---
id: 587d7daf367417b2b2512b7e
title: Elewa sifa ya kijenzi
challengeType: 1
forumTopicId: 301327
dashedName: understand-the-constructor-property
---

# --description--

Kuna sifa maalum ya `constructor` iliyoko kwenye matukio ya vitu `duck` na `beagle` ambayo yaliundwa katika changamoto zilizopita:

```js
let duck = new Bird();
let beagle = new Dog();

console.log(duck.constructor === Bird); 
console.log(beagle.constructor === Dog);
```

Wito zote mbili za `console.log` zingetangaza `true` kwenye konsoli.

Kumbuka kuwa sifa ya `constructor` ni rejea kwa kitendakazi cha kijenzi kilichounda tukio hilo. Faida ya sifa ya `constructor` ni kwamba inawezekana kukagua sifa hii ili kujua ni aina gani ya kitu. Hapa kuna mfano wa jinsi hii inaweza kutumika:

```js
function joinBirdFraternity(candidate) {
  if (candidate.constructor === Bird) {
    return true;
  } else {
    return false;
  }
}
```

**Kumbuka:** Kwa kuwa sifa ya `constructor` inaweza kuandikwa upya (ambayo itajadiliwa katika changamoto mbili zijazo) kwa ujumla ni bora kutumia njia ya `instanceof` kukagua aina ya kitu.

# --instructions--

Andika kitendakazi cha `joinDogFraternity` kinachopokea kigezo cha `candidate` na, kwa kutumia sifa ya `constructor`, rudisha `true` ikiwa mgombea ni `Dog`, vinginevyo rudisha `false`.

# --hints--

`joinDogFraternity` inapaswa kufafanuliwa kama kitendakazi.

```js
assert(typeof joinDogFraternity === 'function');
```

`joinDogFraternity` inapaswa kurudisha `true` ikiwa `candidate` ni tukio la `Dog`.

```js
assert(joinDogFraternity(new Dog('')) === true);
```

`joinDogFraternity` inapaswa kutumia sifa ya `constructor`.

```js
assert(/\.constructor/.test(__helpers.removeJSComments(code)) && !/instanceof/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
function joinDogFraternity(candidate) {

}
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
function joinDogFraternity(candidate) {
  return candidate.constructor === Dog;
}
```
