---
id: 587d7dae367417b2b2512b7a
title: Thibitisha kijenzi cha kitu kwa kutumia instanceof
challengeType: 1
forumTopicId: 301337
dashedName: verify-an-objects-constructor-with-instanceof
---

# --description--

Kila wakati kitendakazi cha kijenzi kinapotengeneza kitu kipya, kitu hicho husemwa kuwa ni <dfn>mfano</dfn> wa kijenzi chake. JavaScript inatoa njia rahisi ya kuthibitisha hili kwa kutumia kielezi `instanceof`. `instanceof` inakuwezesha kulinganisha kitu na kijenzi, ikirudisha `true` au `false` kulingana na kama kitu hicho kilitengenezwa kwa kijenzi hicho au la. Hapa kuna mfano:

```js
let Bird = function(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}

let crow = new Bird("Alexis", "black");

crow instanceof Bird;
```

Njia hii ya `instanceof` itarudisha `true`.

Ikiwa kitu kitatengenezwa bila kutumia kijenzi, `instanceof` itathibitisha kuwa si mfano wa kijenzi hicho:

```js
let canary = {
  name: "Mildred",
  color: "Yellow",
  numLegs: 2
};

canary instanceof Bird;
```

Njia hii ya `instanceof` itarudisha `false`.

# --instructions--

Tengeneza mfano mpya wa kijenzi cha `House`, ukimuita `myHouse` na ukipatia idadi ya vyumba vya kulala. Kisha, tumia `instanceof` kuthibitisha kuwa ni mfano wa `House`.

# --hints--

`myHouse` inapaswa kuwa na sifa ya kipengele `numBedrooms` iliyowekwa kuwa nambari.

```js
assert(typeof myHouse.numBedrooms === 'number');
```

Unapaswa kuthibitisha kuwa `myHouse` ni mfano wa `House` kwa kutumia kielezi cha `instanceof`.

```js
assert(/myHouse\s*instanceof\s*House/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}

// Only change code below this line
```

# --solutions--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}
const myHouse = new House(4);
console.log(myHouse instanceof House);
```
