---
id: 587d7dac367417b2b2512b74
title: Tumia uandishi wa nukta kupata sifa za kitu
challengeType: 1
forumTopicId: 301333
dashedName: use-dot-notation-to-access-the-properties-of-an-object
---

# --description--

Changamoto ya mwisho iliunda kitu chenye sifa mbalimbali. Sasa utaona jinsi ya kupata thamani za sifa hizo. Hapa kuna mfano:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
console.log(duck.name);
```

Uandishi wa nukta hutumika kwenye jina la kitu, `duck`, ikifuatiwa na jina la sifa, `name`, ili kupata thamani ya `Aflac`.

# --instructions--

Chapisha sifa zote mbili za kitu cha `dog` kwenye konsoli yako.

# --hints--

Msimbo wako unapaswa kutumia `console.log` kuchapisha thamani ya sifa ya `name` ya kitu cha `dog`.

```js
assert(/console.log\(.*dog\.name.*\)/g.test(__helpers.removeJSComments(code)));
```

Msimbo wako unapaswa kutumia `console.log` kuchapisha thamani ya sifa ya `numLegs` ya kitu cha `dog`.

```js
assert(/console.log\(.*dog\.numLegs.*\)/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
// Only change code below this line
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
console.log(dog.name);
console.log(dog.numLegs);
```
