---
id: 587d7dae367417b2b2512b7b
title: Elewa vigezo binafsi
challengeType: 1
forumTopicId: 301326
dashedName: understand-own-properties
---

# --description--

Katika mfano ufuatao, kijenzi `Bird` kinafafanua vigezo viwili: `name` na `numLegs`:

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let duck = new Bird("Donald");
let canary = new Bird("Tweety");
```

`name` na `numLegs` huitwa <dfn>vigezo binafsi</dfn>, kwa sababu vinafafanuliwa moja kwa moja kwenye mfano wa kitu. Hii inamaanisha kuwa `duck` na `canary` kila moja ina nakala yake tofauti ya vigezo hivi. Kwa kweli kila mfano wa `Bird` utakuwa na nakala yake ya vigezo hivi. Msimbo ufuatao unaongeza vigezo vyote binafsi vya `duck` kwenye safu ya data `ownProps`:

```js
let ownProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  }
}

console.log(ownProps);
```

Konsoli itaonyesha thamani `["name", "numLegs"]`.

# --instructions--

Ongeza vigezo binafsi vya `canary` kwenye safu ya data `ownProps`.

# --hints--

`ownProps` inapaswa kujumuisha thamani za `numLegs` na `name`.

```js
assert(ownProps.indexOf('name') !== -1 && ownProps.indexOf('numLegs') !== -1);
```

Unapaswa kutatua changamoto hii bila kutumia njia ya kujengwa `Object.keys()`.

```js
assert(!/Object(\.keys|\[(['"`])keys\2\])/.test(__helpers.removeJSComments(code)));
```

Unapaswa kutatua changamoto hii bila kuandika safu ya data `ownProps` moja kwa moja.

```js
assert(
  !/\[\s*(?:'|")(?:name|numLegs)|(?:push|concat)\(\s*(?:'|")(?:name|numLegs)/.test(
    __helpers.removeJSComments(code)
  )
);
```

# --seed--

## --seed-contents--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
let ownProps = [];
// Only change code below this line
```

# --solutions--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
function getOwnProps (obj) {
  const props = [];

  for (let prop in obj) {
    if (obj.hasOwnProperty(prop)) {
      props.push(prop);
    }
  }

  return props;
}

const ownProps = getOwnProps(canary);
```
