---
id: 587d7daf367417b2b2512b7f
title: Badilisha mfano wa awali kuwa kitu kipya
challengeType: 1
forumTopicId: 301316
dashedName: change-the-prototype-to-a-new-object
---

# --description--

Hadi sasa umekuwa ukiongeza sifa kwa `prototype` moja moja:

```js
Bird.prototype.numLegs = 2;
```

Hii huwa ya kuchosha baada ya sifa chache tu.

```js
Bird.prototype.eat = function() {
  console.log("nom nom nom");
}

Bird.prototype.describe = function() {
  console.log("My name is " + this.name);
}
```

Njia bora zaidi ni kuweka `prototype` kuwa kitu kipya ambacho tayari kina sifa hizo. Kwa njia hii, sifa zinaongezwa kwa pamoja mara moja:

```js
Bird.prototype = {
  numLegs: 2, 
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --instructions--

Ongeza sifa `numLegs` na njia mbili `eat()` na `describe()` kwenye `prototype` ya `Dog` kwa kuweka `prototype` kuwa kitu kipya.

# --hints--

`Dog.prototype` inapaswa kuwekwa kuwa kitu kipya.

```js
assert(/Dog\.prototype\s*?=\s*?{/.test(__helpers.removeJSComments(code)));
```

`Dog.prototype` inapaswa kuwa na sifa `numLegs`.

```js
assert(Dog.prototype.numLegs !== undefined);
```

`Dog.prototype` inapaswa kuwa na njia `eat()`.

```js
assert(typeof Dog.prototype.eat === 'function');
```

`Dog.prototype` inapaswa kuwa na njia `describe()`.

```js
assert(typeof Dog.prototype.describe === 'function');
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype = {
  // Only change code below this line

};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
numLegs: 4,
  eat () {
    console.log('nom nom nom');
  },
  describe () {
    console.log('My name is ' + this.name);
  }
};
```
