---
id: 587d7dae367417b2b2512b7c
title: Tumia sifa za prototype kupunguza msimbo unaojirudia
challengeType: 1
forumTopicId: 301336
dashedName: use-prototype-properties-to-reduce-duplicate-code
---

# --description--

Kwa kuwa `numLegs` huenda ikawa na thamani ile ile kwa matukio yote ya `Bird`, kwa msingi huo una variable `numLegs` inayojirudia ndani ya kila tukio la `Bird`.

Hii huenda isakuwa tatizo pale ambapo kuna matukio mawili tu, lakini fikiria kama kuna mamilioni ya matukio. Hiyo itakuwa variable nyingi zinazojirudia.

Njia bora ni kutumia `prototype` ya `Bird`. Sifa katika `prototype` zinashirikiwa kati ya matukio YOTE ya `Bird`. Hapa ni jinsi ya kuongeza `numLegs` kwenye `Bird prototype`:

```js
Bird.prototype.numLegs = 2;
```

Sasa matukio yote ya `Bird` yana sifa ya `numLegs`.

```js
console.log(duck.numLegs);
console.log(canary.numLegs);
```

Kwa kuwa matukio yote yana sifa hizo moja kwa moja kwenye `prototype`, fikiria `prototype` kama "mapishi" ya kuunda vitu. Kumbuka kuwa `prototype` kwa `duck` na `canary` ni sehemu ya kijenzi cha `Bird` kama `Bird.prototype`.

# --instructions--

Ongeza sifa ya `numLegs` kwenye `prototype` ya `Dog`

# --hints--

`beagle` inapaswa kuwa na sifa ya `numLegs`.

```js
assert(beagle.numLegs !== undefined);
```

`beagle.numLegs` inapaswa kuwa nambari.

```js
assert(typeof beagle.numLegs === 'number');
```

`numLegs` inapaswa kuwa sifa ya `prototype` si sifa ya mwenyewe.

```js
assert(beagle.hasOwnProperty('numLegs') === false);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}



// Only change code above this line
let beagle = new Dog("Snoopy");
```

# --solutions--

```js
function Dog (name) {
  this.name = name;
}
Dog.prototype.numLegs = 4;
let beagle = new Dog("Snoopy");
```
