---
id: 587d7db0367417b2b2512b81
title: "Elewa wapi `prototype` ya kitu hutoka"
challengeType: 1
forumTopicId: 301330
dashedName: understand-where-an-objects-prototype-comes-from
---

# --description--

Kama watu wanavyopata vinasaba kutoka kwa wazazi wao, kitu hupata `prototype` yake moja kwa moja kutoka kwa kitendakazi cha kijenzi kilichoumba. Kwa mfano, hapa kijenzi cha `Bird` kinaunda kitu cha `duck`:

```js
function Bird(name) {
  this.name = name;
}

let duck = new Bird("Donald");
```

`duck` hupata `prototype` yake kutoka kwa kitendakazi cha kijenzi cha `Bird`. Unaweza kuonyesha uhusiano huu kwa njia ya `isPrototypeOf`:

```js
Bird.prototype.isPrototypeOf(duck);
```

Hii itarudisha `true`.

# --instructions--

Tumia `isPrototypeOf` kuangalia `prototype` ya `beagle`.

# --hints--

Unapaswa kuonyesha kuwa `Dog.prototype` ni `prototype` ya `beagle`

```js
assert(/Dog\.prototype\.isPrototypeOf\(beagle\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
```
