---
id: 587d7daf367417b2b2512b7d
title: Rudia juu ya vigezo vyote
challengeType: 1
forumTopicId: 301320
dashedName: iterate-over-all-properties
---

# --description--

Sasa umeona aina mbili za vigezo: <dfn>vigezo vya mwenyewe</dfn> na vigezo vya `prototype`. Vigezo vya mwenyewe vinafafanuliwa moja kwa moja kwenye mfano wa kitu hicho. Na vigezo vya prototype vinafafanuliwa kwenye `prototype`.

```js
function Bird(name) {
  this.name = name;  //own property
}

Bird.prototype.numLegs = 2; // prototype property

let duck = new Bird("Donald");
```

Hapa ni jinsi unavyoongeza vigezo vya mwenyewe vya `duck` kwenye safu ya data `ownProps` na vigezo vya `prototype` kwenye safu ya data `prototypeProps`:

```js
let ownProps = [];
let prototypeProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  } else {
    prototypeProps.push(property);
  }
}

console.log(ownProps);
console.log(prototypeProps);
```

`console.log(ownProps)` itaonyesha `["name"]` kwenye konsoli, na `console.log(prototypeProps)` itaonyesha `["numLegs"]`.

# --instructions--

Ongeza vigezo vyote vya mwenyewe vya `beagle` kwenye safu ya data `ownProps`. Ongeza vigezo vyote vya `prototype` vya `Dog` kwenye safu ya data `prototypeProps`.

# --hints--

Safu ya data `ownProps` inapaswa kuwa na `name` tu.

```js
assert.deepEqual(ownProps, ['name']);
```

Safu ya data `prototypeProps` inapaswa kuwa na `numLegs` tu.

```js
assert.deepEqual(prototypeProps, ['numLegs']);
```

Unapaswa kutatua changamoto hii bila kutumia njia iliyojengwa ya `Object.keys()`.

```js
assert(!/\Object.keys/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];
for (let prop in beagle) {
  if (beagle.hasOwnProperty(prop)) {
    ownProps.push(prop);
  } else {
    prototypeProps.push(prop);
  }
}
```
