---
id: 587d7db1367417b2b2512b86
title: Weka upya sifa ya kijenzi iliyorithiwa
challengeType: 1
forumTopicId: 301324
dashedName: reset-an-inherited-constructor-property
---

# --description--

Wakati kitu kinaporithi `prototype` kutoka kwa kitu kingine, pia kinaporithi sifa ya kijenzi ya aina kuu.

Hapa kuna mfano:

```js
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
let duck = new Bird();
duck.constructor
```

Lakini `duck` na matukio yote ya `Bird` yanapaswa kuonyesha kuwa yaliundwa na `Bird` na si `Animal`. Ili kufanya hivyo, unaweza kuweka kwa mkono sifa ya kijenzi ya `Bird` kuwa kitu cha `Bird`:

```js
Bird.prototype.constructor = Bird;
duck.constructor
```

# --instructions--

Rekebisha msimbo ili `duck.constructor` na `beagle.constructor` zirudishe kijenzi chao husika.

# --hints--

`Bird.prototype` inapaswa kuwa mfano wa `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Bird.prototype));
```

`duck.constructor` inapaswa kurudisha `Bird`.

```js
assert(duck.constructor === Bird);
```

`Dog.prototype` inapaswa kuwa mfano wa `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

`beagle.constructor` inapaswa kurudisha `Dog`.

```js
assert(beagle.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Animal() { }
function Bird() { }
function Dog() { }

Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);

// Only change code below this line



let duck = new Bird();
let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
function Bird() { }
function Dog() { }
Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Bird.prototype.constructor = Bird;
let duck = new Bird();
let beagle = new Dog();
```
