---
id: 587d7dad367417b2b2512b75
title: Unda njia kwenye kitu
challengeType: 1
forumTopicId: 301318
dashedName: create-a-method-on-an-object
---

# --description--

Vitu vinaweza kuwa na aina maalum ya sifa, inayoitwa <dfn>njia</dfn>.

Njia ni sifa ambazo ni kitendakazi. Hii huongeza tabia tofauti kwa kitu. Hapa kuna mfano wa `duck` wenye njia:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + duck.name + ".";}
};
duck.sayName();
```

Mfano huu unaongeza njia ya `sayName`, ambayo ni kitendakazi kinachorudisha sentensi inayotoa jina la `duck`. Angalia kwamba njia ilifikilia sifa ya `name` katika taarifa ya return kwa kutumia `duck.name`. Changamoto inayofuata itashughulikia njia nyingine ya kufanya hivi.

# --instructions--

Ukitegemea kitu cha `dog`, kipa njia iitwayo `sayLegs`. Njia hiyo inapaswa kurudisha sentensi `This dog has 4 legs.`

# --hints--

`dog.sayLegs()` inapaswa kuwa kitendakazi.

```js
assert(typeof dog.sayLegs === 'function');
```

`dog.sayLegs()` inapaswa kurudisha mfuatano wa herufi uliotolewa - kumbuka kuwa alama za uandishi na nafasi ni muhimu.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,

};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
