---
id: 6619240f46cec8e04d77e03a
title: Esercizio base sulle funzioni A
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

Crea una funzione che riceve un numero intero. Questa funzione dovrebbe restituire il `integer + 7` dato se l'intero è minore di `10`. Se l'intero è maggiore o uguale a `10`, dovrebbe restituire il `integer - 3` dato.

Il nome della funzione dovrebbe essere `addOrSubtract`.

# --hints--

Dovresti avere una funzione chiamata `addOrSubtract`.

```js
assert.isFunction(addOrSubtract);
```

La tua funzione dovrebbe ricevere un numero intero come argomento.

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

Dovresti restituire l'intero dato + 7 se l'intero è minore di 10.

```js
assert.strictEqual(addOrSubtract(5), 12);
```

Dovresti restituire l'intero dato - 3 se l'intero è maggiore o uguale a 10.

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
