---
id: 661e151f068359c3ccf2f4d7
title: Esercizio base sulle funzioni C
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

Scrivi una funzione, chiamata `capitalize`, che prende una stringa come parametro e restituisce una nuova stringa con la prima lettera maiuscola.

# --hints--

Dovresti avere una funzione chiamata `capitalize`.

```js
assert.isFunction(capitalize);
```

La tua funzione dovrebbe prendere una stringa come parametro.

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

La tua funzione dovrebbe restituire una nuova stringa con la prima lettera maiuscola.

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
