---
id: 661e17c6068359c3ccf2f4d8
title: esercizio base sulle funzioni d
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

Scrivi una funzione, chiamata `lastLetter`, che prende una stringa come parametro e restituisce l'ultima lettera della stringa.

# --hints--

Dovresti avere una funzione chiamata `lastLetter`.

```js
assert.isFunction(lastLetter);
```

La tua funzione dovrebbe prendere una stringa come parametro.

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

Dovresti restituire l'ultima lettera della stringa.

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
