---
id: 661e131f068359c3ccf2f4d6
title: esercizio base sulle funzioni b
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

Scrivi una funzione, chiamata `multiply`, che prende due parametri e restituisce il loro prodotto.

# --hints--

Dovresti avere una funzione chiamata `multiply`.

```js
assert.isFunction(multiply);
```

La tua funzione dovrebbe prendere due numeri interi come argomenti.

```js
assert.match(multiply.toString(), /\s*multiply\(\s*\w+\s*,\s*\w+\s*\)/);
```

Dovresti restituire il prodotto dei due numeri interi.

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js 
function multiply(a, b) {
  return a * b;
}
```
