---
id: 66fe4f33a2cc9b33f4d5cd9b
title: 构建一个排列生成器
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

在本实验中，你将构建一个排列生成器，该生成器将接受一个字串并返回该字串中字符的所有可能排列。例如，字串 `cat` 的可能排列有 `cat`、`cta`、`act`、`atc`、`tac` 和 `tca`。

递归的创建字串排列的方法是通过存储字串的固定起始部分（前缀），并创建其余部分的排列。

例如，考虑字串 `machine`。第一轮创建排列将固定 `m` 作为字串的前缀，然后创建剩余字串 `achine` 的排列。

对于剩余的字串，排列组合以相同的方式继续。一个字母被添加到前缀，可能是 `c`，所以前缀变成 `mc`。然后，将 `ahine` 的每个排列组合连接到前缀。

这将持续，直到前缀包含所有字母且剩余的字串为空，这意味着一个排列已被创建。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `permuteString` 的函数。
2. `permuteString` 函数应接受三个参数：一个字串、一个前缀值和一个用于存储及返回结果的空数组。前缀值应是一个用于累积字符以形成排列的字串。
3. 在函数内部，你应该查看传入字串的长度是否为 `0`。如果是，将当前前缀推入结果并返回结果。
4. 迭代输入字串中的每个字符，并且对于每次迭代，从字串中移除当前字符，递归地调用 `permuteString` 函数，使用更新后的参数来构建剩余的排列。
5. 你应该返回最终结果数组。
6. 你应该通过移除副本来确保排列是唯一的。

# --hints--

你应该有一个函数 `permuteString`。

```js
assert.isFunction(permuteString);
```

你应该在你的 `permuteString` 函数中使用递归。

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` 应该返回 `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`。

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` 应该返回 `[ "fcc", "cfc", "ccf" ]`。

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` 应该返回 `[ "p" ]`。

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` 应该返回 `[""]`。

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` 应该返回 `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`。

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` 应该返回正确的结果。

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
