---
id: 67051431a73f1ca25d9a6f25
title: Construir un Tablero de Ideas de Proyectos
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**Objetivo:** Cumplir las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de Usuario:**

1. Debes definir una constante de objeto llamada `projectStatus` con las tres claves: `PENDING`, `SUCCESS`, y `FAILURE`. A cada estado se le debe asignar un objeto con una clave `description` con el valor `Pending Execution`, `Executed Successfully`, y `Execution Failed`, respectivamente.
2. Debes definir una clase llamada `ProjectIdea` con un `constructor` que tome `title` y `description` como parámetros. Inicializa las propiedades de `title` y `description` con los parámetros proporcionados. La clase también debe tener una propiedad llamada `status` que se establece en el valor `projectStatus.PENDING` de manera predeterminada.
3. Debes definir un método llamado `updateProjectStatus` dentro de la clase `ProjectIdea`. Este método debe aceptar un parámetro `newStatus` y actualizar la propiedad `status` al valor dado.
4. Debes definir una clase `ProjectIdeaBoard` con un `constructor` que acepte un `title` e inicialice un array vacío llamado `ideas` para contener instancias de la clase `ProjectIdea`.
5. Debes definir un método llamado `pin` dentro de la clase `ProjectIdeaBoard` que acepte una instancia de la clase `ProjectIdea` y la agregue al array `ideas`.
6. Debes definir un método llamado `unpin` dentro de la clase `ProjectIdeaBoard`. Este método debe aceptar una instancia de la clase `ProjectIdea` y eliminarla del array `ideas`.
7. Debes definir un método llamado `count` que devuelva el número de ideas de proyectos en el array `ProjectIdeaBoard` dado.
8. Debes definir un método llamado `formatToString` que devuelva el nombre, la descripción y el estado de las ideas de proyectos en el formato:

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

Debes definir una constante de objeto llamada `projectStatus` con los estados y descripciones correctos del proyecto.

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

Debes tener una clase `ProjectIdea`.

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

Tu clase `ProjectIdea` debe inicializar la propiedad `title` y la propiedad `description` según los parámetros pasados. También debe establecer el estado en `projectStatus.PENDING`.

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

Tu clase `ProjectIdea` debe tener un método `updateProjectStatus`.

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

Debes tener una clase `ProjectIdeaBoard`.

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

Tu `ProjectIdeaBoard` debería inicializar la propiedad `title` según el parámetro pasado e inicializar un array vacío llamado `ideas` para contener instancias de la clase `ProjectIdea`.

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

Tu clase `ProjectIdeaBoard` debería tener un método `pin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

Tu clase `ProjectIdeaBoard` debería tener un método `unpin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

Tu clase `ProjectIdeaBoard` debería tener un método `count`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

Tu clase `ProjectIdeaBoard` debería tener un método `formatToString`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")` debería devolver `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }`.

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

Llamar a `updateProjectStatus(projectStatus.SUCCESS)` en `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")` debería actualizar el estado a `{ description: 'Executed Successfully' }`.

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

Llamar a `updateProjectStatus(projectStatus.FAILURE)` en `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")` debería actualizar el estado a `{ description: 'Execution Failed' }`.

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

Llamar a `updateProjectStatus(projectStatus.SUCCESS)` en `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")` debería actualizar el estado a `{ description: 'Executed Successfully' }`.

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

Deberías poder fijar un objeto `ProjectIdea` en tu `ProjectIdeaBoard` usando el método `pin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

Deberías poder desanclar un objeto `ProjectIdea` en tu `ProjectIdeaBoard` usando el método `unpin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

Cuando `new ProjectIdeaBoard("Empty Board")` está vacío, `emptyBoard.formatToString()` debería devolver `Empty Board has 0 idea(s)\n`.

```js

let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

// Create an empty board
let emptyBoard = new ProjectIdeaBoard("Empty Board");

// Attempt to unpin a project from an empty board
emptyBoard.unpin(smartHome);

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

Cuando fijas `new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")` en `new ProjectIdeaBoard("Tech Projects Board")`, el `techProjects.formatToString()` debería devolver `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n`.

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
