---
id: 66f686b8ebdb982fa8e14330
title: Ein Programm zur Datumsumwandlung erstellen
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

In diesem Labor programmieren Sie ein Datumsumwandlungsprogramm, das ein gegebenes Datum in verschiedene Formate umwandelt. Zum Beispiel würde das aktuelle Datum `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` auf folgende 2 Arten formatiert:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Variable namens `currentDate` erstellen und ihr das aktuelle Datum und die aktuelle Uhrzeit mit dem `Date`-Objekt zuweisen.
2. Sie sollten eine Variable namens `currentDateFormat` erstellen und ihr den String `Current Date and Time: [current date]` zuweisen. Ersetzen Sie `[current date]` durch das Ergebnis der `currentDate`-Variablen.
3. Sie sollten den Wert von `currentDateFormat` in der Konsole ausgeben.
4. Sie sollten eine Funktion namens `formatDateMMDDYYYY` erstellen, die das Datumsobjekt als Parameter entgegennimmt. Sie können diesen Parameter beliebig benennen.
5. Ihre `formatDateMMDDYYYY`-Funktion sollte das aktuelle Datum in das `month/day/year`-Format umwandeln und den String `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]` zurückgeben.
6. Sie sollten eine Funktion namens `formatDateLong` erstellen, die das Datumsobjekt als Parameter entgegennimmt.
7. Ihre `formatDateLong`-Funktion sollte das aktuelle Datum in das `Month Day, Year`-Format umwandeln und den String `Formatted Date (Month Day, Year): [formatted date]` zurückgeben.

**Hinweis:** Damit die Tests bestehen, stellen Sie sicher, dass Sie `en-US` als Locale beim Formatieren der Daten verwenden.

# --hints--

Sie sollten eine Variable namens `currentDate` haben, die das aktuelle Datum und die aktuelle Uhrzeit mit dem `Date`-Objekt hält.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

Sie sollten eine Variable namens `currentDateFormat` haben, die das aktuelle Datum im Format `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>` hält.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

Sie sollten den Wert von `currentDateFormat` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

Sie sollten eine Funktion `formatDateMMDDYYYY` haben

```js
assert.isFunction(formatDateMMDDYYYY);
```

Die Funktion `formatDateMMDDYYYY` sollte einen einzelnen Parameter entgegennehmen.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

Wenn das Datumsobjekt `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` enthält, sollte die Funktion `formatDateMMDDYYYY` `Formatted Date (MM/DD/YYYY): 9/27/2024` zurückgeben.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

Sie sollten eine Funktion `formatDateLong` haben

```js
assert.isFunction(formatDateLong);
```

Die Funktion `formatDateLong` sollte einen einzelnen Parameter entgegennehmen.

```js
assert.lengthOf(formatDateLong, 1);
```

Wenn das Datumsobjekt `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` enthält, sollte die Funktion `formatDateLong` `Formatted Date (Month Day, Year): September 27, 2024` zurückgeben.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```

