---
id: 687ad7be1ff45032ccf1d92f
title: テーマスイッチャーを作る
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

このラボでは、異なるテーマを切り替えるアプリを作成します。テーマを切り替えると、背景色が変わり、異なるメッセージが表示されます。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `button` 要素に `Switch Theme` というテキストがあること。
2. `button` 要素に以下の属性があること:
   - `id` 属性が `"theme-switcher-button"` に設定されていること。
   - `aria-haspopup` 属性が `"true"` に設定されていること。
   - `aria-expanded` 属性が `"false"` に設定されていること。
   - `aria-controls` 属性が `"theme-dropdown"` に設定されていること。
3. `ul` 要素に以下の属性があること:
   - `id` 属性が `"theme-dropdown"` に設定されていること。
   - `role` 属性が `"menu"` に設定されていること。
   - `aria-labelledby` 属性が `"theme-switcher-button"` に設定されていること。
   - `hidden` 属性があること。
4. `ul` 要素に少なくとも2つの `li` 要素があり、それぞれに `role` 属性が `"menuitem"` に設定され、異なるテーマを表すテキストがあること。
5. 各 `li` 要素に `id` 属性があり、その値は `theme-` で始まり、`li` 要素のテキストで設定したテーマで終わっていること。例えば、テーマの1つが `Light` の場合、`id` は `theme-light` である必要があります。
6. `aria-live` 属性が `"polite"` に設定された要素があること。
7. `themes` 配列に少なくとも2つのオブジェクトがあり、それぞれに `name` と `message` プロパティがあること。`name` は異なるテーマを表し、`message` はテーマが切り替わったときに表示されます。`name` と `message` の値は自由に設定できますが、`name` の値は `li` 要素のテキスト内容と一致している必要があります。
8. ユーザーが `#theme-switcher-button` をクリックすると、`#theme-dropdown` 要素に利用可能なテーマが表示されること。
9. `#theme-dropdown` 要素が表示されると、`aria-expanded` 属性が `"true"` に設定され、`button` 要素に適用されること。
10. ユーザーが `#theme-switcher-button` をクリックし、`#theme-dropdown` 要素が表示されている間にメニューが閉じられ、`hidden` 属性が追加され、`aria-expanded` 属性が `"false"` に設定されること。
11. ユーザーが `#theme-switcher-button` をクリックしてテーマを選択すると、対応する `theme-<name>` クラスが `body` 要素に追加され、`themes` 配列で設定した関連するテーマメッセージが `aria-live="polite"` 要素に表示されること。
12. ユーザーが `#theme-switcher-button` をクリックして現在のテーマと異なるテーマを選択すると、`body` テーマクラスとステータスメッセージが新しい選択に合わせて更新されること。

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

`button` 要素があること。

```js
assert.exists(document.querySelector("button"));
```

`button` 要素に `Switch Theme` というテキストがあること。

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

`button` 要素に `id` 属性が `"theme-switcher-button"` に設定されていること。

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

`button` 要素に `aria-haspopup` 属性が `"true"` に設定されていること。

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

`button` 要素に `aria-expanded` 属性が `"false"` に設定されていること。

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

`button` 要素に `aria-controls` 属性が `"theme-dropdown"` に設定されていること。

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

`ul` 要素があること。

```js
assert.exists(document.querySelector("ul"));
```

`ul` 要素に `id` 属性が `"theme-dropdown"` に設定されていること。

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

`ul` 要素に `role` 属性が `"menu"` に設定されていること。

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

`ul` 要素に `aria-labelledby` 属性が `"theme-switcher-button"` に設定されていること。

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

`ul` 要素に `hidden` 属性があること。

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

`ul` 要素に少なくとも2つの `li` 要素があること。

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

各 `li` 要素に `role` 属性が `"menuitem"` に設定されていること。

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

各 `li` 要素にテーマを表すテキストがあること。

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

各 `li` 要素に `id` 属性があり、その値は `theme-` で始まり、その要素に対応するテーマで終わっていること。`id` の値はすべて小文字であることを確認してください。

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

`aria-live` 属性が `"polite"` に設定された要素があること。

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

`themes` 配列があること。

```js
assert.isArray(themes);
```

`themes` 配列に少なくとも2つのオブジェクトがあること。

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

`themes` 配列に少なくとも2つのオブジェクトがあり、それぞれに `name` と `message` プロパティがあること。

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

`name` と `message` のプロパティ値は空でない文字列であること。

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

各 `name` プロパティ値は、`li` 要素のテキスト内容のいずれかと一致していること。これらの値はすべて小文字であることを確認してください。

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

ユーザーが `#theme-switcher-button` をクリックしてテーマオプションを表示すると、`hidden` 属性が `#theme-dropdown` 要素から削除されること。

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

`#theme-dropdown` 要素が表示されると、`aria-expanded` 属性が `"true"` に設定され、`button` 要素に適用されること。

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

ユーザーが `#theme-switcher-button` をクリックし、`#theme-dropdown` 要素が表示されている間にドロップダウンメニューが `hidden` に戻されること。

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

ユーザーが `#theme-switcher-button` をクリックし、`#theme-dropdown` 要素が表示されている間に `aria-expanded` 属性が `"false"` に設定されること。

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

ユーザーが `#theme-switcher-button` をクリックしてテーマを選択すると、対応する `theme-<name>` クラスが `body` 要素に追加されること。

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

ユーザーが `#theme-switcher-button` をクリックしてテーマを選択すると、`message` 配列の選択された `theme` に関連する `themes` が `aria-live="polite"` 要素に表示されること。

```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

ユーザーが `#theme-switcher-button` をクリックして2つの異なるテーマを選択すると、`body` クラスとステータスメッセージがそれぞれの選択テーマに合わせて更新されること。

```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[0]?.dispatchEvent(new Event("click", { bubbles: true }));
const firstClass = document.body.className;
const firstMessage = liveRegion?.textContent;

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", { bubbles: true }));
const secondClass = document.body.className;
const secondMessage = liveRegion?.textContent;

assert.equal(firstClass, listItems[0]?.id);
assert.equal(secondClass, listItems[1]?.id);
assert.notEqual(firstClass, secondClass);
assert.notEqual(firstMessage, secondMessage);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
