---
id: 66e45c8140f9fda5c105ae26
title: 创建一个事件传单页面
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你的 `body` 中应包含一个 `header` 元素。
1. 你的 `header` 元素中应包含一个用于你的事件的图像，以及一个包含你的事件标题的 `h1`，顺序如上。如果你愿意，可以使用此图像：`https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. 你应该在 `body` 中有一个 `main` 元素。
1. 你的 `main` 元素内应至少包含两个 `section` 元素，用于展示事件特色。
1. 你的每个 `section` 元素中都应包含一个 `h2`。
1. 你的体应有 `50px` 的顶部和底部内边距。
1. 你的 body 应该有 `0` 的上下边距，以及使其居中的左右边距。
1. 你的 `body` 应该设置相对于 `viewport` 宽度的宽度。
1. 你的体应使用 `calc` 函数将其 `min-height` 属性设置为视口高度的 100% 减去应用于体顶部和底部的所有内边距。
1. 你的传单中应至少包含一个 `hr`。
1. 你应该将所有 `hr` 和 `section` 元素的 `width` 设置为相对于其父元素的百分比值。

**注意：** 确保在 HTML 中链接样式表以应用 CSS。

# --hints--

你应该在 `body` 元素内有一个 `header` 元素。

```js
assert.exists(document.querySelector('body > header'));
```

你应该在 `header` 中将 `img` 元素作为第一个子元素。

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

你应该在 `header` 中将 `h1` 元素作为第二个子元素。

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

你应该只有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `h1` 应该包含你的事件标题。

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

你应该在 `body` 元素内有一个 `main` 元素。

```js
assert.exists(document.querySelector('body > main'));
```

你的 `main` 元素内应至少包含两个 `section` 元素。

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

你的 `section` 元素中每个都应包含一个 `h2`。

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

你的 `h2` 元素不应为空。

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

你的 `body` 元素应有上下各 `50px` 的内边距。

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

你的 `body` 元素应具有上下边距为 `0`。

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

你的 `body` 元素应具有左右 `auto` 外边距。

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

你的 `body` 元素应使用 `vw` 设置其宽度。

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

你的 `body` 应该使用 `calc` 将其 `min-height` 设置为 `100vh - 100px`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

你应该至少有一个 `hr` 元素。

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

你的 `hr` 元素的宽度应使用百分比值设置。

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

你的 `section` 元素的宽度应使用百分比值设置。

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
