---
id: 68ed661e9f463c155bb677a1
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

このワークショップでは、CSS flexboxを使って色付きのボックスをデザインする練習をします。

まず、`body` 要素の中に `header` 要素を追加してください。`h1` 要素の中に `header` 要素をネストし、テキストとして `Colored Boxes Layout` を入れてください。

# --hints--

`header` 要素の中に `body` 要素があるはずです。

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

`header` 要素の中に `h1` 要素があるはずです。

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

`h1` 要素の中にテキスト `Colored Boxes Layout` があるはずです。

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
