---
id: 68ef0b3237fef991c92d9c19
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`div`の下にクラス`flex-container`を持つ`header`要素を追加してください。次に、先ほど作成した`div`の中にクラス`box`を持つ`div`要素を6つ作成してください。

クラス属性はスタイルシートのクラス名を指定するためによく使われます。この場合、すべての`div`要素は`.flex-container`および`.box`のスタイル定義に従って同じようにスタイル付きになります。

# --hints--

`div`要素の下に`header`要素があるはずです。

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

`div`要素の下にある`header`要素はクラス`flex-container`を持つべきです。

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

`div`要素の中に6つの`.flex-container`要素がネストされているはずです。

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

`div`要素の中にある6つの`.flex-container`要素すべてがクラス`box`を持つべきです。

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
