---
id: 68ef0d6d156f86a6ac1824c6
title: ステップ 23
challengeType: 0
dashedName: step-23
---

# --description--

これから`align-content`プロパティを使います。このプロパティは交差軸に沿ったアイテムの配置を制御します。`align-content`のデフォルト値は`stretch`ですが、`align-content`に使える最も関連性の高いオプションは`space-between`、`start`、`center`、`end`、`space-around`、`space-evenly`です。

`align-content`要素内のボックスがどのように揃うかを見るために、値`space-between`で`.flex-container`プロパティを追加してください。

# --hints--

`.flex-container`セレクターに、値`align-content`の`space-between`プロパティを設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('align-content'), 'space-between');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}

.box {
	max-height: 120px;
	color: #000;
	border: 1px solid #000;
	display: flex;
	flex: 1 1 100px;
	flex-direction: column;
	align-items: center;
	margin: 10px;
	font-weight: bold;
	font-size: 1.125rem;
	border-radius: 5px;
	order: 0; 
}

.flex-container {
	display: flex;
	flex-wrap: wrap;
	width: 70%;
	height: 600px;
	padding: 10px;
	margin: 20px auto;
--fcc-editable-region--
	
--fcc-editable-region--
}
```
