---
id: 68f471d46fa193b002a5bfb8
title: ステップ 43
challengeType: 0
dashedName: step-43
---

# --description--

`.box6` セレクターを作成し、`background` プロパティに `#3f51b5` の値を設定してください。

次に、新しいセレクターに `order` プロパティを追加し、値を `5` にしてください。`.box6` 要素が `.box5` 要素の後ろに移動するのが見えます。

さらに、`flex-grow` プロパティを値 `1` で追加してください。これにより、ボックス6がコンテナ内の残りのスペースを埋めるように拡大します。

これで、カラフルなボックスのワークショップは完了です！

# --hints--

`.box6` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box6'));
```

`.box6` セレクターには、`background` プロパティが値 `#3f51b5` で設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.getPropVal('background'), 'rgb(63, 81, 181)');
```

`.box6` セレクターには、`order` プロパティが値 `5` で設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.getPropVal('order'), '5');
```

`.box6` セレクターには、`flex-grow` プロパティが値 `1` で設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box box4">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box box5">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box box6">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2;
  flex-shrink: 3; 
}

.box4 {
  background: #4caf50;
  order: 3; 
}

.box5 {
  background: #457b9d;
  order: 4;
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box box4">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box box5">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box box6">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1; 
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2; 
  flex-shrink: 3; 
}

.box4 {
  background: #4caf50;
  order: 3;
}

.box5 {
  background: #457b9d;
  order: 4; 
}

.box6 {
  background: #3f51b5;
  order: 5; 
  flex-grow: 1; 
}
```
