---
id: 68f46b758eb8ea7043abd73c
title: ステップ 33
challengeType: 0
dashedName: step-33
---

# --description--

これから、フレックスボックスのコンテナ内のボックスを整理していきます。まず、`.box1` セレクターを作成し、`background` プロパティに `#f16e79` の値を設定してください。

次に、`order` プロパティに `1` の値を追加します。これで最初のボックスが最後に移動するのがわかります。これは、`box` クラスにすでに `order` プロパティが `0` の値で設定されており、値が大きい順にアイテムが後ろに表示されるためです。

さらに、`flex-grow` プロパティに `1` の値を追加してください。これにより、そのボックスは行の余白を埋めるように拡大します。

# --hints--

`.box1` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box1'));
```

`.box1` セレクターには、`background` プロパティを `#f16e79` の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('background'), 'rgb(241, 110, 121)');
```

`.box1` セレクターには、`order` プロパティを `1` の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('order'), '1');
```

`.box1` セレクターには、`flex-grow` プロパティを `1` の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--
```
