---
id: 68f91df5776e693049fb8942
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

これで`styles.css`ファイルにCSSルールを追加する時間です。まず`h1`要素のセレクターを作成してください。

`h1`要素の内容を中央に配置するには、その`text-align`プロパティを`center`に設定します。また、`margin-bottom`要素の下に続くHTML要素との間にマージンを追加するために、値を`10px`にした`h1`プロパティも作成してください。

# --hints--

`h1`セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

`h1`セレクターは`text-align`プロパティを`center`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

`h1`セレクターは`margin-bottom`プロパティを`10px`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
