---
id: 692f314d449329ba8bcd5d2e
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

`border-radius` プロパティは要素のボーダーの角を丸くします。次のように指定できます：

- すべての角に適用する1つの値。
- 2つの値で、1つ目は左上／右下、2つ目は右上／左下に適用。
- 3つの値で、左上、右上／左下、右下に対応。
- 4つの値で、左上の角から時計回りにボーダー半径を設定。

`border-radius` セレクターに、`5px` プロパティを値`.box` で追加してください。

# --hints--

`.box` セレクターには、`border-radius` プロパティを値`5px` で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border-radius'), '5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
