---
id: 692f2c20cb7f9e5050b935a7
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`border` セレクターに、値が `1px solid #000` のプロパティ `.box` を追加してください。`border` プロパティは、要素の周りに見えるアウトラインを作成するために使われます。これは `border-width`、`border-style`、`border-color` の点でカスタマイズできます。

この場合、`1px` は `border-width` で、ボーダーの太さを定義します。`solid` は `border-style` で、ボーダーの見た目（実線、破線、点線など）を設定します。そして `#000` は `border-color` で、ボーダーの色を定義します。

# --hints--

`.box` セレクターには、値が `border` のプロパティ `1px solid #000` があるはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border'), '1px solid rgb(0, 0, 0)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
