---
id: 692f318859d5e4bf8fd77ba1
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

`order` プロパティは、同じコンテナ内の他のフレックスアイテムに対するフレックスアイテムの順序を指定します。デフォルトでは、すべてのフレックスアイテムの `order` の値は `0` で、HTMLに書かれている順に表示されます。値が低い `order` を持つアイテムは先に表示され、値が高いアイテムは後に表示されます。

`order` セレクターに、値 `0` の `.box` プロパティを追加してください。

# --hints--

`.box` セレクターには、値 `order` の `0` プロパティが必要です。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('order'), '0');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
