---
id: 692f22933ba1cb86cd75a70c
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

ボックスの色名を作成する時間です。これを行うには、各`p`要素の下に、次の順番で色名を持つ`h2`要素を追加してください：`Red`、`Orange`、`Yellow`、`Green`、`Blue`、`Indigo`。ボックスごとに1つの色です。

# --hints--

最初の`p`要素内の`Red`要素の下に、テキスト`h2`を持つ`.box`要素を追加してください。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

2番目の`p`要素内の`Orange`要素の下に、テキスト`h2`を持つ`.box`要素を追加してください。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

3番目の`p`要素内の`Yellow`要素の下に、テキスト`h2`を持つ`.box`要素を追加してください。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

4番目の`p`要素内の`Green`要素の下に、テキスト`h2`を持つ`.box`要素を追加してください。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

5番目の`p`要素内の`Blue`要素の下に、テキスト`h2`を持つ`.box`要素を追加してください。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

6番目の`p`要素内の`Indigo`要素の下に、テキスト`h2`を持つ`.box`要素を追加してください。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
