---
id: 692f2d37f34ed866ed42fcd6
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`flex-direction` プロパティは、フレックスアイテムがコンテナ内でどのように配置されるかを制御します。デフォルト値は `row` で、アイテムを左から右へ水平方向に並べます。このプロパティは `display: flex` または `display: inline-flex` を持つ要素にのみ適用されます。

`column` の値を使うと、アイテムは上から下へ垂直に積み重なります。他の可能な値には `row`、`row-reverse`、`column-reverse` があります。

今、`flex-direction` セレクターに `column` プロパティを値 `.box` で追加してください。

# --hints--

`.box` セレクターには、プロパティ `flex-direction` を値 `column` で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex-direction'), 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
