---
id: 692f22a2ad44ae887a24c8f9
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

各 `.box` 要素の中に、テキスト `h2` と空白、その後に `Box` から始まる連番を含む `1` 要素を追加してください。

# --hints--

最初の `h2` の中にクラス `Box 1` を持つ `div` 要素があり、その中にテキスト `box` があるはずです。

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

2番目の `h2` の中にクラス `Box 2` を持つ `div` 要素があり、その中にテキスト `box` があるはずです。

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

3番目の `h2` の中にクラス `Box 3` を持つ `div` 要素があり、その中にテキスト `box` があるはずです。

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

4番目の `h2` の中にクラス `Box 4` を持つ `div` 要素があり、その中にテキスト `box` があるはずです。

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

5番目の `h2` の中にクラス `Box 5` を持つ `div` 要素があり、その中にテキスト `box` があるはずです。

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

最後の `h2` の中にクラス `Box 6` を持つ `div` 要素があり、その中にテキスト `box` があるはずです。

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
