---
id: 673339d96c56d36e963e888e
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, costruirai un catalogo per un vivaio mentre pratichi gli oggetti Map e Set di JavaScript.

Per rappresentare ogni pianta nel tuo catalogo, userai oggetti, ciascuno con tre proprietà:

- `commonName`, il nome comune della pianta
- `scientificName`, il nome scientifico (o nome latino) della pianta
- `cultivar`, la varietà della pianta

Inizia creando un oggetto con le tre proprietà sopra per il `Spanish lavender` (nome scientifico `Lavandula stoechas`) di una cultivar chiamata `Ballerina` e assegnalo a una variabile chiamata `ballerina`.

# --hints--

Dovresti avere una variabile chiamata `ballerina`.

```js
assert.exists(ballerina);
```

Dovresti assegnare un oggetto alla tua variabile `ballerina`.

```js
assert.isObject(ballerina);
```

Il tuo oggetto `ballerina` dovrebbe avere una proprietà `commonName` impostata su `"Spanish lavender"`.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

Il tuo oggetto `ballerina` dovrebbe avere una proprietà `scientificName` impostata su `"Lavandula stoechas"`.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

Il tuo oggetto `ballerina` dovrebbe avere una proprietà `cultivar` impostata su `"Ballerina"`.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

Il tuo oggetto `ballerina` dovrebbe avere tre chiavi.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
