---
id: 6733b5e6f210c1bd1afdf8b3
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Gli oggetti delle piante rimanenti sono stati aggiunti per te. Ora è il momento di creare il tuo catalogo.

Come i normali oggetti JavaScript, gli oggetti `Map` contengono una raccolta di coppie chiave-valore. Per creare un oggetto `Map`, puoi usare il costruttore `Map`:

```js
const myMap = new Map();
```

Dichiara una variabile `catalog` e assegnale un oggetto `Map`.

# --hints--

Dovresti avere una variabile chiamata `catalog`.

```js
assert.exists(catalog);
```

Dovresti assegnare un oggetto `Map` alla tua variabile `catalog`.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
