---
id: 6733b5632fb854bc6c58dc0d
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Ora crea un altro oggetto, con le stesse tre proprietà, per rappresentare un'altra cultivar di lavanda spagnola chiamata `Pretty Polly`, e assegnalo a una variabile chiamata `prettyPolly`.

# --hints--

Dovresti avere una variabile chiamata `prettyPolly`.

```js
assert.exists(prettyPolly);
```

Dovresti assegnare un oggetto alla tua variabile `prettyPolly`.

```js
assert.isObject(prettyPolly);
```

Il tuo oggetto `prettyPolly` dovrebbe avere una proprietà `commonName` impostata su `"Spanish lavender"`.

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

Il tuo oggetto `prettyPolly` dovrebbe avere una proprietà `scientificName` impostata su `"Lavandula stoechas"`.

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

Il tuo oggetto `prettyPolly` dovrebbe avere una proprietà `cultivar` impostata su `"Pretty Polly"`.

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

Il tuo oggetto `prettyPolly` dovrebbe avere tre chiavi.

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
