---
id: 6734f0ae9d1cf4dfd43a7805
title: Passaggio 15
challengeType: 1
dashedName: step-15
---

# --description--

Per rimuovere un elemento da un oggetto map, puoi usare il metodo `delete`, che prende una chiave come argomento ed elimina la voce associata a quella chiave. Restituisce un booleano che indica se la rimozione è stata effettuata con successo.

Crea una funzione chiamata `removePlant` che prende un oggetto pianta come argomento e lo rimuove dal catalogo. Fai in modo che restituisca `true` o `false` a seconda che la pianta specificata sia stata eliminata dal catalogo.

Dopodiché, registra il risultato della chiamata alla tua nuova funzione con la stringa `"ballerina"` come argomento.

# --hints--

Dovresti avere una funzione chiamata `removePlant`.

```js
assert.isFunction(removePlant);
```

La tua funzione `removePlant` dovrebbe prendere un solo argomento.

```js
assert.lengthOf(removePlant, 1);
```

La tua funzione `removePlant` dovrebbe chiamare il metodo `delete` su `catalog`.

```js
let flag = false;
const temp = catalog.delete;
try {
    catalog.delete = (p) => {flag = true}
    removePlant("");
    assert.isTrue(flag);
} finally {
    catalog.delete = temp;
}
```

`removePlant(item)` dovrebbe eliminare `item` dal tuo catalogo.

```js
const testArr = [];
const temp = catalog.delete;
try {
    catalog.delete = (p) => {testArr.push(p)}
    removePlant("test");
    assert.deepEqual(testArr, ["test"]);
} finally {
    catalog.delete = temp;
}
```

La tua funzione `removePlant` dovrebbe restituire un valore booleano che indica se la rimozione è stata effettuata con successo.

```js
try {
    assert.isFalse(removePlant("test"));
    assert.isTrue(removePlant(royalCrown));
} finally {
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 })
}
```

Dovresti registrare `removePlant("ballerina")` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*removePlant\s*\(\s*("|')ballerina\1\s*\)\s*\)/);
```



# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

--fcc-editable-region--

--fcc-editable-region--

const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```
