---
id: 67349ba230cff237c0f2bdf2
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Gli oggetti map ricordano l’ordine in cui le chiavi vengono inserite nella mappa. Per impostare una coppia chiave-valore nel tuo oggetto map, puoi usare il metodo `set`. Il seguente codice imposta la chiave `"name"` al valore `"Jerry"` nell’oggetto map `myMap`.

```js
myMap.set("name", "Jerry");
```

Imposta la chiave `"ballerina"` al valore del tuo oggetto `ballerina` nella tua mappa `catalog`. Poi, stampa il tuo catalogo nella console.

Nota che per vedere correttamente il tuo oggetto map dovrai aprire la console del browser.

# --hints--

Dovresti impostare la chiave `"ballerina"` al valore del tuo oggetto `ballerina` nella tua mappa `catalog`.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

La tua mappa `catalog` dovrebbe contenere una singola coppia chiave-valore.

```js
assert.lengthOf(catalog, 1);
```

Dovresti stampare la tua mappa `catalog` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
