---
id: 6734a13a29acdf47f2b8e161
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

A differenza degli oggetti normali, per i quali le chiavi devono sempre essere stringhe, gli oggetti map ti permettono di usare qualsiasi tipo di dato come chiavi. Nel tuo catalogo, userai coppie chiave-valore oggetto-oggetto, associando a ogni tuo oggetto pianta un oggetto contenente il numero di vasi disponibili per diverse dimensioni.

Modifica la tua chiamata `set` per usare il tuo oggetto `ballerina` come chiave. Per il valore, usa un oggetto contenente tre proprietà, `small`, `medium` e `large`, ognuna con il valore di un numero intero positivo.

# --hints--

Dovresti impostare il tuo oggetto `ballerina` come chiave della tua mappa `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Il valore della chiave `ballerina` nella mappa `catalog` dovrebbe essere un oggetto con le proprietà `small`, `medium` e `large`.

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

La proprietà `small` dell'oggetto associato alla chiave `ballerina` nella mappa `catalog` dovrebbe essere un numero intero positivo.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

La proprietà `medium` dell'oggetto associato alla chiave `ballerina` nella mappa `catalog` dovrebbe essere un numero intero positivo.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

La proprietà `large` dell'oggetto associato alla chiave `ballerina` nella mappa `catalog` dovrebbe essere un numero intero positivo.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

La tua mappa `catalog` dovrebbe contenere una singola coppia chiave-valore.

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
