---
id: 6734a606e5968a5a19b77171
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Continua ad aggiungere le piante rimanenti al tuo catalogo come hai fatto nei passaggi precedenti.

# --hints--

Dovresti impostare il tuo oggetto `willowVale` come chiave della tua mappa `catalog`.

```js
assert.isTrue(catalog.has(willowVale));
```

Il valore della chiave `willowVale` nella mappa `catalog` dovrebbe essere un oggetto con le proprietà `small`, `medium` e `large`, ciascuna con un valore intero positivo.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Dovresti impostare il tuo oggetto `hidcote` come chiave della tua mappa `catalog`.

```js
assert.isTrue(catalog.has(hidcote));
```

Il valore della chiave `hidcote` nella mappa `catalog` dovrebbe essere un oggetto con le proprietà `small`, `medium` e `large`, ciascuna con un valore intero positivo.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Dovresti impostare il tuo oggetto `imperialGem` come chiave della tua mappa `catalog`.

```js
assert.isTrue(catalog.has(imperialGem));
```

Il valore della chiave `imperialGem` nella mappa `catalog` dovrebbe essere un oggetto con le proprietà `small`, `medium` e `large`, ciascuna con un valore intero positivo.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Dovresti impostare il tuo oggetto `royalCrown` come chiave della tua mappa `catalog`.

```js
assert.isTrue(catalog.has(royalCrown));
```

Il valore della chiave `royalCrown` nella mappa `catalog` dovrebbe essere un oggetto con le proprietà `small`, `medium` e `large`, ciascuna con un valore intero positivo.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Dovresti ancora avere la chiave `ballerina` nella tua mappa `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

Dovresti ancora avere la chiave `prettyPolly` nella tua mappa `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

La tua mappa `catalog` dovrebbe contenere sei elementi.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
