---
id: 6734a2385c52c54d916a8b86
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Aggiungi un’altra coppia chiave-valore al tuo catalogo. Usa la tua variabile `prettyPolly` come chiave e fai in modo che il valore sia un oggetto con le proprietà `small`, `medium` e `large` i cui valori siano di nuovo numeri interi.

# --hints--

Dovresti impostare il tuo oggetto `prettyPolly` come chiave della tua mappa `catalog`.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Il valore della chiave `prettyPolly` nella mappa `catalog` dovrebbe essere un oggetto con le proprietà `small`, `medium` e `large`.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

La proprietà `small` dell’oggetto associato alla chiave `prettyPolly` nella mappa `catalog` dovrebbe essere un numero intero positivo.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

La proprietà `medium` dell’oggetto associato alla chiave `prettyPolly` nella mappa `catalog` dovrebbe essere un numero intero positivo.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

La proprietà `large` dell’oggetto associato alla chiave `prettyPolly` nella mappa `catalog` dovrebbe essere un numero intero positivo.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Dovresti avere ancora la chiave `ballerina` nella tua mappa `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

La tua mappa `catalog` dovrebbe contenere due elementi.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
