---
id: 6734a772196245601e4b80e0
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Quando hai bisogno di recuperare un elemento da una mappa, puoi usare il metodo `get`. Questo metodo prende una chiave come argomento e restituisce il valore associato.

Usa il metodo `get` per recuperare il valore della chiave `ballerina` dal tuo catalogo e assegnalo a una variabile chiamata `getBallerina`. Poi, stampa la tua nuova variabile nella console.

# --hints--

Dovresti avere una variabile chiamata `getBallerina`.

```js
assert.exists(getBallerina);
```

La tua variabile `getBallerina` dovrebbe avere il valore di `catalog.get(ballerina)`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+getBallerina\s*=\s*catalog\s*\.\s*get\s*\(\s*ballerina\s*\)/);
```

Dovresti stampare la tua variabile `getBallerina` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*getBallerina\s*\)/);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--

--fcc-editable-region--
console.log(catalog);
```
