---
id: 594810f028c0303b75339acb
title: 100 porte
challengeType: 1
forumTopicId: 302217
dashedName: 100-doors
---

# --description--

Ci sono 100 porte in fila, tutte inizialmente chiuse. Fai 100 passaggi davanti alle porte. La prima volta, visita ogni porta e "commuta" la porta (se la porta è chiusa, aprila; se è aperta, chiudila). La seconda volta, visita solo ogni seconda porta (cioè la porta #2, #4, #6, ...) e commutala. La terza volta, visita ogni terza porta (cioè la porta #3, #6, #9, ...), e così via, fino a visitare solo la porta numero 100.

# --instructions--

Implementa una funzione per determinare lo stato delle porte dopo l’ultimo passaggio. Ritorna il risultato finale in un array, includendo nell’array solo il numero della porta se è aperta.

# --hints--

`getFinalOpenedDoors` dovrebbe essere una funzione.

```js
assert(typeof getFinalOpenedDoors === 'function');
```

`getFinalOpenedDoors` dovrebbe restituire un array.

```js
assert(Array.isArray(getFinalOpenedDoors(100)));
```

`getFinalOpenedDoors` dovrebbe produrre il risultato corretto.

```js
const solution = [1, 4, 9, 16, 25, 36, 49, 64, 81, 100];
assert.deepEqual(getFinalOpenedDoors(100), solution);
```

# --seed--

## --seed-contents--

```js
function getFinalOpenedDoors(numDoors) {

}
```

# --solutions--

```js
function getFinalOpenedDoors(numDoors) {
  // this is the final pattern (always squares).
  // thus, the most efficient solution simply returns an array of squares up to numDoors).
  const finalState = [];
  let i = 1;
  while (Math.pow(i, 2) <= numDoors) {
    finalState.push(Math.pow(i, 2));
    i++;
  }
  return finalState;
}
```
