---
id: 594810f028c0303b75339ad4
title: A capo automatico
challengeType: 1
forumTopicId: 302344
dashedName: word-wrap
---

# --description--

Anche oggi, con caratteri tipografici proporzionali e layout complessi, ci sono ancora casi in cui è necessario andare a capo a una colonna specifica. L'attività di base è andare a capo in modo semplice in un paragrafo di testo.

# --instructions--

Scrivi una funzione che possa andare a capo in questo testo a un numero qualsiasi di caratteri. Nota che il testo di input contiene già interruzioni di riga, che la tua funzione deve gestire correttamente. Come esempio, il testo andato a capo a 80 caratteri dovrebbe apparire così:

<pre>
Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX
algorithm. If your language provides this, you get easy extra credit, but you
must reference documentation indicating that the algorithm is something better
than a simple minimum length algorithm.
</pre>

# --before-each--

```js
const text =
`Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX algorithm.
If your language provides this, you get easy extra credit,
but you ''must reference documentation'' indicating that the algorithm
is something better than a simple minimum length algorithm.`;

const wrapped80 = wrap(text, 80);
const wrapped42 = wrap(text, 42);

const firstRow80 =
    'Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX';

const firstRow42 = 'Wrap text using a more sophisticated';
```

# --hints--

`wrap` dovrebbe essere una funzione.

```js
assert.equal(typeof wrap, 'function');
```

`wrap` dovrebbe restituire una stringa.

```js
assert.equal(typeof wrap('abc', 10), 'string');
```

`wrap(text,80)` dovrebbe restituire 4 righe.

```js
assert(wrapped80.split('\n').length === 4);
```

La tua funzione `wrap` dovrebbe restituire il testo atteso.

```js
assert.equal(wrapped80.split('\n')[0], firstRow80);
```

`wrap(text,42)` dovrebbe restituire 7 righe.

```js
assert(wrapped42.split('\n').length === 7);
```

La tua funzione `wrap` dovrebbe restituire il testo atteso.

```js
assert.equal(wrapped42.split('\n')[0], firstRow42);
```

# --seed--

## --seed-contents--

```js
function wrap(text, limit) {
  return text;
}
```

# --solutions--

```js
function wrap(text, limit) {
  const noNewlines = text.replace('\n', '');
  if (noNewlines.length > limit) {
    // find the last space within limit
    const edge = noNewlines.slice(0, limit).lastIndexOf(' ');
    if (edge > 0) {
      const line = noNewlines.slice(0, edge);
      const remainder = noNewlines.slice(edge + 1);
      return line + '\n' + wrap(remainder, limit);
    }
  }
  return text;
}
```
