---
id: 5949b579404977fbaefcd736
title: "9 miliardi di nomi di Dio l'intero"
challengeType: 1
forumTopicId: 302219
dashedName: 9-billion-names-of-god-the-integer
---

# --description--

Questa attività è una variazione del racconto breve di Arthur C. Clarke.

(Chi risolve dovrebbe essere consapevole delle conseguenze del completare questa attività.)

In dettaglio, per specificare cosa si intende per "nome":

<ul>
  <li>L'intero 1 ha 1 nome "1".</li>
  <li>L'intero 2 ha 2 nomi "1+1" e "2".</li>
  <li>L'intero 3 ha 3 nomi "1+1+1", "2+1" e "3".</li>
  <li>L'intero 4 ha 5 nomi "1+1+1+1", "2+1+1", "2+2", "3+1", "4".</li>
  <li>L'intero 5 ha 7 nomi "1+1+1+1+1", "2+1+1+1", "2+2+1", "3+1+1", "3+2", "4+1", "5".</li>
</ul>

Questo può essere visualizzato nella seguente forma:

<pre>          1
        1   1
      1   1   1
    1   2   1   1
  1   2   2   1   1
1   3   3   2   1   1
</pre>

Dove la riga $n$ corrisponde all'intero $n$, e ogni colonna $C$ nella riga $m$ da sinistra a destra corrisponde al numero di nomi che iniziano con $C$.

Nota opzionale che la somma della $n$-esima riga $P(n)$ è la funzione di partizione intera.

# --instructions--

Implementa una funzione che restituisce la somma della $n$-esima riga.

# --hints--

`numberOfNames` dovrebbe essere function.

```js
assert(typeof numberOfNames === 'function');
```

`numberOfNames(5)` dovrebbe essere uguale a 7.

```js
assert.equal(numberOfNames(5), 7);
```

`numberOfNames(12)` dovrebbe essere uguale a 77.

```js
assert.equal(numberOfNames(12), 77);
```

`numberOfNames(18)` dovrebbe essere uguale a 385.

```js
assert.equal(numberOfNames(18), 385);
```

`numberOfNames(23)` dovrebbe essere uguale a 1255.

```js
assert.equal(numberOfNames(23), 1255);
```

`numberOfNames(42)` dovrebbe essere uguale a 53174.

```js
assert.equal(numberOfNames(42), 53174);
```

`numberOfNames(123)` dovrebbe essere uguale a 2552338241.

```js
assert.equal(numberOfNames(123), 2552338241);
```

# --seed--

## --seed-contents--

```js
function numberOfNames(num) {

  return true;
}
```

# --solutions--

```js
function numberOfNames(num) {
  const cache = [
    [1]
  ];
  for (let l = cache.length; l < num + 1; l++) {
    let Aa;
    let Mi;
    const r = [0];
    for (let x = 1; x < l + 1; x++) {
      r.push(r[r.length - 1] + (Aa = cache[l - x < 0 ? cache.length - (l - x) : l - x])[(Mi = Math.min(x, l - x)) < 0 ? Aa.length - Mi : Mi]);
    }
    cache.push(r);
  }
  return cache[num][cache[num].length - 1];
}
```
