---
id: 594faaab4e2a8626833e9c3d
title: Tokenizzare una stringa con escape
challengeType: 1
forumTopicId: 302338
dashedName: tokenize-a-string-with-escaping
---

# --description--

Scrivi una funzione o un programma che possa dividere una stringa in corrispondenza di ogni occorrenza non escape di un carattere separatore.

Deve accettare tre parametri di input:

<ul>
  <li>La <strong>stringa</strong></li>
  <li>Il <strong>carattere separatore</strong></li>
  <li>Il <strong>carattere di escape</strong></li>
</ul>

Deve restituire una lista di stringhe.

Regole per la divisione:

<ul>
  <li>I campi separati dai separatori diventano gli elementi della lista di output.</li>
  <li>I campi vuoti devono essere preservati, anche all'inizio e alla fine.</li>
</ul>

Regole per l'escape:

<ul>
  <li>"Escape" significa preceduto da un'occorrenza del carattere di escape che non è già a sua volta escape.</li>
  <li>Quando il carattere di escape precede un carattere che non ha un significato speciale, conta comunque come escape (ma non fa nulla di speciale).</li>
  <li>Ogni occorrenza del carattere di escape usata per fare escape non deve far parte dell'output.</li>
</ul>

Dimostra che la tua funzione soddisfa il seguente caso di test:

Data la stringa

<pre>one^|uno||three^^^^|four^^^|^cuatro|</pre>

e usando `|` come separatore e `^` come carattere di escape, la tua funzione dovrebbe restituire il seguente array:

<pre>  ['one|uno', '', 'three^^', 'four^|cuatro', '']
</pre>

# --before-each--

```js
const testStr1 = 'one^|uno||three^^^^|four^^^|^cuatro|';
const res1 = ['one|uno', '', 'three^^', 'four^|cuatro', ''];

const testStr2 = 'a@&bcd&ef&&@@hi';
const res2 = ['a&bcd', 'ef', '', '@hi'];

const testStr3 = 'hello^|world|how^are^you^|';
const res3 = ['hello|world', 'howareyou|'];
```

# --hints--

`tokenize` dovrebbe essere una funzione.

```js
assert(typeof tokenize === 'function');
```

`tokenize` dovrebbe restituire un array.

```js
assert(typeof tokenize('a', 'b', 'c') === 'object');
```

`tokenize('one^|uno||three^^^^|four^^^|^cuatro|', '|', '^')` dovrebbe restituire `['one|uno', '', 'three^^', 'four^|cuatro', '']`

```js
assert.deepEqual(tokenize(testStr1, '|', '^'), res1);
```

`tokenize('a@&bcd&ef&&@@hi', '&', '@')` dovrebbe restituire `['a&bcd', 'ef', '', '@hi']`

```js
assert.deepEqual(tokenize(testStr2, '&', '@'), res2);
```

`tokenize('hello^|world|how^are^you^|', '|', '^')` dovrebbe restituire `['hello|world', 'howareyou|']`

```js
assert.deepEqual(tokenize(testStr3, '|', '^'), res3);
```

# --seed--

## --seed-contents--

```js
function tokenize(str, sep, esc) {
  return true;
}
```

# --solutions--

```js
// tokenize :: String -> Character -> Character -> [String]
function tokenize(str, charDelim, charEsc) {
  const dctParse = str.split('')
    .reduce((a, x) => {
      const blnEsc = a.esc;
      const blnBreak = !blnEsc && x === charDelim;
      const blnEscChar = !blnEsc && x === charEsc;

      return {
        esc: blnEscChar,
        token: blnBreak ? '' : (
          a.token + (blnEscChar ? '' : x)
        ),
        list: a.list.concat(blnBreak ? a.token : [])
      };
    }, {
      esc: false,
      token: '',
      list: []
    });

  return dctParse.list.concat(
    dctParse.token
  );
}
```
