---
id: 5951ed8945deab770972ae56
title: Torri di Hanoi
challengeType: 1
forumTopicId: 302341
dashedName: towers-of-hanoi
---

# --description--

Risolvi il problema delle torri di Hanoi. Il numero di oggetti sarà dato come primo parametro, seguito dalle stringhe usate per identificare ogni pila di oggetti. Crea un array annidato contenente la lista delle mosse, `["source", "destination"]`.
 
Per esempio, i parametri `(4, 'A', 'B', 'C')` produrranno un array annidato di mosse `[['A', 'C'], ['B', 'A']]`, che indica che la prima mossa è stata spostare un oggetto dalla pila `A` a `C` e la seconda mossa è stata spostare un oggetto dalla pila `B` a `A`.

# --instructions--

Scrivi una funzione che restituisca le mosse per impilare gli oggetti in un array annidato.

# --before-each--

```js
const res3Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B']];
const res7First10Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B'], ['A', 'C'], ['B', 'C'], ['B', 'A']];

function getTowerResults_() {
  return {
    res3: towerOfHanoi(3, 'A', 'B', 'C'),
    res5: towerOfHanoi(5, 'X', 'Y', 'Z'),
    res7: towerOfHanoi(7, 'A', 'B', 'C')
  };
}
```

# --hints--

`towerOfHanoi` dovrebbe essere una funzione.

```js
assert(typeof towerOfHanoi === 'function');
```

`towerOfHanoi(3, ...)` dovrebbe restituire 7 mosse.

```js
assert(getTowerResults_().res3.length === 7);
```

`towerOfHanoi(3, 'A', 'B', 'C')` dovrebbe restituire `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B']]`.

```js
assert.deepEqual(getTowerResults_().res3, res3Moves);
```

La decima mossa `towerOfHanoi(5, "X", "Y", "Z")` dovrebbe essere Y -> X.

```js
const { res5 } = getTowerResults_();
assert.deepEqual(res5[9], ['Y', 'X']);
```

Le prime dieci mosse `towerOfHanoi(7, 'A', 'B', 'C')` dovrebbero essere `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B'], ['A','C'], ['B','C'], ['B','A']]`

```js
const { res7 } = getTowerResults_();
assert.deepEqual(res7.slice(0, 10), res7First10Moves);
```

# --seed--

## --seed-contents--

```js
function towerOfHanoi(n, a, b, c) {

  return [[]];
}
```

# --solutions--

```js
function towerOfHanoi(n, a, b, c) {
  const res = [];
  towerOfHanoiHelper(n, a, c, b, res);
  return res;
}

function towerOfHanoiHelper(n, a, b, c, res) {
  if (n > 0) {
    towerOfHanoiHelper(n - 1, a, c, b, res);
    res.push([a, c]);
    towerOfHanoiHelper(n - 1, b, a, c, res);
  }
}
```
