---
id: 595608ff8bcd7a50bd490181
title: sequenza hailstone
challengeType: 1
forumTopicId: 302279
dashedName: hailstone-sequence
---

# --description--

La sequenza hailstone di numeri può essere generata da un intero positivo iniziale, `n`, seguendo queste regole:

- Se `n` è `1` allora la sequenza termina
- Se `n` è `even` allora il prossimo `n` della sequenza è `= n/2`
- Se `n` è `odd` allora il prossimo `n` della sequenza è `= (3 * n) + 1`

La congettura di Collatz (non dimostrata) afferma che la sequenza hailstone per qualsiasi numero iniziale termina sempre.

La sequenza hailstone è anche conosciuta come numeri hailstone (perché i valori sono solitamente soggetti a discese e risalite multiple come i chicchi di grandine in una nuvola), o come sequenza di Collatz.

# --instructions--

1. Crea una routine per generare la sequenza hailstone per un numero
2. La tua funzione dovrebbe restituire un array con il numero minore di `limit` che ha la sequenza hailstone più lunga e la lunghezza di quella sequenza. (Ma non mostrare la sequenza effettiva!)

# --hints--

`hailstoneSequence` dovrebbe essere una funzione.

```js
assert(typeof hailstoneSequence === 'function');
```

`hailstoneSequence(30)` dovrebbe restituire un array.

```js
assert(Array.isArray(hailstoneSequence(30)));
```

`hailstoneSequence(30)` dovrebbe restituire `[27, 112]`.

```js
assert.deepEqual(hailstoneSequence(30), [27, 112]);
```

`hailstoneSequence(50000)` dovrebbe restituire `[35655, 324]`.

```js
assert.deepEqual(hailstoneSequence(50000), [35655, 324]);
```

`hailstoneSequence(100000)` dovrebbe restituire `[77031, 351]`.

```js
assert.deepEqual(hailstoneSequence(100000), [77031, 351]);
```

# --seed--

## --seed-contents--

```js
function hailstoneSequence(limit) {
  const res = [];


  return res;
}
```

# --solutions--

```js
function hailstoneSequence (limit) {
  function hailstone(n) {
    const seq = [n];
    while (n > 1) {
      n = n % 2 ? 3 * n + 1 : n / 2;
      seq.push(n);
    }
    return seq;
  }

  let n = 0;
  let max = 0;
  for (let i = limit; --i;) {
    const seq = hailstone(i);
    const sLen = seq.length;

    if (sLen > max) {
      n = i;
      max = sLen;
    }
  }

  return [n, max];
}
```
