---
id: 59667989bf71cf555dd5d2ff
title: S-espressioni
challengeType: 1
forumTopicId: 302303
dashedName: s-expressions
---

# --description--

<a href="https://rosettacode.org/wiki/S-expressions" target="_blank" rel="noopener noreferrer nofollow">Le S-espressioni</a> sono un modo comodo per analizzare e archiviare dati.

# --instructions--

Scrivi un semplice lettore/parser per S-espressioni che gestisca stringhe quotate e non quotate, numeri interi e numeri in virgola mobile.

La funzione deve leggere una singola S-espressione annidata da una stringa e restituirla come un array (annidato).

Le nuove righe e altri spazi bianchi possono essere ignorati a meno che non siano contenuti all’interno di una stringa quotata.

I ``()`` all’interno delle stringhe quotate non devono essere interpretati, ma trattati come parte della stringa.

Gestire le virgolette escape all’interno di una stringa è opzionale; quindi ``(foo"bar)`` può essere trattato come una stringa ``foo"bar`` oppure come un errore.

Per questo, il lettore non deve riconoscere ``\`` per l’escape, ma dovrebbe inoltre riconoscere i numeri se il linguaggio ha tipi di dati appropriati.

Nota che, fatta eccezione per ``()"`` (``\`` se l’escape è supportato) e gli spazi bianchi, non ci sono caratteri speciali. Qualsiasi altra cosa è permessa senza virgolette.

Il lettore dovrebbe essere in grado di leggere il seguente input

<pre>((data "quoted data" 123 4.5)
(data (!@# (4.5) "(more" "data)")))
</pre>

e trasformarlo in una struttura dati nativa.

# --before-each--

```js
const simpleSExpr = '(data1 data2 data3)';
const simpleSolution = ['data1', 'data2', 'data3'];

const basicSExpr = '((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))';
const basicSolution = [["data","\"quoted data\"",123,4.5],["data",["!@#",[4.5],"\"(more\"","\"data)\""]]];
```

# --hints--

``parseSexpr`` dovrebbe essere una funzione.

```js
assert(typeof parseSexpr === 'function');
```

``parseSexpr('(data1 data2 data3)')`` dovrebbe restituire ``['data1', 'data2', 'data3']``

```js
assert.deepEqual(parseSexpr(simpleSExpr), simpleSolution);
```

``parseSexpr('((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))')`` dovrebbe restituire ``[['data', '"quoted data"', 123, 4.5], ['data', ['!@#', [4.5], '"(more"', '"data)"']]]``.

```js
assert.deepEqual(parseSexpr(basicSExpr), basicSolution);
```

# --seed--

## --seed-contents--

```js
function parseSexpr(str) {

  return true;
}
```

# --solutions--

```js
function parseSexpr(str) {
  const t = str.match(/\s*("[^"]*"|\(|\)|"|[^\s()"]+)/g);
  for (var o, c = 0, i = t.length - 1; i >= 0; i--) {
    var n,
      ti = t[i].trim();
    if (ti == '"') return;
    else if (ti == '(') t[i] = '[', c += 1;
    else if (ti == ')') t[i] = ']', c -= 1;
    else if ((n = +ti) == ti) t[i] = n;
    else t[i] = `'${ti.replace('\'', '\\\'')}'`;
    if (i > 0 && ti != ']' && t[i - 1].trim() != '(') t.splice(i, 0, ',');
    if (!c) if (!o) o = true; else return;
  }
  return c ? undefined : eval(t.join(''));
}
```
