---
id: 59669d08d75b60482359409f
title: Formato della data
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

Restituisci un array con due stringhe di data della data corrente con le seguenti specifiche:

- L'ordine della data nella prima stringa deve essere anno, mese e giorno separati da trattini (`-`).
- L'anno nella prima stringa deve essere lungo quattro cifre.
- Il mese e il giorno nella prima stringa non devono contenere zeri iniziali.
- I nomi del giorno della settimana e del mese nella seconda stringa non devono essere abbreviati.
- Il giorno nella seconda stringa non deve contenere zeri iniziali.

Esempi di output:

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

`getDateFormats` dovrebbe essere una funzione.

```js
assert(typeof getDateFormats === 'function');
```

`getDateFormats` dovrebbe restituire un oggetto.

```js
assert(typeof getDateFormats() === 'object');
```

`getDateFormats` dovrebbe restituire un array con 2 elementi.

```js
assert(getDateFormats().length === 2);
```

`getDateFormats` dovrebbe restituire la data corretta nel formato giusto

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
