---
id: 59694356a6e7011f7f1c5f4e
title: Distribuire le carte per FreeCell
challengeType: 1
forumTopicId: 302246
dashedName: deal-cards-for-freecell
---

# --description--

*FreeCell* è il gioco di carte solitario che Paul Alfille ha introdotto nel sistema PLATO nel 1978. Jim Horne, alla Microsoft, ha cambiato il nome in FreeCell e ha reimplementato il gioco per DOS, poi per Windows. Questa versione ha introdotto 32000 distribuzioni numerate.

Con la crescente popolarità del gioco, Jim Horne ha reso pubblico l'algoritmo, e altre implementazioni di FreeCell hanno iniziato a riprodurre le distribuzioni Microsoft. Queste distribuzioni sono numerate da 1 a 32000. Versioni più recenti di Microsoft hanno 1 milione di distribuzioni, numerate da 1 a 1000000; alcune implementazioni permettono numeri fuori da questo intervallo.

L'algoritmo usa questo generatore congruenziale lineare di Microsoft C:

<ul>
  <li>$state_{n + 1} \equiv 214013 \times state_n + 2531011 \pmod{2^{31}}$</li>
  <li>$rand_n = state_n \div 2^{16}$</li>
  <li>$rand_n$ è nell'intervallo da 0 a 32767.</li>
</ul>

L'algoritmo segue questi passi:

<ol>
  <li>Inizializza il generatore di numeri casuali con il numero della distribuzione.
  </li><li>Crea un array di 52 carte: asso di fiori, asso di quadri, asso di cuori, asso di picche, 2 di fiori, 2 di quadri, e così via attraverso i valori: asso, 2, 3, 4, 5, 6, 7, 8, 9, 10, jack, regina, re. Gli indici dell'array vanno da 0 a 51, con asso di fiori a 0 e re di picche a 51.</li>
  <li>Finché l'array non è vuoto:</li>
    <ul>
      <li>Scegli una carta casuale con <i>indice</i> ≡ <i>prossimo numero casuale</i> (mod <i>lunghezza dell'array</i>).</li>
      <li>Scambia questa carta casuale con l'ultima carta dell'array.</li>
      <li>Rimuovi questa carta casuale dall'array. (La lunghezza dell'array diminuisce di 1.)</li>
      <li>Distribuisci questa carta casuale.</li>
    </ul>
  <li>Distribuisci tutte le 52 carte, scoperte, su 8 colonne. Le prime 8 carte vanno nelle 8 colonne, le successive 8 carte vanno sopra le prime 8, e così via.</li>
</ol>

**Esempio:**

**Ordine di distribuzione delle carte**

<pre> 1  2  3  4  5  6  7  8
 9 10 11 12 13 14 15 16
17 18 19 20 21 22 23 24
25 26 27 28 29 30 31 32
33 34 35 36 37 38 39 40
41 42 43 44 45 46 47 48
49 50 51 52</pre>

**Partita #1**

```js
[
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
]
```

**Partita #617**

```js
[
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
]
```

# --instructions--

Scrivi una funzione che prenda un numero di distribuzione e distribuisca le carte nello stesso ordine di questo algoritmo. La funzione deve restituire un array bidimensionale che rappresenta la plancia di FreeCell.

# --before-each--

```js
const replaceThis = 3;
const game1 = [
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
];
const game617 = [
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
];
```

# --hints--

`dealFreeCell` dovrebbe essere una funzione.

```js
assert(typeof dealFreeCell === 'function');
```

`dealFreeCell(seed)` dovrebbe restituire un oggetto.

```js
assert(typeof dealFreeCell(1) === 'object');
```

`dealFreeCell(seed)` dovrebbe restituire un array di lunghezza 7.

```js
assert(dealFreeCell(1).length === 7);
```

`dealFreeCell(1)` dovrebbe restituire un array identico all'esempio "Partita #1"

```js
assert.deepEqual(dealFreeCell(1), game1);
```

`dealFreeCell(617)` dovrebbe restituire un array identico all'esempio "Partita #617"

```js
assert.deepEqual(dealFreeCell(617), game617);
```

# --seed--

## --seed-contents--

```js
function dealFreeCell(seed) {

  return true;
}
```

# --solutions--

```js
// RNG
function FreeCellRNG(seed) {
  return {
    lastNum: seed,
    next() {
      this.lastNum = ((214013 * this.lastNum) + 2531011) % (Math.pow(2, 31));
      return this.lastNum >> 16;
    }
  };
}
// Get cards
function getDeck() {
  const ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', 'T', 'J', 'Q', 'K'];
  const suits = ['C', 'D', 'H', 'S'];
  const cards = [];
  for (let i = 0; i < ranks.length; i += 1) {
    for (let j = 0; j < suits.length; j += 1) {
      cards.push(`${ranks[i]}${suits[j]}`);
    }
  }
  return cards;
}
function dealFreeCell(seed) {
  const rng = FreeCellRNG(seed);
  const deck = getDeck();

  const dealtCards = [[], [], [], [], [], [], []];
  let currentColumn = 0;
  let currentRow = 0;

  let rand;
  let temp;
  let card;
  while (deck.length > 0) {
    // Choose a random card
    rand = rng.next() % deck.length;

    // Swap this random card with the last card in the array
    temp = deck[deck.length - 1];
    deck[deck.length - 1] = deck[rand];
    deck[rand] = temp;

    // Remove this card from the array
    card = deck.pop();

    // Deal this card
    dealtCards[currentRow].push(card);
    currentColumn += 1;
    if (currentColumn === 8) {
      currentColumn = 0;
      currentRow += 1;
    }
  }

  return dealtCards;
}
```
