---
id: 596a8888ab7c01048de257d5
title: Copia profonda
challengeType: 1
forumTopicId: 302247
dashedName: deepcopy
---

# --description--

Scrivi una funzione che restituisca una copia profonda di un oggetto dato. La copia non deve essere lo stesso oggetto fornito.

Questa attività non testerà:

<ul>
  <li>Oggetti con proprietà che sono funzioni</li>
  <li>Oggetti data o oggetti con proprietà che sono oggetti data</li>
  <li>Espressioni regolari o oggetti con proprietà che sono espressioni regolari</li>
  <li>Copia del prototipo</li>
</ul>

# --before-each--

```js
const obj1 = { test: 'test' };
const obj2 = {
  t: 'test',
  a: ['an', 'array']
};
const obj3 = {
  t: 'try',
  o: obj2
};
```

# --hints--

`deepcopy` dovrebbe essere una funzione.

```js
assert(typeof deepcopy === 'function');
```

`deepcopy({test: "test"})` dovrebbe restituire un oggetto.

```js
assert(typeof deepcopy(obj1) === 'object');
```

`deepcopy` non dovrebbe restituire lo stesso oggetto fornito.

```js
assert(deepcopy(obj2) != obj2);
```

Quando riceve un oggetto contenente un array, `deepcopy` dovrebbe restituire una copia profonda dell'oggetto.

```js
assert.deepEqual(deepcopy(obj2), obj2);
```

Quando riceve un oggetto contenente un altro oggetto, `deepcopy` dovrebbe restituire una copia profonda dell'oggetto.

```js
assert.deepEqual(deepcopy(obj3), obj3);
```

# --seed--

## --seed-contents--

```js
function deepcopy(obj) {

  return true;
}
```

# --solutions--

```js
function deepcopy(obj) {
  return JSON.parse(JSON.stringify(obj));
}
```
