---
id: 596e414344c3b2872167f0fe
title: Disputa sulla virgola
challengeType: 1
forumTopicId: 302234
dashedName: comma-quibbling
---

# --description--

<a href="https://rosettacode.org/wiki/Comma_quibbling" target="_blank" rel="noopener noreferrer nofollow">Disputa sulla virgola</a> è un'attività originariamente proposta da Eric Lippert nel suo blog.

# --instructions--

Scrivi una funzione per generare una stringa di output che sia la concatenazione delle parole di input da una lista/sequenza dove:

<ol>
  <li>Un input senza parole produce la stringa di output composta solo dai due caratteri graffa (<code>"{}"</code>)</li>
  <li>Un input con una sola parola, ad esempio <code>["ABC"]</code>, produce la stringa di output con la parola racchiusa tra le due graffe, ad esempio <code>"{ABC}"</code></li>
  <li>Un input con due parole, ad esempio <code>["ABC", "DEF"]</code>, produce la stringa di output con le due parole racchiuse tra le graffe e separate dalla stringa <code>" and "</code>, ad esempio <code>"{ABC and DEF}"</code></li>
  <li>Un input con tre o più parole, ad esempio <code>["ABC", "DEF", "G", "H"]</code>, produce la stringa di output con tutte le parole tranne l'ultima separate da <code>", "</code>, l'ultima parola separata da <code>" and "</code> e tutte racchiuse tra le graffe; ad esempio <code>"{ABC, DEF, G and H}"</code></li>
</ol>

Metti alla prova la tua funzione con la seguente serie di input mostrando il tuo output qui su questa pagina:

<ul>
  <li>[] # (Nessuna parola di input).</li>
  <li>["ABC"]</li>
  <li>["ABC", "DEF"]</li>
  <li>["ABC", "DEF", "G", "H"]</li>
</ul>

**Nota:** Per questa attività, considera che le parole sono stringhe non vuote di caratteri maiuscoli.

# --before-each--

```js
const testCases = [[], ["ABC"], ["ABC", "DEF"], ["ABC", "DEF", "G", "H"]];
const results = ["{}", "{ABC}", "{ABC and DEF}", "{ABC, DEF, G and H}"];
```

# --hints--

`quibble` dovrebbe essere una funzione.

```js
assert(typeof quibble === 'function');
```

`quibble(["ABC"])` dovrebbe restituire una stringa.

```js
assert(typeof quibble(['ABC']) === 'string');
```

`quibble([])` dovrebbe restituire "{}".

```js
assert.equal(quibble(testCases[0]), results[0]);
```

`quibble(["ABC"])` dovrebbe restituire `"{ABC}"`.

```js
assert.equal(quibble(testCases[1]), results[1]);
```

`quibble(["ABC", "DEF"])` dovrebbe restituire `"{ABC and DEF}"`.

```js
assert.equal(quibble(testCases[2]), results[2]);
```

`quibble(["ABC", "DEF", "G", "H"])` dovrebbe restituire `"{ABC, DEF, G and H}"`.

```js
assert.equal(quibble(testCases[3]), results[3]);
```

# --seed--

## --seed-contents--

```js
function quibble(words) {

  return true;
}
```

# --solutions--

```js
function quibble(words) {
  return "{" +
    words.slice(0, words.length - 1).join(", ") +
   (words.length > 1 ? " and " : "") +
   (words[words.length - 1] || '') +
  "}";
}
```
