---
id: 596e457071c35c882915b3e4
title: Confronta una lista di stringhe
challengeType: 1
forumTopicId: 302235
dashedName: compare-a-list-of-strings
---

# --description--

Una lista è un insieme ordinato di valori che può contenere duplicati. Ecco un esempio:

```js
const list = [['AA',  'BB', 'CC'], ['AA', 'ACB', 'AA'], [], ['AA']];
````

Data una lista di un numero arbitrario di stringhe, implementa una funzione per ciascuna delle seguenti condizioni:

<ul>
  <li>verificare se sono tutte lessicalmente uguali</li>
  <li>verificare se ogni stringa è lessicalmente minore di quella che la segue (cioè se la lista è in ordine ascendente stretto)</li>
</ul>

# --before-each--

```js
const testCases = [['AA', 'AA', 'AA', 'AA'], ['AA', 'ACB', 'BB', 'CC'], [], ['AA'], ['BB', 'AA']];
```

# --hints--

`allEqual` dovrebbe essere una funzione.

```js
assert(typeof allEqual === 'function');
```

`azSorted` dovrebbe essere una funzione.

```js
assert(typeof azSorted === 'function');
```

`allEqual(["AA", "AA", "AA", "AA"])` dovrebbe restituire true.

```js
assert(allEqual(testCases[0]));
```

`azSorted(["AA", "AA", "AA", "AA"])` dovrebbe restituire false.

```js
assert(!azSorted(testCases[0]));
```

`allEqual(["AA", "ACB", "BB", "CC"])` dovrebbe restituire false.

```js
assert(!allEqual(testCases[1]));
```

`azSorted(["AA", "ACB", "BB", "CC"])` dovrebbe restituire true.

```js
assert(azSorted(testCases[1]));
```

`allEqual([])` dovrebbe restituire true.

```js
assert(allEqual(testCases[2]));
```

`azSorted([])` dovrebbe restituire true.

```js
assert(azSorted(testCases[2]));
```

`allEqual(["AA"])` dovrebbe restituire true.

```js
assert(allEqual(testCases[3]));
```

`azSorted(["AA"])` dovrebbe restituire true.

```js
assert(azSorted(testCases[3]));
```

`allEqual(["BB", "AA"])` dovrebbe restituire false.

```js
assert(!allEqual(testCases[4]));
```

`azSorted(["BB", "AA"])` dovrebbe restituire false.

```js
assert(!azSorted(testCases[4]));
```

# --seed--

## --seed-contents--

```js
function allEqual(arr) {

  return true;
}

function azSorted(arr) {

  return true;
}
```

# --solutions--

```js
function allEqual(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] === a[i]);
  } return out;
}

function azSorted(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] < a[i]);
  } return out;
}
```
