---
id: 59e0a8df964e4540d5abe599
title: Esegui brain****
challengeType: 1
forumTopicId: 302261
dashedName: execute-brain
---

# --description--

Scrivi una funzione per implementare un interprete Brain\*\*\*\*. La funzione riceverà una stringa come parametro e dovrebbe restituire una stringa come output. Ulteriori dettagli sono forniti di seguito:

RCBF è un insieme di <a href="https://rosettacode.org/wiki/Brainf***" target="_blank" rel="noopener noreferrer nofollow">compilatori e interpreti Brainf\*\*\*</a> scritti per Rosetta Code in diversi linguaggi.

Di seguito ci sono i collegamenti a ciascuna delle versioni di RCBF.

Un'implementazione deve solo implementare correttamente le seguenti istruzioni:

| Comando           | Descrizione                                                                        |
| ----------------- | ---------------------------------------------------------------------------------- |
| <code>></code>    | Sposta il puntatore a destra                                                       |
| <code>&lt;</code> | Sposta il puntatore a sinistra                                                     |
| <code>+</code>    | Incrementa la cella di memoria sotto il puntatore                                 |
| <code>-</code>    | Decrementa la cella di memoria sotto il puntatore                                 |
| <code>.</code>    | Stampa il carattere indicato dalla cella sotto il puntatore                       |
| <code>,</code>    | Inserisci un carattere e salvalo nella cella sotto il puntatore                   |
| <code>\[</code>   | Salta oltre il corrispondente <code>]</code> se la cella sotto il puntatore è 0   |
| <code>]</code>    | Torna al corrispondente <code>\[</code> se la cella sotto il puntatore è diversa da zero |

Qualsiasi dimensione della cella è consentita, il supporto EOF (*E*nd-*O*-*F*ile) è opzionale, così come la memoria limitata o illimitata.

# --hints--

`brain(hello)` dovrebbe restituire una stringa

```js
assert(typeof brain(hello) === 'string');
```

`brain("++++++[>++++++++++<-]>+++++.")` dovrebbe restituire "A"

```js
assert.equal(brain('++++++[>++++++++++<-]>+++++.'), 'A');
```

`brain(hello)` dovrebbe restituire `Hello World!\n`

```js
assert.equal(brain(hello), 'Hello World!\n');
```

`brain(fib)` dovrebbe restituire `1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89`

```js
assert.equal(brain(fib), '1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89');
```

# --seed--

## --before-user-code--

```js
let fib=`+

++

+++

++++

+>+>>

>>++++

+++++++

++++++++

+++++++++

++++++++++

++++++>++++

++++++++++++

+++++++++++++

+++<<<<<<[>[>>

>>>>+>+<<<<<<<-

]>>>>>>>[<<<<<<<

+>>>>>>>-]<[>++++

++++++[-<-[>>+>+<<

<-]>>>[<<<+>>>-]+<[

>[-]<[-]]>[<<[>>>+<<

<-]>>[-]]<<]>>>[>>+>+

<<<-]>>>[<<<+>>>-]+<[>

[-]<[-]]>[<<+>>[-]]<<<<

<<<]>>>>>[++++++++++++++

+++++++++++++++++++++++++

+++++++++.[-]]++++++++++<[

->-<]>+++++++++++++++++++++

+++++++++++++++++++++++++++.

[-]<<<<<<<<<<<<[>>>+>+<<<<-]>

>>>[<<<<+>>>>-]<-[>>.>.<<<[-]]

<<[>>+>+<<<-]>>>[<<<+>>>-]<<[<+

>-]>[<+>-]<<<-]`;
let hello='++++++++[>++++[>++>+++>+++>+<<<<-]>+>+>->>+[<]<-]>>.>---.+++++++..+++.>>.<-.<.+++.------.--------.>>+.>++.'
```

## --seed-contents--

```js
function brain(prog) {

}
```

# --solutions--

```js
function brain(prog){
  var output="";
    var code; // formatted code
  var ip = 0; // current instruction within code
  var nest = 0; // current bracket nesting (for Out button)
  var ahead = []; // locations of matching brackets

  var data = [0]; // data array (mod by +, -)
  var dp = 0; // index into data (mod by <, >)

  var inp = 0; // current input character (fetch with ,)
  var quit = 0;
    var commands = {
    '>':function() { if (++dp >= data.length) data[dp]=0 },
    '<':function() { if (--dp < 0) quit++ },
    '+':function() { ++data[dp] },
    '-':function() { --data[dp] },
    '[':function() { if (!data[dp]) ip = ahead[ip]; else ++nest },
    ']':function() { if ( data[dp]) ip = ahead[ip]; else --nest },
    ',':function() {
        var c = document.getElementById("input").value.charCodeAt(inp++);
        data[dp] = isNaN(c) ? 0 : c; // EOF: other options are -1 or no change
    },
    '.':function() {
            output+=String.fromCharCode(data[dp]);
            /*var s = document.getElementById("output").innerHTML)
             + String.fromCharCode(data[dp]);
            s = s.replace(/\n/g,"<br>").replace(/ /g,"&amp;nbsp;");
            document.getElementById("output").innerHTML = s;*/
        },
    };

    let ar=prog.split('');
    var st = [], back, error = -1;
    for (ip=0; ip<ar.length; ip++) {
        switch(ar[ip]) {
        case '[':
            st.push(ip);
            break;
        case ']':
            if (st.length == 0) error = ip;
            back = st.pop();
            ahead[ip] = back;
            ahead[back] = ip;
            break;
        }
    }

    for(ip=0;ip<ar.length;ip++){
    if(commands.hasOwnProperty(ar[ip]))
          commands[ar[ip]]();
    }

    return output;
}
```
