---
id: 5a23c84252665b21eecc8028
title: Sequenza di Stern-Brocot
challengeType: 1
forumTopicId: 302324
dashedName: stern-brocot-sequence
---

# --description--

Per questa attività, la sequenza di Stern-Brocot deve essere generata da un algoritmo simile a quello usato per generare la <a href="https://rosettacode.org/wiki/Fibonacci_sequence" target="_blank" rel="noopener noreferrer nofollow">sequenza di Fibonacci</a>.

<ol>
  <li>I primi due membri della sequenza sono entrambi 1:</li>
    <ul><li>1, 1</li></ul>
  <li>Inizia considerando il secondo membro della sequenza</li>
  <li>Somma il membro considerato della sequenza e il suo precedente, (1 + 1) = 2, e aggiungilo alla fine della
    sequenza:</li>
    <ul><li>1, 1, 2</li></ul>
  <li>Aggiungi il membro considerato della sequenza alla fine della sequenza:</li>
    <ul><li>1, 1, 2, 1</li></ul>
  <li>Considera il membro successivo della serie, (il terzo membro cioè 2)</li>
  <li>TORNA al punto 3 </li>
    <ul>
      <li></li>
      <li> ─── Espandendo un altro loop otteniamo: ───</li>
      <li></li>
    </ul>
  <li>Somma il membro considerato della sequenza e il suo precedente, (2 + 1) = 3, e aggiungilo alla fine della
    sequenza:</li>
    <ul><li>1, 1, 2, 1, 3</li></ul>
  <li>Aggiungi il membro considerato della sequenza alla fine della sequenza:</li>
    <ul><li>1, 1, 2, 1, 3, 2</li></ul>
  <li>Considera il membro successivo della serie, (il quarto membro cioè 1)</li>
</ol>

# --instructions--

Crea una funzione che restituisca la posizione nella sequenza di Stern-Brocot in cui $ n $ viene incontrato per la prima volta, dove la sequenza è generata con il metodo descritto sopra. Nota che questa sequenza usa un indice basato su 1.

# --hints--

`sternBrocot` dovrebbe essere una funzione.

```js
assert(typeof sternBrocot == 'function');
```

`sternBrocot(2)` dovrebbe restituire un numero.

```js
assert(typeof sternBrocot(2) == 'number');
```

`sternBrocot(2)` dovrebbe restituire `3`.

```js
assert.equal(sternBrocot(2), 3);
```

`sternBrocot(3)` dovrebbe restituire `5`.

```js
assert.equal(sternBrocot(3), 5);
```

`sternBrocot(5)` dovrebbe restituire `11`.

```js
assert.equal(sternBrocot(5), 11);
```

`sternBrocot(7)` dovrebbe restituire `19`.

```js
assert.equal(sternBrocot(7), 19);
```

`sternBrocot(10)` dovrebbe restituire `39`.

```js
assert.equal(sternBrocot(10), 39);
```

# --seed--

## --seed-contents--

```js
function sternBrocot(num) {

}
```

# --solutions--

```js
function sternBrocot(num) {
  function f(n) {
    return n < 2
      ? n
      : n & 1
      ? f(Math.floor(n / 2)) + f(Math.floor(n / 2 + 1))
      : f(Math.floor(n / 2));
  }

  function gcd(a, b) {
    return a ? (a < b ? gcd(b % a, a) : gcd(a % b, b)) : b;
  }
  var n;
  for (n = 1; f(n) != num; n++);
  return n;
}
```
