---
id: 5a23c84252665b21eecc8045
title: Clipping di poligoni con Sutherland-Hodgman
challengeType: 1
forumTopicId: 302336
dashedName: sutherland-hodgman-polygon-clipping
---

# --description--

L'algoritmo di clipping di Sutherland-Hodgman trova il poligono che rappresenta l'intersezione tra un poligono arbitrario (il "poligono soggetto") e un poligono convesso (il "poligono di clipping"). Viene usato nella grafica computerizzata (soprattutto nella grafica 2D) per ridurre la complessità di una scena visualizzata eliminando le parti di un poligono che non devono essere mostrate. Prendi il poligono chiuso definito dai punti:

<pre>[(50, 150), (200, 50), (350, 150), (350, 300), (250, 300), (200, 250), (150, 350), (100, 250), (100, 200)]</pre>

e ritaglialo con il rettangolo definito dai punti:

<pre>[(100, 100), (300, 100), (300, 300), (100, 300)]</pre>

# --instructions--

Scrivi una funzione che prenda 2 array come parametri. Il primo array contiene i punti del poligono soggetto e il secondo array contiene i punti del poligono di clipping. La funzione deve restituire un array contenente i punti del poligono ritagliato. Ogni numero deve essere arrotondato a 3 cifre decimali.

# --before-each--

```js
const subjectPolygon1 = [
  [50, 150],
  [200, 50],
  [350, 150],
  [350, 300],
  [250, 300],
  [200, 250],
  [150, 350],
  [100, 250],
  [100, 200]
];
const clipPolygon1 = [
  [100, 100],
  [300, 100],
  [300, 300],
  [100, 300]
];
const expected1 = [
  [100, 116.667],
  [125, 100],
  [275, 100],
  [300, 116.667],
  [300, 300],
  [250, 300],
  [200, 250],
  [175, 300],
  [125, 300],
  [100, 250]
];

const subjectPolygon2 = [
  [150, 200],
  [400, 450],
  [30, 50]
];
const clipPolygon2 = [
  [10, 10],
  [300, 200],
  [400, 600],
  [100, 300]
];
const expected2 = [
  [150, 200],
  [350, 400],
  [348.611, 394.444],
  [30, 50]
];

const subjectPolygon3 = [
  [250, 200],
  [100, 450],
  [130, 250]
];
const clipPolygon3 = [
  [50, 60],
  [100, 230],
  [400, 600],
  [100, 300]
];
const expected3 = [
  [129.167, 329.167],
  [119.565, 319.565],
  [121.854, 304.305]
];
```

# --hints--

`clip` dovrebbe essere una funzione.

```js
assert(typeof clip == 'function');
```

`clip([[50, 150], [200, 50], [350, 150], [350, 300], [250, 300], [200, 250], [150, 350], [100, 250], [100, 200]], [[100, 100], [300, 100], [300, 300], [100, 300]])` dovrebbe restituire un array.

```js
assert(Array.isArray(clip(subjectPolygon1, clipPolygon1)));
```

`clip([[50, 150], [200, 50], [350, 150], [350, 300], [250, 300], [200, 250], [150, 350], [100, 250], [100, 200]], [[100, 100], [300, 100], [300, 300], [100, 300]])` dovrebbe restituire `[[100, 116.667], [125, 100], [275, 100], [300, 116.667], [300, 300], [250, 300], [200, 250], [175, 300], [125, 300], [100, 250]]`.

```js
assert.deepEqual(clip(subjectPolygon1, clipPolygon1), expected1);
```

`clip([[150, 200], [400, 450], [30, 50]], [[10, 10], [300, 200], [400, 600], [100, 300]])` dovrebbe restituire `[[150, 200], [350, 400], [348.611, 394.444], [30, 50]]`.

```js
assert.deepEqual(clip(subjectPolygon2, clipPolygon2), expected2);
```

`clip([[250, 200], [100, 450], [130, 250]], [[50, 60], [100, 230], [400, 600], [100, 300]])` dovrebbe restituire `[[129.167, 329.167], [119.565, 319.565], [121.854, 304.305]]`.

```js
assert.deepEqual(clip(subjectPolygon3, clipPolygon3), expected3);
```

# --seed--

## --seed-contents--

```js
function clip(subjectPolygon, clipPolygon) {

}
```

# --solutions--

```js
function clip(subjectPolygon, clipPolygon) {
  var cp1, cp2, s, e, i, j;
  var inside = function(p) {
    return (
      (cp2[0] - cp1[0]) * (p[1] - cp1[1]) > (cp2[1] - cp1[1]) * (p[0] - cp1[0])
    );
  };
  var intersection = function() {
    var dc = [cp1[0] - cp2[0], cp1[1] - cp2[1]],
      dp = [s[0] - e[0], s[1] - e[1]],
      n1 = cp1[0] * cp2[1] - cp1[1] * cp2[0],
      n2 = s[0] * e[1] - s[1] * e[0],
      n3 = 1.0 / (dc[0] * dp[1] - dc[1] * dp[0]);
    return [(n1 * dp[0] - n2 * dc[0]) * n3, (n1 * dp[1] - n2 * dc[1]) * n3];
  };
  var outputList = subjectPolygon;
  cp1 = clipPolygon[clipPolygon.length - 1];
  for (j in clipPolygon) {
    var cp2 = clipPolygon[j];
    var inputList = outputList;
    outputList = [];
    s = inputList[inputList.length - 1]; //last on the input list
    for (i in inputList) {
      var e = inputList[i];
      if (inside(e)) {
        if (!inside(s)) {
          outputList.push(intersection());
        }
        outputList.push(e);
      } else if (inside(s)) {
        outputList.push(intersection());
      }
      s = e;
    }
    cp1 = cp2;
  }
  return outputList.map(e => e.map(f => Math.round(f * 1000) / 1000));
}
```
