---
id: 5a23c84252665b21eecc8010
title: Algoritmi di ordinamento/ordinamento shell
challengeType: 1
forumTopicId: 302317
dashedName: sorting-algorithmsshell-sort
---

# --description--

Scrivi una funzione per ordinare un array di elementi usando l'algoritmo di ordinamento Shell, un ordinamento a incremento decrescente. La funzione dovrebbe restituire l'array ordinato.

L'ordinamento Shell (noto anche come Shellsort o metodo di Shell) prende il nome dal suo inventore, Donald Shell, che pubblicò l'algoritmo nel 1959.

L'ordinamento Shell è una sequenza di ordinamenti per inserimento intercalati basati su una sequenza di incrementi. La dimensione dell'incremento viene ridotta dopo ogni passaggio fino a quando l'incremento è 1.

Con un incremento di 1, l'ordinamento è un ordinamento per inserimento base, ma a questo punto i dati sono garantiti quasi ordinati, che è il "miglior caso" per l'ordinamento per inserimento.

Qualsiasi sequenza ordinerà i dati purché termini con 1, ma alcune funzionano meglio di altre.

Studi empirici hanno dimostrato che una sequenza di incrementi geometrica con un rapporto di circa 2,2 funziona bene nella pratica.

# --hints--

`shellSort` dovrebbe essere una funzione.

```js
assert(typeof shellSort == 'function');
```

`shellSort([25, 32, 12, 7, 20])` dovrebbe restituire un array.

```js
assert(Array.isArray(shellSort([25, 32, 12, 7, 20])));
```

`shellSort([25, 32, 12, 7, 20])` dovrebbe restituire `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(shellSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`shellSort([38, 45, 35, 8, 13])` dovrebbe restituire `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(shellSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`shellSort([43, 36, 20, 34, 24])` dovrebbe restituire `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(shellSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`shellSort([12, 33, 26, 18, 1, 16, 38])` dovrebbe restituire `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(shellSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`shellSort([3, 39, 48, 16, 1, 4, 29])` dovrebbe restituire `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(shellSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function shellSort(a) {

}
```

# --solutions--

```js
function shellSort(a) {
  for (var h = a.length; h > 0; h = parseInt(h / 2)) {
    for (var i = h; i < a.length; i++) {
      var k = a[i];
      for (var j = i; j >= h && k < a[j - h]; j -= h) a[j] = a[j - h];
      a[j] = k;
    }
  }
  return a;
}
```
