---
id: 5a23c84252665b21eecc8017
title: Soundex
challengeType: 1
forumTopicId: 302320
dashedName: soundex
---

# --description--

**L'algoritmo Soundex** si occupa delle *intenzioni* delle parole. Crea una rappresentazione per parole che suonano in modo simile.

Viene usato per la ricerca di <em>nomi</em> e <em>indirizzi</em>. Questo significa che la persona che ha inserito il <em>nome</em> può concentrarsi su come suona invece di correggere l'ortografia dei <em>nomi</em>.

Per esempio:

Se senti il nome `Quenci` per la prima volta e lo scrivi male, otterrai il codice **Soundex** `Q520`.

Quando scrivi correttamente il nome `Quincy` la volta successiva, otterrai comunque lo stesso codice `Q520`, il che significa che puoi collegare più pronunce di un nome alla stessa <em>persona</em> senza dover aggiungere ogni ortografia.

Ecco le regole:

<ul>
  <li>Se una vocale (A, E, I, O, U) separa due consonanti che hanno lo stesso codice soundex, viene codificata la consonante a destra della vocale. Tymczak viene codificato come T-522 (T, 5 per la M, 2 per la C, Z ignorata (vedi la regola "Side-by-Side" sopra), 2 per la K). Poiché la vocale "A" separa la Z e la K, la K viene codificata.</li>
  <li>Se "H" o "W" separano due consonanti che hanno lo stesso codice soundex, la consonante a destra della vocale non viene codificata. Esempio: Ashcraft viene codificato A-261 (A, 2 per la S, C ignorata, 6 per la R, 1 per la F). Non viene codificato A-226.</li>
</ul>

# --instructions--

Scrivi una funzione che prende una stringa come parametro e restituisce la stringa codificata.

# --hints--

`soundex` dovrebbe essere una funzione.

```js
assert(typeof soundex == 'function');
```

`soundex("Soundex")` dovrebbe restituire una stringa.

```js
assert(typeof soundex('Soundex') == 'string');
```

`soundex("Soundex")` dovrebbe restituire `"S532"`.

```js
assert.equal(soundex('Soundex'), 'S532');
```

`soundex("Example")` dovrebbe restituire `"E251"`.

```js
assert.equal(soundex('Example'), 'E251');
```

`soundex("Sownteks")` dovrebbe restituire `"S532"`.

```js
assert.equal(soundex('Sownteks'), 'S532');
```

`soundex("Ekzampul")` dovrebbe restituire `"E251"`.

```js
assert.equal(soundex('Ekzampul'), 'E251');
```

`soundex("Euler")` dovrebbe restituire `"E460"`.

```js
assert.equal(soundex('Euler'), 'E460');
```

`soundex("Gauss")` dovrebbe restituire `"G200"`.

```js
assert.equal(soundex('Gauss'), 'G200');
```

`soundex("Hilbert")` dovrebbe restituire `"H416"`.

```js
assert.equal(soundex('Hilbert'), 'H416');
```

`soundex("Knuth")` dovrebbe restituire `"K530"`.

```js
assert.equal(soundex('Knuth'), 'K530');
```

`soundex("Lloyd")` dovrebbe restituire `"L300"`.

```js
assert.equal(soundex('Lloyd'), 'L300');
```

`soundex("Lukasiewicz")` dovrebbe restituire `"L222"`.

```js
assert.equal(soundex('Lukasiewicz'), 'L222');
```

# --seed--

## --seed-contents--

```js
function soundex(s) {

}
```

# --solutions--

```js
function soundex(s) {
  var a = s.toLowerCase().split('');
  var f = a.shift(),
    r = '',
    codes = {
      a: '',
      e: '',
      i: '',
      o: '',
      u: '',
      b: 1,
      f: 1,
      p: 1,
      v: 1,
      c: 2,
      g: 2,
      j: 2,
      k: 2,
      q: 2,
      s: 2,
      x: 2,
      z: 2,
      d: 3,
      t: 3,
      l: 4,
      m: 5,
      n: 5,
      r: 6
    };
  r =
    f +
    a
      .map(function(v, i, a) {
        return codes[v];
      })
      .filter(function(v, i, a) {
        return i === 0 ? v !== codes[f] : v !== a[i - 1];
      })
      .join('');

  return (r + '000').slice(0, 4).toUpperCase();
}
```
