---
id: 5a23c84252665b21eecc7eb0
title: I prima di E tranne dopo C
challengeType: 1
forumTopicId: 302288
dashedName: i-before-e-except-after-c
---

# --description--

**<a href="http://www.rosettacode.org/wiki/I_before_E_except_after_C" target="_blank" rel="noopener noreferrer nofollow">"I prima di E, tranne dopo C"</a>** è una regola generale per l'ortografia della lingua inglese. Se non si è sicuri se una parola si scrive con il digramma `ei` o `ie`, la rima suggerisce che l'ordine corretto è `ie` a meno che la lettera precedente non sia `c`, nel qual caso potrebbe essere `ei`.

Usando le parole fornite, verifica se le due sotto-clausole della frase sono plausibili singolarmente:

<ol>
  <li>
    <i>"I prima di E quando non è preceduto da C".</i>
  </li>
  <li>
    <i>"E prima di I quando è preceduto da C".</i>
  </li>
</ol>

Se entrambe le sotto-frasi sono plausibili, allora si può dire che la frase originale è plausibile.

# --instructions--

Scrivi una funzione che accetta una parola e verifica se la parola segue questa regola. La funzione dovrebbe restituire true se la parola segue la regola e false se non la segue.

# --hints--

`IBeforeExceptC` dovrebbe essere una funzione.

```js
assert(typeof IBeforeExceptC == 'function');
```

`IBeforeExceptC("receive")` dovrebbe restituire un booleano.

```js
assert(typeof IBeforeExceptC('receive') == 'boolean');
```

`IBeforeExceptC("receive")` dovrebbe restituire `true`.

```js
assert.equal(IBeforeExceptC('receive'), true);
```

`IBeforeExceptC("science")` dovrebbe restituire `false`.

```js
assert.equal(IBeforeExceptC('science'), false);
```

`IBeforeExceptC("imperceivable")` dovrebbe restituire `true`.

```js
assert.equal(IBeforeExceptC('imperceivable'), true);
```

`IBeforeExceptC("inconceivable")` dovrebbe restituire `true`.

```js
assert.equal(IBeforeExceptC('inconceivable'), true);
```

`IBeforeExceptC("insufficient")` dovrebbe restituire `false`.

```js
assert.equal(IBeforeExceptC('insufficient'), false);
```

`IBeforeExceptC("omniscient")` dovrebbe restituire `false`.

```js
assert.equal(IBeforeExceptC('omniscient'), false);
```

# --seed--

## --seed-contents--

```js
function IBeforeExceptC(word) {

}
```

# --solutions--

```js
function IBeforeExceptC(word)
{
    if(word.indexOf("c")==-1 && word.indexOf("ie")!=-1)
        return true;
    else if(word.indexOf("cei")!=-1)
        return true;
    return false;
}
```
