---
id: 5e4ce2f5ac708cc68c1df261
title: Generatore congruenziale lineare
challengeType: 1
forumTopicId: 385266
dashedName: linear-congruential-generator
---

# --description--

Un generatore congruenziale lineare (LCG) è un <em>algoritmo</em> che produce una sequenza di numeri pseudo-casuali calcolati con un'equazione lineare a tratti discontinua. Tutti i generatori congruenziali lineari usano questa formula:

$$r_{n + 1} = (a \times r_n + c) \bmod m$$

Dove:

<ul>
<li>$ r_0 $ è un seme.</li>
<li>$r_1$, $r_2$, $r_3$, ..., sono i numeri casuali.</li>
<li>$a$, $c$, $m$ sono costanti.</li>
</ul>

Se si scelgono con cura i valori di $a$, $c$ e $m$, allora il generatore produce una distribuzione uniforme di numeri interi da $0$ a $m - 1$.

I numeri del <abbr title="linear congruential generator">LCG</abbr> sono di scarsa qualità. $r_n$ e $r\_{n + 1}$ non sono indipendenti, come invece sarebbero i veri numeri casuali. Chiunque conosca $r_n$ può prevedere $r\_{n + 1}$, quindi il <abbr title="linear congruential generator">LCG</abbr> non è sicuro dal punto di vista crittografico. Il <abbr title="linear congruential generator">LCG</abbr> è comunque abbastanza valido per attività semplici come il test di primalità di Miller-Rabin o le distribuzioni di FreeCell. Tra i vantaggi del <abbr title="linear congruential generator">LCG</abbr> c’è la facilità di riprodurre una sequenza di numeri a partire dallo stesso $r_0$. Si può anche riprodurre la stessa sequenza con un linguaggio di programmazione diverso, perché la formula è molto semplice.

# --instructions--

Scrivi una funzione che prenda $r_0,a,c,m,n$ come parametri e restituisca $r_n$.

# --hints--

`linearCongGenerator` dovrebbe essere una funzione.

```js
assert(typeof linearCongGenerator == 'function');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` dovrebbe restituire un numero.

```js
assert(typeof linearCongGenerator(324, 1145, 177, 2148, 3) == 'number');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` dovrebbe restituire `855`.

```js
assert.equal(linearCongGenerator(324, 1145, 177, 2148, 3), 855);
```

`linearCongGenerator(234, 11245, 145, 83648, 4)` dovrebbe restituire `1110`.

```js
assert.equal(linearCongGenerator(234, 11245, 145, 83648, 4), 1110);
```

`linearCongGenerator(85, 11, 1234, 214748, 5)` dovrebbe restituire `62217`.

```js
assert.equal(linearCongGenerator(85, 11, 1234, 214748, 5), 62217);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 1)` dovrebbe restituire `12345`.

```js
assert.equal(linearCongGenerator(0, 1103515245, 12345, 2147483648, 1), 12345);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 2)` dovrebbe restituire `1406932606`.

```js
assert.equal(
  linearCongGenerator(0, 1103515245, 12345, 2147483648, 2),
  1406932606
);
```

# --seed--

## --seed-contents--

```js
function linearCongGenerator(r0, a, c, m, n) {

}
```

# --solutions--

```js
function linearCongGenerator(r0, a, c, m, n) {
    for (let i = 0; i < n; i++) {
        r0 = (a * r0 + c) % m;
    }
    return r0;
}
```
