---
id: 5eaf48389ee512d4d103684b
title: Numeri autodescrittivi
challengeType: 1
forumTopicId: 385289
dashedName: self-describing-numbers
---

# --description--

Esistono diversi numeri interi cosiddetti "autodescrittivi".

Un numero intero si dice "autodescrittivo" se ha la proprietà che, quando le posizioni delle cifre sono etichettate da 0 a N-1, la cifra in ogni posizione è uguale al numero di volte che quella cifra appare nel numero.

Per esempio, **2020** è un numero autodescrittivo di quattro cifre:

<ul>
    <li> la posizione 0 ha valore 2 e ci sono due 0 nel numero; </li>
    <li> la posizione 1 ha valore 0 e non ci sono 1 nel numero; </li>
    <li> la posizione 2 ha valore 2 e ci sono due 2; </li>
    <li> la posizione 3 ha valore 0 e non ci sono 3 nel numero; </li>
</ul>

I numeri autodescrittivi &lt; 100.000.000 sono: 1210, 2020, 21200, 3211000, 42101000.

# --instructions--

Scrivi una funzione che prende un numero intero positivo come parametro. Se è autodescrittivo restituisci true. Altrimenti, restituisci false.

# --hints--

`isSelfDescribing` dovrebbe essere una funzione.

```js
assert(typeof isSelfDescribing == 'function');
```

`isSelfDescribing()` dovrebbe restituire un booleano.

```js
assert(typeof isSelfDescribing(2020) == 'boolean');
```

`isSelfDescribing(2020)` dovrebbe restituire `true`.

```js
assert.equal(isSelfDescribing(2020), true);
```

`isSelfDescribing(3021)` dovrebbe restituire `false`.

```js
assert.equal(isSelfDescribing(3021), false);
```

`isSelfDescribing(3211000)` dovrebbe restituire `true`.

```js
assert.equal(isSelfDescribing(3211000), true);
```

# --seed--

## --seed-contents--

```js
function isSelfDescribing(n) {

}
```

# --solutions--

```js
function isSelfDescribing(n) {
    let digits = String(n).split("");
    digits = digits.map(function(e) {return parseInt(e)});
    let count = digits.map((x) => {return 0})
    digits.forEach((d) =>{
        if (d >= count.length) {
            return false
        }
        count[d] += 1;
    });

     if (digits === count) {
        return true;
    }
    if (digits.length != count.length) {
        return false;
    }
    
    for (let i=0; i< digits.length; i++){
      if (digits[i] !== count[i]) {
        return false;
      }
    }
    return true;
}
```
