---
id: 6a89c69cbe1bb69845992d8b
title: 数値の範囲を生成するジェネレーターを作成する
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

このラボでは、指定された範囲内の数字のリストを生成するために再帰を使う `range_of_numbers` 関数を作成します。

例えば：

- `range_of_numbers(3, 9)` は `[3, 4, 5, 6, 7, 8, 9]` を返すべきです。
- `range_of_numbers(5, 5)` は `[5]` を返すべきです。

以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `range_of_numbers` という名前の関数を作成し、`start_num` と `end_num` という2つのパラメータを持つ必要があります。
1. 関数は、`start_num` から始まり `end_num` で終わる連続した整数のリストを返す必要があり、両方の値を含みます。
1. `start_num` の値は常に `end_num` 以下です。
1. 関数は自分自身を呼び出すことで再帰を使う必要があります。
1. 関数は `for` や `while` ループ、内包表記、または `range()` 関数を使ってはいけません。
1. `start_num` が `end_num` と等しい場合、関数は `start_num` を含むリストを返す必要があります。これが基本ケースです。
1. 再帰ケースでは、関数は基本ケースに向かう引数で自分自身を呼び出し、現在の数字を返されたリストに追加する必要があります。
1. 関数は返されたリストを保存するためにグローバル変数を使ってはいけません。

# --hints--

`start_num` と `end_num` というパラメータを持つ `range_of_numbers` という名前の関数を用意してください。

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` は `[1, 2, 3, 4, 5]` を返すべきです。

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` は `[6, 7, 8, 9]` を返すべきです。

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` は `[4]` を返すべきです。

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` は `[10, 11, 12, 13, 14, 15]` を返すべきです。

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

`range_of_numbers` 関数は自分自身を呼び出す必要があります。

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

`range_of_numbers` 関数はループ、内包表記、または `range()` 関数を使ってはいけません。

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

`range_of_numbers(1, 3)` を呼び出しても後の呼び出しに影響を与えてはいけません。その後の `range_of_numbers(6, 9)` は引き続き `[6, 7, 8, 9]` を返す必要があります。

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
