---
id: 681cb1b0dab50c87ddb2e51b
title: "挑戰 10：3 次機會"
challengeType: 29
dashedName: challenge-10
---

# --description--

給定一個介於 1 和 10,000 之間的整數，傳回從 1 到該整數中，其平方數包含至少一個數字 3 的數字個數。

# --hints--

`squares_with_three(1)` 應該傳回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1), 0)`)
}})
```

`squares_with_three(10)` 應該傳回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10), 1)`)
}})
```

`squares_with_three(100)` 應該傳回 `19`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(100), 19)`)
}})
```

`squares_with_three(1000)` 應該傳回 `326`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1000), 326)`)
}})
```

`squares_with_three(10000)` 應該傳回 `4531`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10000), 4531)`)
}})
```

# --seed--

## --seed-contents--

```py
def squares_with_three(n):

    return n
```

# --solutions--

```py
def squares_with_three(n):
    count = 0
    for i in range(1, n + 1):
        square = i * i
        if '3' in str(square):
            count += 1
    return count
```
