---
id: 681cb1afdab50c87ddb2e515
title: "挑戰 4：S  P  A  C  E  J  A  M"
challengeType: 29
dashedName: challenge-4
---

# --description--

給定一個字串，移除字串中所有空格，並在每個字元之間插入兩個空格，將所有英文字母轉換為大寫，然後傳回結果。

- 非字母字元應保持不變（空格除外）。

# --hints--

`space_jam("freeCodeCamp")` 應該傳回 `"F  R  E  E  C  O  D  E  C  A  M  P"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("   free   Code   Camp   ")` 應該傳回 `"F  R  E  E  C  O  D  E  C  A  M  P"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("Hello World?!")` 應該傳回 `"H  E  L  L  O  W  O  R  L  D  ?  !"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !")`)
}})
```

`space_jam("C@t$ & D0g$")` 應該傳回 `"C  @  T  $  &  D  0  G  $"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $")`)
}})
```

`space_jam("allyourbase")` 應該傳回 `"A  L  L  Y  O  U  R  B  A  S  E"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E")`)
}})
```

# --seed--

## --seed-contents--

```py
def space_jam(s):

    return s
```

# --solutions--

```py
def space_jam(s):
    s = s.replace(" ", "")
    s = "  ".join(s)
    return s.upper()
```
