---
id: 681cb1afdab50c87ddb2e516
title: "挑戰 5：Jbelmud 文字"
challengeType: 29
dashedName: challenge-5
---

# --description--

給定一個字串，傳回該字串的亂序版本，其中每個字都依照以下條件進行轉換：

- 字的第一個和最後一個字母保持原位。
- 第一個和最後一個字母之間的所有字母皆按字母順序排序。
- 輸入字串將不包含標點符號，且全部為小寫。

# --hints--

`jbelmu("hello world")` 應該傳回 `"hello wlord"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("hello world"), "hello wlord")`)
}})
```

`jbelmu("i love jumbled text")` 應該傳回 `"i love jbelmud text"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("i love jumbled text"), "i love jbelmud text")`)
}})
```

`jbelmu("freecodecamp is my favorite place to learn to code")` 應該傳回 `"faccdeeemorp is my faiortve pacle to laern to cdoe"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe")`)
}})
```

`jbelmu("the quick brown fox jumps over the lazy dog")` 應該傳回 `"the qciuk borwn fox jmpus oevr the lazy dog"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def jbelmu(text):

    return text
```

# --solutions--

```py
def jbelmu(text):
    words = text.split()
    jumbled = []

    for word in words:
        if len(word) <= 3:
            jumbled.append(word)
        else:
            first = word[0]
            last = word[-1]
            middle = ''.join(sorted(word[1:-1]))
            jumbled.append(first + middle + last)
    
    return ' '.join(jumbled)
```
