---
id: 681cb1b0dab50c87ddb2e518
title: "挑戰 7：標的總和"
challengeType: 29
dashedName: challenge-7
---

# --description--

給定一個數字陣列和一個整數標的，找出陣列中兩個相加等於標的值的唯一數字。傳回一個包含這兩個數字索引的陣列，或在沒有兩個數字加總等於標的時傳回 `"Target not found"`。

- 傳回的陣列應該具有遞增順序的索引。

# --hints--

`find_target([2, 7, 11, 15], 9)` 應該傳回 `[0, 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([2, 7, 11, 15], 9), [0, 1])`)
}})
```

`find_target([3, 2, 4, 5], 6)` 應該傳回 `[1, 2]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([3, 2, 4, 5], 6), [1, 2])`)
}})
```

`find_target([1, 3, 5, 6, 7, 8], 15)` 應該傳回 `[4, 5]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 6, 7, 8], 15), [4, 5])`)
}})
```

`find_target([1, 3, 5, 7], 14)` 應該傳回 `'Target not found'`。

```js

({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 7], 14), "Target not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_target(arr, target):

    return arr
```

# --solutions--

```py
def find_target(arr, target):
    for i in range(len(arr)):
        for j in range(i + 1, len(arr)):
            if arr[i] + arr[j] == target:
                return [i, j]
    return 'Target not found'
```
