---
id: 681cb1afdab50c87ddb2e517
title: "挑戰 6：字母重組檢查器"
challengeType: 29
dashedName: challenge-6
---

# --description--

給定兩個字串，判斷它們是否為彼此的變位字（包含相同字元但順序可不同）。

- 忽略大小寫和空白。

# --hints--

`are_anagrams("listen", "silent")` 應該傳回 `true`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("listen", "silent"), True)`)
}})
```

`are_anagrams("School master", "The classroom")` 應該傳回 `true`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("School master", "The classroom"), True)`)
}})
```

`are_anagrams("A gentleman", "Elegant man")` 應該傳回 `true`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("A gentleman", "Elegant man"), True)`)
}})
```

`are_anagrams("Hello", "World")` 應該傳回 `false`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("Hello", "World"), False)`)
}})
```

`are_anagrams("apple", "banana")` 應該傳回 `false`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("apple", "banana"), False)`)
}})
```

`are_anagrams("cat", "dog")` 應該傳回 `false`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("cat", "dog"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def are_anagrams(str1, str2):

    return str1
```

# --solutions--

```py
def are_anagrams(str1, str2):
    def clean(s):
        return sorted(s.replace(" ", "").lower())
    
    return clean(str1) == clean(str2)
```
