---
id: 681cb1a2dab50c87ddb2e514
title: "挑戰 3：費氏數列"
challengeType: 29
dashedName: challenge-3
---

# --description--

費波那契數列是一串數字，其中每個數字都是前兩個數字的總和。當從 `0` 和 `1` 開始時，數列中的前 10 個數字是 `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34`。

給定一個包含費波那契數列前兩個數字的陣列，以及一個表述數列長度的整數，傳回一個包含該長度數列的陣列。

- 你的函式（程式）應該控制代碼長度大於或相等於零的序列。
- 如果長度為 0，傳回空陣列。
- 請注意，起始數字是序列的零件。

# --hints--


`fibonacci_sequence([0, 1], 20)` 應該傳回 `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181])`)
}})
```

`fibonacci_sequence([21, 32], 1)` 應該傳回 `[21]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([21, 32], 1), [21])`)
}})
```

`fibonacci_sequence([0, 1], 0)` 應該傳回 `[]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 0), [])`)
}})
```

`fibonacci_sequence([10, 20], 2)` 應該傳回 `[10, 20]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([10, 20], 2), [10, 20])`)
}})
```

`fibonacci_sequence([123456789, 987654321], 5)` 應該傳回 `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541])`)
}})
```

# --seed--

## --seed-contents--

```py
def fibonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def fibonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2]
        sequence.append(next_value)

    return sequence
```
