---
id: 681cb1b0dab50c87ddb2e51a
title: "挑戰 9：平方和"
challengeType: 29
dashedName: challenge-9
---

# --description--

給定一個最大為 1,000 的正整數，傳回從 1 到該數字所有整數平方的總和。

# --hints--

`sum_of_squares(5)` 應該傳回 `55`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(5), 55)`)
}})
```

`sum_of_squares(10)` 應該傳回 `385`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(10), 385)`)
}})
```

`sum_of_squares(25)` 應該傳回 `5525`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(25), 5525)`)
}})
```

`sum_of_squares(500)` 應該傳回 `41791750`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(500), 41791750)`)
}})
```

`sum_of_squares(1000)` 應該傳回 `333833500`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(1000), 333833500)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_of_squares(n):

    return n
```

# --solutions--

```py
def sum_of_squares(n):
    sum = ((n)*(n+1)*(2*n+1))//6
    return sum
```
