---
id: 681cb1b0dab50c87ddb2e519
title: "挑戰 8：階乘器"
challengeType: 29
dashedName: challenge-8
---

# --description--

給定一個介於 0 到 20 的整數，傳回該數字的階乘。階乘是從 1 到該數字之間所有數字的乘積。

- 零的階乘是 1。

# --hints--

`factorial(0)` 應該傳回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(0), 1)`)
}})
```

`factorial(5)` 應該傳回 `120`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(5), 120)`)
}})
```

`factorial(20)` 應該傳回 `2432902008176640000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(20), 2432902008176640000)`)
}})
```

# --seed--

## --seed-contents--

```py
def factorial(n):

    return n
```

# --solutions--

```py
def factorial(n):
    result = 1
    for i in range(1, n + 1):
        result *= i
    return result
```
