---
id: 6821ec02237de8297eaee79a
title: "挑戰 24：全字母句"
challengeType: 29
dashedName: challenge-24
---

# --description--

給定一個字或句子以及一串小寫字母，判斷該字或句子是否至少使用了給定集合中的所有字母一次，且沒有使用其他字母。

- 忽略字或句子中非字母字元。
- 忽略字母大小寫。

# --hints--

`is_pangram("hello", "helo")` 應該傳回 `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "helo"), True)`)
}})
```

`is_pangram("hello", "hel")` 應該傳回 `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "hel"), False)`)
}})
```

`is_pangram("hello", "helow")` 應該傳回 `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "helow"), False)`)
}})
```

`is_pangram("hello world", "helowrd")` 應該傳回 `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello world", "helowrd"), True)`)
}})
```

`is_pangram("Hello World!", "helowrd")` 應該傳回 `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("Hello World!", "helowrd"), True)`)
}})
```

`is_pangram("Hello World!", "heliowrd")` 應該傳回 `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("Hello World!", "heliowrd"), False)`)
}})
```

`is_pangram("freeCodeCamp", "frcdmp")` 應該傳回 `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("freeCodeCamp", "frcdmp"), False)`)
}})
```

`is_pangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz")` 應該傳回 `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_pangram(sentence, letters):

    return sentence
```

# --solutions--

```py
import re
def is_pangram(sentence, letters):
    used_letters = []
    for char in sentence.lower():
        if re.match(r'[a-z]', char) and char not in used_letters:
            used_letters.append(char)

    sorted_letters = ''.join(sorted(letters.lower()))
    sorted_used_letters = ''.join(sorted(used_letters))

    return sorted_letters == sorted_used_letters
```
