---
id: 6821ebf8237de8297eaee798
title: "挑戰 22：三角數列"
challengeType: 29
dashedName: challenge-22
---

# --description--

Tribonacci 數列是一串數字，其中每個數字都是前三個數字的總和。當以 `0`、`0` 和 `1` 開始時，數列中的前 10 個數字是 `0`、`0`、`1`、`1`、`2`、`4`、`7`、`13`、`24`、`44`。

給定一個包含 Tribonacci 數列前三個數字的陣列，以及一個表述數列長度的整數，傳回一個包含該長度數列的陣列。

- 你的函式（程式）應該控制代碼長度大於或相等於零的序列。
- 如果長度為 0，傳回空陣列。
- 請注意，起始數字是序列的零件。

# --hints--

`tribonacci_sequence([0, 0, 1], 20)` 應該傳回 `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513])`)
}})
```

`tribonacci_sequence([21, 32, 43], 1)` 應該傳回 `[21]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([21, 32, 43], 1), [21])`)
}})
```

`tribonacci_sequence([0, 0, 1], 0)` 應該傳回 `[]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 0), [])`)
}})
```

`tribonacci_sequence([10, 20, 30], 2)` 應該傳回 `[10, 20]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 2), [10, 20])`)
}})
```

`tribonacci_sequence([10, 20, 30], 3)` 應該傳回 `[10, 20, 30]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 3), [10, 20, 30])`)
}})
```

`tribonacci_sequence([123, 456, 789], 8)` 應該傳回 `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134])`)
}})
```

# --seed--

## --seed-contents--

```py
def tribonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def tribonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return [start_sequence[0], start_sequence[1]]
    if length == 3:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2] + sequence[-3]
        sequence.append(next_value)

    return sequence
```
