---
id: 6821ebce237de8297eaee790
title: "挑戰 14：字元戰鬥"
challengeType: 29
dashedName: challenge-14
---

# --description--

給定兩個字串，分別表述你的軍隊和敵方軍隊，你的軍隊中每個字元會依照以下規則與敵方軍隊中相同位置的字元戰鬥：

- 字元 `a-z` 分別具有 `1-26` 的強度。
- 字元 `A-Z` 分別具有 `27-52` 的強度。
- 數字 `0-9` 的強度等於其面值。
- 所有其他字元的值為零。
- 每個字元只能參加一場戰鬥。

每場戰鬥中，較強的字元獲勝。勝利次數較多的軍隊贏得戰爭。傳回以下數值：

- 如果你的軍隊有比對方軍隊更多的字元，則顯示 `"Opponent retreated"`。
- 「We retreated」如果敵軍的字元數比你的多。
- 如果你的軍隊贏得更多戰役，則顯示「We won」。
- 如果敵軍贏得更多戰役，則「We lost」。
- 如果雙方軍隊贏得相同數量的戰役，則為「平手」。

# --hints--

`battle("Hello", "World")` 應該傳回 `"We lost"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` 應該傳回 `"We won"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` 應該傳回 `"We won"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` 應該傳回 `"Opponent retreated"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` 應該傳回 `"We retreated"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` 應該傳回 `"It was a tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` 應該傳回 `"It was a tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
