---
id: 6821ebda237de8297eaee792
title: "挑戰 16：反轉括號"
challengeType: 29
dashedName: challenge-16
---

# --description--

給定一個包含正確巢狀小括弧的字串，使用以下規則傳回該字串的解碼版本：

- 每一對小括弧內的所有字元都應該反轉。
- 最終結果中應該移除小括弧。
- 如果小括弧巢狀，應該先反轉最內層的配對，然後將其結果包含在外層配對的反轉中。
- 假設所有小括弧均勻平衡且正確巢狀。

# --hints--

`decode("(f(b(dc)e)a)")` 應該傳回 `"abcdef"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` 應該傳回 `"Can you read this?"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` 應該傳回 `"freeCodeCamp"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
