---
id: 6821ebd4237de8297eaee791
title: "挑戰 15：camelCase"
challengeType: 29
dashedName: challenge-15
---

# --description--

給定一個字串，使用以下規則傳回其駝峰式大小寫版本：

- 字串引數中的字、單字由以下集合中的一個或多個字元分隔：空格（` `）、破折號（`-`）或底線（`_`）。將這些字元的任意序列視為字、單字的分隔。
- 第一個字應該全部是小寫。
- 每個後續字、單字的首字母應為大寫，其餘字母則為小寫。
- 所有空格和分隔符應該被移除。

# --hints--

`to_camel_case("hello world")` 應該傳回 `"helloWorld"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("hello world"), "helloWorld")`)
}})
```

`to_camel_case("HELLO WORLD")` 應該傳回 `"helloWorld"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("HELLO WORLD"), "helloWorld")`)
}})
```

`to_camel_case("secret agent-X")` 應該傳回 `"secretAgentX"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("secret agent-X"), "secretAgentX")`)
}})
```

`to_camel_case("FREE cODE cAMP")` 應該傳回 `"freeCodeCamp"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("FREE cODE cAMP"), "freeCodeCamp")`)
}})
```

`to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` 應該傳回 `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_camel_case(s):

    return s
```

# --solutions--

```py
import re
def to_camel_case(s):
    words = re.split(r'[_\- ]+', s)

    camel = [
        words[0].lower() if words else ''
    ] + [
        word.capitalize() for word in words[1:]
    ]

    return ''.join(camel)
```
