---
id: 68216eb60f957572e7c340c4
title: "挑戰 11：英里配速"
challengeType: 29
dashedName: challenge-11
---

# --description--

給定跑了多少英里，以及跑這些英里所花費的時間（格式為 `"MM:SS"`（分鐘：秒）），傳回一個字串，表示跑每英里的平均時間，格式為 `"MM:SS"`。

- 在需要時添加前導零。

# --hints--

`mile_pace(3, "24:00")` 應該傳回 `"08:00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(3, "24:00"), "08:00")`)
}})
```

`mile_pace(1, "06:45")` 應該傳回 `"06:45"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(1, "06:45"), "06:45")`)
}})
```

`mile_pace(2, "07:00")` 應該傳回 `"03:30"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(2, "07:00"), "03:30")`)
}})
```

`mile_pace(26.2, "120:35")` 應該傳回 `"04:36"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(26.2, "120:35"), "04:36")`)
}})
```

# --seed--

## --seed-contents--

```py
def mile_pace(miles, duration):

    return miles
```

# --solutions--

```py
def mile_pace(miles, duration):
    minutes, seconds = map(int, duration.split(":"))
    total_seconds = minutes * 60 + seconds
    avg_seconds_per_mile = total_seconds / miles

    avg_minutes = int(avg_seconds_per_mile // 60)
    avg_seconds = round(avg_seconds_per_mile % 60)

    return f"{avg_minutes:02d}:{avg_seconds:02d}"
```
