---
id: 6821ebfd237de8297eaee799
title: "挑戰 23：RGB 轉 Hex"
challengeType: 29
dashedName: challenge-23
---

# --description--

給定一個 CSS `rgb(r, g, b)` 顏色字串，傳回其十六進位等價。

以下是給定輸入的一些範例輸出：

| Input   | Output   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- 將所有字母轉為小寫。
- 傳回一個 `#`，後面接著六個字元。不要使用任何縮寫值。

# --hints--

`rgb_to_hex("rgb(255, 255, 255)")` 應該傳回 `"#ffffff"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(255, 255, 255)"), "#ffffff")`)
}})
```

`rgb_to_hex("rgb(1, 11, 111)")` 應該傳回 `"#010b6f"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(1, 11, 111)"), "#010b6f")`)
}})
```

`rgb_to_hex("rgb(173, 216, 230)")` 應該傳回 `"#add8e6"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(173, 216, 230)"), "#add8e6")`)
}})
```

`rgb_to_hex("rgb(79, 123, 201)")` 應該傳回 `"#4f7bc9"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(79, 123, 201)"), "#4f7bc9")`)
}})
```

# --seed--

## --seed-contents--

```py
def rgb_to_hex(rgb):

    return rgb
```

# --solutions--

```py
def rgb_to_hex(rgb):
    import re
    match = re.findall(r'\d+', rgb)
    r, g, b = [max(0, min(255, int(x))) for x in match[:3]]
    return f'#{r:02x}{g:02x}{b:02x}'
```
