---
id: 6821ebea237de8297eaee795
title: "挑戰 19：燭光"
challengeType: 29
dashedName: challenge-19
---

# --description--

給定一個整數表示你開始時擁有的蠟燭數量，以及一個整數表示燃燒多少根蠟燭可以創建一根新的蠟燭，傳回在創建並燃燒盡可能多的蠟燭後你將使用的蠟燭數量。

例如，如果給定 7 根蠟燭且需要 2 根燃燒過的蠟燭來製作一根新的：

1. 燃燒 7 根蠟燭以獲得 7 個剩餘物，
2. 將 6 個剩餘物回收成 3 根新蠟燭（剩餘 1 個剩餘物），
3. 燃燒 3 根蠟燭以獲得 3 個剩餘物（共 4 個），
4. 將 4 個剩餘物回收成 2 根新蠟燭，
5. 燃燒 2 根蠟燭以獲得 2 個剩餘物，
6. 將 2 個剩餘物回收成 1 支新蠟燭，
7. 點燃 1 根蠟燭。

在範例中你將總共點燃 13 根蠟燭。

# --hints--

`burn_candles(7, 2)` 應該傳回 `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

`burn_candles(10, 5)` 應該傳回 `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

`burn_candles(20, 3)` 應該傳回 `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

`burn_candles(17, 4)` 應該傳回 `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

`burn_candles(2345, 3)` 應該傳回 `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```
