---
id: 68adce01c0e1144d0a902956
title: "挑戰 25：母音重複器"
challengeType: 29
dashedName: challenge-25
---

# --description--

給定一個字串，傳回該字串的新版本，其中每個母音的重複次數比你遇到的前一個母音多一次。例如，句子中的第一個母音應保持不變。第二個母音應連續出現兩次。第三個母音應連續出現三次，依此類推。

- 字母 `a`、`e`、`i`、`o` 和 `u`，無論是大寫還是小寫，都被視為母音。
- 原始母音應保持其大小寫。
- 重複的母音應該是小寫。
- 所有非母音字元應保持其原始大小寫。

# --hints--

`repeat_vowels("hello world")` 應該傳回 `"helloo wooorld"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` 應該傳回 `"freeeCooodeeeeCaaaaamp"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` 應該傳回 `"AEeIiiOoooUuuuu"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` 應該傳回 `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
