---
id: 68adce01c0e1144d0a902958
title: "挑戰 26：IPv4 驗證器"
challengeType: 29
dashedName: challenge-26
---

# --description--

給定一個字串，判斷它是否為有效的 IPv4 位址。有效的 IPv4 位址由四個以點（`.`）分隔的整數所組成。每個數字必須符合以下條件：

- 數值介於 0 和 255（含）之間。
- 它沒有前導零（例如 0 是允許的，01 則不行）。
- 只允許數字字元。

# --hints--

`is_valid_ipv4("192.168.1.1")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.1.1"), True)`)
}})
```

`is_valid_ipv4("0.0.0.0")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("0.0.0.0"), True)`)
}})
```

`is_valid_ipv4("255.01.50.111")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("255.01.50.111"), False)`)
}})
```

`is_valid_ipv4("255.00.50.111")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("255.00.50.111"), False)`)
}})
```

`is_valid_ipv4("256.101.50.115")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("256.101.50.115"), False)`)
}})
```

`is_valid_ipv4("192.168.101.")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.101."), False)`)
}})
```

`is_valid_ipv4("192168145213")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192168145213"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_ipv4(ipv4):

    return ipv4
```

# --solutions--

```py
def is_valid_ipv4(ipv4):
    parts = ipv4.split(".")    

    if len(parts) != 4:
        return False

    for part in parts:
        if not part.isdigit():
            return False

        num = int(part)
        if num < 0 or num > 255:
            return False

        if len(part) > 1 and part.startswith("0"):
            return False

    return True
```
