---
id: 68adce01c0e1144d0a90295a
title: "挑戰 27：矩陣旋轉"
challengeType: 29
dashedName: challenge-27
---

# --description--

給定一個矩陣（一個陣列的陣列），將矩陣順時針旋轉 90 度並傳回它。例如，給定 `[[1, 2], [3, 4]]`，它看起來像這樣：

| 1 | 2 |
|---|---|
| 3 | 4 |

你應該傳回 `[[3, 1], [4, 2]]`，它看起來像這樣：

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

`rotate([[1]])` 應該傳回 `[[1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1]]), [[1]])`)
}})
```

`rotate([[1, 2], [3, 4]])` 應該傳回 `[[3, 1], [4, 2]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]])`)
}})
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` 應該傳回 `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]])`)
}})
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` 應該傳回 `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def rotate(matrix):

    return matrix
```

# --solutions--

```py
def rotate(matrix):
    n = len(matrix)
    result = [[0] * n for _ in range(n)]

    for i in range(n):
        for j in range(n):
            result[j][n - 1 - i] = matrix[i][j]

    return result
```
