---
id: 68adce01c0e1144d0a90295e
title: "挑戰 29：縮寫詞建置器"
challengeType: 29
dashedName: challenge-29
---

# --description--

給定一個包含一個或多個字的字串，使用以下條件傳回這些字的縮寫詞：

- 縮寫應由每個字的第一個字母大寫組成，除非另有說明。
- 縮寫應忽略這些字的第一個字母，除非它們是給定字串的第一個字：`a`、`for`、`an`、`and`、`by` 和 `of`。
- 縮寫字母應以給定的順序傳回。
- 縮寫不應包含任何空格。

# --hints--

`build_acronym("Search Engine Optimization")` 應該傳回 `"SEO"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Search Engine Optimization"), "SEO")`)
}})
```

`build_acronym("Frequently Asked Questions")` 應該傳回 `"FAQ"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Frequently Asked Questions"), "FAQ")`)
}})
```

`build_acronym("National Aeronautics and Space Administration")` 應該傳回 `"NASA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("National Aeronautics and Space Administration"), "NASA")`)
}})
```

`build_acronym("Federal Bureau of Investigation")` 應該傳回 `"FBI"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Federal Bureau of Investigation"), "FBI")`)
}})
```

`build_acronym("For your information")` 應該傳回 `"FYI"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("For your information"), "FYI")`)
}})
```

`build_acronym("By the way")` 應該傳回 `"BTW"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("By the way"), "BTW")`)
}})
```

`build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` 應該傳回 `"AUHWPOTIMSH"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH")`)
}})
```

# --seed--

## --seed-contents--

```py
def build_acronym(s):

    return s
```

# --solutions--

```py
def build_acronym(s):
    small_words = {"a", "for", "an", "and", "by", "of"}
    words = s.split()
    acronym = ""

    for i, word in enumerate(words):
        if i == 0 or word.lower() not in small_words:
            acronym += word[0].upper()

    return acronym
```
