---
id: 68af0687ef34c76c28ffa54d
title: "挑戰 33：螢幕時間"
challengeType: 29
dashedName: challenge-33
---

# --description--

給定一個包含七個整數的輸入陣列，表述一週的時間，其中每個整數是該天使用手機的時數，根據這些條件判斷是否螢幕使用時間過長：

- 如果任何單一天有 10 小時或以上，這就太多了。
- 如果連續 3 天中任意一天的平均值大於或相等於 8 小時，則表示過多。
- 如果七天的平均值大於或相等於 6 小時，則太多了。

# --hints--

`too_much_screen_time([1, 2, 3, 4, 5, 6, 7])` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 4, 5, 6, 7]), False)`)
}})
```

`too_much_screen_time([7, 8, 8, 4, 2, 2, 3])` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([7, 8, 8, 4, 2, 2, 3]), False)`)
}})
```

`too_much_screen_time([5, 6, 6, 6, 6, 6, 6])` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([5, 6, 6, 6, 6, 6, 6]), False)`)
}})
```

`too_much_screen_time([1, 2, 3, 11, 1, 3, 4])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 11, 1, 3, 4]), True)`)
}})
```

`too_much_screen_time([1, 2, 3, 10, 2, 1, 0])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 10, 2, 1, 0]), True)`)
}})
```

`too_much_screen_time([3, 3, 5, 8, 8, 9, 4])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([3, 3, 5, 8, 8, 9, 4]), True)`)
}})
```

`too_much_screen_time([3, 9, 4, 8, 5, 7, 6])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([3, 9, 4, 8, 5, 7, 6]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def too_much_screen_time(hours):

    return hours
```

# --solutions--

```py
def too_much_screen_time(hours):
    for h in hours:
        if h >= 10:
            return True

    for i in range(5):
        avg3 = sum(hours[i:i+3]) / 3
        if avg3 >= 8:
            return True

    weekly_avg = sum(hours) / 7
    if weekly_avg >= 6:
        return True

    return False
```
