---
id: 68af0687ef34c76c28ffa54b
title: "挑戰 32：反轉句子"
challengeType: 29
dashedName: challenge-32
---

# --description--

給定一個字串的字，傳回一個字串，其字的順序為反向。 例如，第一個字應該位於傳回字串的末尾，而最後一個字應該位於傳回字串的開頭。

- 在給定的字串中，字可以由一個或多個空格分隔。
- 傳回的字串中，字與字之間應該只有一個空格。

# --hints--

`reverse_sentence("world hello")` 應該傳回 `"hello world"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` 應該傳回 `"git commit push"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm  install   apt    sudo")` 應該傳回 `"sudo apt install npm"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import    default   function  export")` 應該傳回 `"export function default import"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```
