---
id: 68af0687ef34c76c28ffa547
title: "挑戰 30：獨特字元"
challengeType: 29
dashedName: challenge-30
---

# --description--

給定一個字串，判斷字串中的所有字元是否唯一。

- 大寫和小寫字母應視為不同的字元。

# --hints--

`all_unique("abc")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("abc"), True)`)
}})
```

`all_unique("aA")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("aA"), True)`)
}})
```

`all_unique("QwErTy123!@")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("QwErTy123!@"), True)`)
}})
```

`all_unique("~!@#$%^&*()_+")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("~!@#$%^&*()_+"), True)`)
}})
```

`all_unique("hello")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("hello"), False)`)
}})
```

`all_unique("freeCodeCamp")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("freeCodeCamp"), False)`)
}})
```

`all_unique("!@#*$%^&*()aA")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("!@#*$%^&*()aA"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def all_unique(s):

    return s
```

# --solutions--

```py
def all_unique(s):
    seen = set()
    for ch in s:
        if ch in seen:
            return False
        seen.add(ch)
    return True
```
