---
id: 68b06e589bf227324381476f
title: "挑戰 35：字頻率"
challengeType: 29
dashedName: challenge-35
---

# --description--

給定一段文字，傳回出現頻率最高的三個字陣列。

- 段落中的字會以空格分隔。
- 忽略給定段落中的大小寫。例如，將 `Hello` 和 `hello` 視為相同的字、單字。
- 忽略給定段落中的標點符號。標點符號包括逗號（`,`）、句號（`.`）和驚嘆號（`!`）。
- 傳回的陣列應該包含所有小寫字。
- 傳回的陣列應該以降冪排序，最常出現的字放在最前面。

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` 應該傳回 `["coding", "python", "in"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` 應該傳回 `["i", "like", "coding"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` 應該傳回 `["debug", "test", "deploy"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
