---
id: 68b06e589bf2273243814777
title: "挑戰 39：填滿油箱"
challengeType: 29
dashedName: challenge-39
---

# --description--

給定油箱的大小、目前的燃料層次和每加侖的價格，傳回將油箱加滿的費用。

- `tankSize` 是油箱的總容量，單位為加侖。
- `fuelLevel` 是油箱中目前的燃料量，單位為加侖。
- `pricePerGallon` 是一加侖燃料的成本。
- 傳回的值應該以 `"$d.dd"` 格式四捨五入到小數點後兩位。

# --hints--

`cost_to_fill(20, 0, 4.00)` 應該傳回 `"$80.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` 應該傳回 `"$17.50"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` 應該傳回 `"$29.25"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` 應該傳回 `"$0.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` 應該傳回 `"$21.89"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
