---
id: 68b06e589bf2273243814771
title: "挑戰 36：恆溫器調節器"
challengeType: 29
dashedName: challenge-36
---

# --description--

給定房間的當前溫度和目標溫度，根據這些條件傳回一個字串，指示如何調整房間溫度：

- 如果目前溫度低於標的，傳回 `"heat"`。
- 如果目前溫度高於標的，傳回 `"cool"`。
- 如果目前溫度與標的相等，傳回 `"hold"`。

# --hints--

`adjust_thermostat(68, 72)` 應該傳回 `"heat"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(68, 72), "heat")`)
}})
```

`adjust_thermostat(75, 72)` 應該傳回 `"cool"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(75, 72), "cool")`)
}})
```

`adjust_thermostat(72, 72)` 應該傳回 `"hold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 72), "hold")`)
}})
```

`adjust_thermostat(-20.5, -10.1)` 應該傳回 `"heat"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(-20.5, -10.1), "heat")`)
}})
```

`adjust_thermostat(100, 99.9)` 應該傳回 `"cool"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(100, 99.9), "cool")`)
}})
```

`adjust_thermostat(0.0, 0.0)` 應該傳回 `"hold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(0.0, 0.0), "hold")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(temp, target):

    return temp
```

# --solutions--

```py
def adjust_thermostat(temp, target):
    if temp < target:
        return "heat"
    elif temp > target:
        return "cool"
    else:
        return "hold"
```
