---
id: 68b06e589bf2273243814775
title: "挑戰 38：Slug 產生器"
challengeType: 29
dashedName: challenge-38
---

# --description--

給定一個字串，使用以下條件傳回一個適合 URL 的字串版本：

- 所有字母應該是小寫。
- 所有非字母、數字或空格的字元都應該被移除。
- 所有空格都應該被取代為 URL 編碼的空格代碼 `%20`。
- 連續空格應該取代為單一 `%20`。
- 傳回的字串不應該有前置或後置的 `%20`。

# --hints--

`generate_slug("helloWorld")` 應該傳回 `"helloworld"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("helloWorld"), "helloworld")`)
}})
```

`generate_slug("hello world!")` 應該傳回 `"hello%20world"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello world!"), "hello%20world")`)
}})
```

`generate_slug(" hello-world ")` 應該傳回 `"helloworld"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug(" hello-world "), "helloworld")`)
}})
```

`generate_slug("hello  world")` 應該傳回 `"hello%20world"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello  world"), "hello%20world")`)
}})
```

`generate_slug("  ?H^3-1*1]0! W[0%R#1]D  ")` 應該傳回 `"h3110%20w0r1d"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_slug(str):

    return str
```

# --solutions--

```py
import re
def generate_slug(s):
    cleaned = re.sub(r'[^a-zA-Z0-9 ]+', '', s)
    cleaned = re.sub(r'\s+', ' ', cleaned).strip()
    return cleaned.lower().replace(' ', '%20')
```
