---
id: 68b1f72371a5ac895ac70a06
title: "挑戰 42：影片儲存"
challengeType: 29
dashedName: challenge-42
---

# --description--

給定一個影片大小、影片大小的單位、一個硬碟容量，以及硬碟的單位，傳回硬碟可以儲存的影片數量，並使用以下條件：

- 影片大小的單位可以是位元組（`"B"`）、千位元組（`"KB"`）、百萬位元組（`"MB"`）或千兆位元組（`"GB"`）。
- 如果沒有給定上述其中一個影片單位，傳回 `"Invalid video unit"`。
- 硬碟容量的單位可以是 gigabytes (`"GB"`) 或 terabytes (`"TB"`)。
- 如果沒有給出上述硬碟單位之一，傳回 `"Invalid drive unit"`。
- 傳回磁碟機可容納的完整影片數量。
- 使用以下轉換：

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |
| 1 TB |   1000 GB  |

例如，給定 `500`、`"MB"`、`100` 和 `"GB"` 作為引數，判斷在 100 GB 硬碟中可以容納多少個 500 MB 影片。

# --hints--

`number_of_videos(500, "MB", 100, "GB")` 應該傳回 `200`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` 應該傳回 `"Invalid video unit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` 應該傳回 `"Invalid drive unit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` 應該傳回 `4000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` 應該傳回 `1466`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```
