---
id: 68b1f72371a5ac895ac70a08
title: "挑戰 43：數字對字母"
challengeType: 29
dashedName: challenge-43
---

# --description--

給定一個字串，如果字串中數字多於字母，傳回 `"digits"`，如果字母多於數字，傳回 `"letters"`，如果數字和字母數量相同，傳回 `"tie"`。

- 數字由 `0-9` 組成。
- 字母由大寫或小寫的 `a-z` 組成。
- 忽略任何其他字元。

# --hints--

`digits_or_letters("abc123")` 應該傳回 `"tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123"), "tie")`)
}})
```

`digits_or_letters("a1b2c3d")` 應該傳回 `"letters"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("a1b2c3d"), "letters")`)
}})
```

`digits_or_letters("1a2b3c4")` 應該傳回 `"digits"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("1a2b3c4"), "digits")`)
}})
```

`digits_or_letters("abc123!@#DEF")` 應該傳回 `"letters"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123!@#DEF"), "letters")`)
}})
```

`digits_or_letters("H3110 W0R1D")` 應該傳回 `"digits"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("H3110 W0R1D"), "digits")`)
}})
```

`digits_or_letters("P455W0RD")` 應該傳回 `"tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("P455W0RD"), "tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def digits_or_letters(s):

    return s
```

# --solutions--

```py
def digits_or_letters(s):
    digit_count = 0
    letter_count = 0

    for char in s:
        if char.isdigit():
            digit_count += 1
        elif char.isalpha():
            letter_count += 1

    if digit_count > letter_count:
        return "digits"
    if letter_count > digit_count:
        return "letters"
    return "tie"
```
