---
id: 68b1f72371a5ac895ac70a04
title: "挑戰 41：檔案儲存"
challengeType: 29
dashedName: challenge-41
---

# --description--

給定一個檔案大小、一個檔案大小的單位，以及以 GB 為單位的硬碟容量，傳回硬碟可以儲存的檔案數量，並使用以下條件：

- 檔案大小的單位可以是位元組（`"B"`）、千位元組（`"KB"`）或兆位元組（`"MB"`）。
- 傳回磁碟機可容納的完整檔案數量。
- 使用以下轉換：

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

例如，給定 `500`、`"KB"` 和 `1` 作為引數，判斷在 1 GB 硬碟中可以容納多少個 500 KB 的檔案。

# --hints--

`number_of_files(500, "KB", 1)` 應該傳回 `2000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(500, "KB", 1), 2000)`)
}})
```

`number_of_files(50000, "B", 1)` 應該傳回 `20000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(50000, "B", 1), 20000)`)
}})
```

`number_of_files(5, "MB", 1)` 應該傳回 `200`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(5, "MB", 1), 200)`)
}})
```

`number_of_files(4096, "B", 1.5)` 應該傳回 `366210`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4096, "B", 1.5), 366210)`)
}})
```

`number_of_files(220.5, "KB", 100)` 應該傳回 `453514`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(220.5, "KB", 100), 453514)`)
}})
```

`number_of_files(4.5, "MB", 750)` 應該傳回 `166666`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4.5, "MB", 750), 166666)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_files(file_size, file_unit, drive_size_gb):

    return file_size
```

# --solutions--

```py
def number_of_files(file_size, file_unit, drive_size_gb):
    drive_size_bytes = drive_size_gb * 1000 * 1000 * 1000

    if file_unit == "B":
        file_size_bytes = file_size
    elif file_unit == "KB":
        file_size_bytes = file_size * 1000
    else:
        file_size_bytes = file_size * 1000 * 1000

    return int(drive_size_bytes // file_size_bytes)
```
