---
id: 68b1f72371a5ac895ac70a02
title: "挑戰 40：照片儲存"
challengeType: 29
dashedName: challenge-40
---

# --description--

給定以百萬位元組（MB）為單位的照片大小，以及以吉位元組（GB）為單位的硬碟容量，傳回硬碟可以儲存的照片數量，並使用以下條件：

- 1 gigabyte 等於 1000 megabytes。
- 傳回磁碟機可以儲存的完整照片數量。

# --hints--

`number_of_photos(1, 1)` 應該傳回 `1000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(1, 1), 1000)`)
}})
```

`number_of_photos(2, 1)` 應該傳回 `500`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(2, 1), 500)`)
}})
```

`number_of_photos(4, 256)` 應該傳回 `64000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(4, 256), 64000)`)
}})
```

`number_of_photos(3.5, 750)` 應該傳回 `214285`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 750), 214285)`)
}})
```

`number_of_photos(3.5, 5.5)` 應該傳回 `1571`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 5.5), 1571)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_photos(photo_size_mb, drive_size_gb):

    return photo_size_mb
```

# --solutions--

```py
def number_of_photos(photo_size_mb, drive_size_gb):
    drive_size_mb = drive_size_gb * 1000
    return drive_size_mb // photo_size_mb
```
