---
id: 68b1f72371a5ac895ac70a0a
title: "挑戰 44：字串鏡像"
challengeType: 29
dashedName: challenge-44
---

# --description--

給定兩個字串，判斷第二個字串是否為第一個字串的鏡像。

- 如果一個字串包含相同的字母且順序相反，則該字串被視為鏡像字串。
- 將大寫字母和小寫字母視為不同。
- 忽略所有非字母字元。

# --hints--

`is_mirror("helloworld", "helloworld")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("helloworld", "helloworld"), False)`)
}})
```

`is_mirror("Hello World", "dlroW olleH")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Hello World", "dlroW olleH"), True)`)
}})
```

`is_mirror("RaceCar", "raCecaR")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("RaceCar", "raCecaR"), True)`)
}})
```

`is_mirror("RaceCar", "RaceCar")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("RaceCar", "RaceCar"), False)`)
}})
```

`is_mirror("Mirror", "rorrim")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Mirror", "rorrim"), False)`)
}})
```

`is_mirror("Hello World", "dlroW-olleH")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Hello World", "dlroW-olleH"), True)`)
}})
```

`is_mirror("Hello World", "!dlroW !olleH")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Hello World", "!dlroW !olleH"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_mirror(str1, str2):

    return str1
```

# --solutions--

```py
def is_mirror(str1, str2):
    clean1 = "".join(c for c in str1 if c.isalpha())
    clean2 = "".join(c for c in str2 if c.isalpha())
    return clean1[::-1] == clean2
```
