---
id: 68b7687dded630607aceccb3
title: "挑戰 49：CSV 標頭解析器"
challengeType: 29
dashedName: challenge-49
---

# --description--

給定逗號分隔值（CSV）`file` 的第一行，傳回包含標題的陣列。

- CSV 檔案的第一行包含以逗號分隔的標題。
- 移除每個標題的前後空白。

# --hints--

`get_headings("name,age,city")` 應該傳回 `["name", "age", "city"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("name,age,city"), ["name", "age", "city"])`)
}})
```

`get_headings("first name,last name,phone")` 應該傳回 `["first name", "last name", "phone"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("first name,last name,phone"), ["first name", "last name", "phone"])`)
}})
```

`get_headings("username , email , signup date ")` 應該傳回 `["username", "email", "signup date"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("username , email , signup date "), ["username", "email", "signup date"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_headings(csv):

    return csv
```

# --solutions--

```py
def get_headings(csv):
    return [h.strip() for h in csv.split(",")]
```
