---
id: 68b7cadffed0e75a517da677
title: "挑戰 54：P@ssw0rd Str3ngth！"
challengeType: 29
dashedName: challenge-54
---

# --description--

給定一個密碼字串，根據密碼強度傳回 `"weak"`、`"medium"` 或 `"strong"`。

密碼會根據以下規則進行評估：

- 長度至少為 8 個字元。
- 它包含大寫和小寫字母。
- 它至少包含 1 個數字。
- 它至少包含來自此集合的 1 個特殊字元：`!`、`@`、`#`、`$`、`%`、`^`、`&` 或 `*`。

如果密碼符合少於兩項規則，傳回 `"weak"`。  
如果密碼符合 2 或 3 項規則，傳回 `"medium"`。  
如果密碼符合全部 4 項規則，傳回 `"strong"`。

# --hints--

`check_strength("123456")` 應該傳回 `"weak"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("123456"), "weak")`)
}})
```

`check_strength("pass!!!")` 應該傳回 `"weak"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("pass!!!"), "weak")`)
}})
```

`check_strength("Qwerty")` 應該傳回 `"weak"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("Qwerty"), "weak")`)
}})
```

`check_strength("PASSWORD")` 應該傳回 `"weak"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD"), "weak")`)
}})
```

`check_strength("PASSWORD!")` 應該傳回 `"medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("PassWord%^!")` 應該傳回 `"medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PassWord%^!"), "medium")`)
}})
```

`check_strength("qwerty12345")` 應該傳回 `"medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("qwerty12345"), "medium")`)
}})
```

`check_strength("S3cur3P@ssw0rd")` 應該傳回 `"strong"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("S3cur3P@ssw0rd"), "strong")`)
}})
```

`check_strength("C0d3&Fun!")` 應該傳回 `"strong"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("C0d3&Fun!"), "strong")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_strength(password):

    return password
```

# --solutions--

```py
import re
def check_strength(password):
    rules_met = 0

    if len(password) >= 8:
        rules_met += 1
    if re.search(r'[a-z]', password) and re.search(r'[A-Z]', password):
        rules_met += 1
    if re.search(r'\d', password):
        rules_met += 1
    if re.search(r'[!@#$%^&*]', password):
        rules_met += 1

    if rules_met < 2:
        return "weak"
    elif rules_met < 4:
        return "medium"
    else:
        return "strong"
```
