---
id: 68c1a929005bf54d342aa8d6
title: "挑戰 59：太空週第 5 天：適居帶"
challengeType: 29
dashedName: challenge-59
---

# --description--

在太空週的第 5 天，你將計算恆星的「適居帶」——恆星周圍條件「恰到好處」以維持液態水存在的區域。

給定恆星的質量，傳回一個包含其黃金圈區域起始和結束距離（以天文單位為單位）的陣列。

計算 Goldilocks Zone：

1. 透過將恆星的質量提升到 3.5 次方來求出其光度。
2. 該區域的起點是其亮度平方根的 0.95 倍。
3. 該區域的末端是其亮度平方根的 1.37 倍。

- 傳回四捨五入到小數點後兩位的距離。

例如，給定 `1` 作為質量，傳回 `[0.95, 1.37]`。

# --hints--

`goldilocks_zone(1)` 應該傳回 `[0.95, 1.37]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` 應該傳回 `[0.28, 0.41]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` 應該傳回 `[21.85, 31.51]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` 應該傳回 `[9.38, 13.52]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` 應該傳回 `[179.69, 259.13]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
