---
id: 68c497f3aaefc9fd9f1b0e25
title: "挑戰 61：太空週第 7 天：發射燃料"
challengeType: 29
dashedName: challenge-61
---

# --description--

在太空週的最後一天，你將會得到你想送入軌道的有效載荷質量（公斤，kg）。請根據以下規則計算將你的有效載荷送入軌道所需的燃料量：

- 火箭每提升 5 公斤的質量需要 1 公斤的燃料。
- 燃料本身具有質量。所以當你添加燃料時，需提升的質量會增加，這需要更多燃料，進而增加質量，如此循環。
- 要計算所需的總燃料：從有效載荷質量開始，計算所需的燃料，將該燃料添加到總質量中，然後再次計算。重複此處理程序，直到額外所需的燃料少於 1 公斤，然後停止。
- 忽略火箭本身的質量。只計算提升有效載荷及其自身燃料所需的燃料量。

例如，給定 50 kg 的有效載荷質量，你需要 10 kg 的燃料來將其提升（有效載荷 / 5），這會使總質量增加到 60 kg，需要 12 kg 來提升（額外 2 kg），這會使總質量增加到 62 kg，需要 12.4 kg 來提升－額外 0.4 kg－這比額外 1 kg 少，所以我們在此停止。總提升質量為 62.4 kg，其中 50 kg 是初始有效載荷，12.4 kg 是燃料。

- 傳回所需燃料量，四捨五入至小數點後 1 位。

# --hints--

`launch_fuel(50)` 應該傳回 `12.4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(50), 12.4)`)
}})
```

`launch_fuel(500)` 應該傳回 `124.8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(500), 124.8)`)
}})
```

`launch_fuel(243)` 應該傳回 `60.7`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(243), 60.7)`)
}})
```

`launch_fuel(11000)` 應該傳回 `2749.8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(11000), 2749.8)`)
}})
```

`launch_fuel(6214)` 應該傳回 `1553.4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(6214), 1553.4)`)
}})
```

# --seed--

## --seed-contents--

```py
def launch_fuel(payload):

    return payload
```

# --solutions--

```py
def launch_fuel(payload):
    total_mass = payload
    additional_fuel = total_mass / 5
    total_fuel = additional_fuel

    while additional_fuel >= 1:
        total_mass += additional_fuel
        additional_fuel = (total_mass / 5) - total_fuel
        total_fuel += additional_fuel

    return round(total_fuel, 1)
```
