---
id: 68c497f3aaefc9fd9f1b0e24
title: "挑戰 60：太空週第 6 天：月相"
challengeType: 29
dashedName: challenge-60
---

# --description--

在太空週的第 6 天，你將會收到一個格式為 `"YYYY-MM-DD"` 的日期，並且需要使用以下規則判斷該日期的月相：

使用 28 天的簡化農曆週期，分為四個相等的階段：

- 「New」：第 1 天至第 7 天
- 「蠟狀期」：第 8 天至第 14 天
- 「Full」：第 15 天至第 21 天
- 「Waning」：第 22 天至第 28 天

在第 28 天之後，週期會從第 1 天（新月）重新開始。

- 使用 `"2000-01-06"` 作為參考新月（週期的第 1 天）來判斷指定日期的月相。
- 你不會被給予任何早於引用日期的日期。
- 傳回正確階段的字串。

**注意：** 第 1 天表述新月當天，表示自上一次新月以來尚未過任何天數。

# --hints--

`moon_phase("2000-01-12")` 應該傳回 `"New"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` 應該傳回 `"Waxing"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` 應該傳回 `"Full"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` 應該傳回 `"Waning"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` 應該傳回 `"New"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
