---
id: 68cae5b538ff798bbd4da001
title: "挑戰 63：字之戰"
challengeType: 29
dashedName: challenge-63
---

# --description--

給定兩個句子，分別表述你的團隊和對手團隊，其中你團隊的每個字與對手團隊對應的字進行對戰，根據以下規則判斷哪個團隊獲勝：

- 給定的句子將始終包含相同數量的字。
- 字、單字 由單一空格分隔，且只包含字母。
- 每個字的值是其字母的總和。
- 字母 `a` 到 `z` 對應值為 `1` 到 `26`。例如，`a` 是 `1`，而 `z` 是 `26`。
- 大寫字母會將該字母的值加倍。 例如，`A` 是 `2`，而 `Z` 是 `52`。
- 字、單字依序對戰：你隊伍的第一個字、單字與對方隊伍的第一個字、單字對戰，依此類推。
- 如果一個字的值大於對方字的值，該字即獲勝。
- 獲勝字數較多的隊伍即為贏家。

傳回 `"We win"` 如果你的隊伍是贏家，傳回 `"We lose"` 如果你的隊伍輸了，並且如果雙方隊伍的勝場數相同，傳回 `"Draw"`。

# --hints--

`battle("hello world", "hello word")` 應該傳回 `"We win"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` 應該傳回 `"We win"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` 應該傳回 `"We lose"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` 應該傳回 `"Draw"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` 應該傳回 `"We win"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` 應該傳回 `"We lose"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` 應該傳回 `"Draw"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
