---
id: 68cae5b538ff798bbd4da002
title: "挑戰 64：24 變 12"
challengeType: 29
dashedName: challenge-64
---

# --description--

給定一個以 24 小時制 `"HHMM"` 表述的時間字串，傳回其等價的 12 小時制 `"H:MM AM"` 或 `"H:MM PM"`。

- 給定的輸入將永遠是 24 小時制時間格式中，從 `"0000"` 到 `"2359"` 的四位數字串。

# --hints--

`to_12("1124")` 應該傳回 `"11:24 AM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1124"), "11:24 AM")`)
}})
```

`to_12("0900")` 應該傳回 `"9:00 AM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0900"), "9:00 AM")`)
}})
```

`to_12("1455")` 應該傳回 `"2:55 PM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1455"), "2:55 PM")`)
}})
```

`to_12("2346")` 應該傳回 `"11:46 PM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("2346"), "11:46 PM")`)
}})
```

`to_12("0030")` 應該傳回 `"12:30 AM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0030"), "12:30 AM")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_12(time):

    return time
```

# --solutions--

```py
def convert_hours(hours):
    if hours == 0:
        return 12
    elif hours > 12:
        return hours - 12
    else:
        return hours

def to_12(time):
    hours = int(time[:2])
    minutes = time[2:]
    period = "AM" if hours < 12 else "PM"

    return f"{convert_hours(hours)}:{minutes} {period}"
```
