---
id: 68cae5b538ff798bbd4da005
title: "挑戰 67：電子郵件驗證器"
challengeType: 29
dashedName: challenge-67
---

# --description--

給定一個字串，判斷它是否為有效的電子郵件位址，使用以下條件：

- 它必須包含恰好一個 `@` 符號。
- 區域的零件（`@` 之前）：
  - 只能包含字母（`a-z`、`A-Z`）、數字（`0-9`）、點（`.`）、底線（`_`）或連字號（`-`）。
  - 不能以點號開始或結束。
- `@` 之後的網域零件
  - 必須包含至少一個點。
  - 必須以句點後接至少兩個字母結尾。
- 區域的或網域零件都不能有連續兩個點。

# --hints--

`validate("a@b.cd")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("a@b.cd"), True)`)
}})
```

`validate("hell.-w.rld@example.com")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hell.-w.rld@example.com"), True)`)
}})
```

`validate(".b@sh.rc")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate(".b@sh.rc"), False)`)
}})
```

`validate("example@test.c0")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("example@test.c0"), False)`)
}})
```

`validate("freecodecamp.org")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("freecodecamp.org"), False)`)
}})
```

`validate("develop.ment_user@c0D!NG.R.CKS")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("develop.ment_user@c0D!NG.R.CKS"), True)`)
}})
```

`validate("hello.@wo.rld")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hello.@wo.rld"), False)`)
}})
```

`validate("hello@world..com")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hello@world..com"), False)`)
}})
```

`validate("develop..ment_user@c0D!NG.R.CKS")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("develop..ment_user@c0D!NG.R.CKS"), False)`)
}})
```

`validate("git@commit@push.io")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("git@commit@push.io"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def validate(email):

    return email
```

# --solutions--

```py
import re
def validate(email):
    if '..' in email:
        return False

    parts = email.split('@')
    if len(parts) != 2:
        return False

    local, domain = parts

    if local.startswith('.') or local.endswith('.'):
        return False
    if not re.match(r'^[a-zA-Z0-9._-]+$', local):
        return False

    if '.' not in domain:
        return False

    tld = domain.split('.')[-1]
    if len(tld) < 2 or not tld.isalpha():
        return False

    return True
```
