---
id: 68cae5b538ff798bbd4da008
title: "挑戰 70：HTML 屬性擷取器"
challengeType: 29
dashedName: challenge-70
---

# --description--

給定一個有效 HTML 元素的字串，根據以下條件傳回該元素的屬性：

- 你只會被給予一個元素。
- 屬性將採用 `attribute="value"` 格式。
- 傳回一個字串陣列，每個屬性屬性和值以逗號分隔，格式為：`["attribute1, value1", "attribute2, value2"]`。
- 以給定的順序傳回屬性。
- 如果找不到任何屬性，傳回空陣列。

# --hints--

`extract_attributes('<span class="red"></span>')` 應該傳回 `["class, red"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<span class="red"></span>'), ["class, red"])`)
}})
```

`extract_attributes('<meta charset="UTF-8" />')` 應該傳回 `["charset, UTF-8"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<meta charset="UTF-8" />'), ["charset, UTF-8"])`)
}})
```

`extract_attributes("<p>Lorem ipsum dolor sit amet</p>")` 應該傳回 `[]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes("<p>Lorem ipsum dolor sit amet</p>"), [])`)
}})
```

`extract_attributes('<input name="email" type="email" required="true" />')` 應該傳回 `["name, email", "type, email", "required, true"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"])`)
}})
```

`extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>')` 應該傳回 `["id, submit", "class, btn btn-primary"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"])`)
}})
```

# --seed--

## --seed-contents--

```py
def extract_attributes(element):

    return element
```

# --solutions--

```py
import re
def extract_attributes(element):
    pattern = r'([\w-]+)="([^"]*)"'
    matches = re.findall(pattern, element)
    return [f"{attr}, {val}" for attr, val in matches]
```
