---
id: 68cae5b538ff798bbd4da004
title: "挑戰 66：HTML 頁籤剝除器"
challengeType: 29
dashedName: challenge-66
---

# --description--

給定一個 HTML `code` 的 `string`，移除 `tag` 並傳回純文字內容。

- 輸入字串將只包含有效的 HTML。
- HTML 頁籤可能會巢狀排列。
- 移除頁籤和任何屬性。

例如，`'<a href="#">Click here</a>'` 應該傳回 `"Click here"`。

# --hints--

`strip_tags('<a href="#">Click here</a>')` 應該傳回 `"Click here"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` 應該傳回 `"Hello World!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` 應該傳回空字串（`""`）。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` 應該傳回 `sectionsection`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
