---
id: 68cae5b538ff798bbd4da00a
title: "挑戰 72：恆溫器調整器 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

給定房間目前的華氏溫度和攝氏的目標溫度，傳回一個字串，根據這些條件指示如何調整房間溫度：

- 傳回 `"Heat: X degrees Fahrenheit"`，如果目前溫度低於標的。`X` 是將房間加熱到標的的華氏溫度數，四捨五入到 1 位小數。
- 傳回 `"Cool: X degrees Fahrenheit"`，如果目前溫度高於標的。`X` 是將房間冷卻到標的所需的華氏溫度數，四捨五入到 1 位小數。
- 如果目前溫度與標的相等，傳回 `"Hold"`。

要將攝氏溫度轉換為華氏溫度，請將攝氏溫度乘以 1.8，然後將 32 加到結果中（`F = (C * 1.8) + 32`）。

# --hints--

`adjust_thermostat(32, 0)` 應該傳回 `"Hold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` 應該傳回 `"Heat: 7.0 degrees Fahrenheit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` 應該傳回 `"Cool: 7.6 degrees Fahrenheit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` 應該傳回 `"Hold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` 應該傳回 `"Heat: 12.6 degrees Fahrenheit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
