---
id: 68d2ba1468508398389487d0
title: "挑戰 75：隱藏寶藏"
challengeType: 29
dashedName: challenge-75
---

# --description--

給定一個表述海底地圖的二維陣列，其中包含一個隱藏的寶藏，以及一個包含你下一次尋寶潛水座標（`[row, column]`）的陣列，請根據以下規則傳回 `"Empty"`、`"Found"` 或 `"Recovered"`：

- 給定的二維陣列將包含恰好一個未被尋獲的寶藏，該寶藏將佔據多個儲存格。
- 二維陣列中的每個儲存格將包含以下值之一：
  - `"-"`：沒有寶藏。
  - "O"：尚未找到的寶藏零件。
  - 「X」：已經找到的寶藏零件。

- 如果潛水地點沒有寶藏，傳回 `"Empty"`。
- 如果潛水地點找到寶藏，但至少有一個其他零件仍未被找到，傳回 `"Found"`。
- 如果潛水地點找到寶藏最後一個未找到的零件，傳回 `"Recovered"`。

例如，給定：

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

而 `[2, 1]` 為潛水地點的座標，傳回 `"Recovered"`，因為潛水找到了寶藏最後一個未找到的零件。

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` 應該傳回 `"Recovered"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered")`)
}})
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` 應該傳回 `"Empty"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty")`)
}})
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` 應該傳回 `"Found"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` 應該傳回 `"Found"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` 應該傳回 `"Recovered"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` 應該傳回 `"Empty"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty")`)
}})
```

# --seed--

## --seed-contents--

```py
def dive(map, coordinates):

    return map
```

# --solutions--

```py
def dive(map, coordinates):
    row, col = coordinates
    target = map[row][col]

    if target == "-":
        return "Empty"

    if target == "O":
        for r in range(len(map)):
            for c in range(len(map[r])):
                if (r != row or c != col) and map[r][c] == "O":
                    return "Found"
        return "Recovered"

    return "Found"
```
