---
id: 68d2ba1468508398389487cf
title: "挑戰 74：最愛歌曲"
challengeType: 29
dashedName: challenge-74
---

# --description--

還記得 iPod 嗎？第一款機型在 2001 年 10 月 23 日，也就是今天，推出，距今已經 24 年。

給定一個表述你 iPod 播放列表的歌曲物件陣列，傳回一個包含播放次數最多的兩首歌曲標題的陣列，並以播放次數最多的歌曲為首。

- 每個物件都會有一個 `"title"` 屬性（字串），以及一個 `"plays"` 屬性（整數）。

# --hints--

`favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` 應該傳回 `["Sync or Swim", "Earbud Blues"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"])`)
}})
```

`favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` 應該傳回 `["Clickwheel Love", "99 Downloads"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"])`)
}})
```

`favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` 應該傳回 `["Song B", "Song C"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"])`)
}})
```

# --seed--

## --seed-contents--

```py
def favorite_songs(playlist):

    return playlist
```

# --solutions--

```py
def favorite_songs(playlist):
    sorted_songs = sorted(playlist, key=lambda x: x["plays"], reverse=True)
    return [song["title"] for song in sorted_songs[:2]]
```
