---
id: 68d30845cc08266018fc46bc
title: "挑戰 76：互補 DNA"
challengeType: 29
dashedName: challenge-76
---

# --description--

給定一個表述 DNA 序列的字串，使用以下規則傳回其互補鏈：

- DNA 由字母 `"A"`、`"C"`、`"G"` 和 `"T"` 組成。
- 字母 `"A"` 和 `"T"` 互補。
- 字母 `"C"` 和 `"G"` 互補。

例如，給定 `"ACGT"`，傳回 `"TGCA"`。

# --hints--

`complementary_dna("ACGT")` 應該傳回 `"TGCA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ACGT"), "TGCA")`)
}})
```

`complementary_dna("ATGCGTACGTTAGC")` 應該傳回 `"TACGCATGCAATCG"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ATGCGTACGTTAGC"), "TACGCATGCAATCG")`)
}})
```

`complementary_dna("GGCTTACGATCGAAG")` 應該傳回 `"CCGAATGCTAGCTTC"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC")`)
}})
```

`complementary_dna("GATCTAGCTAGGCTAGCTAG")` 應該傳回 `"CTAGATCGATCCGATCGATC"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC")`)
}})
```

# --seed--

## --seed-contents--

```py
def complementary_dna(strand):

    return strand
```

# --solutions--

```py
def complementary_dna(strand):
    complements = { "A": "T", "T": "A", "C": "G", "G": "C" }
    return "".join(complements[n] for n in strand)
```
