---
id: 68d30845cc08266018fc46bd
title: "挑戰 77：持續時間格式化器"
challengeType: 29
dashedName: challenge-77
---

# --description--

給定一個整數秒數，傳回一個字串，表述相同的持續時間，格式為 `"H:MM:SS"`，其中 `"H"` 是小時數，`"MM"` 是分鐘數，`"SS"` 是秒數。請依照以下規則傳回時間：

- 秒數：應始終為兩位數。
- 分鐘：當不需要時應省略前導零。如果持續時間少於 1 分鐘，則使用 `"0"`。
- 小時：只有當其大於零時才應包含。

# --hints--

`format(500)` 應該傳回 `"8:20"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(500), "8:20")`)
}})
```

`format(4000)` 應該傳回 `"1:06:40"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(4000), "1:06:40")`)
}})
```

`format(1)` 應該傳回 `"0:01"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(1), "0:01")`)
}})
```

`format(5555)` 應該傳回 `"1:32:35"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(5555), "1:32:35")`)
}})
```

`format(99999)` 應該傳回 `"27:46:39"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(99999), "27:46:39")`)
}})
```

# --seed--

## --seed-contents--

```py
def format(seconds):

    return seconds
```

# --solutions--

```py
def format(seconds):
    h = seconds // 3600
    m = (seconds % 3600) // 60
    s = seconds % 60

    seconds_str = f"{s:02d}"
    minutes_str = str(m)

    if h > 0:
        hours_str = str(h)
        return f"{hours_str}:{minutes_str.zfill(2)}:{seconds_str}"
    else:
        return f"{minutes_str}:{seconds_str}"
```
