---
id: 68d30845cc08266018fc46be
title: "挑戰 78：整數序列"
challengeType: 29
dashedName: challenge-78
---

# --description--

給定一個正整數，傳回一個字串，該字串包含從 `1` 到該數字（包含該數字）所有整數，且依數字順序排列。

例如，給定 `5`，傳回 `"12345"`。

# --hints--

`sequence(5)` 應該傳回 `"12345"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(5), "12345")`)
}})
```

`sequence(10)` 應該傳回 `"12345678910"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(10), "12345678910")`)
}})
```

`sequence(1)` 應該傳回 `"1"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(1), "1")`)
}})
```

`sequence(27)` 應該傳回 `"123456789101112131415161718192021222324252627"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(27), "123456789101112131415161718192021222324252627")`)
}})
```

# --seed--

## --seed-contents--

```py
def sequence(n):

    return n
```

# --solutions--

```py
def sequence(n):
    result = ""
    for i in range(1, n + 1):
        result += str(i)
    return result
```
