---
id: 68e39ed6106dac2f0a98fd66
title: "挑戰 84：感染"
challengeType: 29
dashedName: challenge-84
---

# --description--

1988 年 11 月 2 日，第一個大型網際網路蠕蟲被釋出，在僅僅 1 天內感染了約 10％ 連接到網際網路的電腦。

在這個挑戰中，你會得到自網際網路蠕蟲釋出以來經過的天數，並且你需要根據以下規則判斷有多少台電腦被感染：

- 在第 0 天，第一台電腦被感染。
- 每隔一天，感染的電腦數量就會翻倍。
- 每隔 3 天，在病毒擴散後會套用一次修補程式，並將受感染的電腦數量減少 20%。將已修補的電腦數量向上取整到最接近的整數。

例如，在：

- 第 0 天：共有 1 台電腦被感染。
- 第 1 天：共有 2 台電腦被感染。
- 第 2 天：共有 4 台電腦被感染。
- 第 3 天：共有 8 台電腦被感染。然後，套用補丁：8 台被感染 * 20％ = 1.6 台已修補。將 1.6 進位為 2。8 台被感染的電腦 - 2 台已修補 = 第 3 天後共有 6 台電腦被感染。

傳回經過指定天數後感染的電腦總數。

# --hints--

`infected(1)` 應該傳回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(1), 2)`)
}})
```

`infected(3)` 應該傳回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(3), 6)`)
}})
```

`infected(8)` 應該傳回 `152`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(8), 152)`)
}})
```

`infected(17)` 應該傳回 `39808`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(17), 39808)`)
}})
```

`infected(25)` 應該傳回 `5217638`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(25), 5217638)`)
}})
```

# --seed--

## --seed-contents--

```py
def infected(days):

    return days
```

# --solutions--

```py
import math
def infected(days):
    infected = 1
    for day in range(1, days + 1):
        infected *= 2
        if day % 3 == 0:
            patched = math.ceil(infected * 0.2)
            infected -= patched

    return infected
```
