---
id: 68e39ed6106dac2f0a98fd64
title: "挑戰 82：SpOoKy~CaSe"
challengeType: 29
dashedName: challenge-82
---

# --description--

給定一個表述變數名稱的字串，使用以下條件將其轉換為「spooky case」：

- 將所有底線（`_`）和連字號（`-`）取代為波浪號（`~`）。
- 將字串的第一個字元及之後每隔一個字元的大寫。計數時忽略波浪號字元。其他字元皆轉為小寫。

例如，給定 `hello_world`，傳回 `HeLlO~wOrLd`。

# --hints--

`spookify("hello_world")` 應該傳回 `"HeLlO~wOrLd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("hello_world"), "HeLlO~wOrLd")`)
}})
```

`spookify("Spooky_Case")` 應該傳回 `"SpOoKy~CaSe"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("Spooky_Case"), "SpOoKy~CaSe")`)
}})
```

`spookify("TRICK-or-TREAT")` 應該傳回 `"TrIcK~oR~tReAt"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt")`)
}})
```

`spookify("c_a-n_d-y_-b-o_w_l")` 應該傳回 `"C~a~N~d~Y~~b~O~w~L"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L")`)
}})
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` 應該傳回 `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS")`)
}})
```

# --seed--

## --seed-contents--

```py
def spookify(boo):

    return boo
```

# --solutions--

```py
def spookify(boo):
    replaced = boo.replace("_", "~").replace("-", "~")
    
    result = []
    capitalize = True

    for char in replaced:
        if char == "~":
            result.append(char)
        else:
            result.append(char.upper() if capitalize else char.lower())
            capitalize = not capitalize

    return "".join(result)
```
