---
id: 68e39ed6106dac2f0a98fd63
title: "挑戰 81：第 N 個質數"
challengeType: 29
dashedName: challenge-81
---

# --description--

質數是大於 1 的正整數，且只能被 1 和它本身整除。前五個質數是 `2`、`3`、`5`、`7` 和 `11`。

給定一個正整數 `n`，傳回第 `n` 個質數。例如，給定 `5` 傳回第 5 個質數：`11`。

# --hints--

`nth_prime(5)` 應該傳回 `11`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(5), 11)`)
}})
```

`nth_prime(10)` 應該傳回 `29`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(10), 29)`)
}})
```

`nth_prime(16)` 應該傳回 `53`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(16), 53)`)
}})
```

`nth_prime(99)` 應該傳回 `523`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(99), 523)`)
}})
```

`nth_prime(1000)` 應該傳回 `7919`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(1000), 7919)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_prime(n):

    return n
```

# --solutions--

```py
def nth_prime(n):
    primes = []
    num = 2

    while len(primes) < n:
        is_prime = True
        for i in range(2, int(num**0.5) + 1):
            if num % i == 0:
                is_prime = False
                break
        if is_prime:
            primes.append(num)
        num += 1

    return primes[n - 1]
```
