---
id: 68e39ed6106dac2f0a98fd62
title: "挑戰 80：電子郵件分類器"
challengeType: 29
dashedName: challenge-80
---

# --description--

在 1971 年 10 月 29 日，第一封電子郵件被發送，介紹了我們至今仍在使用的 `username@domain` 格式。現在，有數十億個電子郵件位址。

在這個挑戰中，你會得到一個電子郵件位址列表，並且需要先依照網域名稱（`@` 之後的部分）字母順序排序，再依照使用者名稱（`@` 之前的部分）排序。

- 排序應該不區分大小寫。
- 如果多於 1 個電子郵件具有相同的網域，請依其使用者名稱排序。
- 傳回排序後的位址陣列。
- 傳回的位址應保留其原始大小寫。

例如，給定 `["jill@mail.com", "john@example.com", "jane@example.com"]`，傳回 `["jane@example.com", "john@example.com", "jill@mail.com"]`。

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` 應該傳回 `["jane@example.com", "john@example.com", "jill@mail.com"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"])`)
}})
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` 應該傳回 `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"])`)
}})
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` 應該傳回 `["user@x.com", "user@y.com", "user@z.com"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"])`)
}})
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` 應該傳回 `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"])`)
}})
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` 應該傳回 `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"])`)
}})
```

# --seed--

## --seed-contents--

```py
def sort(emails):

    return emails
```

# --solutions--

```py
def sort(emails):
    return sorted(
        emails,
        key=lambda email: (email.split('@')[1].lower(), email.split('@')[0].lower())
    )
```
