---
id: 68ee9e3066cfd4eb2328e8a6
title: "挑戰 87：矩陣建置器"
challengeType: 29
dashedName: challenge-87
---

# --description--

給定兩個整數（列數和行數），傳回一個指定大小、以零（`0`）填滿的矩陣（陣列的陣列）。

例如，給定 `2` 和 `3`，傳回：

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`build_matrix(2, 3)` 應該傳回 `[[0, 0, 0], [0, 0, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(2, 3), [[0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(3, 2)` 應該傳回 `[[0, 0], [0, 0], [0, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(3, 2), [[0, 0], [0, 0], [0, 0]])`)
}})
```

`build_matrix(4, 3)` 應該傳回 `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(9, 1)` 應該傳回 `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def build_matrix(rows, cols):

    return rows
```

# --solutions--

```py
def build_matrix(rows, cols):

    return [[0 for _ in range(cols)] for _ in range(rows)]
```
