---
id: 68f6587287ad1f4ad39b0c85
title: "挑戰 99：指紋測試"
challengeType: 29
dashedName: challenge-99
---

# --description--

給定兩個表述指紋的字串，使用以下規則判斷它們是否匹配：

- 每個指紋將僅由小寫字母（`a-z`）組成。
- 如果兩個指紋符合以下條件，則視為匹配：
  - 它們長度相同。
  - 不同字元的數量不超過指紋長度的 10％。

# --hints--

`is_match("helloworld", "helloworld")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("helloworld", "helloworld"), True)`)
}})
```

`is_match("helloworld", "helloworlds")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("helloworld", "helloworlds"), False)`)
}})
```

`is_match("helloworld", "jelloworld")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("helloworld", "jelloworld"), True)`)
}})
```

`is_match("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"), True)`)
}})
```

`is_match("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"), True)`)
}})
```

`is_match("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_match(fingerprint_a, fingerprint_b):

    return fingerprint_a
```

# --solutions--

```py
def is_match(fingerprint_a, fingerprint_b):
    if len(fingerprint_a) != len(fingerprint_b):
        return False

    mismatches = 0

    for c1, c2 in zip(fingerprint_a, fingerprint_b):
        if c1 != c2:
            mismatches += 1
            if mismatches > len(fingerprint_a) // 10:
                return False

    return True
```
