---
id: 68f6587287ad1f4ad39b0c81
title: "挑戰 95：陣列位移"
challengeType: 29
dashedName: challenge-95
---

# --description--

給定一個陣列和一個整數，表述要將陣列位移多少個位置，傳回位移後的陣列。

- 正整數將陣列向左移動。
- 負整數會將陣列向右移動。
- 位移會在陣列中環繞。

例如，給定 `[1, 2, 3]` 和 `1`，將陣列向左移動 1 個位置，傳回 `[2, 3, 1]`。

# --hints--

`shift_array([1, 2, 3], 1)` 應該傳回 `[2, 3, 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], 1), [2, 3, 1])`)
}})
```

`shift_array([1, 2, 3], -1)` 應該傳回 `[3, 1, 2]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], -1), [3, 1, 2])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], 5)` 應該傳回 `["charlie", "alpha", "bravo"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], -11)` 應該傳回 `["bravo", "charlie", "alpha"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"])`)
}})
```

`shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` 應該傳回 `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_array(arr, n):

    return arr
```

# --solutions--

```py
def shift_array(arr, n):
    length = len(arr)
    n = n % length
    if n < 0:
        n += length

    return arr[n:] + arr[:n]
```
