---
id: 68f6587287ad1f4ad39b0c7c
title: "挑戰 90：字元限制"
challengeType: 29
dashedName: challenge-90
---

# --description--

在這個挑戰中，你會得到一個字串，並且需要判斷它是否適合用於社群媒體貼文。根據給定的規則傳回以下字串：

- 「短貼文」如果符合 40 字元限制。
- 「long post」如果它超過 40 個字元且符合 80 個字元限制內。
- 「invalid post」如果它太長而無法符合任一限制。

# --hints--

`can_post("Hello world")` 應該傳回 `"short post"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` 應該傳回 `"long post"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` 應該傳回 `"invalid post"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
