---
id: 68f6587287ad1f4ad39b0c83
title: "挑戰 97：GCD"
challengeType: 29
dashedName: challenge-97
---

# --description--

給定兩個正整數，傳回它們的最大公因數（GCD）。

- 兩個整數的最大公因數是能夠整除兩個數字且不留餘數的最大數字。

例如，`4` 的因數是 `1`、`2` 和 `4`。`6` 的因數是 `1`、`2`、`3` 和 `6`。所以給定 `4` 和 `6`，傳回 `2`，即出現在兩組因數中最大的數字。

# --hints--

`gcd(4, 6)` 應該傳回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(4, 6), 2)`)
}})
```

`gcd(20, 15)` 應該傳回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(20, 15), 5)`)
}})
```

`gcd(13, 17)` 應該傳回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(13, 17), 1)`)
}})
```

`gcd(654, 456)` 應該傳回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(654, 456), 6)`)
}})
```

`gcd(3456, 4320)` 應該傳回 `864`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(3456, 4320), 864)`)
}})
```

# --seed--

## --seed-contents--

```py
def gcd(x, y):

    return x
```

# --solutions--

```py
def gcd(x, y):
    while y != 0:
        x, y = y, x % y
    return x
```
