---
id: 68f6587287ad1f4ad39b0c82
title: "挑戰 96：今天是週末嗎？"
challengeType: 29
dashedName: challenge-96
---

# --description--

給定格式為 `"YYYY-MM-DD"` 的日期，傳回距離週末剩餘的天數。

- 週末從星期六開始。
- 如果給定的日期是星期六或星期日，傳回 `"It's the weekend!"`。
- 否則，傳回 `"X days until the weekend."`，其中 `X` 是距離星期六的天數。
- 如果 `X` 是 `1`，請使用 `"day"`（單數）而非 `"days"`（複數）。
- 請確保計算忽略你的區域的時區。

# --hints--

`days_until_weekend("2025-11-14")` 應該傳回 `"1 day until the weekend."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-11-14"), "1 day until the weekend.")`)
}})
```

`days_until_weekend("2025-01-01")` 應該傳回 `"3 days until the weekend."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-01-01"), "3 days until the weekend.")`)
}})
```

`days_until_weekend("2025-12-06")` 應該傳回 `"It's the weekend!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-12-06"), "It's the weekend!")`)
}})
```

`days_until_weekend("2026-01-27")` 應該傳回 `"4 days until the weekend."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-01-27"), "4 days until the weekend.")`)
}})
```

`days_until_weekend("2026-09-07")` 應該傳回 `"5 days until the weekend."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-09-07"), "5 days until the weekend.")`)
}})
```

`days_until_weekend("2026-11-29")` 應該傳回 `"It's the weekend!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-11-29"), "It's the weekend!")`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_weekend(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime

def days_until_weekend(date_string):
    date = datetime.strptime(date_string, "%Y-%m-%d").date()
    day_of_week = date.weekday()

    if day_of_week == 5 or day_of_week == 6:
        return "It's the weekend!"

    days_until_saturday = 5 - day_of_week
    return f"{days_until_saturday} day{'s' if days_until_saturday != 1 else ''} until the weekend."
```
