---
id: 68f6587287ad1f4ad39b0c7e
title: "挑戰 92：擴充套件擷取器"
challengeType: 29
dashedName: challenge-92
---

# --description--

給定一個表述檔名的字串，傳回該檔案的擴充套件。

- 擴充套件是檔名中最後一個句點（`.`）之後的零件。
- 如果檔案名稱不包含句點或以句點結尾，傳回 `"none"`。
- 擴充套件應保持原樣傳回，並保留大小寫。

# --hints--

`get_extension("document.txt")` 應該傳回 `"txt"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("document.txt"), "txt")`)
}})
```

`get_extension("README")` 應該傳回 `"none"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("README"), "none")`)
}})
```

`get_extension("image.PNG")` 應該傳回 `"PNG"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("image.PNG"), "PNG")`)
}})
```

`get_extension(".gitignore")` 應該傳回 `"gitignore"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension(".gitignore"), "gitignore")`)
}})
```

`get_extension("archive.tar.gz")` 應該傳回 `"gz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("archive.tar.gz"), "gz")`)
}})
```

`get_extension("final.draft.")` 應該傳回 `"none"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("final.draft."), "none")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_extension(filename):

    return filename
```

# --solutions--

```py
def get_extension(filename):
    last_dot = filename.rfind('.')
    
    if last_dot == -1 or last_dot == len(filename) - 1:
        return "none"
    
    return filename[last_dot + 1:]
```
