---
id: 68f6587287ad1f4ad39b0c84
title: "挑戰 98：矩形計數"
challengeType: 29
dashedName: challenge-98
---

# --description--

給定兩個正整數，表述矩形的寬度和高度，判斷在該矩形中可以容納多少個矩形。

- 只計算寬度和高度為整數的矩形。

例如，給定 `1` 和 `3`，傳回 `6`。三個 1x1 矩形、兩個 1x2 矩形和一個 1x3 矩形。

# --hints--

`count_rectangles(1, 3)` 應該傳回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 3), 6)`)
}})
```

`count_rectangles(3, 2)` 應該傳回 `18`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(3, 2), 18)`)
}})
```

`count_rectangles(1, 2)` 應該傳回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 2), 3)`)
}})
```

`count_rectangles(5, 4)` 應該傳回 `150`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(5, 4), 150)`)
}})
```

`count_rectangles(11, 19)` 應該傳回 `12540`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(11, 19), 12540)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_rectangles(width, height):

    return width
```

# --solutions--

```py
def count_rectangles(width, height):
    total = 0
    for w in range(1, width + 1):
        for h in range(1, height + 1):
            total += (width - w + 1) * (height - h + 1)
    return total
```
