---
id: 68f6587287ad1f4ad39b0c7f
title: "挑戰 93：母音和子音"
challengeType: 29
dashedName: challenge-93
---

# --description--

給定一個字串，傳回一個陣列，該陣列包含字串中母音和子音的數量。

- 母音由任何大小寫的 `a`、`e`、`i`、`o`、`u` 組成。
- 子音由所有其他字母（不論大小寫）組成。
- 忽略任何非字母字元。

例如，給定 `"Hello World"`，傳回 `[3, 7]`。

# --hints--

`count("Hello World")` 應該傳回 `[3, 7]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("JavaScript")` 應該傳回 `[3, 7]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("JavaScript"), [3, 7])`)
}})
```

`count("Python")` 應該傳回 `[1, 5]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Python"), [1, 5])`)
}})
```

`count("freeCodeCamp")` 應該傳回 `[5, 7]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("freeCodeCamp"), [5, 7])`)
}})
```

`count("Hello, World!")` 應該傳回 `[3, 7]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("The quick brown fox jumps over the lazy dog.")` 應該傳回 `[11, 24]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

# --seed--

## --seed-contents--

```py
def count(s):

    return s
```

# --solutions--

```py
def count(s):
    vowels = "aeiou"
    consonants = "bcdfghjklmnpqrstvwxyz"
    v = 0
    c = 0

    for char in s:
        ch = char.lower()
        if ch in vowels:
            v += 1
        if ch in consonants:
            c += 1

    return [v, c]
```
