---
id: 68ffb91507a5b645769328c4
title: "挑戰 101：Markdown 標題轉換器"
challengeType: 29
dashedName: challenge-101
---

# --description--

給定一個表述 Markdown 標題的字串，傳回等價的 HTML 標題。

有效的 Markdown 標題必須：

- 以零個或多個空格開始，接著
- 連續 1 到 6 個井字號（`#`），然後
- 至少一個空格。最後，
- 標題文字。

雜湊符號的數量決定標題的層級。例如，一個雜湊符號對應 `h1` 頁籤，六個雜湊符號對應 `h6` 頁籤。

如果給定的字串沒有上述的正確格式，傳回 `"Invalid format"`。

例如，給定 `"# My level 1 heading"`，傳回 `"<h1>My level 1 heading</h1>"`。

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤。請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`convert("# My level 1 heading")` 應該傳回 `"<h1>My level 1 heading</h1>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>")`)
}})
```

`convert("My heading")` 應該傳回 `"Invalid format"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("My heading"), "Invalid format")`)
}})
```

`convert("##### My level 5 heading")` 應該傳回 `"<h5>My level 5 heading</h5>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>")`)
}})
```

`convert("#My heading")` 應該傳回 `"Invalid format"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("#My heading"), "Invalid format")`)
}})
```

`convert("  ###  My level 3 heading")` 應該傳回 `"<h3>My level 3 heading</h3>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>")`)
}})
```

`convert("####### My level 7 heading")` 應該傳回 `"Invalid format"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("####### My level 7 heading"), "Invalid format")`)
}})
```

`convert("## My #2 heading")` 應該傳回 `"<h2>My #2 heading</h2>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("## My #2 heading"), "<h2>My #2 heading</h2>")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert(heading):

    return heading
```

# --solutions--

```py
import re
def convert(heading):
    trimmed = heading.strip()

    if not re.match(r'^#{1,6}\s+\S+', trimmed):
        return "Invalid format"

    hash_count = trimmed.find(' ')
    text = trimmed[hash_count:].lstrip()

    return f"<h{hash_count}>{text}</h{hash_count}>"
```
