---
id: 68ffb91507a5b645769328cc
title: "挑戰 109：我到底幾歲了？"
challengeType: 29
dashedName: challenge-109
---

# --description--

給定某人生日的日期格式為 `YYYY-MM-DD`，傳回該人在 2025 年 11 月 27 日的年齡。

- 假設所有生日皆為 2025 年 11 月 27 日之前的有效日期。
- 傳回年齡為整數。
- 請務必考慮該人在 2025 年是否已經過生日。

# --hints--

`calculate_age("2000-11-20")` 應該傳回 `25`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-11-20"), 25)`)
}})
```

`calculate_age("2000-12-01")` 應該傳回 `24`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-12-01"), 24)`)
}})
```

`calculate_age("2014-10-25")` 應該傳回 `11`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2014-10-25"), 11)`)
}})
```

`calculate_age("1994-01-06")` 應該傳回 `31`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-01-06"), 31)`)
}})
```

`calculate_age("1994-12-14")` 應該傳回 `30`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-12-14"), 30)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_age(birthday):

    return birthday
```

# --solutions--

```py
from datetime import datetime
def calculate_age(birthday):
    today = datetime(2025, 11, 27)
    birth_date = datetime.strptime(birthday, "%Y-%m-%d")

    age = today.year - birth_date.year

    if (today.month, today.day) < (birth_date.month, birth_date.day):
        age -= 1

    return age
```
