---
id: 68ffb91507a5b645769328c6
title: "挑戰 103：最小公倍數"
challengeType: 29
dashedName: challenge-103
---

# --description--

給定兩個整數，傳回這兩個數字的最小公倍數（LCM）。

兩個數字的最小公倍數是同時為這兩個數字的倍數中最小的正整數。 例如，給定 `4` 和 `6`，傳回 `12`，因為：

- `4` 的倍數是 `4`、`8`、`12`，依此類推。
- `6` 的倍數有 `6`、`12`、`18` 等等。
- `12` 是同時為兩者倍數的最小數字。

# --hints--

`lcm(4, 6)` 應該傳回 `12`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(4, 6), 12)`)
}})
```

`lcm(9, 6)` 應該傳回 `18`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(9, 6), 18)`)
}})
```

`lcm(10, 100)` 應該傳回 `100`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(10, 100), 100)`)
}})
```

`lcm(13, 17)` 應該傳回 `221`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(13, 17), 221)`)
}})
```

`lcm(45, 70)` 應該傳回 `630`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(45, 70), 630)`)
}})
```

# --seed--

## --seed-contents--

```py
def lcm(a, b):

    return a
```

# --solutions--

```py
def lcm(a, b):
    def gcd(x, y):
        return x if y == 0 else gcd(y, x % y)

    return abs(a * b) // gcd(a, b)
```
