---
id: 68ffb91507a5b645769328c5
title: "挑戰 102：最長字、單字"
challengeType: 29
dashedName: challenge-102
---

# --description--

給定一個句子字串，傳回句子中最長的單字。

- 字、單字 以單一空格分隔。
- 只有字母（`a-z`，不分大小寫）計入字的長度。
- 如果有多個相同長度的字，傳回第一個出現的字。
- 傳回該字串中出現的字，並移除標點符號。

# --hints--

`longest_word("The quick red fox")` 應該傳回 `"quick"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("The quick red fox"), "quick")`)
}})
```

`longest_word("Hello coding challenge.")` 應該傳回 `"challenge"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Hello coding challenge."), "challenge")`)
}})
```

`longest_word("Do Try This At Home.")` 應該傳回 `"This"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Do Try This At Home."), "This")`)
}})
```

`longest_word("This sentence... has commas, ellipses, and an exclamation point!")` 應該傳回 `"exclamation"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation")`)
}})
```

`longest_word("A tie? No way!")` 應該傳回 `"tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("A tie? No way!"), "tie")`)
}})
```

`longest_word("Wouldn't you like to know.")` 應該傳回 `"Wouldnt"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Wouldn't you like to know."), "Wouldnt")`)
}})
```

# --seed--

## --seed-contents--

```py
def longest_word(sentence):

    return sentence
```

# --solutions--

```py
import re
def longest_word(sentence):
    words = sentence.split(" ")
    longest = ""

    for word in words:
        cleaned = re.sub(r'[^a-zA-Z]', '', word)
        if len(cleaned) > len(longest):
            longest = cleaned

    return longest
```
