---
id: 69162d64f96574d9bb629eff
title: "挑戰 114：駝峰式轉蛇形"
challengeType: 29
dashedName: challenge-114
---

# --description--

給定一個駝峰式字串，請使用以下規則傳回該字串的蛇底線式版本：

- 輸入字串將只包含字母（`A-Z` 和 `a-z`），且總是以小寫字母開頭。
- 駝峰式字串中的每個大寫字母都代表一個新字。
- 將所有字母轉換為小寫。
- 以底線（`_`）分隔字、單字。

# --hints--

`to_snake("helloWorld")` 應該傳回 `"hello_world"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` 應該傳回 `"my_variable_name"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` 應該傳回 `"freecodecamp_daily_challenges"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
