---
id: 69162d64f96574d9bb629efe
title: "挑戰 113：英里轉公里"
challengeType: 29
dashedName: challenge-113
---

# --description--

給定以英里為單位的距離數字，傳回等價的公里距離。

- 輸入將永遠是非負數字。
- 1 英里相等於 1.60934 公里。
- 將結果四捨五入到兩位小數。
- 從四捨五入的結果中移除不必要的尾隨零。

# --hints--

`convert_to_km(1)` 應該傳回 `1.61`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(1), 1.61)`)
}})
```

`convert_to_km(21)` 應該傳回 `33.8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(21), 33.8)`)
}})
```

`convert_to_km(3.5)` 應該傳回 `5.63`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(3.5), 5.63)`)
}})
```

`convert_to_km(0)` 應該傳回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0), 0)`)
}})
```

`convert_to_km(0.621371)` 應該傳回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0.621371), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_km(miles):

    return miles
```

# --solutions--

```py
def convert_to_km(miles):
    km = miles * 1.60934
    return round(km, 2)
```
