---
id: 69162d64f96574d9bb629efb
title: "挑戰 110：字猜測器"
challengeType: 29
dashedName: challenge-110
---

# --description--

給定兩個長度相同的字串，一個是秘密字、一個是猜測字，使用以下規則比較猜測字與秘密字：

- 秘密字和猜測將僅包含大寫字母（`"A"` 到 `"Z"`）；
- 對於猜測中的每個字母，根據它與秘密字的匹配方式將其取代為數字：
  - 「2」如果該字母在秘密字中且位置正確。
  - 「1」如果該字母在秘密字中但位置錯誤。
  - 如果該字母不在秘密字中，則為 `"0"`。
- 秘密字中的每個字、單字最多只能使用一次。
- 完全符合 (`"2"`) 會先被指定，然後剩餘字母會從左到右指定部分符合 (`"1"`)。
- 如果一個字母在猜測中出現多次，它只能匹配秘密字中出現的次數。

例如，給定一個秘密字為 `"APPLE"` 且猜測為 `"POPPA"`，傳回 `"10201"`：

第一個 `"P"` 不在正確位置（`"1"`），`"O"` 不在秘密字中（`"0"`），第二個 `"P"` 在正確位置（`"2"`），第三個 `"P"` 是零（`"0"`），因為秘密字中的兩個 `"P"` 已被使用，而 `"A"` 不在正確位置（`"1"`）。

# --hints--

`compare("APPLE", "POPPA")` 應該傳回 `"10201"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("APPLE", "POPPA"), "10201")`)
}})
```

`compare("REACT", "TRACE")` 應該傳回 `"11221"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("REACT", "TRACE"), "11221")`)
}})
```

`compare("DEBUGS", "PYTHON")` 應該傳回 `"000000"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("DEBUGS", "PYTHON"), "000000")`)
}})
```

`compare("JAVASCRIPT", "TYPESCRIPT")` 應該傳回 `"0000222222"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222")`)
}})
```

`compare("ORANGE", "ROUNDS")` 應該傳回 `"110200"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("ORANGE", "ROUNDS"), "110200")`)
}})
```

`compare("WIRELESS", "ETHERNET")` 應該傳回 `"10021000"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("WIRELESS", "ETHERNET"), "10021000")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare(word, guess):

    return word
```

# --solutions--

```py
def compare(word, guess):
    n = len(word)
    result = ["0"] * n
    used_secret = [False] * n

    for i in range(n):
        if guess[i] == word[i]:
            result[i] = "2"
            used_secret[i] = True

    for i in range(n):
        if result[i] == "0":
            for j in range(n):
                if not used_secret[j] and guess[i] == word[j]:
                    result[i] = "1"
                    used_secret[j] = True
                    break

    return "".join(result)
```
