---
id: 69162d64f96574d9bb629efd
title: "挑戰 112：AI 偵測器"
challengeType: 29
dashedName: challenge-112
---

# --description--

今天的挑戰靈感來自於 2022 年 11 月 30 日 ChatGPT 的發布。

給定一個包含一個或多個句子的字串，判斷它是否很可能是使用以下規則由 AI 產生。

- 它包含兩個或更多的破折號（`-`）。
- 它包含兩組或以上的括號（`()`）。文字可以在括號內。
- 它包含三個或更多字、單字且長度為 7 個或以上字母。

- 字以單一空格分隔，且僅由字母（`A-Z`）組成。不要將標點符號或其他非字母視為字的一部分。

如果給定的句子符合上述任何規則，請傳回 `"AI"`，否則，請傳回 `"Human"`。

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` 應該傳回 `"Human"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` 應該傳回 `"Human"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` 應該傳回 `"AI"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` 應該傳回 `"AI"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` 應該傳回 `"AI"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
