---
id: 69162d64f96574d9bb629f01
title: "挑戰 116：排列數量"
challengeType: 29
dashedName: challenge-116
---

# --description--

給定一個字串，傳回可以由其字元組成的不同排列數量。

- 排列是字串中字元的任何重新排序。
- 不要計算重複的排列組合。
- 如果字串包含重複的字元，重複的排列組合應該只計算一次。
- 該字串將只包含字母（`A-Z`、`a-z`）。

例如，給定 `"abb"`，傳回 `3`，因為有三種不同的排列方式：`"abb"`、`"bab"` 和 `"bba"`。

# --hints--

`count_permutations("abb")` 應該傳回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` 應該傳回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` 應該傳回 `630`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` 應該傳回 `39916800`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
