---
id: 69162d64f96574d9bb629f04
title: "挑戰 119：字串壓縮"
challengeType: 29
dashedName: challenge-119
---

# --description--

給定一個字串句子，傳回句子的壓縮版本，其中連續重複的字被該字及其重複次數（小括弧內）所取代。

- 只有連續重複項會被壓縮。
- 字、單字由單一空格分隔。

例如，給定 `"yes yes yes please"`，傳回 `"yes(3) please"`。

# --hints--

`compress_string("yes yes yes please")` 應該傳回 `"yes(3) please"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("yes yes yes please"), "yes(3) please")`)
}})
```

`compress_string("I have have have apples")` 應該傳回 `"I have(3) apples"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("I have have have apples"), "I have(3) apples")`)
}})
```

`compress_string("one one three and to the the the the")` 應該傳回 `"one(2) three and to the(4)"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("one one three and to the the the the"), "one(2) three and to the(4)")`)
}})
```

`compress_string("route route route route route route tee tee tee tee tee tee")` 應該傳回 `"route(6) tee(6)"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)")`)
}})
```

# --seed--

## --seed-contents--

```py
def compress_string(sentence):

    return sentence
```

# --solutions--

```py
def compress_string(sentence):
    words = sentence.split(" ")
    result = []
    count = 1

    for i in range(len(words)):
        if i < len(words) - 1 and words[i] == words[i + 1]:
            count += 1
        else:
            if count > 1:
                result.append(f"{words[i]}({count})")
            else:
                result.append(words[i])
            count = 1

    return " ".join(result)
```
