---
id: 69162d64f96574d9bb629f02
title: "挑戰 117：對稱差集"
challengeType: 29
dashedName: challenge-117
---

# --description--

給定兩個陣列，傳回一個包含它們對稱差集的新陣列。

- 兩個集合的對稱差集是出現在任一集合中但不在兩者皆有的值的集合。
- 傳回值的順序為它們在輸入陣列中首次出現的順序。

# --hints--

`difference([1, 2, 3], [3, 4, 5])` 應該傳回 `[1, 2, 4, 5]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5])`)
}})
```

`difference(["a", "b"], ["c", "b"])` 應該傳回 `["a", "c"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"])`)
}})
```

`difference([1, "a", 2], [2, "b", "a"])` 應該傳回 `[1, "b"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"])`)
}})
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` 應該傳回 `[2, 4, 6, 8]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8])`)
}})
```

# --seed--

## --seed-contents--

```py
def difference(arr1, arr2):

    return arr1
```

# --solutions--

```py
def difference(arr1, arr2):
    diff = []

    for v in arr1:
        if v not in arr2 and v not in diff:
            diff.append(v)

    for v in arr2:
        if v not in arr1 and v not in diff:
            diff.append(v)

    return diff
```
