---
id: 691b559495c5cb5a37b9b480
title: "挑戰 120：磅轉公斤"
challengeType: 29
dashedName: challenge-120
---

# --description--

給定一個以磅為單位的重量數字，傳回字串 `"(lbs) pounds equals (kgs) kilograms."`。

- 將 `"(lbs)"` 取代為輸入的數字。
- 將 `"(kgs)"` 取代為轉換成公斤的輸入值，四捨五入至小數點後兩位，且數值始終包含兩位小數。
- 1 磅相等於 0.453592 公斤。
- 如果輸入是 `1`，請使用 `"pound"` 而非 `"pounds"`。
- 如果轉換後的值是 `1`，請使用 `"kilogram"` 而非 `"kilograms"`。

# --hints--

`convert_to_kgs(1)` 應該傳回 `"1 pound equals 0.45 kilograms."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` 應該傳回 `"0 pounds equals 0.00 kilograms."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` 應該傳回 `"100 pounds equals 45.36 kilograms."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` 應該傳回 `"2.5 pounds equals 1.13 kilograms."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` 應該傳回 `"2.20462 pounds equals 1.00 kilogram."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
