---
id: 691b559495c5cb5a37b9b489
title: "挑戰 129：Markdown 區塊引用解析器"
challengeType: 29
dashedName: challenge-129
---

# --description--

給定一個包含 Markdown 中區塊引用的字串，傳回等價的 HTML 字串。

Markdown 中的區塊引言是任何符合以下條件的行：

- 以零開始
- 後面接著大於號（`>`）
- 然後，一個或多個空格
- 最後，是區塊引用的文字。

傳回被開啟和關閉的 HTML `blockquote` 標籤包圍的區塊引用文字。

例如，給定 `"> This is a quote"`，傳回 `<blockquote>This is a quote</blockquote>`。

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤。請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`parse_blockquote("> This is a quote")` 應該傳回 `"<blockquote>This is a quote</blockquote>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` 應該傳回 `"<blockquote>This is also a quote</blockquote>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` 應該傳回 `"<blockquote>So  Is  This</blockquote>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
