---
id: 691b559495c5cb5a37b9b483
title: "挑戰 123：羅馬數字建置器"
challengeType: 29
dashedName: challenge-123
---

# --description--

給定一個整數，傳回其等價的羅馬數字值。

羅馬數字使用這些符號：

| Symbol | Value |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

羅馬數字從最大到最小、由左至右書寫，遵循以下規則：

- 當符號後面跟著一個相等或較小的值時，會使用加法。 例如，`18` 寫成 `XVIII`（10 ＋ 5 ＋ 1 ＋ 1 ＋ 1 ＝ 18）。
- 當較小的符號出現在較大的符號之前時，會使用減法來表述任何位值中的 4 或 9。例如，19 寫作 `XIX`（10 + (10 - 1)）。
- 同一個符號連續出現不得超過三次。當你必須連續書寫同一個符號超過三次時，會使用減法表示法。因此你能書寫的最大數字是 3999。

這裡有另一個範例：給定 `1464`，傳回 `"MCDLXIV"`（1000 + (500 - 100) + 50 + 10 + (5 - 1)）。

# --hints--

`to_roman(18)` 應該傳回 `"XVIII"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` 應該傳回 `"XIX"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` 應該傳回 `"MCDLXIV"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` 應該傳回 `"MMXXV"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` 應該傳回 `"MMMCMXCIX"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```
