---
id: 691b559495c5cb5a37b9b486
title: "挑戰 126：將首字母大寫"
challengeType: 29
dashedName: challenge-126
---

# --description--

給定一個字串 title，傳回一個使用以下規則格式化為標題大小寫的新字串：

- 將每個字的第一個字母大寫。
- 將每個字中所有其他字母轉為小寫。
- 字、單字總是以單一空格分隔。

# --hints--

`title_case("hello world")` 應該傳回 `"Hello World"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("hello world"), "Hello World")`)
}})
```

`title_case("the quick brown fox")` 應該傳回 `"The Quick Brown Fox"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("the quick brown fox"), "The Quick Brown Fox")`)
}})
```

`title_case("JAVASCRIPT AND PYTHON")` 應該傳回 `"Javascript And Python"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("JAVASCRIPT AND PYTHON"), "Javascript And Python")`)
}})
```

`title_case("AvOcAdO tOAst fOr brEAkfAst")` 應該傳回 `"Avocado Toast For Breakfast"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast")`)
}})
```

# --seed--

## --seed-contents--

```py
def title_case(title):

    return title
```

# --solutions--

```py
def title_case(title):
    return " ".join(
        w[:1].upper() + w[1:].lower()
        for w in title.split(" ")
    )
```
