---
id: 691b559495c5cb5a37b9b485
title: "挑戰 125：生命遊戲"
challengeType: 29
dashedName: challenge-125
---

# --description--

給定一個矩陣（陣列的陣列）表述康威生命遊戲中的當前狀態，使用這些規則傳回矩陣的下一個狀態：

- 每個儲存格要麼是 `1`（存活），要麼是 `0`（死亡）。
- 一個儲存格的鄰居是最多八個周圍的儲存格（垂直、水平和對角線方向）。
- 邊緣的儲存格鄰居少於八個。

更新每個儲存格的規則：

- 任何鄰近活細胞少於 2 個的活細胞會死亡（人口過少）。
- 任何有兩個或三個活鄰居的活細胞會繼續存活。
- 任何鄰近活細胞數超過 3 個的活細胞會死亡（過度繁殖）。
- 任何恰好有三個活鄰居的死細胞會變成活細胞（繁殖）。

例如，給定：

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

傳回：

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

每個格子會根據鄰近活著的實體數量更新。例如，`[0][0]` 保持死亡（2 個活著的鄰近實體）、`[0][1]` 保持存活（2 個活著的鄰近實體）、`[0][2]` 死亡（3 個活著的鄰近實體），依此類推。

# --hints--

`game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` 應該傳回 `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]])`)
}})
```

`game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` 應該傳回 `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]])`)
}})
```

`game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` 應該傳回 `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]])`)
}})
```

`game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` 應該傳回 `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def game_of_life(grid):

    return grid
```

# --solutions--

```py
def game_of_life(grid):
    rows, cols = len(grid), len(grid[0])

    def count_live_neighbors(r, c):
        count = 0
        for i in range(r-1, r+2):
            for j in range(c-1, c+2):
                if 0 <= i < rows and 0 <= j < cols and (i != r or j != c):
                    count += grid[i][j]
        return count

    next_state = [row[:] for row in grid]

    for r in range(rows):
        for c in range(cols):
            live_neighbors = count_live_neighbors(r, c)

            if grid[r][c] == 1:
                if live_neighbors < 2 or live_neighbors > 3:
                    next_state[r][c] = 0
            else:
                if live_neighbors == 3:
                    next_state[r][c] = 1

    return next_state
```
