---
id: 691b559495c5cb5a37b9b487
title: "挑戰 127：速度檢查"
challengeType: 29
dashedName: challenge-127
---

# --description--

根據你以英里每小時（MPH）行駛的速度，以及以公里每小時（KPH）為單位的速限，判斷你是否超速，以及你會收到警告還是罰單。

- 1 英里相等於 1.60934 公里。
- 如果你行駛速度小於或相等於速限，傳回 `"Not Speeding"`。
- 如果你超速不超過 5 KPH，傳回 `"Warning"`。
- 如果你超速超過 5 KPH，請傳回 `"Ticket"`。

# --hints--

`speed_check(30, 70)` 應該傳回 `"Not Speeding"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` 應該傳回 `"Warning"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` 應該傳回 `"Not Speeding"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` 應該傳回 `"Ticket"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` 應該傳回 `"Warning"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` 應該傳回 `"Ticket"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
