---
id: 6925e2068081f40f549ced1d
title: "挑戰 139：剪刀、石頭、布"
challengeType: 29
dashedName: challenge-139
---

# --description--

給定兩個字串，第一個表述玩家 1，第二個表述玩家 2，判斷石頭、剪刀、布比賽的勝者。

- 輸入字串將永遠是 `"Rock"`、`"Paper"` 或 `"Scissors"`。
- 「Rock」贏過「Scissors」。
- 「Paper」贏過「Rock」。
- 「Scissors」贏過「Paper」。

傳回：

- 如果 Player 1 獲勝，則顯示「Player 1 wins」。
- 如果 Player 2 獲勝，則顯示「Player 2 wins」。
- 「Tie」如果雙方玩家選擇相同的選項。

# --hints--

`rock_paper_scissors("Rock", "Rock")` 應該傳回 `"Tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Rock"), "Tie")`)
}})
```

`rock_paper_scissors("Rock", "Paper")` 應該傳回 `"Player 2 wins"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Paper"), "Player 2 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Paper")` 應該傳回 `"Player 1 wins"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Paper"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Rock", "Scissors")` 應該傳回 `"Player 1 wins"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Scissors"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Scissors")` 應該傳回 `"Tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Scissors"), "Tie")`)
}})
```

`rock_paper_scissors("Scissors", "Rock")` 應該傳回 `"Player 2 wins"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Rock"), "Player 2 wins")`)
}})
```

# --seed--

## --seed-contents--

```py
def rock_paper_scissors(player1, player2):

    return player1
```

# --solutions--

```py
def rock_paper_scissors(player1, player2):
    if player1 == player2:
        return "Tie"
    
    if (player1 == "Rock" and player2 == "Scissors") or \
       (player1 == "Paper" and player2 == "Rock") or \
       (player1 == "Scissors" and player2 == "Paper"):
        return "Player 1 wins"
    else:
        return "Player 2 wins"
```
