---
id: 6925e2068081f40f549ced1c
title: "挑戰 138：除數總和"
challengeType: 29
dashedName: challenge-138
---

# --description--

給定一個正整數，傳回它所有因數的總和。

- 除數是任何能夠整除該數字（餘數為 `0`）的整數。
- 每個除數只計算一次。

例如，給定 `6`，傳回 `12`，因為 `6` 的因數是 `1`、`2`、`3` 和 `6`，這些數字的總和是 `12`。

# --hints--

`sum_divisors(6)` 應該傳回 `12`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` 應該傳回 `14`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` 應該傳回 `56`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` 應該傳回 `224`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` 應該傳回 `806`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` 應該傳回 `23520`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
