---
id: 6925e2068081f40f549ced1b
title: "挑戰 137：Snowflake 產生器"
challengeType: 29
dashedName: challenge-137
---

# --description--

給定一個使用換行字元（`\n`）表述換行的多行字串，傳回一個新字串，其中每一行都水平鏡像並附加到原始行的末尾。

- 透過反轉所有字元（包括空格）來鏡像一行。

例如，給定 `"* \n *\n* "`，它會記錄到主控台為：

```sh
* 
 *
* 
```

傳回 `"*  *\n ** \n*  *"`，該字串會在主控台記錄為：

```sh
*  *
 ** 
*  *
```

請仔細注意給定字串和傳回字串中的空白字元。務必不要刪除任何空白字元。

# --hints--

`generate_snowflake("* \n *\n* ")` 應該傳回 `"*  *\n ** \n*  *"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` 應該傳回 `"X=~~=X"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` 應該傳回 `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` 應該傳回 `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` 應該傳回 `"*  --  *\n * -- * \n*  --  *"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
