---
id: 69272dcf1c24b44fd79137c6
title: "挑戰 143：Markdown 斜體字解析器"
challengeType: 29
dashedName: challenge-143
---

# --description--

給定一個可能包含部分 Markdown 斜體文字的字串，傳回等價的 HTML 字串。

- Markdown 中的斜體文字是任何以單一星號（`*`）或單一底線（`_`）開始和結束的文字。
- 文字和星號或底線之間不能有任何空格，但文字本身可以有空格。
- 將所有斜體字改為 HTML 的 `i` 標籤並傳回該字串。

例如，給定 `"*This is italic*"`，傳回 `"<i>This is italic</i>"`。

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤。請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`parse_italics("*This is italic*")` 應該傳回 `"<i>This is italic</i>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is italic*"), "<i>This is italic</i>")`)
}})
```

`parse_italics("_This is also italic_")` 應該傳回 `"<i>This is also italic</i>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_This is also italic_"), "<i>This is also italic</i>")`)
}})
```

`parse_italics("*This is not italic *")` 應該傳回 `"*This is not italic *"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is not italic *"), "*This is not italic *")`)
}})
```

`parse_italics("_ This is also not italic_")` 應該傳回 `"_ This is also not italic_"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_ This is also not italic_"), "_ This is also not italic_")`)
}})
```

`parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog.")` 應該傳回 `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_italics(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_italics(markdown):
    pattern = r'(\*|_)([^\s][^]*?[^\s])\1'
    pattern = r'(\*|_)([^\s].*?[^\s])\1'
    return re.sub(pattern, r'<i>\2</i>', markdown)
```
