---
id: 69272dcf1c24b44fd79137c5
title: "挑戰 142：字串求和"
challengeType: 29
dashedName: challenge-142
---

# --description--

給定一個包含數字和其他字元的字串，傳回字串中所有數字的總和。

- 將連續的數字視為一個整數。例如，`"13"` 計算為 13，而不是 1 + 3。
- 忽略任何非數字字元。

# --hints--

`string_sum("3apples2bananas")` 應該傳回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("3apples2bananas"), 5)`)
}})
```

`string_sum("10cats5dogs2birds")` 應該傳回 `17`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("10cats5dogs2birds"), 17)`)
}})
```

`string_sum("125344")` 應該傳回 `125344`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("125344"), 125344)`)
}})
```

`string_sum("a1b20c300")` 應該傳回 `321`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a1b20c300"), 321)`)
}})
```

`string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` 應該傳回 `1653`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653)`)
}})
```

# --seed--

## --seed-contents--

```py
def string_sum(s):

    return s
```

# --solutions--

```py
import re
def string_sum(s):
    numbers = re.findall(r'\d+', s)
    return sum(int(num) for num in numbers)
```
