---
id: 69272dcf1c24b44fd79137c3
title: "挑戰 140：SCREAMING_SNAKE_CASE"
challengeType: 29
dashedName: challenge-140
---

# --description--

給定一個表述變數名稱的字串，傳回轉換為 SCREAMING_SNAKE_CASE 的變數名稱。

給定的變數名稱將以以下其中一種格式書寫：

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

在上述格式中，字、單字由底線（`_`）、連字號（`-`）分隔，或新字以大寫字母開頭。

要轉換為 SCREAMING_SNAKE_CASE：

- 將所有字母轉為大寫
- 用底線（`_`）分隔字、單字

# --hints--

`to_screaming_snake_case("userEmail")` 應該傳回 `"USER_EMAIL"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("userEmail"), "USER_EMAIL")`)
}})
```

`to_screaming_snake_case("UserPassword")` 應該傳回 `"USER_PASSWORD"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("UserPassword"), "USER_PASSWORD")`)
}})
```

`to_screaming_snake_case("user_id")` 應該傳回 `"USER_ID"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user_id"), "USER_ID")`)
}})
```

`to_screaming_snake_case("user-address")` 應該傳回 `"USER_ADDRESS"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user-address"), "USER_ADDRESS")`)
}})
```

`to_screaming_snake_case("username")` 應該傳回 `"USERNAME"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("username"), "USERNAME")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_screaming_snake_case(variable_name):

    return variable_name
```

# --solutions--

```py
import re
def to_screaming_snake_case(variable_name):
    temp = re.sub(r'[-_]+', ' ', variable_name)
    temp = re.sub(r'([a-z0-9])([A-Z])', r'\1 \2', temp)
    words = temp.strip().split()
    return '_'.join(words).upper()
```
