---
id: 69306364df283fcaff2e1ad8
title: "挑戰 147：閏年計算器"
challengeType: 29
dashedName: challenge-147
---

# --description--

給定一個整數年份，判斷它是否為閏年。

如果符合以下規則，該年即為閏年：

- 該年份能被 4 整除，且
- 該年份不能被 100 整除，除非
- 該年份可被 400 整除。

# --hints--

`is_leap_year(2024)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2024), True)`)
}})
```

`is_leap_year(2023)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2023), False)`)
}})
```

`is_leap_year(2100)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2100), False)`)
}})
```

`is_leap_year(2000)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2000), True)`)
}})
```

`is_leap_year(1999)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(1999), False)`)
}})
```

`is_leap_year(2040)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2040), True)`)
}})
```

`is_leap_year(2026)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2026), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_leap_year(year):

    return year
```

# --solutions--

```py
def is_leap_year(year):
    if (year % 4 == 0 and year % 100 != 0) or (year % 400 == 0):
        return True
    else:
        return False
```
