---
id: 69306364df283fcaff2e1ad6
title: "挑戰 145：第 N 個費氏數列數字"
challengeType: 29
dashedName: challenge-145
---

# --description--

給定一個整數 `n`，傳回費氏數列中的第 `n` 個數字。

費波那契數列是一串數字，其中每個數字都是前兩個數字的總和。數列中的前 10 個數字是 `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34`。

# --hints--

`nth_fibonacci(4)` 應該傳回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(4), 2)`)
}})
```

`nth_fibonacci(10)` 應該傳回 `34`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(10), 34)`)
}})
```

`nth_fibonacci(15)` 應該傳回 `377`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(15), 377)`)
}})
```

`nth_fibonacci(40)` 應該傳回 `63245986`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(40), 63245986)`)
}})
```

`nth_fibonacci(75)` 應該傳回 `1304969544928657`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(75), 1304969544928657)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_fibonacci(n):

    return n
```

# --solutions--

```py
def nth_fibonacci(n):
    if n == 1:
        return 0
    if n == 2:
        return 1

    a, b = 0, 1
    for _ in range(3, n + 1):
        a, b = b, a + b
    return b
```
