---
id: 69306364df283fcaff2e1ad7
title: "挑戰 146：桌上的左撇子座位"
challengeType: 29
dashedName: challenge-146
---

# --description--

給定一個 4x2 矩陣（陣列的陣列）表述用餐的座位安排，判斷左撇子可以坐的座位數量。

- 左撇子無法坐在右撇子正左側的座位上。

在給定的矩陣中：

- 「R」是由右撇子佔用的座位。
- 「L」是由左撇子佔用的座位。
- 「U」是未被佔用的座位。
- 只有未被佔用的座位可供坐下。
- 最上列的座位面向「下方」，最下列的座位面向「上方」（像桌子一樣），所以左右是相對於座位的方向。
- 角落座位旁邊只有一個座位。

例如，在給定的矩陣中：

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

左上角的座位無法坐，因為左邊有一個慣用右手的人。其他兩個空位可以坐，因為左邊沒有慣用右手的人。

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` 應該傳回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` 應該傳回 `8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` 應該傳回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` 應該傳回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` 應該傳回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
