---
id: 69306364df283fcaff2e1ada
title: "挑戰 149：vOwElcAsE"
challengeType: 29
dashedName: challenge-149
---

# --description--

給定一個字串，傳回一個新字串，其中所有母音都轉換為大寫，所有其他字母字元都轉換為小寫。

- 母音是任何大小寫的 `"a"`、`"e"`、`"i"`、`"o"` 和 `"u"`。
- 非字母字元應保持不變。

# --hints--

`vowel_case("vowelcase")` 應該傳回 `"vOwElcAsE"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("vowelcase"), "vOwElcAsE")`)
}})
```

`vowel_case("coding is fun")` 應該傳回 `"cOdIng Is fUn"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("coding is fun"), "cOdIng Is fUn")`)
}})
```

`vowel_case("HELLO, world!")` 應該傳回 `"hEllO, wOrld!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HELLO, world!"), "hEllO, wOrld!")`)
}})
```

`vowel_case("git cherry-pick")` 應該傳回 `"gIt chErry-pIck"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("git cherry-pick"), "gIt chErry-pIck")`)
}})
```

`vowel_case("HEAD~1")` 應該傳回 `"hEAd~1"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HEAD~1"), "hEAd~1")`)
}})
```

# --seed--

## --seed-contents--

```py
def vowel_case(s):

    return s
```

# --solutions--

```py
def vowel_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char in vowels:
            result += char.upper()
        elif char.isalpha():
            result += char.lower()
        else:
            result += char

    return result
```
