---
id: 69373793f5a867f769cde13a
title: "挑戰 155：種植作物"
challengeType: 29
dashedName: challenge-155
---

# --description--

給定一個整數代表你的農場欄的大小，以及 `"acres"` 或 `"hectares"` 代表你的農場欄大小的單位，還有一種作物的型別，判斷你可以在你的欄中種植多少該型別的植物。

- 1 英畝相等於 4046.86 平方公尺。
- 1 公頃相等於 10,000 平方公尺。

以下是將作為輸入的作物列表以及單株植物佔用的空間：

|Crop|Space per plant|
|-|-|
|`"corn"`|1 square meter|
|`"wheat"`|0.1 square meters|
|`"soybeans"`|0.5 square meters|
|`"tomatoes"`|0.25 square meters|
|`"lettuce"`|0.2 square meters|

傳回可放入欄中的植物數量，並向下取整至最接近的整株植物。

# --hints--

`get_number_of_plants(1, "acres", "corn")` 應該傳回 `4046`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

`get_number_of_plants(2, "hectares", "lettuce")` 應該傳回 `100000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

`get_number_of_plants(20, "acres", "soybeans")` 應該傳回 `161874`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

`get_number_of_plants(3.75, "hectares", "tomatoes")` 應該傳回 `150000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

`get_number_of_plants(16.75, "acres", "tomatoes")` 應該傳回 `271139`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
