---
id: 69373793f5a867f769cde136
title: "挑戰 151：排序過的陣列？"
challengeType: 29
dashedName: challenge-151
---

# --description--

給定一個數字陣列，判斷這些數字是以升冪排序、降冪排序，還是都不是。

如果給定的陣列是：

- 以遞增順序（由最低到最高）傳回 `"Ascending"`。
- 以遞減順序（由高到低）傳回 `"Descending"`。
- 未依遞增或遞減順序排序，傳回 `"Not sorted"`。

# --hints--

`is_sorted([1, 2, 3, 4, 5])` 應該傳回 `"Ascending"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([1, 2, 3, 4, 5]), "Ascending")`)
}})
```

`is_sorted([10, 8, 6, 4, 2])` 應該傳回 `"Descending"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([10, 8, 6, 4, 2]), "Descending")`)
}})
```

`is_sorted([1, 3, 2, 4, 5])` 應該傳回 `"Not sorted"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([1, 3, 2, 4, 5]), "Not sorted")`)
}})
```

`is_sorted([3.14, 2.71, 1.61, 0.57])` 應該傳回 `"Descending"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([3.14, 2.71, 1.61, 0.57]), "Descending")`)
}})
```

`is_sorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])` 應該傳回 `"Ascending"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending")`)
}})
```

`is_sorted([0.4, 0.5, 0.3])` 應該傳回 `"Not sorted"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([0.4, 0.5, 0.3]), "Not sorted")`)
}})
```

# --seed--

## --seed-contents--

```py
def is_sorted(arr):

    return arr
```

# --solutions--

```py
def is_sorted(arr):
    ascending = True
    descending = True

    for i in range(1, len(arr)):
        if arr[i] < arr[i - 1]:
            ascending = False
        if arr[i] > arr[i - 1]:
            descending = False

    if ascending:
        return "Ascending"
    if descending:
        return "Descending"
    return "Not sorted"
```
