---
id: 69373793f5a867f769cde139
title: "挑戰 154：標準桿數"
challengeType: 29
dashedName: challenge-154
---

# --description--

給定兩個整數，分別為高爾夫球洞的標準桿數和高爾夫球手在該洞的擊球次數，傳回使用高爾夫術語的球手得分。

傳回：

- 「一桿進洞！」如果只用了 1 桿。
- 如果比標準桿少打兩桿，則為「Eagle」。
- 「Birdie」如果比標準桿少打一桿。
- 「Par」如果所用的擊球數與標準桿相同。
- 「Bogey」是指比標準桿多打一桿。
- 「Double bogey」是指比標準桿多打兩桿。

# --hints--

`golf_score(3, 3)` 應該傳回 `"Par"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 3), "Par")`)
}})
```

`golf_score(4, 3)` 應該傳回 `"Birdie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 3), "Birdie")`)
}})
```

`golf_score(3, 1)` 應該傳回 `"Hole in one!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 1), "Hole in one!")`)
}})
```

`golf_score(5, 7)` 應該傳回 `"Double bogey"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 7), "Double bogey")`)
}})
```

`golf_score(4, 5)` 應該傳回 `"Bogey"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 5), "Bogey")`)
}})
```

`golf_score(5, 3)` 應該傳回 `"Eagle"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 3), "Eagle")`)
}})
```

# --seed--

## --seed-contents--

```py
def golf_score(par, strokes):

    return par
```

# --solutions--

```py
def golf_score(par, strokes):
    if strokes == 1:
        return "Hole in one!"

    diff = strokes - par

    if diff == -2:
        return "Eagle"
    if diff == -1:
        return "Birdie"
    if diff == 0:
        return "Par"
    if diff == 1:
        return "Bogey"
    if diff == 2:
        return "Double bogey"
```
