---
id: 6939b873185d8e00d453563d
title: "挑戰 158：陣列交換"
challengeType: 29
dashedName: challenge-158
---

# --description--

給定一個包含兩個值的陣列，傳回一個值已交換的陣列。

例如，給定 `["A", "B"]` 傳回 `["B", "A"]`。

# --hints--

`array_swap(["A", "B"])` 應該傳回 `["B", "A"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["A", "B"]), ["B", "A"])`)
}})
```

`array_swap([25, 20])` 應該傳回 `[20, 25]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([25, 20]), [20, 25])`)
}})
```

`array_swap([True, False])` 應該傳回 `[False, True]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([True, False]), [False, True])`)
}})
```

`array_swap(["1", 1])` 應該傳回 `[1, "1"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["1", 1]), [1, "1"])`)
}})
```

# --seed--

## --seed-contents--

```py
def array_swap(arr):

    return arr
```

# --solutions--

```py
def array_swap(arr):

    return [arr[1], arr[0]]
```
