---
id: 6939b873185d8e00d453563b
title: "挑戰 156：奇數或偶數？"
challengeType: 29
dashedName: challenge-156
---

# --description--

給定一個正整數，如果它是奇數，傳回 `"Odd"`，如果是偶數，傳回 `"Even"`。

# --hints--

`odd_or_even(1)` 應該傳回 `"Odd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(1), "Odd")`)
}})
```

`odd_or_even(2)` 應該傳回 `"Even"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(2), "Even")`)
}})
```

`odd_or_even(13)` 應該傳回 `"Odd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(13), "Odd")`)
}})
```

`odd_or_even(196)` 應該傳回 `"Even"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(196), "Even")`)
}})
```

`odd_or_even(123456789)` 應該傳回 `"Odd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(123456789), "Odd")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even(n):

    return n
```

# --solutions--

```py
def odd_or_even(n):
    return "Even" if n % 2 == 0 else "Odd"
```
