---
id: 6939b873185d8e00d453563c
title: "挑戰 157：Markdown 聯結解析器"
challengeType: 29
dashedName: challenge-157
---

# --description--

給定 Markdown 中的聯結字串，傳回等價的 HTML 字串。

Markdown 影像、圖片具有以下格式：`"[link_text](link_url)"`。傳回 HTML `a` 頁籤的字串，其 `href` 設為 `link_url`，並以 `link_text` 作為頁籤內容。

例如，給定 `"[freeCodeCamp](https://freecodecamp.org/)"` 傳回 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`；

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤 — 請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` 應該傳回 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` 應該傳回 `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` 應該傳回 `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
