---
id: 694596b0585c11170ac7c7fa
title: "挑戰 162：能量消耗"
challengeType: 29
dashedName: challenge-162
---

# --description--

根據運動期間燃燒的卡路里數量，以及你的電子裝置在該運動期間使用的瓦特時數量，判斷哪一個使用的能量較多。

要比較它們，請使用以下轉換將兩個值都轉換為焦耳：

- 1 卡路里相等於 4184 焦耳。
- 1 瓦特時相等於 3600 焦耳。

傳回：

- 「Workout」如果該 workout 使用了更多能量。
- 「裝置」如果該裝置使用了更多能量。
- 「相等」如果兩者使用了相同量的能量。

# --hints--

`compare_energy(250, 50)` 應該傳回 `"Workout"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` 應該傳回 `"Devices"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` 應該傳回 `"Equal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` 應該傳回 `"Workout"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` 應該傳回 `"Devices"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` 應該傳回 `"Equal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
