---
id: 694596b0585c11170ac7c7fb
title: "挑戰 163：子音大小寫"
challengeType: 29
dashedName: challenge-163
---

# --description--

給定一個表述變數名稱的字串，使用以下規則將其轉換為子音大小寫：

- 所有子音字母應該轉換為大寫。
- 所有母音（任何大小寫的 `a`、`e`、`i`、`o`、`u`）都應該轉換成小寫。
- 所有連字號（`-`）都應該轉換為底線（`_`）。

# --hints--

`to_consonant_case("helloworld")` 應該傳回 `"HeLLoWoRLD"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("helloworld"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("HELLOWORLD")` 應該傳回 `"HeLLoWoRLD"`。


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("HELLOWORLD"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("_hElLO-WOrlD-")` 應該傳回 `"_HeLLo_WoRLD_"`。


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_hElLO-WOrlD-"), "_HeLLo_WoRLD_")`)
}})
```

`to_consonant_case("_~-generic_~-variable_~-name_~-here-~_")` 應該傳回 `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`。


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_consonant_case(s):

    return s
```

# --solutions--

```py
def to_consonant_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char == "-":
            result += "_"
        elif char.isalpha():
            if char in vowels:
                result += char.lower()
            else:
                result += char.upper()
        else:
            result += char

    return result
```
