---
id: 696655d24b614176d4c9b78c
title: "挑戰 169：FizzBuzz 迷你"
challengeType: 29
dashedName: challenge-169
---

# --description--

給定一個整數，根據以下規則傳回一個字串：

- 如果數字能被 3 整除，傳回 `"Fizz"`。
- 如果數字能被 5 整除，傳回 `"Buzz"`。
- 如果數字同時可被 3 和 5 整除，傳回 `"FizzBuzz"`。
- 否則，將給定的數字傳回為字串。

# --hints--

`fizz_buzz_mini(3)` 應該傳回 `"Fizz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(3), "Fizz")`)
}})
```

`fizz_buzz_mini(4)` 應該傳回 `"4"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(4), "4")`)
}})
```

`fizz_buzz_mini(35)` 應該傳回 `"Buzz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(35), "Buzz")`)
}})
```

`fizz_buzz_mini(75)` 應該傳回 `"FizzBuzz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(75), "FizzBuzz")`)
}})
```

`fizz_buzz_mini(98)` 應該傳回 `"98"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(98), "98")`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_mini(n):

    return n
```

# --solutions--

```py
def fizz_buzz_mini(n):
    if n % 3 == 0 and n % 5 == 0:
        return "FizzBuzz"
    if n % 3 == 0:
        return "Fizz"
    if n % 5 == 0:
        return "Buzz"
    return str(n)
```
