---
id: 696655d24b614176d4c9b78b
title: "挑戰 168：縮放影像"
challengeType: 29
dashedName: challenge-168
---

# --description--

給定一個表述影像寬度和高度的字串，以及一個用來縮放影像的數字，傳回縮放後的寬度和高度。

- 輸入字串的格式為 `"WxH"`。例如，`"800x600"`。
- 比例是一個用來乘以寬度和高度的數字。

以相同的 `"WxH"` 格式傳回縮放後的尺寸。

# --hints--

`scale_image("800x600", 2)` 應該傳回 `"1600x1200"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("800x600", 2), "1600x1200")`)
}})
```

`scale_image("100x100", 10)` 應該傳回 `"1000x1000"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("100x100", 10), "1000x1000")`)
}})
```

`scale_image("1024x768", 0.5)` 應該傳回 `"512x384"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("1024x768", 0.5), "512x384")`)
}})
```

`scale_image("300x200", 1.5)` 應該傳回 `"450x300"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("300x200", 1.5), "450x300")`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_image(size, scale):

    return size
```

# --solutions--

```py
def scale_image(size, scale):
    width, height = map(int, size.split("x"))

    new_width = round(width * scale)
    new_height = round(height * scale)

    return f"{new_width}x{new_height}"
```
