---
id: 696655d24b614176d4c9b78d
title: "挑戰 170：奇數日或偶數日"
challengeType: 29
dashedName: challenge-170
---

# --description--

給定一個時間戳記（自 Unix 紀元以來的毫秒數），傳回：

- 如果該時間戳的日期是奇數，則為 `"odd"`。
- 如果該時間戳的日期是偶數，則為 "even"。

例如，給定 `1769472000000`，一個代表 2026 年 1 月 27 日的時間戳，傳回 `"odd"`，因為日期數字（`27`）是奇數。

注意：時間戳記是以毫秒為單位，且你應該使用 UTC 時區的日期，而非你的區域的時間。

# --hints--

`odd_or_even_day(1769472000000)` 應該傳回 `"odd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` 應該傳回 `"even"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` 應該傳回 `"odd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` 應該傳回 `"odd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` 應該傳回 `"even"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
