---
id: 69738771fb5a7b8b24cca29d
title: "挑戰 171：展平陣列"
challengeType: 29
dashedName: challenge-171
---

# --description--

給定一個包含巢狀陣列的陣列，傳回一個新的陣列，將所有值攤平成單一的一維陣列。保留陣列中條款的原始順序。

# --hints--

`flatten([1, [2, 3], 4])` 應該傳回 `[1, 2, 3, 4]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([1, [2, 3], 4]), [1, 2, 3, 4])`)
}})
```

`flatten([5, [4, [3, 2]], 1])` 應該傳回 `[5, 4, 3, 2, 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([5, [4, [3, 2]], 1]), [5, 4, 3, 2, 1])`)
}})
```

`flatten(["A", [[[["B"]]]], "C"])` 應該傳回 `["A", "B", "C"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten(["A", [[[["B"]]]], "C"]), ["A", "B", "C"])`)
}})
```

`flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]])` 應該傳回 `["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]]), ["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"])`)
}})
```

`flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]])` 應該傳回 `["red","blue","green","yellow","purple","orange","pink","brown"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]]), ["red","blue","green","yellow","purple","orange","pink","brown"])`)
}})
```

# --seed--

## --seed-contents--

```py
def flatten(arr):

    return arr
```

# --solutions--

```py
def flatten(arr):
    result = []
    for item in arr:
        if isinstance(item, list):
            result.extend(flatten(item))
        else:
            result.append(item)
    return result
```
