---
id: 69738771fb5a7b8b24cca2a5
title: "挑戰 179：零錢"
challengeType: 29
dashedName: challenge-179
---

# --description--

給定一個整數陣列，表述你口袋中硬幣的數量，每個整數為硬幣的分值，傳回總金額，格式為 `"$D.CC"`。

- 100 分相等於 1 元。
- 在返回值中，對於小於 1 元的金額包含前導零，並且分 Always exactly two digits for the cents。

# --hints--

`count_change([25, 10, 5, 1])` 應該傳回 `"$0.41"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1]), "$0.41")`)
}})
```

`count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` 應該傳回 `"$1.43"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43")`)
}})
```

`count_change([100, 25, 100, 1000, 5, 500, 2000, 25])` 應該傳回 `"$37.55"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55")`)
}})
```

`count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` 應該傳回 `"$0.70"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70")`)
}})
```

`count_change([1])` 應該傳回 `"$0.01"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([1]), "$0.01")`)
}})
```

`count_change([25, 25, 25, 25])` 應該傳回 `"$1.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 25, 25, 25]), "$1.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_change(change):

    return change
```

# --solutions--

```py
def count_change(change):
    total_cents = sum(change)
    dollars = total_cents // 100
    cents = total_cents % 100
    return f"${dollars}.{cents:02d}"
```
