---
id: 69738771fb5a7b8b24cca2a3
title: "挑戰 177：字串鏡像"
challengeType: 29
dashedName: challenge-177
---

# --description--

給定一個字串，傳回一個新字串，該字串由給定字串與其反轉副本附加在末尾所組成。

# --hints--

`mirror("freeCodeCamp")` 應該傳回 `"freeCodeCamppmaCedoCeerf"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf")`)
}})
```

`mirror("RaceCar")` 應該傳回 `"RaceCarraCecaR"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("RaceCar"), "RaceCarraCecaR")`)
}})
```

`mirror("helloworld")` 應該傳回 `"helloworlddlrowolleh"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("helloworld"), "helloworlddlrowolleh")`)
}})
```

`mirror("The quick brown fox...")` 應該傳回 `"The quick brown fox......xof nworb kciuq ehT"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT")`)
}})
```

# --seed--

## --seed-contents--

```py
def mirror(s):

    return s
```

# --solutions--

```py
def mirror(s):
    reversed_s = s[::-1]
    return s + reversed_s
```
