---
id: 69738771fb5a7b8b24cca2a4
title: "挑戰 178：截斷 `Text`"
challengeType: 29
dashedName: challenge-178
---

# --description--

給定一個字串，如果它的長度在 20 個字元以內，則原樣傳回。如果超過 20 個字元，則將它截斷為前 17 個字元，並在結尾附加 `"..."`（總共 20 個字元），然後傳回結果。

# --hints--

`truncate_text("Hello, world!")` 應該傳回 `"Hello, world!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Hello, world!"), "Hello, world!")`)
}})
```

`truncate_text("This string should get truncated.")` 應該傳回 `"This string shoul..."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("This string should get truncated."), "This string shoul...")`)
}})
```

`truncate_text("Exactly twenty chars")` 應該傳回 `"Exactly twenty chars"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Exactly twenty chars"), "Exactly twenty chars")`)
}})
```

`truncate_text(".....................")` 應該傳回 `"...................."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("....................."), "....................")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(text):

    return text
```

# --solutions--

```py
def truncate_text(text):
    if len(text) <= 20:
        return text
    return text[:17] + "..."
```
