---
id: 69738771fb5a7b8b24cca2a2
title: "挑戰 176：土撥鼠日"
challengeType: 29
dashedName: challenge-176
---

# --description--

今天是土撥鼠日，土撥鼠會根據是否看到自己的影子來預測天氣。

給定一個表述土撥鼠外觀的值，傳回正確的預測：

- 如果給定的值是布林值 `true`（土撥鼠看到了牠的影子），傳回 `"Looks like we'll have six more weeks of winter."`。
- 如果值是布林值 `false`（土撥鼠沒有看到牠的影子），傳回 `"It's going to be an early spring."`。
- 如果值是其他任何東西（土撥鼠沒有出現），傳回 `"No prediction this year."`。

# --hints--

`groundhog_day_prediction(True)` 應該傳回 `"Looks like we'll have six more weeks of winter."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

`groundhog_day_prediction(False)` 應該傳回 `"It's going to be an early spring."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

`groundhog_day_prediction(None)` 應該傳回 `"No prediction this year."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` 應該傳回 `"No prediction this year."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

`groundhog_day_prediction("True")` 應該傳回 `"No prediction this year."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
