---
id: 69738771fb5a7b8b24cca29f
title: "挑戰 173：有效的兵的移動"
challengeType: 29
dashedName: challenge-173
---

# --description--

給定你其中一個棋子在棋盤上的位置，傳回一個包含它可以移動到的所有有效格子且按升冪排序的陣列。

標準的棋盤是 8x8，行從左到右標記為 `A` 到 `H`，列從下到上標記為 `1` 到 `8`。它看起來像這樣：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

對於這個挑戰：

- 你是棋盤底部的玩家。
- 兵只能「向上」移動一格。
- 除非兵在起始列（第 2 列），否則它只能向上移動一格或兩格。

例如，給定 `"D4"`，傳回 `["D5"]`，這是你的兵唯一能移動到的格子。  
給定 `"B2"`，傳回 `["B3", "B4"]`，因為它在起始列且需要以遞增順序排列。

# --hints--

`find_pawn_moves("D4")` 應該傳回 `["D5"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("D4"), ["D5"])`)
}})
```

`find_pawn_moves("B2")` 應該傳回 `["B3", "B4"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("B2"), ["B3", "B4"])`)
}})
```

`find_pawn_moves("A7")` 應該傳回 `["A8"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("A7"), ["A8"])`)
}})
```

`find_pawn_moves("G2")` 應該傳回 `["G3", "G4"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("G2"), ["G3", "G4"])`)
}})
```

`find_pawn_moves("E3")` 應該傳回 `["E4"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("E3"), ["E4"])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_pawn_moves(position):

    return position
```

# --solutions--

```py
def find_pawn_moves(position):
    column = position[0].upper()
    row = int(position[1])
    moves = []

    if row >= 8:
        return moves

    moves.append(f"{column}{row + 1}")

    if row == 2:
        moves.append(f"{column}{row + 2}")

    return moves
```
