---
id: 69738771fb5a7b8b24cca29e
title: "挑戰 172：字母-數字"
challengeType: 29
dashedName: challenge-172
---

# --description--

給定一個只包含字母和數字的字串，傳回一個新字串，當字串從字母切換到數字，或從數字切換到字母時，每次插入一個連字號（`-`）。

# --hints--

`separate_letters_and_numbers("ABC123")` 應該傳回 `"ABC-123"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(separate_letters_and_numbers("ABC123"), "ABC-123")`)
}})
```

`separate_letters_and_numbers("Route66")` 應該傳回 `"Route-66"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(separate_letters_and_numbers("Route66"), "Route-66")`)
}})
```

`separate_letters_and_numbers("H3LL0W0RLD")` 應該傳回 `"H-3-LL-0-W-0-RLD"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(separate_letters_and_numbers("H3LL0W0RLD"), "H-3-LL-0-W-0-RLD")`)
}})
```

`separate_letters_and_numbers("a1b2c3d4")` 應該傳回 `"a-1-b-2-c-3-d-4"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(separate_letters_and_numbers("a1b2c3d4"), "a-1-b-2-c-3-d-4")`)
}})
```

# --seed--

## --seed-contents--

```py
def separate_letters_and_numbers(s):

    return s
```

# --solutions--

```py
def separate_letters_and_numbers(s):
    if not s:
        return s

    result = s[0]

    for i in range(1, len(s)):
        prev = s[i - 1]
        curr = s[i]

        prev_is_letter = prev.isalpha()
        curr_is_letter = curr.isalpha()

        if prev_is_letter != curr_is_letter:
            result += "-"

        result += curr

    return result
```
