---
id: 697a49e6ff50d756c9b69366
title: "挑戰 189：2026 冬季遊戲第 10 天：高山滑雪"
challengeType: 29
dashedName: challenge-189
---

# --description--

給定滑雪坡的垂直落差、水平距離和型別，判斷該坡的難度評級。

要判定評等：

- 計算坡度的方法是將落差除以距離。
- 然後，根據山丘的 `type` 計算調整後的陡度：
  - 「Downhill」將陡峭度乘以 1.2
  - 「Slalom」：將陡峭度乘以 0.9
  - 「Giant Slalom」：將陡峭度乘以 1.0

傳回：

- 「Green」如果調整後的陡度小於或相等於 0.1
- 「Blue」如果調整後的陡度大於 0.1 且小於或相等於 0.25
- 當調整後的陡度大於 0.25 時，顯示「Black」


# --hints--

`get_hill_rating(95, 900, "Slalom")` 應該傳回 `"Green"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(95, 900, "Slalom"), "Green")`)
}})
```

`get_hill_rating(620, 2800, "Downhill")` 應該傳回 `"Black"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(620, 2800, "Downhill"), "Black")`)
}})
```

`get_hill_rating(420, 1680, "Giant Slalom")` 應該傳回 `"Blue"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(420, 1680, "Giant Slalom"), "Blue")`)
}})
```

`get_hill_rating(250, 3000, "Downhill")` 應該傳回 `"Green"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(250, 3000, "Downhill"), "Green")`)
}})
```

`get_hill_rating(110, 900, "Slalom")` 應該傳回 `"Blue"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(110, 900, "Slalom"), "Blue")`)
}})
```

`get_hill_rating(380, 1500, "Giant Slalom")` 應該傳回 `"Black"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(380, 1500, "Giant Slalom"), "Black")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_hill_rating(drop, distance, hill_type):

    return drop
```

# --solutions--

```py
def get_hill_rating(drop, distance, hill_type):
    steepness = drop / distance

    if hill_type == "Downhill":
        steepness *= 1.2
    elif hill_type == "Slalom":
        steepness *= 0.9
    elif hill_type == "Giant Slalom":
        steepness *= 1.0

    if steepness <= 0.1:
        return "Green"
    elif steepness <= 0.25:
        return "Blue"
    else:
        return "Black"
```
