---
id: 697a49e6ff50d756c9b69365
title: "挑戰 188：2026 冬季遊戲第 9 天：Skeleton"
challengeType: 29
dashedName: challenge-188
---

# --description--

給定一個表述骨架賽道曲線的字串，判斷該賽道的難度。

- 給定的字串將只包含以下字母：
  - 「L」代表左轉
  - 「R」表示右轉
  - 「S」代表直線段

- 每次方向改變會增加 15 分（`"L"` 後接 `"R"` 或反之亦然）。
- 所有其他曲線各加 5 點（所有其他 `"L"` 或 `"R"` 字元）。
- 直線段加 0 分。

賽道的難度是根據總分計算。傳回：

- 「Easy」如果總數是 0 - 100
- 「Medium」如果總數在 101-200 之間
- 如果總數超過 200，則為「Hard」

# --hints--

`get_difficulty("SLSLLSRRLSRLRL")` 應該傳回 `"Easy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLSLLSRRLSRLRL"), "Easy")`)
}})
```

`get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` 應該傳回 `"Hard"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard")`)
}})
```

`get_difficulty("SRRRRLSLLRLRSSRLSRL")` 應該傳回 `"Medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRRRRLSLLRLRSSRLSRL"), "Medium")`)
}})
```

`get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` 應該傳回 `"Hard"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard")`)
}})
```

`get_difficulty("SLLSSLRLSLSLRSLSSLRL")` 應該傳回 `"Medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium")`)
}})
```

`get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` 應該傳回 `"Easy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_difficulty(track):

    return track
```

# --solutions--

```py
def get_difficulty(track):
  score = 0

  for i in range(len(track)):
    current = track[i]
    previous = track[i - 1] if i > 0 else None

    if current == 'S':
      score += 0
    elif current == 'L' or current == 'R':
      if previous and previous != 'S' and previous != current:
        score += 15
      else:
        score += 5

  if score <= 100:
    return 'Easy'
  elif score <= 200:
    return 'Medium'
  else:
    return 'Hard'
```
