---
id: 697a49e6ff50d756c9b6935d
title: "挑戰 180：冬季運動會第 1 天：開幕日"
challengeType: 29
dashedName: challenge-180
---

# --description--

接下來的 17 天將帶給你受冬季運動會啟發的程式碼挑戰。

對於第一個，你會得到一個兩個字母的國家代碼，並且需要傳回該國家的旗標表情符號。

使用此列表：

| 國家 | 代碼 | 旗標 |
| - | -: | :--: |
| 阿爾巴尼亞                    |            `"AL"` | `"🇦🇱"` |
| 安道爾                    |            `"AD"` | `"🇦🇩"` |
| 阿根廷                  |            `"AR"` | `"🇦🇷"` |
| 亞美尼亞                    |            `"AM"` | `"🇦🇲"` |
| 澳洲                  |            `"AU"` | `"🇦🇺"` |
| 奧地利                    |            `"AT"` | `"🇦🇹"` |
| 亞塞拜然                 |            `"AZ"` | `"🇦🇿"` |
| 比利時                    |            `"BE"` | `"🇧🇪"` |
| 貝寧                      |            `"BJ"` | `"🇧🇯"` |
| 玻利維亞                    |            `"BO"` | `"🇧🇴"` |
| 波士尼亞與赫塞哥維納     |            `"BA"` | `"🇧🇦"` |
| 巴西                     |            `"BR"` | `"🇧🇷"` |
| 保加利亞                   |            `"BG"` | `"🇧🇬"` |
| 加拿大                     |            `"CA"` | `"🇨🇦"` |
| 智利                      |            `"CL"` | `"🇨🇱"` |
| 中國                      |            `"CN"` | `"🇨🇳"` |
| 哥倫比亞                   |            `"CO"` | `"🇨🇴"` |
| 克羅埃西亞                    |            `"HR"` | `"🇭🇷"` |
| 賽普勒斯                     |            `"CY"` | `"🇨🇾"` |
| 捷克共和國             |            `"CZ"` | `"🇨🇿"` |
| 丹麥                    |            `"DK"` | `"🇩🇰"` |
| 厄瓜多爾                    |            `"EC"` | `"🇪🇨"` |
| 厄立特里亞                    |            `"ER"` | `"🇪🇷"` |
| 愛沙尼亞                    |            `"EE"` | `"🇪🇪"` |
| 芬蘭                    |            `"FI"` | `"🇫🇮"` |
| 法國                     |            `"FR"` | `"🇫🇷"` |
| 喬治亞                    |            `"GE"` | `"🇬🇪"` |
| 德國                    |            `"DE"` | `"🇩🇪"` |
| 大不列顛              |            `"GB"` | `"🇬🇧"` |
| 希臘                     |            `"GR"` | `"🇬🇷"` |
| 幾內亞比紹              |            `"GW"` | `"🇬🇼"` |
| 海地                      |            `"HT"` | `"🇭🇹"` |
| 香港                  |            `"HK"` | `"🇭🇰"` |
| 匈牙利                    |            `"HU"` | `"🇭🇺"` |
| 冰島                    |            `"IS"` | `"🇮🇸"` |
| 印度                      |            `"IN"` | `"🇮🇳"` |
| 伊朗                       |            `"IR"` | `"🇮🇷"` |
| 愛爾蘭                    |            `"IE"` | `"🇮🇪"` |
| 以色列                     |            `"IL"` | `"🇮🇱"` |
| 義大利                      |            `"IT"` | `"🇮🇹"` |
| 牙買加                    |            `"JM"` | `"🇯🇲"` |
| 日本                      |            `"JP"` | `"🇯🇵"` |
| 哈薩克斯坦                 |            `"KZ"` | `"🇰🇿"` |
| 肯亞                      |            `"KE"` | `"🇰🇪"` |
| 科索沃                     |            `"XK"` | `"🇽🇰"` |
| 吉爾吉斯斯坦                 |            `"KG"` | `"🇰🇬"` |
| 拉脫維亞                     |            `"LV"` | `"🇱🇻"` |
| 黎巴嫩                    |            `"LB"` | `"🇱🇧"` |
| 列支敦斯登              |            `"LI"` | `"🇱🇮"` |
| 立陶宛                  |            `"LT"` | `"🇱🇹"` |
| 盧森堡                 |            `"LU"` | `"🇱🇺"` |
| 馬達加斯加                 |            `"MG"` | `"🇲🇬"` |
| 馬來西亞                   |            `"MY"` | `"🇲🇾"` |
| 馬爾他                      |            `"MT"` | `"🇲🇹"` |
| 墨西哥                     |            `"MX"` | `"🇲🇽"` |
| 摩爾多瓦                    |            `"MD"` | `"🇲🇩"` |
| 摩納哥                     |            `"MC"` | `"🇲🇨"` |
| 蒙古                   |            `"MN"` | `"🇲🇳"` |
| 黑山                 |            `"ME"` | `"🇲🇪"` |
| 摩洛哥                    |            `"MA"` | `"🇲🇦"` |
| 荷蘭                |            `"NL"` | `"🇳🇱"` |
| 紐西蘭                |            `"NZ"` | `"🇳🇿"` |
| 奈及利亞                    |            `"NG"` | `"🇳🇬"` |
| 北馬其頓            |            `"MK"` | `"🇲🇰"` |
| 挪威                     |            `"NO"` | `"🇳🇴"` |
| 巴基斯坦                   |            `"PK"` | `"🇵🇰"` |
| 菲律賓                |            `"PH"` | `"🇵🇭"` |
| 波蘭                     |            `"PL"` | `"🇵🇱"` |
| 葡萄牙                   |            `"PT"` | `"🇵🇹"` |
| 波多黎各                |            `"PR"` | `"🇵🇷"` |
| 羅馬尼亞                    |            `"RO"` | `"🇷🇴"` |
| 聖馬利諾                 |            `"SM"` | `"🇸🇲"` |
| 沙烏地阿拉伯               |            `"SA"` | `"🇸🇦"` |
| 塞爾維亞                     |            `"RS"` | `"🇷🇸"` |
| 新加坡                  |            `"SG"` | `"🇸🇬"` |
| 斯洛伐克                   |            `"SK"` | `"🇸🇰"` |
| 斯洛維尼亞                   |            `"SI"` | `"🇸🇮"` |
| 南非               |            `"ZA"` | `"🇿🇦"` |
| 南韓                |            `"KR"` | `"🇰🇷"` |
| 西班牙                      |            `"ES"` | `"🇪🇸"` |
| 瑞典                     |            `"SE"` | `"🇸🇪"` |
| 瑞士                |            `"CH"` | `"🇨🇭"` |
| 泰國                   |            `"TH"` | `"🇹🇭"` |
| 千里達及托巴哥          |            `"TT"` | `"🇹🇹"` |
| 土耳其                     |            `"TR"` | `"🇹🇷"` |
| 烏克蘭                    |            `"UA"` | `"🇺🇦"` |
| 阿拉伯聯合大公國       |            `"AE"` | `"🇦🇪"` |
| 美國              |            `"US"` | `"🇺🇸"` |
| 烏拉圭                    |            `"UY"` | `"🇺🇾"` |
| 烏茲別克                 |            `"UZ"` | `"🇺🇿"` |
| 委內瑞拉                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`get_flag("AL")` 應該傳回 `"🇦🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "🇦🇱")`)
}})
```

`get_flag("AD")` 應該傳回 `"🇦🇩"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "🇦🇩")`)
}})
```

`get_flag("AR")` 應該傳回 `"🇦🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "🇦🇷")`)
}})
```

`get_flag("AM")` 應該傳回 `"🇦🇲"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "🇦🇲")`)
}})
```

`get_flag("AU")` 應該傳回 `"🇦🇺"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "🇦🇺")`)
}})
```

`get_flag("AT")` 應該傳回 `"🇦🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "🇦🇹")`)
}})
```

`get_flag("AZ")` 應該傳回 `"🇦🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "🇦🇿")`)
}})
```

`get_flag("BE")` 應該傳回 `"🇧🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "🇧🇪")`)
}})
```

`get_flag("BJ")` 應該傳回 `"🇧🇯"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "🇧🇯")`)
}})
```

`get_flag("BO")` 應該傳回 `"🇧🇴"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "🇧🇴")`)
}})
```

`get_flag("BA")` 應該傳回 `"🇧🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "🇧🇦")`)
}})
```

`get_flag("BR")` 應該傳回 `"🇧🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "🇧🇷")`)
}})
```

`get_flag("BG")` 應該傳回 `"🇧🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "🇧🇬")`)
}})
```

`get_flag("CA")` 應該傳回 `"🇨🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "🇨🇦")`)
}})
```

`get_flag("CL")` 應該傳回 `"🇨🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "🇨🇱")`)
}})
```

`get_flag("CN")` 應該傳回 `"🇨🇳"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "🇨🇳")`)
}})
```

`get_flag("CO")` 應該傳回 `"🇨🇴"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "🇨🇴")`)
}})
```

`get_flag("HR")` 應該傳回 `"🇭🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "🇭🇷")`)
}})
```

`get_flag("CY")` 應該傳回 `"🇨🇾"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "🇨🇾")`)
}})
```

`get_flag("CZ")` 應該傳回 `"🇨🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "🇨🇿")`)
}})
```

`get_flag("DK")` 應該傳回 `"🇩🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "🇩🇰")`)
}})
```

`get_flag("EC")` 應該傳回 `"🇪🇨"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "🇪🇨")`)
}})
```

`get_flag("ER")` 應該傳回 `"🇪🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "🇪🇷")`)
}})
```

`get_flag("EE")` 應該傳回 `"🇪🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "🇪🇪")`)
}})
```

`get_flag("FI")` 應該傳回 `"🇫🇮"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "🇫🇮")`)
}})
```

`get_flag("FR")` 應該傳回 `"🇫🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "🇫🇷")`)
}})
```

`get_flag("GE")` 應該傳回 `"🇬🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "🇬🇪")`)
}})
```

`get_flag("DE")` 應該傳回 `"🇩🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "🇩🇪")`)
}})
```

`get_flag("GB")` 應該傳回 `"🇬🇧"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "🇬🇧")`)
}})
```

`get_flag("GR")` 應該傳回 `"🇬🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "🇬🇷")`)
}})
```

`get_flag("GW")` 應該傳回 `"🇬🇼"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "🇬🇼")`)
}})
```

`get_flag("HT")` 應該傳回 `"🇭🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "🇭🇹")`)
}})
```

`get_flag("HK")` 應該傳回 `"🇭🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "🇭🇰")`)
}})
```

`get_flag("HU")` 應該傳回 `"🇭🇺"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "🇭🇺")`)
}})
```

`get_flag("IS")` 應該傳回 `"🇮🇸"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "🇮🇸")`)
}})
```

`get_flag("IN")` 應該傳回 `"🇮🇳"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "🇮🇳")`)
}})
```

`get_flag("IR")` 應該傳回 `"🇮🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "🇮🇷")`)
}})
```

`get_flag("IE")` 應該傳回 `"🇮🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "🇮🇪")`)
}})
```

`get_flag("IL")` 應該傳回 `"🇮🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "🇮🇱")`)
}})
```

`get_flag("IT")` 應該傳回 `"🇮🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "🇮🇹")`)
}})
```

`get_flag("JM")` 應該傳回 `"🇯🇲"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "🇯🇲")`)
}})
```

`get_flag("JP")` 應該傳回 `"🇯🇵"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "🇯🇵")`)
}})
```

`get_flag("KZ")` 應該傳回 `"🇰🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "🇰🇿")`)
}})
```

`get_flag("KE")` 應該傳回 `"🇰🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "🇰🇪")`)
}})
```

`get_flag("XK")` 應該傳回 `"🇽🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "🇽🇰")`)
}})
```

`get_flag("KG")` 應該傳回 `"🇰🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "🇰🇬")`)
}})
```

`get_flag("LV")` 應該傳回 `"🇱🇻"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "🇱🇻")`)
}})
```

`get_flag("LB")` 應該傳回 `"🇱🇧"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "🇱🇧")`)
}})
```

`get_flag("LI")` 應該傳回 `"🇱🇮"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "🇱🇮")`)
}})
```

`get_flag("LT")` 應該傳回 `"🇱🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "🇱🇹")`)
}})
```

`get_flag("LU")` 應該傳回 `"🇱🇺"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "🇱🇺")`)
}})
```

`get_flag("MG")` 應該傳回 `"🇲🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "🇲🇬")`)
}})
```

`get_flag("MY")` 應該傳回 `"🇲🇾"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "🇲🇾")`)
}})
```

`get_flag("MT")` 應該傳回 `"🇲🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "🇲🇹")`)
}})
```

`get_flag("MX")` 應該傳回 `"🇲🇽"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "🇲🇽")`)
}})
```

`get_flag("MD")` 應該傳回 `"🇲🇩"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "🇲🇩")`)
}})
```

`get_flag("MC")` 應該傳回 `"🇲🇨"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "🇲🇨")`)
}})
```

`get_flag("MN")` 應該傳回 `"🇲🇳"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "🇲🇳")`)
}})
```

`get_flag("ME")` 應該傳回 `"🇲🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "🇲🇪")`)
}})
```

`get_flag("MA")` 應該傳回 `"🇲🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "🇲🇦")`)
}})
```

`get_flag("NL")` 應該傳回 `"🇳🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "🇳🇱")`)
}})
```

`get_flag("NZ")` 應該傳回 `"🇳🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "🇳🇿")`)
}})
```

`get_flag("NG")` 應該傳回 `"🇳🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "🇳🇬")`)
}})
```

`get_flag("MK")` 應該傳回 `"🇲🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "🇲🇰")`)
}})
```

`get_flag("NO")` 應該傳回 `"🇳🇴"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "🇳🇴")`)
}})
```

`get_flag("PK")` 應該傳回 `"🇵🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "🇵🇰")`)
}})
```

`get_flag("PH")` 應該傳回 `"🇵🇭"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "🇵🇭")`)
}})
```

`get_flag("PL")` 應該傳回 `"🇵🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "🇵🇱")`)
}})
```

`get_flag("PT")` 應該傳回 `"🇵🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "🇵🇹")`)
}})
```

`get_flag("PR")` 應該傳回 `"🇵🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "🇵🇷")`)
}})
```

`get_flag("RO")` 應該傳回 `"🇷🇴"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "🇷🇴")`)
}})
```

`get_flag("SM")` 應該傳回 `"🇸🇲"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "🇸🇲")`)
}})
```

`get_flag("SA")` 應該傳回 `"🇸🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "🇸🇦")`)
}})
```

`get_flag("RS")` 應該傳回 `"🇷🇸"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "🇷🇸")`)
}})
```

`get_flag("SG")` 應該傳回 `"🇸🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "🇸🇬")`)
}})
```

`get_flag("SK")` 應該傳回 `"🇸🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "🇸🇰")`)
}})
```

`get_flag("SI")` 應該傳回 `"🇸🇮"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "🇸🇮")`)
}})
```

`get_flag("ZA")` 應該傳回 `"🇿🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "🇿🇦")`)
}})
```

`get_flag("KR")` 應該傳回 `"🇰🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "🇰🇷")`)
}})
```

`get_flag("ES")` 應該傳回 `"🇪🇸"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "🇪🇸")`)
}})
```

`get_flag("SE")` 應該傳回 `"🇸🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "🇸🇪")`)
}})
```

`get_flag("CH")` 應該傳回 `"🇨🇭"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "🇨🇭")`)
}})
```

`get_flag("TH")` 應該傳回 `"🇹🇭"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "🇹🇭")`)
}})
```

`get_flag("TT")` 應該傳回 `"🇹🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "🇹🇹")`)
}})
```

`get_flag("TR")` 應該傳回 `"🇹🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "🇹🇷")`)
}})
```

`get_flag("UA")` 應該傳回 `"🇺🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "🇺🇦")`)
}})
```

`get_flag("AE")` 應該傳回 `"🇦🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "🇦🇪")`)
}})
```

`get_flag("US")` 應該傳回 `"🇺🇸"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "🇺🇸")`)
}})
```

`get_flag("UY")` 應該傳回 `"🇺🇾"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "🇺🇾")`)
}})
```

`get_flag("UZ")` 應該傳回 `"🇺🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "🇺🇿")`)
}})
```

`get_flag("VE")` 應該傳回 `"🇻🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "🇻🇪")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
