---
id: 697a49e6ff50d756c9b6935f
title: "挑戰 182：2026 冬季遊戲第 3 天：冬季兩項"
challengeType: 29
dashedName: challenge-182
---

# --description--

給定一個整數陣列，其中每個值表述在一次冬季兩項比賽回合中擊中的標的數量，傳回運動員必須滑行的總罰距離。

- 每回合包含 5 個標的。
- 每錯過一個標的，將導致 150 公尺的懲罰迴圈。

# --hints--

`calculate_penalty_distance([4, 4])` 應該傳回 `300`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 4]), 300)`)
}})
```

`calculate_penalty_distance([5, 5])` 應該傳回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 5]), 0)`)
}})
```

`calculate_penalty_distance([4, 5, 3, 5])` 應該傳回 `450`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 5, 3, 5]), 450)`)
}})
```

`calculate_penalty_distance([5, 4, 5, 5])` 應該傳回 `150`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 4, 5, 5]), 150)`)
}})
```

`calculate_penalty_distance([4, 3, 0, 3])` 應該傳回 `1500`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 3, 0, 3]), 1500)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_penalty_distance(rounds):

    return rounds
```

# --solutions--

```py
def calculate_penalty_distance(rounds):
    total_penalty = 0

    for hits in rounds:
        misses = 5 - hits
        total_penalty += misses * 150

    return total_penalty
```
