---
id: 697a49e9860d24853adef67d
title: "挑戰 192：2026 冬季遊戲第 13 天：北歐兩項"
challengeType: 29
dashedName: challenge-192
---

# --description--

給定一個運動員跳躍分數的陣列，計算他們在北歐兩項賽事越野部分的起跑延遲時間。

跳躍分數最高的運動員先開始（延遲 0 秒）。所有其他運動員根據他們的跳躍分數與最佳跳躍分數的差距，稍後開始。

要計算每位運動員的延遲，請用最佳整體跳躍分數減去該運動員的跳躍分數，然後將結果乘以 1.5。將延遲向上取整到最接近的整數。

# --hints--

`calculate_start_delays([120, 110, 125])` 應該傳回 `[8, 23, 0]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([120, 110, 125]), [8, 23, 0])`)
}})
```

`calculate_start_delays([118, 125, 122, 120])` 應該傳回 `[11, 0, 5, 8]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([118, 125, 122, 120]), [11, 0, 5, 8])`)
}})
```

`calculate_start_delays([100, 105, 95, 110, 120, 115, 108])` 應該傳回 `[30, 23, 38, 15, 0, 8, 18]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18])`)
}})
```

`calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` 應該傳回 `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_start_delays(jump_scores):

    return jump_scores
```

# --solutions--

```py
import math
def calculate_start_delays(jump_scores):
    best_jump = max(jump_scores)
    return [math.ceil((best_jump - score) * 1.5) for score in jump_scores]
```
