---
id: 697a49e6ff50d756c9b69360
title: "挑戰 183：2026 冬季遊戲第 4 天：跳台滑雪"
challengeType: 29
dashedName: challenge-183
---

# --description--

給定距離分數、風力補償值和 K 點獎勵值，計算你的跳台滑雪分數並判斷你是否獲得獎牌。

- 你的分數是將上述四個數值相加計算而得。

其他跳躍者目前的總分為：

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- 如果你的分數是最佳，傳回 `"Gold"`
- 如果是第二名，傳回 `"Silver"`
- 如果是第三名，傳回 `"Bronze"`
- 否則，傳回 `"No Medal"`

# --hints--

`ski_jump_medal(125.0, 58.0, 0.0, 6.0)` 應該傳回 `"Gold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(125.0, 58.0, 0.0, 6.0), "Gold")`)
}})
```

`ski_jump_medal(119.0, 50.0, 1.0, 4.0)` 應該傳回 `"Bronze"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(119.0, 50.0, 1.0, 4.0), "Bronze")`)
}})
```

`ski_jump_medal(122.0, 52.0, -1.0, 4.0)` 應該傳回 `"Silver"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(122.0, 52.0, -1.0, 4.0), "Silver")`)
}})
```

`ski_jump_medal(118.0, 50.5, -1.5, 4.0)` 應該傳回 `"No Medal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(118.0, 50.5, -1.5, 4.0), "No Medal")`)
}})
```

`ski_jump_medal(124.0, 50.5, 2.0, 5.0)` 應該傳回 `"Gold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(124.0, 50.5, 2.0, 5.0), "Gold")`)
}})
```

`ski_jump_medal(119.0, 49.5, 0.0, 3.0)` 應該傳回 `"No Medal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(119.0, 49.5, 0.0, 3.0), "No Medal")`)
}})
```

# --seed--

## --seed-contents--

```py
def ski_jump_medal(distance_points, style_points, wind_comp, k_point_bonus):

    return distance_points
```

# --solutions--

```py
def ski_jump_medal(distance_points, style_points, wind_comp, k_point_bonus):
    my_score = distance_points + style_points + wind_comp + k_point_bonus
    other_scores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0]
    all_scores = other_scores + [my_score]
    all_scores.sort(reverse=True)
    rank = all_scores.index(my_score) + 1

    if rank == 1:
        return "Gold"
    elif rank == 2:
        return "Silver"
    elif rank == 3:
        return "Bronze"
    else:
        return "No Medal"
```
