---
id: 697a49e6ff50d756c9b6935e
title: "挑戰 181：2026 冬季遊戲第 2 天：單板滑雪"
challengeType: 29
dashedName: challenge-181
---

# --description--

給定滑雪板手的起始姿勢和旋轉角度，判斷他們的著陸姿勢。

- 滑雪板手的站姿是 `"Regular"` 或 `"Goofy"`。
- 技巧旋轉是 90 度的倍數。正數顯示順時針旋轉，負數顯示逆時針旋轉。
- 著陸姿勢每旋轉 180 度翻轉一次。

例如，給定 `"Regular"` 和 `90`，傳回 `"Regular"`。給定 `"Regular"` 和 `180` 度，傳回 `"Goofy"`。

# --hints--

`get_landing_stance("Regular", 90)` 應該傳回 `"Regular"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` 應該傳回 `"Goofy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` 應該傳回 `"Regular"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` 應該傳回 `"Goofy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` 應該傳回 `"Goofy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` 應該傳回 `"Regular"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
