---
id: 697a49e6ff50d756c9b69364
title: "挑戰 187：2026 冬季遊戲第 8 天：魯格競速"
challengeType: 29
dashedName: challenge-187
---

# --description--

給定一個包含五個數字的陣列，每個數字表述完成賽道一段所花費的時間（以秒為單位），判斷哪一段的速度最快以及該速度是多少。

賽道被分成以下區段：

- 第 1 段：320 公尺
- 第 2 段：280 公尺
- 第 3 段：350 公尺
- 第 4 段：300 公尺
- 第 5 段：250 公尺

給定陣列中的第一個值對應於區段 1 的時間，第二個值對應於區段 2，依此類推。

要計算一段的速度（以公尺／秒為單位），請將距離除以時間。

傳回 `"The luger's fastest speed was X m/s on segment Y."`。其中 `X` 是最快速度，四捨五入到小數點後兩位，`Y` 是發生最快速度的區段編號。

# --hints--

`get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128])` 應該傳回 `"The luger's fastest speed was 35.07 m/s on segment 5."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.")`)
}})
```

`get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284])` 應該傳回 `"The luger's fastest speed was 37.75 m/s on segment 2."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.")`)
}})
```

`get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912])` 應該傳回 `"The luger's fastest speed was 38.49 m/s on segment 3."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.")`)
}})
```

`get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840])` 應該傳回 `"The luger's fastest speed was 37.69 m/s on segment 1."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.")`)
}})
```

`get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508])` 應該傳回 `"The luger's fastest speed was 39.24 m/s on segment 4."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_fastest_speed(times):

    return times
```

# --solutions--

```py
def get_fastest_speed(times):
    distances = [320, 280, 350, 300, 250]
    fastest_speed = 0
    fastest_segment = 0

    for i, t in enumerate(times):
        speed = distances[i] / t
        if speed > fastest_speed:
            fastest_speed = speed
            fastest_segment = i + 1  # 1-indexed

    return f"The luger's fastest speed was {fastest_speed:.2f} m/s on segment {fastest_segment}."
```
