---
id: 698a1a73ade5ac0e19180fa3
title: "挑戰 200：字母與數字計數"
challengeType: 29
dashedName: challenge-200
---

# --description--

給定一個字串，傳回一則包含其中有多少字母和數字的訊息。

- 字母是 `A-Z` 和 `a-z`。
- 數字是 `0-9`。
- 忽略所有其他字元。

傳回 `"The string has X letters and Y numbers."`，其中 `"X"` 是字母的數量，`"Y"` 是數字的數量。如果任一數量為 1，則該項目使用單數形式。例如：使用 `"1 letter"` 而非 `"1 letters"`，以及使用 `"1 number"` 而非 `"1 numbers"`。

# --hints--

`count_letters_and_numbers("helloworld123")` 應該傳回 `"The string has 10 letters and 3 numbers."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("helloworld123"), "The string has 10 letters and 3 numbers.")`)
}})
```

`count_letters_and_numbers("Catch 22")` 應該傳回 `"The string has 5 letters and 2 numbers."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("Catch 22"), "The string has 5 letters and 2 numbers.")`)
}})
```

`count_letters_and_numbers("A1!")` 應該傳回 `"The string has 1 letter and 1 number."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("A1!"), "The string has 1 letter and 1 number.")`)
}})
```

`count_letters_and_numbers("12345")` 應該傳回 `"The string has 0 letters and 5 numbers."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("12345"), "The string has 0 letters and 5 numbers.")`)
}})
```

`count_letters_and_numbers("password")` 應該傳回 `"The string has 8 letters and 0 numbers."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("password"), "The string has 8 letters and 0 numbers.")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_letters_and_numbers(s):

    return s
```

# --solutions--

```py
def count_letters_and_numbers(s):
    letter_count = 0
    number_count = 0

    for char in s:
        if ('A' <= char <= 'Z') or ('a' <= char <= 'z'):
            letter_count += 1
        elif '0' <= char <= '9':
            number_count += 1

    letter_word = "letter" if letter_count == 1 else "letters"
    number_word = "number" if number_count == 1 else "numbers"

    return f"The string has {letter_count} {letter_word} and {number_count} {number_word}."
```
