---
id: 698a1a73ade5ac0e19180fa6
title: "挑戰 203：扁平化"
challengeType: 29
dashedName: challenge-203
---

# --description--

給定一個陣列，判斷它是否為平坦。

- 如果陣列中沒有任何元素是陣列，則該陣列是平坦的。

# --hints--

`is_flat([1, 2, 3, 4])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, 2, 3, 4]), True)`)
}})
```

`is_flat([1, [2, 3], 4])` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, [2, 3], 4]), False)`)
}})
```

`is_flat([1, 0, False, True, "a", "b"])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, 0, False, True, "a", "b"]), True)`)
}})
```

`is_flat(["a", [0], "b", True])` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat(["a", [0], "b", True]), False)`)
}})
```

`is_flat([1, [2, [3, [4, [5]]]], 6])` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, [2, [3, [4, [5]]]], 6]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_flat(arr):

    return arr
```

# --solutions--

```py
def is_flat(arr):
    for el in arr:
        if isinstance(el, list):
            return False
    return True
```
