---
id: 698a1a863194f1f4e63f645e
title: "挑戰 207：最小間距"
challengeType: 29
dashedName: challenge-207
---

# --description--

給定一個字串，傳回兩個相同字元之間字元數最少（最小間距）的子字串。

- 總會至少有一組匹配的字元配對。
- 傳回的子字串應該排除匹配的字元。
- 如果兩個或以上的空格長度相同，則傳回第一個空格中的字元。

例如，給定 `"ABCDAC"`，傳回 `"DA"`。

- 只有 `"A"` 和 `"C"` 在字串中重複。
- 兩個 `"A"` 字元之間的字元數是 3，兩個 `"C"` 字元之間的字元數是 2。
- 所以傳回兩個 `"C"` 字元之間的字串。

# --hints--

`smallest_gap("ABCDAC")` 應該傳回 `"DA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` 應該傳回 `"e"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` 應該傳回 `"#q6e&rkf(p"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` 應該傳回 `""`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` 應該傳回 `"fox"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
