---
id: 698a1a73ade5ac0e19180fa9
title: "挑戰 206：撲克牌點數"
challengeType: 29
dashedName: challenge-206
---

# --description--

給定一個撲克牌的陣列，傳回一個包含每張牌數值的新陣列。

卡牌數值：

- 一張 Ace (`"A"`) 的值為 1。
- 數字牌（`"2"` - `"10"`）的點數即為其面值：分別為 2 - 10。
- 人頭牌：Jack (`"J"`)、Queen (`"Q"`) 和 King (`"K"`) 各值 10。

適用於：

- 每張牌都有一個花色：黑桃（`"S"`）、梅花（`"C"`）、方塊（`"D"`）或紅心（`"H"`）。

卡片格式：

- 每張牌以字串表示：`"valueSuit"`。例如：`"AS"` 是黑桃 A，`"10H"` 是紅心 10，`"QC"` 是梅花 Q。

# --hints--

`card_values(["3H", "4D", "5S"])` 應該傳回 `[3, 4, 5]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` 應該傳回 `[1, 10, 10, 6, 7]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` 應該傳回 `[8, 10, 2, 10, 9]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` 應該傳回 `[1, 10]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` 應該傳回 `[10, 10, 10, 10, 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```
