---
id: 698a1a863194f1f4e63f645f
title: "挑戰 208：路徑遍歷"
challengeType: 29
dashedName: challenge-208
---

# --description--

給定一個表述你路線圖的字串陣列，傳回一個到達目標所需移動的字串。

給定的字串將包含以下值：

- `"C"`：你目前的位置
- `"G"`：你的目標
- `"T"`：路徑中可遍歷的零件
- `"-"`：地圖中不可通行的零件

傳回一個字串，內容是從你的位置到目標所需的移動步驟，條件如下：

- `"R"` 是向右移動
- `"D"` 是向下移動
- `"L"` 是向左移動
- `"U"` 是向上移動

- 從你目前的位置到你的目標，將始終有一條單一且連續的路徑，且不會有任何分支。
- 每個路徑位置最多會有兩個可通行的位置相鄰。

例如，給定：

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

傳回 `"RDDRDD"`。

# --hints--

`navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` 應該傳回 `"RDDRDD"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD")`)
}})
```

`navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` 應該傳回 `"RRUUURR"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR")`)
}})
```

`navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` 應該傳回 `"DLDDRRRRD"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD")`)
}})
```

`navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` 應該傳回 `"RRRDDL"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL")`)
}})
```

`navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` 應該傳回 `"ULLLUUURRRRRRDDDR"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate_trail(map):

    return map
```

# --solutions--

```py
def navigate_trail(map):
    rows = len(map)
    cols = len(map[0])

    directions = [
        (0, 1, "R"),
        (1, 0, "D"),
        (0, -1, "L"),
        (-1, 0, "U"),
    ]

    for r in range(rows):
        for c in range(cols):
            if map[r][c] == "C":
                row, col = r, c

    result = ""
    prev = None

    while map[row][col] != "G":
        for dr, dc, move in directions:
            new_row = row + dr
            new_col = col + dc

            if (
                0 <= new_row < rows and
                0 <= new_col < cols and
                (map[new_row][new_col] == "T" or
                 map[new_row][new_col] == "G") and
                (not prev or (new_row, new_col) != prev)
            ):
                result += move
                prev = (row, col)
                row, col = new_row, new_col
                break

    return result
```
