---
id: 698a1a863194f1f4e63f6460
title: "挑戰 209：元素大小"
challengeType: 29
dashedName: challenge-209
---

# --description--

給定一個視窗大小、一個元素以視口寬度單位 `"vw"` 表示的寬度，以及一個元素以視口高度單位 `"vh"` 表示的高度，判斷該元素的大小（以像素為單位）。

- 給定的視窗大小和傳回的元素大小是以 `"width x height"` 格式的字串，例如 `"1200 x 800"`。

- `"vw"` 單位是視窗寬度的百分比。`"50vw"` 例如，是視窗寬度的 50%。
- `"vh"` 單位是視窗高度的百分比。`"50vh"` 例如，是視窗高度的 50%。

# --hints--

`get_element_size("1200 x 800", "50vw", "50vh")` 應該傳回 `"600 x 400"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "50vw", "50vh"), "600 x 400")`)
}})
```

`get_element_size("320 x 480", "25vw", "50vh")` 應該傳回 `"80 x 240"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("320 x 480", "25vw", "50vh"), "80 x 240")`)
}})
```

`get_element_size("1000 x 500", "7vw", "3vh")` 應該傳回 `"70 x 15"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1000 x 500", "7vw", "3vh"), "70 x 15")`)
}})
```

`get_element_size("1920 x 1080", "95vw", "100vh")` 應該傳回 `"1824 x 1080"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1920 x 1080", "95vw", "100vh"), "1824 x 1080")`)
}})
```

`get_element_size("1200 x 800", "0vw", "0vh")` 應該傳回 `"0 x 0"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "0vw", "0vh"), "0 x 0")`)
}})
```

`get_element_size("1440 x 900", "100vw", "114vh")` 應該傳回 `"1440 x 1026"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1440 x 900", "100vw", "114vh"), "1440 x 1026")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_element_size(window_size, element_vw, element_vh):

    return window_size
```

# --solutions--

```py
def get_element_size(window_size, element_vw, element_vh):
    window_width, window_height = map(int, window_size.split(" x "))

    vw = int(element_vw.replace("vw", ""))
    vh = int(element_vh.replace("vh", ""))

    width_px = round((vw / 100) * window_width)
    height_px = round((vh / 100) * window_height)

    return f"{width_px} x {height_px}"
```
