---
id: 6994cff2290543b3aec9f510
title: "挑戰 211：陣列總和"
challengeType: 29
dashedName: challenge-211
---

# --description--

給定一個數字的陣列，傳回所有數字的總和。

# --hints--

`sum_array([1, 2, 3, 4, 5])` 應該傳回 `15`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([1, 2, 3, 4, 5]), 15)`)
}})
```

`sum_array([42])` 應該傳回 `42`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([42]), 42)`)
}})
```

`sum_array([5, -2, 7, -3])` 應該傳回 `7`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([5, -2, 7, -3]), 7)`)
}})
```

`sum_array([203, 145, -129, 6293, 523, -919, 845, 2434])` 應該傳回 `9395`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395)`)
}})
```

`sum_array([0, 0])` 應該傳回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([0, 0]), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_array(numbers):

    return numbers
```

# --solutions--

```py
def sum_array(numbers):

    return sum(numbers)
```
