---
id: 699c8e045ee7cb94ed2322d6
title: "挑戰 215：停車費計算器"
challengeType: 29
dashedName: challenge-215
---

# --description--

給定兩個表示你停車時間和取車時間的字串，計算停車費用。

- 給定的字串將會是使用 24 小時制的 `"HH:MM"` 格式。例如，`"14:00"` 是下午 2 點。
- 第一個字串將是你停車的時間，第二個將是你取車的時間。
- 如果取車時間早於進場時間，表示停車時間跨越午夜進入隔天。

費用規則：

- 每停車 1 小時收費 3 美元。
- 部分小時會進位到下一個整點小時。
- 如果停車時間跨越過夜（過午夜），將額外收取 10 美元過夜費。
- 最低費用為 5 美元（僅在總額低於 5 美元時使用）。

以 `"$cost"` 格式傳回總費用，`"$5"` 例如。

# --hints--

`calculate_parking_fee("09:00", "11:00")` 應該傳回 `"$6"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("09:00", "11:00"), "$6")`)
}})
```

`calculate_parking_fee("10:00", "10:30")` 應該傳回 `"$5"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("10:00", "10:30"), "$5")`)
}})
```

`calculate_parking_fee("08:10", "10:45")` 應該傳回 `"$9"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("08:10", "10:45"), "$9")`)
}})
```

`calculate_parking_fee("14:40", "23:10")` 應該傳回 `"$27"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("14:40", "23:10"), "$27")`)
}})
```

`calculate_parking_fee("18:15", "01:30")` 應該傳回 `"$34"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("18:15", "01:30"), "$34")`)
}})
```

`calculate_parking_fee("11:11", "11:10")` 應該傳回 `"$82"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("11:11", "11:10"), "$82")`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_parking_fee(park_time, pickup_time):

    return park_time
```

# --solutions--

```py
import math

def calculate_parking_fee(park_time, pickup_time):
    def to_minutes(time_str):
        hours, minutes = map(int, time_str.split(":"))
        return hours * 60 + minutes

    entry_minutes = to_minutes(park_time)
    exit_minutes = to_minutes(pickup_time)

    overnight = False

    if exit_minutes < entry_minutes:
        overnight = True
        total_minutes = (24 * 60 - entry_minutes) + exit_minutes
    else:
        total_minutes = exit_minutes - entry_minutes

    hours = math.ceil(total_minutes / 60)

    cost = hours * 3

    if overnight:
        cost += 10

    if cost < 5:
        cost = 5

    return f"${cost}"
```
