---
id: 699c8e045ee7cb94ed2322d9
title: "挑戰 218：均勻可整除"
challengeType: 29
dashedName: challenge-218
---

# --description--

給定兩個整數，判斷你是否能將第一個整數整除以第二個整數。

# --hints--

`is_evenly_divisible(4, 2)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4, 2), True)`)
}})
```

`is_evenly_divisible(7, 3)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(7, 3), False)`)
}})
```

`is_evenly_divisible(5, 10)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(5, 10), False)`)
}})
```

`is_evenly_divisible(48, 6)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(48, 6), True)`)
}})
```

`is_evenly_divisible(3186, 9)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(3186, 9), True)`)
}})
```

`is_evenly_divisible(4192, 11)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4192, 11), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_evenly_divisible(a, b):

    return a
```

# --solutions--

```py
def is_evenly_divisible(a, b):
    return a % b == 0
```
