---
id: 699c8e045ee7cb94ed2322dd
title: "挑戰 222：晝夜平分影子"
challengeType: 29
dashedName: challenge-222
---

# --description--

今天是晝夜平分點，太陽正好位於赤道上方，並在中午時分垂直照射。給定一個時間，判斷一根 4 英尺直立柱所投射的影子。

- 時間將會是 `"HH:MM"` 24 小時制格式的字串（例如，`"15:00"` 是下午 3 點）。
- 你只會被給予以 30 分鐘為增量的時間。

規則：

- 太陽在早上 6 點直接從 `"east"` 升起，並在晚上 6 點直接從 `"west"` 落下。
- 陰影總是指向與太陽相反的方向。
- 影子的長度（以英尺為單位）是距離中午小時數的立方。
- 日出前（早上 6 點前）、日落後（下午 6 點或之後）或正午時沒有陰影。

傳回：

- 如果存在陰影，傳回 `"(length)ft (direction)"`。例如，`"8ft west"`。
- 否則，傳回 `"No shadow"`。

例如，給定 `"10:00"`，傳回 `"8ft west"`，因為上午 10 點距離中午有 2 小時，所以 2<sup>3</sup> = 8 英尺，且影子指向西方，因為上午 10 點太陽在東方。

# --hints--

`get_shadow("10:00")` 應該傳回 `"8ft west"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` 應該傳回 `"27ft east"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` 應該傳回 `"No shadow"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` 應該傳回 `"166.375ft east"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` 應該傳回 `"No shadow"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` 應該傳回 `"216ft west"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` 應該傳回 `"No shadow"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` 應該傳回 `"91.125ft west"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` 應該傳回 `"No shadow"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
