---
id: 69a890af247de743333bd4cd
title: "挑戰 224：咖啡烘焙檢測器"
challengeType: 29
dashedName: challenge-224
---

# --description--

給定一個表示用來製作一杯咖啡的豆子的字串，判斷該杯咖啡的烘焙程度。

- 給定的字串將包含以下字元，每個字元表述一種豆子的型別：

  - 撇號（`'`）是一種輕度烘焙豆，每顆價值 1 點。
  - 破折號（`-`）是一種中度烘焙豆，每顆價值 2 點。
  - 句點（`.`）是每顆價值 3 點的深焙咖啡豆。

- 烘焙層次由所有豆子的平均值決定。

傳回：

- `"Light"` 如果平均值小於 1.75。
- `"Medium"` 如果平均值在 1.75 到 2.5 之間。
- 如果平均值大於 2.5，則為 `"Dark"`。

# --hints--

`detect_roast("''-''''''-'-''--''''")` 應該傳回 `"Light"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("''-''''''-'-''--''''"), "Light")`)
}})
```

`detect_roast(".'-''-''..'''.-.-''-")` 應該傳回 `"Medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".'-''-''..'''.-.-''-"), "Medium")`)
}})
```

`detect_roast("--.''--'-''.--..-.--")` 應該傳回 `"Medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("--.''--'-''.--..-.--"), "Medium")`)
}})
```

`detect_roast("-...'-......-..-...-")` 應該傳回 `"Dark"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-...'-......-..-...-"), "Dark")`)
}})
```

`detect_roast(".--.-..-......----.'")` 應該傳回 `"Medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".--.-..-......----.'"), "Medium")`)
}})
```

`detect_roast("..-..-..-..-....-.-.")` 應該傳回 `"Dark"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("..-..-..-..-....-.-."), "Dark")`)
}})
```

`detect_roast("-'-''''''..-'.''-'.'")` 應該傳回 `"Light"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-'-''''''..-'.''-'.'"), "Light")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_roast(beans):

    return beans
```

# --solutions--

```py
def detect_roast(beans):
    total = 0

    for bean in beans:
        if bean == "'":
            total += 1
        elif bean == "-":
            total += 2
        elif bean == ".":
            total += 3

    avg = total / len(beans)

    if avg < 1.75:
        return "Light"
    elif avg <= 2.5:
        return "Medium"
    else:
        return "Dark"
```
