---
id: 69a890af247de743333bd4d1
title: "挑戰 228：電影之夜"
challengeType: 29
dashedName: challenge-228
---

# --description--

給定一個表示星期幾的字串、一個表示場次時間的字串，以及一個整數票數，傳回該場次電影票的總費用。

給定的日期將會是以下之一：

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

演出時間將以 `"H:MMam"` 或 `"H:MMpm"` 格式給出。例如 `"10:00am"` 或 `"10:00pm"`。

以 `"$D.CC"` 格式傳回總成本，並使用以下規則：

- 週末（星期五－星期日）：每張票 $12.00。
- 平日（星期一至星期四）：每張票 10.00 美元。
- Matinee（下午 5:00 前）：每張票減 2.00 美元（星期二除外）。
- 星期二：所有票券每張 $5.00。

# --hints--

`get_movie_night_cost("Saturday", "10:00pm", 1)` 應該傳回 `"$12.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Saturday", "10:00pm", 1), "$12.00")`)
}})
```

`get_movie_night_cost("Sunday", "10:00am", 1)` 應該傳回 `"$10.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Sunday", "10:00am", 1), "$10.00")`)
}})
```

`get_movie_night_cost("Tuesday", "7:20pm", 2)` 應該傳回 `"$10.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "7:20pm", 2), "$10.00")`)
}})
```

`get_movie_night_cost("Wednesday", "5:40pm", 3)` 應該傳回 `"$30.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Wednesday", "5:40pm", 3), "$30.00")`)
}})
```

`get_movie_night_cost("Monday", "11:50am", 4)` 應該傳回 `"$32.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Monday", "11:50am", 4), "$32.00")`)
}})
```

`get_movie_night_cost("Friday", "4:30pm", 5)` 應該傳回 `"$50.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Friday", "4:30pm", 5), "$50.00")`)
}})
```

`get_movie_night_cost("Tuesday", "11:30am", 1)` 應該傳回 `"$5.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "11:30am", 1), "$5.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):

    return day
```

# --solutions--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):
    if day == "Tuesday":
        price = 5
    else:
        if day in ["Friday", "Saturday", "Sunday"]:
            price = 12
        else:
            price = 10

        hour = int(showtime.split(":")[0])
        period = showtime[-2:]

        hour24 = hour

        if period == "pm" and hour != 12:
            hour24 += 12
        if period == "am" and hour == 12:
            hour24 = 0

        if hour24 < 17:
            price -= 2

    total = price * number_of_tickets

    return f"${total:.2f}"
```
