---
id: 69a890af247de743333bd4ce
title: "挑戰 225：禁止連續重複"
challengeType: 29
dashedName: challenge-225
---

# --description--

給定一個字串，判斷它是否沒有重複的字元。

- 如果一個字串沒有相同的字元連續出現兩次或以上，則該字串沒有重複。

# --hints--

`has_no_repeats("hi world")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hi world"), True)`)
}})
```

`has_no_repeats("hello world")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hello world"), False)`)
}})
```

`has_no_repeats("abcdefghijklmnopqrstuvwxyz")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("abcdefghijklmnopqrstuvwxyz"), True)`)
}})
```

`has_no_repeats("freeCodeCamp")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("freeCodeCamp"), False)`)
}})
```

`has_no_repeats("The quick brown fox jumped over the lazy dog.")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("The quick brown fox jumped over the lazy dog."), True)`)
}})
```

`has_no_repeats("Mississippi")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("Mississippi"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_no_repeats(s):

    return s
```

# --solutions--

```py
def has_no_repeats(s):
    for i in range(1, len(s)):
        if s[i] == s[i-1]:
            return False
    return True
```
