---
id: 69b1028d6e265413d0198a2c
title: "挑戰 233：起床鬧鐘"
challengeType: 29
dashedName: challenge-233
---

# --description--

給定一個字串，代表你設定鬧鐘的時間，以及另一個字串，代表你實際起床的時間，判斷你是早起、準時還是遲到。

- 兩個時間都會以 `"HH:MM"` 24 小時制格式給出。

傳回：

- `"early"` 如果你在鬧鐘時間之前起床。
- `"on time"` 如果你在鬧鐘時間起床，或在鬧鐘時間後 10 分鐘的貪睡時間內起床。
- `"late"` 如果你在鬧鐘時間後超過 10 分鐘才起床。

兩個時間都在同一天。

# --hints--

`alarm_check("07:00", "06:45")` 應該傳回 `"early"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("07:00", "06:45"), "early")`)
}})
```

`alarm_check("06:30", "06:30")` 應該傳回 `"on time"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("06:30", "06:30"), "on time")`)
}})
```

`alarm_check("08:10", "08:15")` 應該傳回 `"on time"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:10", "08:15"), "on time")`)
}})
```

`alarm_check("09:30", "09:45")` 應該傳回 `"late"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("09:30", "09:45"), "late")`)
}})
```

`alarm_check("08:15", "08:25")` 應該傳回 `"on time"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:15", "08:25"), "on time")`)
}})
```

`alarm_check("05:45", "05:56")` 應該傳回 `"late"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("05:45", "05:56"), "late")`)
}})
```

`alarm_check("04:30", "04:00")` 應該傳回 `"early"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("04:30", "04:00"), "early")`)
}})
```

# --seed--

## --seed-contents--

```py
def alarm_check(alarm_time, wake_time):

    return alarm_time
```

# --solutions--

```py
def alarm_check(alarm_time, wake_time):
    alarm_h, alarm_m = map(int, alarm_time.split(":"))
    wake_h, wake_m = map(int, wake_time.split(":"))

    alarm_minutes = alarm_h * 60 + alarm_m
    wake_minutes = wake_h * 60 + wake_m
    diff = wake_minutes - alarm_minutes

    if diff < 0:
        return "early"
    if diff <= 10:
        return "on time"
    return "late"
```
