---
id: 69b1028d6e265413d0198a2f
title: "挑戰 236：瀏覽器歷史紀錄"
challengeType: 29
dashedName: challenge-236
---

# --description--

給定一個瀏覽器指令的陣列，傳回一個包含兩個值的陣列：歷史紀錄（以 URL 陣列表示）和目前頁面的索引。

有效的指令有：

- `"URL"` - 其中 URL 是一個網頁位址（`"freecodecamp.org"` 例如）。導向指定的 URL，將它加入歷史紀錄的下一個位置，並捨棄任何前進的歷史紀錄。
- `"Back"` - 移動到歷史紀錄中的上一頁，如果沒有上一頁則停留在目前頁面。
- `"Forward"` - 移動到歷史紀錄中的下一頁，如果沒有下一頁則停留在目前頁面。

例如，給定 `["freecodecamp.org", "freecodecamp.org/learn", "Back"]`，傳回 `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`。

# --hints--

`get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` 應該傳回 `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` 應該傳回 `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"])` 應該傳回 `[["example.com", "example.com/contact", "example.com/blog"], 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` 應該傳回 `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "Back"])` 應該傳回 `[["example.com", "example.com/about"], 0]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Forward"])` 應該傳回 `[["example.com", "example.com/about"], 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_browser_history(commands):

    return commands
```

# --solutions--

```py
def get_browser_history(commands):
    stack = []
    index = -1

    for command in commands:
        if command == "Back":
            if index > 0:
                index -= 1
        elif command == "Forward":
            if index < len(stack) - 1:
                index += 1
        else:
            del stack[index + 1:]
            stack.append(command)
            index += 1

    return [stack, index]
```
