---
id: 69b58ce40693f140c84c8559
title: "挑戰 240：回文字元"
challengeType: 29
dashedName: challenge-240
---

# --description--

給定一個字串，判斷它是否為回文，並傳回中間的字元（如果長度為奇數）或中間兩個字元（如果長度為偶數）。

- 回文是指正著讀和反著讀都相同的字串。
- 如果不是回文，傳回 `"none"`。

# --hints--

`palindrome_locator("racecar")` 應該傳回 `"e"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("racecar"), "e")`)
}})
```

`palindrome_locator("level")` 應該傳回 `"v"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("level"), "v")`)
}})
```

`palindrome_locator("freecodecamp")` 應該傳回 `"none"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("freecodecamp"), "none")`)
}})
```

`palindrome_locator("noon")` 應該傳回 `"oo"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("noon"), "oo")`)
}})
```

`palindrome_locator("11100111")` 應該傳回 `"00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("11100111"), "00")`)
}})
```

# --seed--

## --seed-contents--

```py
def palindrome_locator(s):

    return s
```

# --solutions--

```py
def palindrome_locator(s):
    if s != s[::-1]:
        return "none"
    mid = len(s) // 2
    return s[mid] if len(s) % 2 == 1 else s[mid - 1] + s[mid]
```
