---
id: 69b58ce40693f140c84c855b
title: "挑戰 242：下一個賓果號碼"
challengeType: 29
dashedName: challenge-242
---

# --description--

給定一個賓果號碼，傳回下一個依序的賓果號碼。

賓果號碼由一個字母接著一個數字組成，數字在下表所示的範圍內：

| 字母 | 數字範圍 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

例如，給定 `"B10"`，傳回下一個賓果號碼 `"B11"`。如果給定的是最後一個賓果號碼，則傳回 `"B1"`。

# --hints--

`get_next_bingo_number("B10")` 應該傳回 `"B11"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("B10"), "B11")`)
}})
```

`get_next_bingo_number("N33")` 應該傳回 `"N34"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("N33"), "N34")`)
}})
```

`get_next_bingo_number("I30")` 應該傳回 `"N31"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("I30"), "N31")`)
}})
```

`get_next_bingo_number("G60")` 應該傳回 `"O61"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("G60"), "O61")`)
}})
```

`get_next_bingo_number("O75")` 應該傳回 `"B1"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("O75"), "B1")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_next_bingo_number(n):

    return n
```

# --solutions--

```py
def get_next_bingo_number(n):
    num = int(n[1:])
    next_num = 1 if num == 75 else num + 1

    if next_num <= 15:
        return "B" + str(next_num)
    if next_num <= 30:
        return "I" + str(next_num)
    if next_num <= 45:
        return "N" + str(next_num)
    if next_num <= 60:
        return "G" + str(next_num)
    return "O" + str(next_num)
```
