---
id: 69b58ce40693f140c84c855a
title: "挑戰 241：FizzBuzz 驗證器"
challengeType: 29
dashedName: challenge-241
---

# --description--

給定一個連續整數的陣列，其中 3 和 5 的倍數已被替換，判斷它是否為有效的 FizzBuzz 序列。

在有效的 FizzBuzz 序列中：

- 3 的倍數會被替換成「Fizz」。
- 5 的倍數會被替換成「Buzz」。
- 同時是 3 和 5 的倍數會被替換成「FizzBuzz」。
- 其他所有數字保持為整數。

# --hints--

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([13, 14, "FizzBuzz", 16, 17])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([13, 14, "FizzBuzz", 16, 17]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, 5])` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, 5]), False)`)
}})
```

`is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", "Buzz", 5])` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", "Buzz", 5]), False)`)
}})
```

`is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]), False)`)
}})
```

`is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_fizz_buzz(arr):

    return arr
```

# --solutions--

```py
def is_fizz_buzz(arr):
    start = None
    for i, val in enumerate(arr):
        if isinstance(val, int):
            start = val - i
            break

    if start is None:
        return False

    for i in range(len(arr)):
        n = start + i
        if n % 15 == 0:
            expected = "FizzBuzz"
        elif n % 3 == 0:
            expected = "Fizz"
        elif n % 5 == 0:
            expected = "Buzz"
        else:
            expected = n
        if arr[i] != expected:
            return False

    return True
```
