---
id: 69b5b2be76ec8135a7fbe973
title: "挑戰 244：車與象的攻擊"
challengeType: 29
dashedName: challenge-244
---

# --description--

給定一個字串表示車在棋盤上的位置，另一個字串表示象的位置，判斷其中一個棋子是否能攻擊另一個。

標準棋盤是 8x8，行標籤從 `A` 到 `H`（由左至右），列標籤從 `1` 到 `8`（由下至上）。棋盤如下所示：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- 車可以在水平或垂直方向移動任意多格。
- 象可以在任意斜角方向移動任意多格。
- 如果一個棋子能移動到另一個棋子的位置，則該棋子可以攻擊另一個。

傳回：

- `"rook"` 如果車能攻擊象。
- `"bishop"` 如果象能攻擊車。
- `"neither"` 如果兩個棋子都無法攻擊對方。

# --hints--

`rook_bishop_attack("A1", "A5")` 應該傳回 `"rook"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("A1", "A5"), "rook")`)
}})
```

`rook_bishop_attack("C3", "F6")` 應該傳回 `"bishop"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("C3", "F6"), "bishop")`)
}})
```

`rook_bishop_attack("D4", "D7")` 應該傳回 `"rook"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("D4", "D7"), "rook")`)
}})
```

`rook_bishop_attack("B7", "H1")` 應該傳回 `"bishop"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B7", "H1"), "bishop")`)
}})
```

`rook_bishop_attack("B3", "C5")` 應該傳回 `"neither"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B3", "C5"), "neither")`)
}})
```

`rook_bishop_attack("G3", "E8")` 應該傳回 `"neither"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("G3", "E8"), "neither")`)
}})
```

# --seed--

## --seed-contents--

```py
def rook_bishop_attack(rook, bishop):

    return rook
```

# --solutions--

```py
def rook_bishop_attack(rook, bishop):
    col_diff = abs(ord(rook[0]) - ord(bishop[0]))
    row_diff = abs(int(rook[1]) - int(bishop[1]))

    if rook[0] == bishop[0] or rook[1] == bishop[1]:
        return "rook"
    if col_diff == row_diff:
        return "bishop"
    return "neither"
```
