---
id: 69b559d2903b9e4afe9075f8
title: "Challenge 238: Digit Rotation Escape"
challengeType: 29
dashedName: challenge-238
---

# --description--

給定一個正整數，判斷它本身或其任一旋轉後的數字，是否能被其位數整除。

旋轉指的是將第一個數字移到最後面。例如，旋轉一次後，123 會變成 231。

- 先檢查旋轉 `0`（給定的數字）。
- 給定的數字不會包含任何零。
- 如果找到符合條件的旋轉數字，傳回該數字；否則傳回 `"none"`。

# --hints--

`get_rotation(123)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(123), 0)`)
}})
```

`get_rotation(13579)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(13579), 3)`)
}})
```

`get_rotation(24681)` should return `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(24681), "none")`)
}})
```

`get_rotation(84138789345)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(84138789345), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rotation(n):

    return n
```

# --solutions--

```py
def get_rotation(n):
    s = str(n)
    digit_count = len(s)
    current = s

    for i in range(digit_count):
        if int(current) % digit_count == 0:
            return i
        current = current[1:] + current[0]

    return "none"
```
