---
id: 69c5f3d787b1725d5f00c8b9
title: "Challenge 255: Earth Day Cleanup Crew"
challengeType: 29
dashedName: challenge-255
---

# --description--

今天是地球日。給定一個你清理過的項目陣列，根據以下規則傳回你的總清理分數。

給定的項目將是以下之一：

| 項目 | 基本分數 |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

稀有項目以 `["rare", value]` 表示。舉例來說，`["rare", 80]`。稀有項目不會獲得連續加分。

- 連續加分：如果相同項目連續出現，會獲得遞增的加分。
  - 第一次連續出現：基本分數
  - 第二次：基本分數 + 1
  - 第三次：基本分數 + 2
  - 依此類推。
- 第五項目倍數：每收集到第 5 項目會有倍數加成。
  - 第五項目：乘以 2
  - 第十項目：乘以 3
  - 依此類推。

- 先計算任何加分後，再套用倍數。

# --hints--

`get_cleanup_score(["bottle", "straw", "shoe", "battery"])` should return `44`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["bottle", "straw", "shoe", "battery"]), 44)`)
}})
```

`get_cleanup_score(["electronics", "straw", "newspaper", "bottle", "bag"])` should return `58`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["electronics", "straw", "newspaper", "bottle", "bag"]), 58)`)
}})
```

`get_cleanup_score(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` should return `79`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79)`)
}})
```

`get_cleanup_score(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` should return `358`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358)`)
}})
```

`get_cleanup_score(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` should return `383`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_cleanup_score(items):

    return items
```

# --solutions--

```py
def get_cleanup_score(items):
    values = {
        'bottle': 10, 'can': 6, 'bag': 8, 'tire': 35,
        'straw': 4, 'cardboard': 3, 'newspaper': 3,
        'shoe': 12, 'electronics': 25, 'battery': 18, 'mattress': 38
    }
    total = 0
    prev_item = None
    streak = 0
    for i, item in enumerate(items):
        is_rare = isinstance(item, list) and item[0] == 'rare'
        if is_rare:
            base = item[1]
            prev_item = None
            streak = 0
        else:
            base = values[item]
            streak = streak + 1 if item == prev_item else 0
            prev_item = item
        points = base if is_rare else base + streak
        if (i + 1) % 5 == 0:
            multiplier = (i + 1) // 5 + 1
            points *= multiplier
        total += points
    return total
```
