---
id: 69c5f3d787b1725d5f00c8ba
title: "挑戰 256：最接近的時間方向"
challengeType: 29
dashedName: challenge-256
---

# --description--

給定兩個時間，判斷從第一個時間到第二個時間，向前走還是向後走比較快。

- 時間以 24 小時制表示 (`"HH:MM"`)
- 時鐘會環繞（向前走時 23:59 會到 00:00，向後走時 00:00 會到 23:59）

傳回：

- 如果向前走較短，傳回 `"forward"`
- 如果向後走較短，傳回 `"backward"`
- 如果兩個方向所需時間相同，傳回 `"equal"`

# --hints--

`get_direction("10:00", "12:00")` 應該傳回 `"forward"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("10:00", "12:00"), "forward")`)
}})
```

`get_direction("11:00", "05:00")` 應該傳回 `"backward"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("11:00", "05:00"), "backward")`)
}})
```

`get_direction("00:00", "12:00")` 應該傳回 `"equal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("00:00", "12:00"), "equal")`)
}})
```

`get_direction("15:45", "01:10")` 應該傳回 `"forward"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("15:45", "01:10"), "forward")`)
}})
```

`get_direction("03:30", "19:50")` 應該傳回 `"backward"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("03:30", "19:50"), "backward")`)
}})
```

`get_direction("06:30", "18:30")` 應該傳回 `"equal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("06:30", "18:30"), "equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_direction(time1, time2):

    return time1
```

# --solutions--

```py
def get_direction(time1, time2):
    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    start = to_minutes(time1)
    end = to_minutes(time2)
    forward = (end - start + 1440) % 1440
    backward = (start - end + 1440) % 1440

    if forward < backward:
        return 'forward'
    if backward < forward:
        return 'backward'
    return 'equal'
```
