---
id: 69c5f3d787b1725d5f00c8b7
title: "Challenge 253: Acronym Finder"
challengeType: 29
dashedName: challenge-253
---

# --description--

給定一個代表縮寫詞的字串，從以下列表中傳回該縮寫詞所屬組織的完整名稱：

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

給定縮寫詞中的每個字母，應該依序對應該組織中每個字的第一個字母。

# --hints--

`find_org("NASA")` should return `"National Avocado Storage Authority"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("NASA"), "National Avocado Storage Authority")`)
}})
```

`find_org("CIA")` should return `"Cats Infiltration Agency"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("CIA"), "Cats Infiltration Agency")`)
}})
```

`find_org("FBI")` should return `"Fluffy Beanbag Inspectors"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("FBI"), "Fluffy Beanbag Inspectors")`)
}})
```

`find_org("DOJ")` should return `"Department Of Jelly"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("DOJ"), "Department Of Jelly")`)
}})
```

`find_org("WHO")` should return `"Wild Honey Organization"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("WHO"), "Wild Honey Organization")`)
}})
```

`find_org("EPA")` should return `"Eating Pancakes Administration"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_org("EPA"), "Eating Pancakes Administration")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_org(acronym):

    return acronym
```

# --solutions--

```py
def find_org(acronym):
    orgs = [
        "National Avocado Storage Authority",
        "Cats Infiltration Agency",
        "Fluffy Beanbag Inspectors",
        "Department Of Jelly",
        "Wild Honey Organization",
        "Eating Pancakes Administration",
    ]
    for org in orgs:
        initials = ''.join(word[0] for word in org.split())
        if initials.upper() == acronym.upper():
            return org
    return "not found"
```
