---
id: 69c5f3d787b1725d5f00c8bd
title: "挑戰 259：FizzBuzz 大爆炸"
challengeType: 29
dashedName: challenge-259
---

# --description--

給定一個整數，傳回將字串 `"fizzbuzz"` 轉換成至少包含指定數量 `"z"` 的字串所需的步驟數，規則如下：

- 從字串 `"fizzbuzz"` 開始。
- 每一步，根據字串中字母的位置套用標準 FizzBuzz 規則（第一個 `"f"` 的位置為 1）。
  - 如果字母位置能被 3 整除，將該字母取代為 `"fizz"`
  - 如果能被 5 整除，將該字母取代為 `"buzz"`
  - 如果能同時被 3 和 5 整除，將該字母取代為 `"fizzbuzz"`

所以經過 1 步後，`"fizzbuzz"` 會變成 `"fifizzzbuzzfizzzz"`，其中包含 9 個 `"z"`。

# --hints--

`explode_fizzbuzz(9)` 應該傳回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` 應該傳回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` 應該傳回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` 應該傳回 `4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` 應該傳回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` 應該傳回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` 應該傳回 `11`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` 應該傳回 `14`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
