---
id: 69c6ff713a52713463aa7927
title: "挑戰 261：數字單字"
challengeType: 29
dashedName: challenge-261
---

# --description--

給定一個介於 0 到 99 的整數，傳回它的英文單字表述。

- `0` 傳回 `"zero"`。
- 數字 1-19 有獨特的名稱（`"one"`、`"two"`、…、`"ten"`、`"eleven"`、…、`"eighteen"`、`"nineteen"`）。
- 20-90 的 10 的倍數有自己的名稱（`"twenty"`、`"thirty"`、…、`"eighty"`、`"ninety"`）。
- 21-99 中非 10 的倍數的數字，寫成兩個單字並用連字號連接。例如 `"forty-two"` 和 `"fifty-three"`。

# --hints--

`get_number_words(0)` 應該傳回 `"zero"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(0), "zero")`)
}})
```

`get_number_words(10)` 應該傳回 `"ten"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(10), "ten")`)
}})
```

`get_number_words(19)` 應該傳回 `"nineteen"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(19), "nineteen")`)
}})
```

`get_number_words(30)` 應該傳回 `"thirty"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(30), "thirty")`)
}})
```

`get_number_words(53)` 應該傳回 `"fifty-three"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(53), "fifty-three")`)
}})
```

`get_number_words(7)` 應該傳回 `"seven"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(7), "seven")`)
}})
```

`get_number_words(12)` 應該傳回 `"twelve"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(12), "twelve")`)
}})
```

`get_number_words(60)` 應該傳回 `"sixty"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(60), "sixty")`)
}})
```

`get_number_words(67)` 應該傳回 `"sixty-seven"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(67), "sixty-seven")`)
}})
```

`get_number_words(98)` 應該傳回 `"ninety-eight"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(98), "ninety-eight")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_words(n):

    return n
```

# --solutions--

```py
def get_number_words(n):
    ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
            "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
            "seventeen","eighteen","nineteen"]
    tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"]

    if n < 20:
        return ones[n]
    if n % 10 == 0:
        return tens[n // 10]
    return tens[n // 10] + "-" + ones[n % 10]
```
