---
id: 69c6ff713a52713463aa7928
title: "挑戰 262：URL 查詢字串解析器"
challengeType: 29
dashedName: challenge-262
---

# --description--

給定一個包含查詢字串的 URL，將查詢字串解析成一個包含鍵值對的物件（或字典）。

- 查詢字串從 `"?"` 之後開始，
- 每個參數由 `"&"` 分隔，
- 每個鍵/值配對由 `"="` 分隔

例如，給定 `"https://example.com/search?name=Alice&age=30"`，傳回：

```json
{
  "name": "Alice",
  "age": "30"
}
```

所有值都應該以字串形式傳回。

# --hints--

`parse_url_query("https://example.com/search?name=Alice&age=30")` 應該傳回 `{"name": "Alice", "age": "30"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"})`)
}})
```

`parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python")` 應該傳回 `{"skill": "programming", "language": "python"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"})`)
}})
```

`parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2")` 應該傳回 `{"category": "books", "sort": "asc", "page": "2"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"})`)
}})
```

`parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now")` 應該傳回 `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_url_query(url):

    return url
```

# --solutions--

```py
def parse_url_query(url):
    query_string = url.split('?')[1]
    return dict(param.split('=') for param in query_string.split('&'))
```
