---
id: 69cfca90e8a0a6d4d6871c51
title: "挑戰 267：Parsec 轉換器"
challengeType: 29
dashedName: challenge-267
---

# --description--

在遙遠的銀河系中，parsec 用來測量時間和距離。給定一個整數 parsec，傳回其等價的時間或距離。

- 如果給定的整數是奇數，代表時間。如果是偶數，代表距離。

使用以下轉換率：

| Parsecs | 時間／距離 |
| - | - |
| 1 | 2 小時 |
| 2 | 6 光年 |

傳回轉換後的整數值。

# --hints--

`convert_parsecs(1)` 應該傳回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(1), 2)`)
}})
```

`convert_parsecs(2)` 應該傳回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(2), 6)`)
}})
```

`convert_parsecs(31)` 應該傳回 `62`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(31), 62)`)
}})
```

`convert_parsecs(88)` 應該傳回 `264`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(88), 264)`)
}})
```

`convert_parsecs(17)` 應該傳回 `34`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(17), 34)`)
}})
```

`convert_parsecs(14)` 應該傳回 `42`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(14), 42)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_parsecs(parsecs):

    return parsecs
```

# --solutions--

```py
def convert_parsecs(parsecs):
    if parsecs % 2 != 0:
        return parsecs * 2
    else:
        return parsecs * 3
```
