---
id: 69cfca90e8a0a6d4d6871c52
title: "挑戰 268：自戀數字"
challengeType: 29
dashedName: challenge-268
---

# --description--

給定一個正整數，判斷它是否為自戀數字。

- 如果一個數字的每個位數的數字，分別以位數的總數為次方後相加，結果等於該數字本身，則該數字為自戀數字。

例如，153 有 3 位數，且 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153，因此它是自戀數字。

# --hints--

`is_narcissistic(153)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(153), True)`)
}})
```

`is_narcissistic(154)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(154), False)`)
}})
```

`is_narcissistic(371)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(371), True)`)
}})
```

`is_narcissistic(512)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(512), False)`)
}})
```

`is_narcissistic(9)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9), True)`)
}})
```

`is_narcissistic(11)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(11), False)`)
}})
```

`is_narcissistic(9474)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9474), True)`)
}})
```

`is_narcissistic(6549)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(6549), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_narcissistic(n):

    return n
```

# --solutions--

```py
def is_narcissistic(n):
    digits = str(n)
    power = len(digits)
    return sum(int(d) ** power for d in digits) == n
```
