---
id: 69cfca90e8a0a6d4d6871c54
title: "挑戰 270：最長共同子字串"
challengeType: 29
dashedName: challenge-270
---

# --description--

給定一個字串，傳回出現超過一次的最長子字串。

- 子字串可以重疊。

# --hints--

`get_longest_substring("abracadabra")` 應該傳回 `"abra"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("abracadabra"), "abra")`)
}})
```

`get_longest_substring("hello world hello")` 應該傳回 `"hello"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("hello world hello"), "hello")`)
}})
```

`get_longest_substring("mississippi")` 應該傳回 `"issi"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("mississippi"), "issi")`)
}})
```

`get_longest_substring("ha ha ha ha ha ha ha")` 應該傳回 `"ha ha ha ha ha ha"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("ha ha ha ha ha ha ha"), "ha ha ha ha ha ha")`)
}})
```

`get_longest_substring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over")` 應該傳回 `"the quick brown fox jumped over"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over"), "the quick brown fox jumped over")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_longest_substring(s):

    return s
```

# --solutions--

```py
def get_longest_substring(s):
    longest = ''

    for length in range(len(s) - 1, 0, -1):
        for i in range(len(s) - length + 1):
            sub = s[i:i + length]
            if s.find(sub) != s.rfind(sub):
                return sub

    return longest
```
