---
id: 69cfca90e8a0a6d4d6871c50
title: "Challenge 266: Good Day"
challengeType: 29
dashedName: challenge-266
---

# --description--

給定一個 `"HH:MM"` 格式（24 小時制）的時間字串，傳回：

- `"Good morning"`，時間介於 `05:00` 到 `11:59` 之間
- `"Good afternoon"`，時間介於 `12:00` 到 `17:59` 之間
- `"Good evening"`，時間介於 `18:00` 到 `21:59` 之間
- `"Good night"`，時間介於 `22:00` 到 `04:59` 之間

# --hints--

`get_greeting("06:30")` should return `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("06:30"), "Good morning")`)
}})
```

`get_greeting("12:00")` should return `"Good afternoon"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("12:00"), "Good afternoon")`)
}})
```

`get_greeting("21:59")` should return `"Good evening"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("21:59"), "Good evening")`)
}})
```

`get_greeting("00:01")` should return `"Good night"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("00:01"), "Good night")`)
}})
```

`get_greeting("11:30")` should return `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("11:30"), "Good morning")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_greeting(s):

    return s
```

# --solutions--

```py
def get_greeting(s):
    hours, minutes = map(int, s.split(':'))
    total = hours * 60 + minutes

    if 300 <= total < 720:
        return 'Good morning'
    if 720 <= total < 1080:
        return 'Good afternoon'
    if 1080 <= total < 1320:
        return 'Good evening'
    return 'Good night'
```
