---
id: 69e2383af7832c8032603b90
title: "Challenge 274: Oldest Person"
challengeType: 29
dashedName: challenge-274
---

# --description--

給定一個物件陣列，每個物件都有 `"name"` 和 `"age"` 屬性，傳回一個包含最年長者姓名的陣列。

如果有多個人擁有相同的最高年齡，請依照他們在輸入中出現的順序，傳回所有人的姓名。

# --hints--

`get_oldest([{ "name": "Brenda", "age": 40 }])` should return `["Brenda"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "Brenda", "age": 40}]), ["Brenda"])`)
}})
```

`get_oldest([{ "name": "Alice", "age": 30 }, { "name": "Bob", "age": 25 }])` should return `["Alice"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "Alice", "age": 30}, {"name": "Bob", "age": 25}]), ["Alice"])`)
}})
```

`get_oldest([{ "name": "Allison", "age": 25 }, { "name": "Bill", "age": 30 }, { "name": "Carol", "age": 30 }])` should return `["Bill", "Carol"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "Allison", "age": 25}, {"name": "Bill", "age": 30}, {"name": "Carol", "age": 30}]), ["Bill", "Carol"])`)
}})
```

`get_oldest([{ "name": "George", "age": 50 }, { "name": "Shirley", "age": 42 }, { "name": "Beth", "age": 48 }, { "name": "Holly", "age": 50 }, { "name": "Kevin", "age": 44 }, { "name": "Frank", "age": 47 }, { "name": "Zach", "age": 50 }, { "name": "Jennifer", "age": 43 }])` should return `["George", "Holly", "Zach"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "George", "age": 50}, {"name": "Shirley", "age": 42}, {"name": "Beth", "age": 48}, {"name": "Holly", "age": 50}, {"name": "Kevin", "age": 44}, {"name": "Frank", "age": 47}, {"name": "Zach", "age": 50}, {"name": "Jennifer", "age": 43}]), ["George", "Holly", "Zach"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_oldest(people):

    return people
```

# --solutions--

```py
def get_oldest(people):
    max_age = max(p["age"] for p in people)
    return [p["name"] for p in people if p["age"] == max_age]
```
