---
id: 69e2383af7832c8032603b94
title: "挑戰 278：咖啡訂單解析器"
challengeType: 29
dashedName: challenge-278
---

# --description--

給定一個咖啡訂單的字串，識別任何選單項目並傳回格式化的訂單。

使用以下選單項目和價格：

| 項目、條款 | 價格 |
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

傳回一個字串，將符合的項目用 `" + "` 連接，後面接冒號和空格（`": "`），以及總價格。

例如，給定 `"I'd like an oat latte with vanilla syrup and an extra shot please."`，傳回 `"oat latte + vanilla syrup + extra shot: $6.25"`

項目應依照它們在選單中出現的順序排列，且總價格應始終保留兩位小數。

# --hints--

`format_coffee_order("I'd like an oat latte with vanilla syrup and an extra shot please.")` 應傳回 `"oat latte + vanilla syrup + extra shot: $6.25"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25")`)
}})
```

`format_coffee_order("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` 應傳回 `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85")`)
}})
```

`format_coffee_order("Can I get a cold brew with some cream and an extra shot.")` 應傳回 `"cold brew + extra shot + cream: $5.75"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75")`)
}})
```

`format_coffee_order("Just an espresso please.")` 應傳回 `"espresso: $3.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Just an espresso please."), "espresso: $3.00")`)
}})
```

`format_coffee_order("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` 應傳回 `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_coffee_order(order):

    return order
```

# --solutions--

```py
def format_coffee_order(order):
    menu = [
        ("cold brew", 450),
        ("oat latte", 500),
        ("cappuccino", 475),
        ("espresso", 300),
        ("vanilla syrup", 75),
        ("caramel drizzle", 60),
        ("extra shot", 50),
        ("oat milk", 75),
        ("cream", 75),
    ]
    found = [(name, price) for name, price in menu if name in order.lower()]
    total = sum(price for _, price in found)
    items = " + ".join(name for name, _ in found)
    return f"{items}: ${total / 100:.2f}"
```
