---
id: 69f35a5bb823ed620fcb7cbd
title: "Challenge 284: I Before E"
challengeType: 29
dashedName: challenge-284
---

# --description--

給定一個字或句子，傳回一個修正版本，其中每個字都遵守「I 在 E 之前，除非在 C 之後」的規則。

- 如果一個字包含 `"ei"`，且前面沒有 `"c"`，就將它取代為 `"ie"`。
- 如果一個字包含 `"ie"`，且前面有 `"c"`，就將它取代為 `"ei"`。
- 其他所有字都保持不變。

# --hints--

`i_before_e("beleive")` should return `"believe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("beleive"), "believe")`)
}})
```

`i_before_e("recieve")` should return `"receive"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("recieve"), "receive")`)
}})
```

`i_before_e("we recieved a breif")` should return `"we received a brief"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("we recieved a breif"), "we received a brief")`)
}})
```

`i_before_e("she beleived the friendly niece could percieve the greif")` should return `"she believed the friendly niece could perceive the grief"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief")`)
}})
```

`i_before_e("we recieved relief after the theif gave us a breif piece of feirce deceit")` should return `"we received relief after the thief gave us a brief piece of fierce deceit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit")`)
}})
```

# --seed--

## --seed-contents--

```py
def i_before_e(sentence):

    return sentence
```

# --solutions--

```py
import re

def i_before_e(sentence):
    words = []
    for word in sentence.split(" "):
        word = re.sub(r"([^c])ei", r"\1ie", word)
        word = re.sub(r"cie", "cei", word)
        words.append(word)
    return " ".join(words)
```
