---
id: 69f8c998d78ad3171a0713b9
title: "挑戰 286：尚未解決的議題"
challengeType: 29
dashedName: challenge-286
---

# --description--

給定一個議題編號陣列和另一個拉取請求（PR）編號陣列，傳回所有 PR 合併後仍然開啟的議題陣列。

- 如果 PR 編號的數字是議題編號數字的旋轉，則該 PR 會關閉該議題。例如，議題 123 會被 PR 231 或 312 關閉。
- PR 不會關閉與其編號完全相同的議題。例如，PR 123 不會關閉議題 123。因此，所有數字都相同的議題無法被關閉。
- 兩者的編號都可能會去除前導零。例如，PR 201 會關閉議題 12（012，是 201 的旋轉）。同理，議題 201 會被 PR 12 關閉。

請依照原本給定的順序，傳回仍然開啟的議題。

# --hints--

`get_open_issues([123, 234], [231])` 應該傳回 `[234]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` 應該傳回 `[345, 16]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` 應該傳回 `[456, 332, 15]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` 應該傳回 `[115, 296, 24]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` 應該傳回 `[95, 395, 754, 296, 709, 237, 1802]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```
