---
id: 69f8c998d78ad3171a0713bf
title: "挑戰 292：更寬的寬高比"
challengeType: 29
dashedName: challenge-292
---

# --description--

給定兩個代表不同影像尺寸的字串，傳回寬高比較大的影像的寬高比。

- 給定的字串格式為 `"WxH"`，例如 `"1920x1080"`。
- 寬高比是寬度與高度的比例，需約分為最簡整數比。例如 `"1920x1080"` 約分後為 `"16:9"`。
- 傳回格式為 `"W:H"` 的字串，例如 `"16:9"`。

# --hints--

`get_wider_aspect_ratio("1920x1080", "800x600")` 應該傳回 `"16:9"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1920x1080", "800x600"), "16:9")`)
}})
```

`get_wider_aspect_ratio("1080x1350", "2048x1536")` 應該傳回 `"4:3"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1080x1350", "2048x1536"), "4:3")`)
}})
```

`get_wider_aspect_ratio("640x480", "2440x1220")` 應該傳回 `"2:1"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("640x480", "2440x1220"), "2:1")`)
}})
```

`get_wider_aspect_ratio("360x640", "1080x1920")` 應該傳回 `"9:16"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("360x640", "1080x1920"), "9:16")`)
}})
```

`get_wider_aspect_ratio("3440x1440", "2048x858")` 應該傳回 `"43:18"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("3440x1440", "2048x858"), "43:18")`)
}})
```

`get_wider_aspect_ratio("12345x61234", "12534x51234")` 應該傳回 `"2089:8539"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("12345x61234", "12534x51234"), "2089:8539")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_wider_aspect_ratio(a, b):

    return a
```

# --solutions--

```py
from math import gcd

def get_wider_aspect_ratio(a, b):
    def parse(s):
        w, h = map(int, s.split("x"))
        g = gcd(w, h)
        return w / h, f"{w // g}:{h // g}"

    ratio_a, simplified_a = parse(a)
    ratio_b, simplified_b = parse(b)

    return simplified_a if ratio_a >= ratio_b else simplified_b
```
