---
id: 6a0dcd03ee4e68698080ef67
title: "挑戰 303：室友分配"
challengeType: 29
dashedName: challenge-303
---

# --description--

給定一個包含人員及其室友群組的陣列，根據以下規則傳回飯店住宿的房間分配：

- 每個人都有 `name` 和 `group` 屬性：

```json
[
  { "name": "Alice", "group": "A" },
  { "name": "Bob", "group": "B" },
  { "name": "Carol", "group": "A" }
]
```

- 人員只能與同一群組的人共用房間，且配對順序依照給定的順序。
- 傳回一個字串陣列，若是共用房間，名字間以 `" and "` 分隔；若是單人房，則只傳回名字。名字必須依照配對的順序出現。以上述範例為例，傳回 `["Alice and Carol", "Bob"]`。

# --hints--

`get_roommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }])` 應該傳回 `["Alice and Carol", "Bob"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "Alice", "group": "A"}, {"name": "Bob", "group": "B"}, {"name": "Carol", "group": "A"}])), sorted(["Alice and Carol", "Bob"]))`)
}})
```

`get_roommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }])` 應該傳回 `["John and Julia", "Jim"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "John", "group": "C"}, {"name": "Julia", "group": "C"}, {"name": "Jim", "group": "C"}])), sorted(["John and Julia", "Jim"]))`)
}})
```

`get_roommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }])` 應該傳回 `["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "Adam", "group": "D"}, {"name": "Abraham", "group": "E"}, {"name": "Austin", "group": "E"}, {"name": "Augustus", "group": "D"}, {"name": "Angelica", "group": "D"}, {"name": "Aaron", "group": "E"}])), sorted(["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"]))`)
}})
```

`get_roommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }])` 應該傳回 `["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "Frank", "group": "A"}, {"name": "Emitt", "group": "B"}, {"name": "Daria", "group": "F"}, {"name": "Charles", "group": "D"}, {"name": "Bailey", "group": "A"}, {"name": "Albert", "group": "F"}])), sorted(["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"]))`)
}})
```

`get_roommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }])` 應該傳回 `["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_roommates([{"name": "Kevin", "group": "A"}, {"name": "Yuri", "group": "A"}, {"name": "Hugo", "group": "B"}, {"name": "Violet", "group": "A"}, {"name": "Brett", "group": "A"}, {"name": "Wayne", "group": "B"}])), sorted(["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def get_roommates(people):

    return people
```

# --solutions--

```py
def get_roommates(people):
    groups = {}
    for person in people:
        groups.setdefault(person["group"], []).append(person["name"])
    result = []
    for members in groups.values():
        for i in range(0, len(members), 2):
            result.append(" and ".join(members[i:i + 2]))
    return result
```
