---
id: 6a0dcd03ee4e68698080ef6b
title: "挑戰 307：分區規定"
challengeType: 29
dashedName: challenge-307
---

# --description--

給定一個二維格子（陣列的陣列），代表一個城市的建築配置，傳回所有違反分區規定的建築座標。

格子中的每個單元格包含下表中的一個標籤。如果建築的（最多）4 個鄰居（水平或垂直方向）中有任何一個是該建築不能相鄰的型別，該建築即為違規。

| 標籤 | 型別 | 不能相鄰的型別 |
| - | - | - |
| `"i"` | industrial | `"R"`, `"I"` |
| `"A"` | Agricultural | `"C"` |
| `"R"` | Residential | `"i"`, `"C"` |
| `"I"` | Institutional | `"i"` |
| `"C"` | Commercial | `"R"`, `"A"` |
| `""` (空字串) | undeveloped | 無限制 |

傳回所有違規單元格的座標，格式為 `[row, col]` 對組的陣列，順序不限。如果沒有違規，傳回空陣列。

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` 應該傳回 `[[0, 0], [0, 1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` 應該傳回 `[[0, 1], [1, 1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` 應該傳回 `[]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` 應該傳回 `[[0, 1], [0, 2], [0, 3]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` 應該傳回 `[[2, 3], [2, 4], [3, 1], [3, 2]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
