---
id: 6a0dcd03ee4e68698080ef68
title: "挑戰 304：行程安排"
challengeType: 29
dashedName: challenge-304
---

# --description--

給定一個至少包含兩個選用停靠點的日間旅遊陣列，傳回有效行程安排的數量。

行程總是包含 `"breakfast"`、`"lunch"` 和 `"dinner"`，這些不會作為引數傳入。選用停靠點可以放在行程中的任何位置，但必須遵守以下規則：

- `"breakfast"` 總是第一個，且在 `"lunch"` 之前至少有一個停靠點。
- `"lunch"` 必須出現在 `"dinner"` 之前，且兩者之間至少有一個停靠點。
- 最多只能有一個選用停靠點出現在 `"dinner"` 之後。

傳回有效安排的數量。

# --hints--

`get_itinerary_count(["library", "park"])` 應該傳回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` 應該傳回 `18`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` 應該傳回 `120`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` 應該傳回 `840`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` 應該傳回 `55440`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
