---
id: 6a0dcc730cb92a616f86f0c5
title: "挑戰 301：最後的洗衣精用量"
challengeType: 29
dashedName: challenge-301
---

# --description--

給定你剩餘的洗衣精勺數以及一個陣列，該陣列記錄你過去幾天每天使用的勺數，請傳回你還能使用多少整天的洗衣精。

從使用紀錄中計算你的每日平均用量，並假設未來每天都使用這個用量。

# --hints--

`last_load_date(10, [2, 2, 2, 2, 2, 2, 2])` 應該傳回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(10, [2, 2, 2, 2, 2, 2, 2]), 5)`)
}})
```

`last_load_date(16, [2, 3, 0, 3, 4, 2, 1])` 應該傳回 `7`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(16, [2, 3, 0, 3, 4, 2, 1]), 7)`)
}})
```

`last_load_date(33, [5, 0, 4, 3, 3, 2])` 應該傳回 `11`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(33, [5, 0, 4, 3, 3, 2]), 11)`)
}})
```

`last_load_date(50, [2, 0, 2, 9, 12, 0, 2])` 應該傳回 `12`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(50, [2, 0, 2, 9, 12, 0, 2]), 12)`)
}})
```

`last_load_date(20, [13, 9, 12, 10, 8])` 應該傳回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(20, [13, 9, 12, 10, 8]), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def last_load_date(scoops, usage):

    return scoops
```

# --solutions--

```py
def last_load_date(scoops, usage):
    avg = sum(usage) / len(usage)
    return int(scoops / avg)
```
