---
id: 6a0dcd03ee4e68698080ef6c
title: "挑戰 308：信用卡驗證器"
challengeType: 29
dashedName: challenge-308
---

# --description--

給定一組信用卡號碼的數字字串，使用以下方法判斷它是否為有效的卡號：

- 從倒數第二個數字開始，向左每隔一個數字將其加倍。
- 如果加倍後的數字大於 9，則減去 9。
- 將所有數字（加倍和未加倍的）相加。
- 如果總和能被 10 整除，則該號碼有效。

# --hints--

`is_valid_card("4532015112830366")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830366"), True)`)
}})
```

`is_valid_card("5425233430109903")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("5425233430109903"), True)`)
}})
```

`is_valid_card("371449635398431")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("371449635398431"), True)`)
}})
```

`is_valid_card("6011111111111117")` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("6011111111111117"), True)`)
}})
```

`is_valid_card("4532015112830367")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830367"), False)`)
}})
```

`is_valid_card("1234567890123456")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("1234567890123456"), False)`)
}})
```

`is_valid_card("4532015112830368")` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830368"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_card(number):

    return number
```

# --solutions--

```py
def is_valid_card(number):
    total = 0
    for i, digit in enumerate(reversed(number)):
        d = int(digit)
        if i % 2 == 1:
            d *= 2
            if d > 9:
                d -= 9
        total += d
    return total % 10 == 0
```
