---
id: 6a151d32d772271dd2448c2d
title: "挑戰 311：魔法施放者"
challengeType: 29
dashedName: challenge-311
---

# --description--

給定一串你正在施放的魔法程式碼字串，計算總分數。

字串中的每個字元代表一個魔法：

| 程式碼 | 魔法 | 類別 | 基本分數 |
| - | - | - | - |
| `"f"` | 火焰 | 破壞 | 3 |
| `"l"` | 閃電 | 破壞 | 3 |
| `"i"` | 冰霜 | 控制元件 | 2 |
| `"w"` | 風 | 控制元件 | 2 |
| `"h"` | 治療 | 恢復 | 1 |
| `"s"` | 盾牌 | 恢復 | 1 |

根據連續施放的魔法類別不同，會套用連擊倍數：

- 第一個魔法總是以基本分數計算。
- 每當施放的魔法類別與前一個不同，連擊倍數會增加 1。
- 若施放的魔法類別與前一個相同，連擊倍數會重設為 1。
- 每個魔法的分數為其基本分數乘以當前連擊倍數。

傳回整串魔法序列的總分數。

# --hints--

`cast("fihwl")` 應該傳回 `33`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` 應該傳回 `45`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` 應該傳回 `37`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` 應該傳回 `50`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` 應該傳回 `101`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
