---
id: 6a15cadf5f240d05a2649556
title: "挑戰 312：串流費用"
challengeType: 29
dashedName: challenge-312
---

# --description--

給定一個陣列，表述你串流服務購物車中的電影，以及一個字串表示你的訂閱等級，傳回電影的總費用。

購物車中的每個條款都是一個物件，包含 `"format"`（`"HD"` 或 `"4K"`）和 `"type"`（`"rent"` 或 `"buy"`）。它們的費用如下：

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

套用以下訂閱等級折扣：

- `"none"`：全價
- `"basic"`：九折
- `"premium"`：七五折

傳回總費用，四捨五入到小數點後兩位，格式為 `"$D.CC"`。

# --hints--

`get_streaming_bill([{ format: "HD", type: "rent" }], "none")` 應該傳回 `"$3.99"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}], "none"), "$3.99")`)
}})
```

`get_streaming_bill([{ format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "premium")` 應該傳回 `"$12.73"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "HD", "type": "buy"}], "premium"), "$12.73")`)
}})
```

`get_streaming_bill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic")` 應該傳回 `"$18.87"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "HD", "type": "rent"}, {"format": "HD", "type": "buy"}], "basic"), "$18.87")`)
}})
```

`get_streaming_bill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium")` 應該傳回 `"$29.98"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "4K", "type": "buy"}, {"format": "4K", "type": "buy"}], "premium"), "$29.98")`)
}})
```

`get_streaming_bill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none")` 應該傳回 `"$42.96"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}], "none"), "$42.96")`)
}})
```

`get_streaming_bill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }, { format: "HD", type: "buy" }], "basic")` 應該傳回 `"$50.36"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}, {"format": "HD", "type": "buy"}], "basic"), "$50.36")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_streaming_bill(cart, subscription):

    return cart
```

# --solutions--

```py
def get_streaming_bill(cart, subscription):
    pricing = {
        "HD": {"rent": 3.99, "buy": 12.99},
        "4K": {"rent": 5.99, "buy": 19.99},
    }
    discounts = {"none": 1, "basic": 0.9, "premium": 0.75}
    multiplier = discounts[subscription]
    total = sum(pricing[item["format"]][item["type"]] * multiplier for item in cart)
    return f"${total:.2f}"
```
