---
id: 6a15cadf5f240d05a2649557
title: "挑戰 313：租賃費用"
challengeType: 29
dashedName: challenge-313
---

# --description--

給定租賃時間戳、歸還時間戳和租賃等級，傳回包含任何逾期費用的租賃總費用。

- 給定的時間戳是 UTC ISO 字串，例如：`"2026-06-18T18:30:00Z"`。
- 租賃等級是指租賃到期前的天數：`1`、`3` 或 `7`。
- 租賃必須在租期最後一天的 UTC 中午 12:00 或之前歸還。例如，3 月 15 日任何時間租出的 1 天租賃，必須在 3 月 16 日 UTC 中午 12:00 前歸還。
- 超過到期日期和時間的每一天都會產生延遲費用。

價格如下：

| 等級 | 基本費用 | 每日逾期費用 |
| - | - | - |
| 1 天 | $4.99 | $3.99 |
| 3 天 | $3.99 | $2.99 |
| 7 天 | $2.99 | $0.99 |

傳回以兩位小數四捨五入的總費用，格式為 `"$D.CC"`。

# --hints--

`get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` 應傳回 `"$4.99"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99")`)
}})
```

`get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` 應傳回 `"$12.97"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97")`)
}})
```

`get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` 應傳回 `"$3.99"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99")`)
}})
```

`get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` 應傳回 `"$9.97"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97")`)
}})
```

`get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` 應傳回 `"$2.99"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99")`)
}})
```

`get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` 應傳回 `"$358.40"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rental_cost(rented, returned, tier):

    return rented
```

# --solutions--

```py
from datetime import datetime, timedelta
import math

def get_rental_cost(rented, returned, tier):
    pricing = {
        1: {"base": 4.99, "late": 3.99},
        3: {"base": 3.99, "late": 2.99},
        7: {"base": 2.99, "late": 0.99},
    }
    rented_dt = datetime.fromisoformat(rented.replace("Z", "+00:00"))
    returned_dt = datetime.fromisoformat(returned.replace("Z", "+00:00"))
    due = rented_dt + timedelta(days=tier)
    due = due.replace(hour=12, minute=0, second=0, microsecond=0)
    base = pricing[tier]["base"]
    late = pricing[tier]["late"]
    if returned_dt <= due:
        return f"${base:.2f}"
    days_late = math.ceil((returned_dt - due).total_seconds() / 86400)
    return f"${base + days_late * late:.2f}"
```
