---
id: 6a15cadf5f240d05a264955e
title: "挑戰 320：血庫"
challengeType: 29
dashedName: challenge-320
---

# --description--

給定一個血庫庫存的陣列和一個病人血型需求的陣列，傳回格式為 `"X of Y patients served"` 的字串。其中 `X` 是能從血庫庫存中供應血液的最大病人數，`Y` 是病人總數。

兩個陣列中的每個項目都是以下血型之一：`"AB"`、`"A"`、`"B"` 或 `"O"`。

相容性規則：

- `"AB"` 可以接受任何血型。
- `"A"` 可以接受 `"A"` 和 `"O"`。
- `"B"` 可以接受 `"B"` 和 `"O"`。
- `"O"` 只能接受 `"O"`。

給定陣列中的重複項目代表數量。

# --hints--

`triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` 應該傳回 `"4 of 4 patients served"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` 應該傳回 `"3 of 5 patients served"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served")`)
}})
```

`triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` 應該傳回 `"4 of 5 patients served"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served")`)
}})
```

`triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` 應該傳回 `"4 of 4 patients served"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` 應該傳回 `"8 of 11 patients served"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served")`)
}})
```

`triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` 應該傳回 `"13 of 14 patients served"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_blood(bank, patients):

    return bank
```

# --solutions--

```py
def triage_blood(bank, patients):
    b = {}
    p = {}
    for t in bank:
        b[t] = b.get(t, 0) + 1
    for t in patients:
        p[t] = p.get(t, 0) + 1

    saved = 0

    def serve(patient, donors):
        nonlocal saved
        for donor in donors:
            n = min(p.get(patient, 0), b.get(donor, 0))
            p[patient] = p.get(patient, 0) - n
            b[donor] = b.get(donor, 0) - n
            saved += n

    serve("O", ["O"])
    serve("A", ["A", "O"])
    serve("B", ["B", "O"])
    serve("AB", ["AB", "A", "B", "O"])
    return f"{saved} of {len(patients)} patients served"
```
