---
id: 6a15cadf5f240d05a264955d
title: "挑戰 319：Frontmatter 解析器"
challengeType: 29
dashedName: challenge-319
---

# --description--

給定一個代表 frontmatter 區段的字串，解析它並傳回一個包含鍵和值的物件（JavaScript）或字典（Python）。

Frontmatter 包裹在 `---` 分隔符中，並包含每行一組的 `key: value` 配對。舉例來說：

```md
---
title: My Post
draft: false
views: 100
---
```

應該傳回：

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- 數字、布林值和字串都應該以它們各自的型別傳回。
- 給定的字串會以換行字元（`"\n"`）分隔新行。上述範例會以 `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"` 的形式給出。

# --hints--

`parse_frontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` 應該傳回 `{ title: "My Post", draft: False, views: 100 }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\ntitle: My Post\\ndraft: false\\nviews: 100\\n---"), {"title": "My Post", "draft": False, "views": 100})`)
}})
```

`parse_frontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` 應該傳回 `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: False }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nid: 6a174db57256a112f932195c\\ntitle: My Book\\nlocale: en\\nwordCount: 10000\\npublished: false\\n---"), {"id": "6a174db57256a112f932195c", "title": "My Book", "locale": "en", "wordCount": 10000, "published": False})`)
}})
```

`parse_frontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` 應該傳回 `{ version: "1.0.0", url: "https://example.com", private: True }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nversion: 1.0.0\\nurl: https://example.com\\nprivate: true\\n---"), {"version": "1.0.0", "url": "https://example.com", "private": True})`)
}})
```

`parse_frontmatter("---\nrating: 4.5\nprice: 9.99\n---")` 應該傳回 `{ rating: 4.5, price: 9.99 }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nrating: 4.5\\nprice: 9.99\\n---"), {"rating": 4.5, "price": 9.99})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_frontmatter(s):

    return s
```

# --solutions--

```py
def parse_frontmatter(s):
    lines = [line for line in s.split("\n") if line != "---" and line != ""]
    result = {}
    for line in lines:
        key, *rest = line.split(":")
        value = ":".join(rest).strip()
        if value == "true":
            result[key.strip()] = True
        elif value == "false":
            result[key.strip()] = False
        else:
            try:
                n = float(value)
                result[key.strip()] = int(n) if n == int(n) else n
            except ValueError:
                result[key.strip()] = value
    return result
```
