---
id: 6a1d9f98e819ed70a0e994e1
title: "挑戰 335：五顆骰子"
challengeType: 29
dashedName: challenge-335
---

# --description--

給定一個包含五顆骰子（數值介於 1 到 6）的陣列，傳回最佳可能的牌型。

以下是牌型由低到高的等級排序：

| 牌型 | 說明 |
|------|-------------|
| `"no pair"` | 沒有配對或更好的牌型 |
| `"pair"` | 兩顆骰子數值相同 |
| `"two pair"` | 兩組不同的配對 |
| `"three of a kind"` | 三顆骰子數值相同 |
| `"small straight"` | 四個連續數字 |
| `"large straight"` | 五個連續數字 |
| `"full house"` | 三顆相同加一對 |
| `"four of a kind"` | 四顆骰子數值相同 |
| `"five of a kind"` | 五顆骰子數值全部相同 |

# --hints--

`five_dice([1, 1, 1, 1, 1])` 應該傳回 `"five of a kind"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 1, 1, 1, 1]), "five of a kind")`)
}})
```

`five_dice([5, 5, 5, 6, 5])` 應該傳回 `"four of a kind"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([5, 5, 5, 6, 5]), "four of a kind")`)
}})
```

`five_dice([2, 5, 6, 4, 3])` 應該傳回 `"large straight"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 5, 6, 4, 3]), "large straight")`)
}})
```

`five_dice([4, 3, 3, 3, 1])` 應該傳回 `"three of a kind"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 3, 3, 3, 1]), "three of a kind")`)
}})
```

`five_dice([4, 6, 2, 6, 5])` 應該傳回 `"pair"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 6, 2, 6, 5]), "pair")`)
}})
```

`five_dice([1, 4, 5, 6, 2])` 應該傳回 `"no pair"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 4, 5, 6, 2]), "no pair")`)
}})
```

`five_dice([1, 3, 4, 6, 2])` 應該傳回 `"small straight"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 3, 4, 6, 2]), "small straight")`)
}})
```

`five_dice([2, 2, 5, 2, 5])` 應該傳回 `"full house"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 2, 5, 2, 5]), "full house")`)
}})
```

`five_dice([6, 4, 5, 6, 4])` 應該傳回 `"two pair"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([6, 4, 5, 6, 4]), "two pair")`)
}})
```

# --seed--

## --seed-contents--

```py
def five_dice(dice):

    return dice
```

# --solutions--

```py
def five_dice(dice):
    from collections import Counter
    counts = Counter(dice)
    values = sorted(counts.values(), reverse=True)
    keys = sorted(counts.keys())

    is_large_straight = len(keys) == 5 and keys[-1] - keys[0] == 4
    is_small_straight = not is_large_straight and any(
        keys[i + 3] - keys[i] == 3 and len(set(keys[i:i + 4])) == 4
        for i in range(len(keys) - 3)
    )

    if values[0] == 5: return "five of a kind"
    if values[0] == 4: return "four of a kind"
    if values[0] == 3 and values[1] == 2: return "full house"
    if is_large_straight: return "large straight"
    if is_small_straight: return "small straight"
    if values[0] == 3: return "three of a kind"
    if values[0] == 2 and values[1] == 2: return "two pair"
    if values[0] == 2: return "pair"
    return "no pair"
```
