---
id: 6a1d9f98e819ed70a0e994d9
title: "挑戰 327：資料庫遷移"
challengeType: 29
dashedName: challenge-327
---

# --description--

給定兩個資料庫物件，傳回第二個物件並填入第一個物件中任何缺少的屬性。

- 記錄中已存在的欄位不應被覆寫。

# --hints--

`migrate_record({ "username": "", "posts": 0 }, { "verified": True })` 應該傳回 `{ "username": "", "posts": 0, "verified": True }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "verified": True }), { "username": "", "posts": 0, "verified": True })`)
}})
```

`migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "posts": 5 })` 應該傳回 `{ "username": "camper", "posts": 5 }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "posts": 5 }), { "username": "camper", "posts": 5 })`)
}})
```

`migrate_record({ "username": "", "posts": 0, "verified": False }, { "username": "camper" })` 應該傳回 `{ "username": "camper", "posts": 0, "verified": False }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0, "verified": False }, { "username": "camper" }), { "username": "camper", "posts": 0, "verified": False })`)
}})
```

`migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "role": "admin" })` 應該傳回 `{ "username": "camper", "role": "admin", "posts": 0 }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "role": "admin" }), { "username": "camper", "role": "admin", "posts": 0 })`)
}})
```

`migrate_record({ "username": "", "email": "", "posts": 0, "verified": False, "role": "user", "banned": False }, { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin" })` 應該傳回 `{ "username": "camper", "email": "camper@freecodecamp.org", "role": "admin", "posts": 0, "verified": False, "banned": False }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "email": "", "posts": 0, "verified": False, "role": "user", "banned": False }, { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin" }), { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin", "posts": 0, "verified": False, "banned": False })`)
}})
```

# --seed--

## --seed-contents--

```py
def migrate_record(schema, record):

    return schema
```

# --solutions--

```py
def migrate_record(schema, record):
    return {**schema, **record}
```
