---
id: 6a1d9f98e819ed70a0e994df
title: "挑戰 333：議題分類 2"
challengeType: 29
dashedName: challenge-333
---

# --description--

給定一個議題標題和一個目前標籤的陣列，根據以下規則傳回更新後的標籤陣列：

如果該議題沒有任何標籤，則添加：

- 如果標題包含 `"error"` 或 `"bug"`，則添加 `"bug"` 和 `"needs triage"`
- 如果標題包含 `"feature"` 或 `"add"`，則添加 `"enhancement"` 和 `"discussing"`

否則，如果給定的標籤包含：

- 如果包含 `"needs triage"` 且標題包含 `"simple"` 或 `"easy"`，則移除 `"needs triage"` 並添加 `"good first issue"`
- 如果包含 `"discussing"` 且標題包含 `"planned"` 或 `"next"`，則移除 `"discussing"` 並添加 `"on the roadmap"`
- 否則，如果存在 `"needs triage"` 或 `"discussing"`，則移除它並添加 `"help wanted"`

如果標題包含：

- `"security"`，則添加 `"critical"` 標籤

# --hints--

`triage_issue("app crashes with error", [])` 應該傳回 `["bug", "needs triage"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", [])), sorted(["bug", "needs triage"]))`)
}})
```

`triage_issue("app crashes with error", ["bug", "needs triage"])` 應該傳回 `["bug", "help wanted"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", ["bug", "needs triage"])), sorted(["bug", "help wanted"]))`)
}})
```

`triage_issue("add dark mode", [])` 應該傳回 `["enhancement", "discussing"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", [])), sorted(["enhancement", "discussing"]))`)
}})
```

`triage_issue("add dark mode", ["enhancement", "discussing"])` 應該傳回 `["enhancement", "help wanted"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted"]))`)
}})
```

`triage_issue("xss security bug", [])` 應該傳回 `["bug", "needs triage", "critical"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("xss security bug", [])), sorted(["bug", "needs triage", "critical"]))`)
}})
```

`triage_issue("security vulnerability in auth", [])` 應該傳回 `["critical"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("security vulnerability in auth", [])), sorted(["critical"]))`)
}})
```

`triage_issue("easy a11y fix", ["bug", "needs triage"])` 應該傳回 `["bug", "good first issue"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("easy a11y fix", ["bug", "needs triage"])), sorted(["bug", "good first issue"]))`)
}})
```

`triage_issue("planned api migration", ["enhancement", "discussing"])` 應該傳回 `["enhancement", "on the roadmap"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("planned api migration", ["enhancement", "discussing"])), sorted(["enhancement", "on the roadmap"]))`)
}})
```

`triage_issue("improve security", ["enhancement", "discussing"])` 應該傳回 `["enhancement", "help wanted", "critical"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("improve security", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted", "critical"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(title, labels):

    return title
```

# --solutions--

```py
def triage_issue(title, labels):
    result = list(labels)
    t = title.lower()

    def has(label):
        return label in result

    def add(label):
        if not has(label):
            result.append(label)

    def remove(label):
        if label in result:
            result.remove(label)

    if len(result) == 0:
        if "error" in t or "bug" in t:
            add("bug")
            add("needs triage")
        elif "feature" in t or "add" in t:
            add("enhancement")
            add("discussing")
    else:
        if has("needs triage") and ("simple" in t or "easy" in t):
            remove("needs triage")
            add("good first issue")
        elif has("discussing") and ("planned" in t or "next" in t):
            remove("discussing")
            add("on the roadmap")
        elif has("needs triage") or has("discussing"):
            remove("needs triage")
            remove("discussing")
            add("help wanted")

    if "security" in t:
        add("critical")

    return result
```
