---
id: 6a2037a68a0bc2aef0006002
title: "挑戰 340：Pig Latin 轉換器"
challengeType: 29
dashedName: challenge-340
---

# --description--

給定一個字串，使用以下規則將它轉換成 Pig Latin：

- 如果一個字以母音（`"a"`、`"e"`、`"i"`、`"o"` 或 `"u"`）開頭，則在字尾加上 `"way"`。例如，`"universe"` 會轉換成 `"universeway"`。
- 如果一個字以一個或多個子音開頭，將這些子音移到字尾並加上 `"ay"`。例如，`"hello"` 會轉換成 `"ellohay"`。
- 保持第一個字母的大小寫。例如，`"Hello"` 會轉換成 `"Ellohay"`。

# --hints--

`pig_latin("universe")` 應該傳回 `"universeway"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("universe"), "universeway")`)
}})
```

`pig_latin("hello")` 應該傳回 `"ellohay"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello"), "ellohay")`)
}})
```

`pig_latin("hello universe")` 應該傳回 `"ellohay universeway"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello universe"), "ellohay universeway")`)
}})
```

`pig_latin("Hello universe")` 應該傳回 `"Ellohay universeway"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Hello universe"), "Ellohay universeway")`)
}})
```

`pig_latin("Pig Latin is fun")` 應該傳回 `"Igpay Atinlay isway unfay"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Pig Latin is fun"), "Igpay Atinlay isway unfay")`)
}})
```

`pig_latin("The quick brown fox jumped over the lazy dog")` 應該傳回 `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday")`)
}})
```

# --seed--

## --seed-contents--

```py
def pig_latin(s):

    return s
```

# --solutions--

```py
def pig_latin(s):
    vowels = "aeiou"

    def convert_word(word):
        lower_word = word.lower()
        if lower_word[0] in vowels:
            return word + "way"
        i = 0
        while i < len(word) and lower_word[i] not in vowels:
            i += 1
        consonants = word[:i]
        rest = word[i:]
        result = rest + consonants + "ay"
        if word[0].isupper():
            return result[0].upper() + result[1:].lower()
        return result

    return " ".join(convert_word(word) for word in s.split(" "))
```
