---
id: 6a2037a68a0bc2aef0005fff
title: "挑戰 337：計數器"
challengeType: 29
dashedName: challenge-337
---

# --description--

給定一串計數記號字串，傳回所表述的總數。

- 每個豎線 `"|"` 代表一個計數。
- 每第五個記號用斜線 `"/"` 表示，完成一個五個一組的群組（`"||||/"`）。
- 群組之間以空格分隔。

# --hints--

`get_tally_count("||||")` 應該傳回 `4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||"), 4)`)
}})
```

`get_tally_count("||||/")` 應該傳回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/"), 5)`)
}})
```

`get_tally_count("||||/ |||")` 應該傳回 `8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ |||"), 8)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||")` 應該傳回 `17`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||"), 17)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` 應該傳回 `41`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_tally_count(s):

    return s
```

# --solutions--

```py
def get_tally_count(s):
    return len(s.replace(" ", ""))
```
