---
id: 6a2037a68a0bc2aef0006006
title: "挑戰 344：黃金比例"
challengeType: 29
dashedName: challenge-344
---

# --description--

給定兩個數字，判斷它們的比例是否接近黃金比例。

- 使用黃金比例 `1.618`
- 容許誤差為 `0.01`

# --hints--

`is_golden_ratio(21, 34)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(21, 34), True)`)
}})
```

`is_golden_ratio(15, 20)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(15, 20), False)`)
}})
```

`is_golden_ratio(8, 13)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(8, 13), True)`)
}})
```

`is_golden_ratio(10, 16)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(10, 16), False)`)
}})
```

`is_golden_ratio(1618, 1000)` 應該傳回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(1618, 1000), True)`)
}})
```

`is_golden_ratio(88, 55)` 應該傳回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(88, 55), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_golden_ratio(a, b):

    return a
```

# --solutions--

```py
def is_golden_ratio(a, b):
    golden_ratio = 1.618
    ratio = max(a, b) / min(a, b)
    return abs(ratio - golden_ratio) <= 0.01
```
