---
id: 6a2037a68a0bc2aef0006001
title: "挑戰 339：陣列分塊"
challengeType: 29
dashedName: challenge-339
---

# --description--

給定一個陣列和分塊大小，傳回將陣列拆分成該大小子陣列的結果。

- 如果陣列無法平均分割，最後一個分塊可能會比較小。

# --hints--

`chunk_array([1, 2, 3, 4, 5, 6], 3)` 應該傳回 `[[1, 2, 3], [4, 5, 6]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]])`)
}})
```

`chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` 應該傳回 `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]])`)
}})
```

`chunk_array([1, 2, 3, 4, 5], 3)` 應該傳回 `[[1, 2, 3], [4, 5]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]])`)
}})
```

`chunk_array(["a", "b", "c", "d", "e"], 1)` 應該傳回 `[["a"], ["b"], ["c"], ["d"], ["e"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]])`)
}})
```

`chunk_array([1, 2, 3], 5)` 應該傳回 `[[1, 2, 3]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3], 5), [[1, 2, 3]])`)
}})
```

# --seed--

## --seed-contents--

```py
def chunk_array(arr, size):

    return arr
```

# --solutions--

```py
def chunk_array(arr, size):
    return [arr[i:i + size] for i in range(0, len(arr), size)]
```
