---
id: 6a2037a68a0bc2aef0006005
title: "挑戰 343：電梯停靠順序"
challengeType: 29
dashedName: challenge-343
---

# --description--

給定電梯目前所在樓層的數字和一個請求樓層的陣列，傳回電梯應該拜訪這些樓層的順序陣列，以最小化行駛樓層數。

- 若有平手，先往上
- 電梯第一次經過有請求的樓層時必須停靠

# --hints--

`elevator_stops(5, [2, 8, 3, 9])` 應該傳回 `[3, 2, 8, 9]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(5, [2, 8, 3, 9]), [3, 2, 8, 9])`)
}})
```

`elevator_stops(6, [2, 10, 8, 3, 1, 9])` 應該傳回 `[8, 9, 10, 3, 2, 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(6, [2, 10, 8, 3, 1, 9]), [8, 9, 10, 3, 2, 1])`)
}})
```

`elevator_stops(1, [4, 8, 3, 6, 9])` 應該傳回 `[3, 4, 6, 8, 9]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(1, [4, 8, 3, 6, 9]), [3, 4, 6, 8, 9])`)
}})
```

`elevator_stops(12, [6, 10, 7, 3, 1, 4])` 應該傳回 `[10, 7, 6, 4, 3, 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(12, [6, 10, 7, 3, 1, 4]), [10, 7, 6, 4, 3, 1])`)
}})
```

`elevator_stops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19])` 應該傳回 `[10, 9, 8, 6, 5, 2, 12, 19, 23]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19]), [10, 9, 8, 6, 5, 2, 12, 19, 23])`)
}})
```

# --seed--

## --seed-contents--

```py
def elevator_stops(current_floor, stops):

    return current_floor
```

# --solutions--

```py
def elevator_stops(current_floor, stops):
    stops = sorted(stops)
    above = [s for s in stops if s > current_floor]
    below = [s for s in stops if s < current_floor][::-1]
    up_distance = above[-1] - current_floor if above else float('inf')
    down_distance = current_floor - below[-1] if below else float('inf')
    return below + above if down_distance < up_distance else above + below
```
