---
id: 6a22d77ddf034bc4e35b1d54
title: "挑戰 347：博弈論"
challengeType: 29
dashedName: challenge-347
---

# --description--

給定兩個相同長度的字串，分別表述兩位玩家在一個遊戲中的策略，傳回分數陣列 `[player1, player2]`。

- 給定的字串只會包含兩個字母之一：`"C"`（合作）或 `"D"`（背叛）。
- 每個字元代表一回合，分數計算如下：
  - 若雙方都合作，雙方各得 3 分。
  - 若雙方都背叛，雙方各得 1 分。
  - 若一方背叛另一方合作，背叛者得 5 分，合作方得 0 分。

# --hints--

`play_game("CCCC", "CCCC")` 應該傳回 `[12, 12]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` 應該傳回 `[4, 4]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` 應該傳回 `[5, 10]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` 應該傳回 `[24, 34]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` 應該傳回 `[66, 21]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
