---
id: 6a22d77ddf034bc4e35b1d59
title: "挑戰 352：對比評級 1"
challengeType: 29
dashedName: challenge-352
---

# --description--

給定一個對比度比率和一個布林值表示文字是否為大字，請根據下表傳回 WCAG 評級：

| 評級 | 一般文字 | 大字文字 |
|--------|-------------|------------|
| `"AAA"` | 7.0 以上 | 4.5 以上 |
| `"AA"` | 4.5 以上 | 3.0 以上 |
| `"Fail"` | 低於 4.5 | 低於 3.0 |

# --hints--

`get_contrast_rating("7.5", False)` 應該傳回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("7.5", False), "AAA")`)
}})
```

`get_contrast_rating("4.8", False)` 應該傳回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.8", False), "AA")`)
}})
```

`get_contrast_rating("4.2", False)` 應該傳回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.2", False), "Fail")`)
}})
```

`get_contrast_rating("4.5", True)` 應該傳回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.5", True), "AAA")`)
}})
```

`get_contrast_rating("3.0", True)` 應該傳回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("3.0", True), "AA")`)
}})
```

`get_contrast_rating("2.7", False)` 應該傳回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("2.7", False), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(ratio, is_large_text):

    return ratio
```

# --solutions--

```py
def get_contrast_rating(ratio, is_large_text):
    ratio = float(ratio)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
