---
id: 6a22d77ddf034bc4e35b1d5b
title: "挑戰 354：對比評級 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

給定兩個代表 RGB 值的陣列和一個布林值表示文字是否為大字，使用以下方法傳回 WCAG 對比評級：

首先，將每個 RGB 值轉換為相對亮度：

- 將每個通道 `[R, G, B]` 除以 255，取得介於 `0` 和 `1` 之間的值
- 對每個通道套用伽瑪校正公式：
  - 如果通道值小於或等於 `0.04045`：`channel / 12.92`
  - 否則：`((channel + 0.055) / 1.055) ^ 2.4`
- 計算亮度：`0.2126 * R + 0.7152 * G + 0.0722 * B`

接著，將每個亮度值加上 0.05，再用較亮的亮度除以較暗的亮度。較亮的亮度值永遠是第一個引數。

根據對比比率，使用下表傳回評級：

| 評級 | 一般文字 | 大字文字 |
|--------|-------------|------------|
| `"AAA"` | 7.0 以上 | 4.5 以上 |
| `"AA"` | 4.5 以上 | 3.0 以上 |
| `"Fail"` | 低於 4.5 | 低於 3.0 |

# --hints--

`get_contrast_rating([255, 255, 255], [0, 0, 0], False)` 應該傳回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

`get_contrast_rating([215, 188, 188], [55, 55, 55], False)` 應該傳回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

`get_contrast_rating([143, 144, 210], [46, 47, 61], False)` 應該傳回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

`get_contrast_rating([167, 167, 210], [53, 10, 53], True)` 應該傳回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

`get_contrast_rating([135, 147, 155], [60, 70, 90], True)` 應該傳回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

`get_contrast_rating([125, 210, 195], [105, 130, 90], True)` 應該傳回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
