---
id: 6a26df95efa55a2524399744
title: "挑戰 363：桶子填色 2"
challengeType: 29
dashedName: challenge-363
---

# --description--

給定一個由單一字母顏色字串組成的二維網格和一個標的顏色，傳回使整個網格變成標的顏色所需的最少洪水填色「點擊」次數。

- 每次點擊會將被點擊的格子顏色和整個與該格子顏色相連的區域改成標的顏色。
- 格子只會水平和垂直相連（不包含斜角相連）。

# --hints--

`bucket_fill([["R", "R"], ["R", "R"]], "G")` 應該傳回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["R", "R"], ["R", "R"]], "G"), 1)`)
}})
```

`bucket_fill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B")` 應該傳回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B"), 0)`)
}})
```

`bucket_fill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R")` 應該傳回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R"), 3)`)
}})
```

`bucket_fill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P")` 應該傳回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P"), 5)`)
}})
```

`bucket_fill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y")` 應該傳回 `12`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y"), 12)`)
}})
```

# --seed--

## --seed-contents--

```py
def bucket_fill(grid, target_color):

    return grid
```

# --solutions--

```py
def bucket_fill(grid, target_color):
    rows = len(grid)
    cols = len(grid[0])
    visited = set()

    def dfs(r, c, color):
        if r < 0 or r >= rows or c < 0 or c >= cols:
            return
        if (r, c) in visited:
            return
        if grid[r][c] != color:
            return
        visited.add((r, c))
        dfs(r + 1, c, color)
        dfs(r - 1, c, color)
        dfs(r, c + 1, color)
        dfs(r, c - 1, color)

    clicks = 0
    for r in range(rows):
        for c in range(cols):
            if (r, c) not in visited and grid[r][c] != target_color:
                dfs(r, c, grid[r][c])
                clicks += 1

    return clicks
```
