---
id: 6a26df3e2988bcdded204897
title: "挑戰 361：口語時間"
challengeType: 29
dashedName: challenge-361
---

# --description--

給定一個指針式時鐘的時針和分針角度（以 12 點方向順時針為單位的度數），傳回口語化的時間。

將分針角度轉換成分鐘（360° = 60 分鐘），然後使用以下規則：

| 分鐘 | 口語說法 |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29（不含 15） | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59（不含 45） | `"X minutes to Z"`（其中 `X` 是 60 減去分鐘數）|

`Y` 是目前的時數，`Z` 是下一個時數，兩者皆由時針角度推算（360° = 12 小時）。

注意：指針角度可能不會剛好落在整數上，請考慮以某種方式進行四捨五入。

# --hints--

`get_spoken_time(90, 0)` 應該傳回 `"3 o'clock"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(90, 0), "3 o'clock")`)
}})
```

`get_spoken_time(160, 120)` 應該傳回 `"20 minutes past 5"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(160, 120), "20 minutes past 5")`)
}})
```

`get_spoken_time(255, 180)` 應該傳回 `"half past 8"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(255, 180), "half past 8")`)
}})
```

`get_spoken_time(67.5, 92)` 應該傳回 `"quarter past 2"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(67.5, 92), "quarter past 2")`)
}})
```

`get_spoken_time(200, 240)` 應該傳回 `"20 minutes to 7"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(200, 240), "20 minutes to 7")`)
}})
```

`get_spoken_time(322.5, 273)` 應該傳回 `"quarter to 11"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(322.5, 273), "quarter to 11")`)
}})
```

`get_spoken_time(117.5, 335)` 應該傳回 `"5 minutes to 4"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(117.5, 335), "5 minutes to 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_time(hour_angle, minute_angle):

    return hour_angle
```

# --solutions--

```py
def get_spoken_time(hour_angle, minute_angle):
    minutes = int(minute_angle / 6)
    hour = int(hour_angle / 30) or 12
    next_hour = (hour % 12) + 1

    if minutes == 0:
        return f"{hour} o'clock"
    if minutes == 15:
        return f"quarter past {hour}"
    if minutes == 30:
        return f"half past {hour}"
    if minutes == 45:
        return f"quarter to {next_hour}"
    if minutes < 30:
        return f"{minutes} minutes past {hour}"
    return f"{60 - minutes} minutes to {next_hour}"
```
