---
id: 6a26df3e2988bcdded204896
title: "挑戰 360：口語化時間長度"
challengeType: 29
dashedName: challenge-360
---

# --description--

給定一個秒數，傳回口語化的英文時間長度。

- 將時間長度拆分成小時、分鐘和秒數。
- 跳過任何為零的數值。
- 適當使用單數或複數（`"1 hour"`、`"2 hours"`）。
- 如果有多個單位，最後兩個單位用 `"and"` 連接，倒數第二和倒數第三個單位用逗號分隔（`"1 hour, 2 minutes and 3 seconds"`）。

# --hints--

`get_spoken_duration(3723)` 應該傳回 `"1 hour, 2 minutes and 3 seconds"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(3723), "1 hour, 2 minutes and 3 seconds")`)
}})
```

`get_spoken_duration(7295)` 應該傳回 `"2 hours, 1 minute and 35 seconds"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(7295), "2 hours, 1 minute and 35 seconds")`)
}})
```

`get_spoken_duration(8521)` 應該傳回 `"2 hours, 22 minutes and 1 second"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(8521), "2 hours, 22 minutes and 1 second")`)
}})
```

`get_spoken_duration(435)` 應該傳回 `"7 minutes and 15 seconds"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(435), "7 minutes and 15 seconds")`)
}})
```

`get_spoken_duration(14455)` 應該傳回 `"4 hours and 55 seconds"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(14455), "4 hours and 55 seconds")`)
}})
```

`get_spoken_duration(72000)` 應該傳回 `"20 hours"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(72000), "20 hours")`)
}})
```

`get_spoken_duration(1)` 應該傳回 `"1 second"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(1), "1 second")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_duration(seconds):

    return seconds
```

# --solutions--

```py
def get_spoken_duration(seconds):
    hours = seconds // 3600
    minutes = (seconds % 3600) // 60
    secs = seconds % 60

    parts = []
    if hours:
        parts.append(f"{hours} {'hour' if hours == 1 else 'hours'}")
    if minutes:
        parts.append(f"{minutes} {'minute' if minutes == 1 else 'minutes'}")
    if secs:
        parts.append(f"{secs} {'second' if secs == 1 else 'seconds'}")

    if len(parts) == 1:
        return parts[0]
    if len(parts) == 2:
        return f"{parts[0]} and {parts[1]}"
    return f"{parts[0]}, {parts[1]} and {parts[2]}"
```
