---
id: 7a97b79c096ecb000d03bd36
title: Construa um calculador do n-ésimo número de Fibonacci
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `fibonacci`.
2. Você deve definir um array chamado `sequence` dentro da função `fibonacci`, e ele deve ser inicializado com os valores `[0, 1]`.
3. A função `fibonacci` deve aceitar um parâmetro, um inteiro não negativo `n`.
4. Chamar `fibonacci(n)` deve usar uma abordagem de programação dinâmica para calcular e retornar o `n`-ésimo número da sequência de Fibonacci, onde cada número é a soma dos dois números anteriores.
5. Cada número calculado da sequência de Fibonacci deve ser anexado ao array `sequence`.

# --hints--

Você deve ter uma função chamada `fibonacci`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.fibonacci);
```

Sua função `fibonacci` deve receber um único parâmetro.

```js
const explorer = await __helpers.Explorer(code);
const { fibonacci } = explorer.allFunctions;
assert.equal(fibonacci?.parameters.length, 1);
```

Você deve ter um array chamado `sequence` dentro da função `fibonacci` inicializado com `[0, 1]`.

```js
const explorer = await __helpers.Explorer(code);
const { fibonacci } = explorer.allFunctions;
const { sequence } = fibonacci?.variables ?? {};
assert.isTrue(sequence?.value.matches('[0, 1]'));
```

`fibonacci(0)` deve retornar `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` deve retornar `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` deve retornar `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` deve retornar `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` deve retornar `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` deve retornar `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` deve retornar `610`.

```js
assert.equal(fibonacci(15), 610);
```

Você não deve usar recursão no seu código.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
