---
id: 7a97b79c096ecb000d03bd36
title: Build an Nth Fibonacci Number Calculator
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Objetivo:** Atenda às user stories abaixo e faça todos os testes passarem para concluir o laboratório.

**User Stories:**

1. Você deve criar uma função chamada `fibonacci`.
2. Você deve definir um array chamado `sequence` dentro da função `fibonacci`, e ele deve ser inicializado com os valores `[0, 1]`.
3. A função `fibonacci` deve aceitar um parâmetro, um inteiro positivo `n`.
4. Chamar `fibonacci(n)` deve usar uma abordagem de programação dinâmica para calcular e retornar o `n`-ésimo número da sequência de Fibonacci, onde cada número é a soma dos dois números anteriores.
5. Cada número calculado na posição `n` na sequência de Fibonacci deve ser armazenado no array `sequence` no índice `n - 1`.

# --hints--

Você deve ter uma função chamada `fibonacci`.

```js
assert.isFunction(fibonacci);
```

Sua função `fibonacci` deve receber um único parâmetro.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

Você deve ter um array chamado `sequence` dentro da função `fibonacci` inicializado com `[0, 1]`.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` should return `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` should return `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` should return `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` should return `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` should return `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` should return `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` should return `610`.

```js
assert.equal(fibonacci(15), 610);
```

Você não deve usar recursão no seu código.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
