---
id: 6619240f46cec8e04d77e03a
title: تمرين الدوال الأساسية أ
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

أنشئ دالة تستقبل عددًا صحيحًا. يجب أن تُرجع هذه الدالة القيمة `integer + 7` إذا كان العدد أقل من `10`. وإذا كان العدد أكبر من أو يساوي `10`، يجب أن تُرجع القيمة `integer - 3`.

يجب أن يكون اسم الدالة `addOrSubtract`.

# --hints--

يجب أن يكون لديك دالة باسم `addOrSubtract`.

```js
assert.isFunction(addOrSubtract);
```

يجب أن تستقبل دالتك عددًا صحيحًا كمعلمة.

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

يجب أن تُرجع العدد الصحيح زائد 7 إذا كان العدد أقل من 10.

```js
assert.strictEqual(addOrSubtract(5), 12);
```

يجب أن تُرجع العدد الصحيح ناقص 3 إذا كان العدد أكبر من أو يساوي 10.

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
