---
id: 661e17c6068359c3ccf2f4d8
title: تمرين الدوال الأساسية د
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

اكتب دالة باسم `lastLetter` تأخذ سلسلة نصية كمعلمة وتُرجع الحرف الأخير من السلسلة النصية.

# --hints--

يجب أن يكون لديك دالة باسم `lastLetter`.

```js
assert.isFunction(lastLetter);
```

يجب أن تأخذ دالتك سلسلة نصية كمعلمة.

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

يجب أن تُرجع الحرف الأخير من السلسلة النصية.

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
