---
id: 661e131f068359c3ccf2f4d6
title: تمرين الدوال الأساسية ب
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

اكتب دالة باسم `multiply` تأخذ معلمتين وتُرجع حاصل ضربهما.

# --hints--

يجب أن يكون لديك دالة باسم `multiply`.

```js
assert.isFunction(multiply);
```

يجب أن تأخذ دالتك عددين صحيحين كمعلمات.

```js
assert.match(multiply.toString(), /\s*multiply\(\s*\w+\s*,\s*\w+\s*\)/);
```

يجب أن تُرجع حاصل ضرب العددين الصحيحين.

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js 
function multiply(a, b) {
  return a * b;
}
```
