---
id: 661e151f068359c3ccf2f4d7
title: تمرين الدوال الأساسية ج
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

اكتب دالة باسم `capitalize` تأخذ سلسلة نصية كمعلمة وتُرجع سلسلة نصية جديدة مع جعل الحرف الأول كبيرًا.

# --hints--

يجب أن يكون لديك دالة باسم `capitalize`.

```js
assert.isFunction(capitalize);
```

يجب أن تأخذ دالتك سلسلة نصية كمعلمة.

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

يجب أن تُرجع دالتك سلسلة نصية جديدة مع جعل الحرف الأول كبيرًا.

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
