---
id: 67fccc2463253b213a000142
title: マルチメディアプレーヤーを作成する
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

前のレッスンでは、`audio` と `video` 要素の操作を学びました。このラボでは、`audio` トラックとトランスクリプト付きの `video` を表示するマルチメディアプレーヤーを作成します。

`audio` 要素には、使用するメディアを指定するための `source` 要素を含める必要があります。

例を示します：

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

`source` 要素は、次のように `video` 要素内でも使用できます：

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. ページのメインタイトル用に `h1` 要素を用意してください。
2. 3つの `section` 要素を用意してください。
3. 最初の `section` 要素の中に、再生中の曲のタイトル用の `h2` 要素を含めてください。
4. `h2` 要素の下に、`audio` 属性と `controls` 属性を持つ `aria-label` 要素を用意してください。
5. `audio` 要素の中に、オーディオファイルを指す `source` 属性と `src` 属性を持つ `type` 要素を含めてください。必要に応じてこのオーディオURLを使っても構いません：`https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. 2番目の `section` 要素の中に、再生中のビデオのタイトル用の `h2` 要素を含めてください。
7. `h2` 要素の下に、`video`、`controls` 属性と `width` 属性を持つ `aria-label` 要素を用意してください。
8. `video` 要素の中に、ビデオファイルを指す `source` 属性と `src` 属性を持つ `type` 要素を含めてください。必要に応じてこのビデオURLを使っても構いません：`https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. `source` 要素の下に、字幕ファイルを指す `track` 属性と `src` 属性、`kind` 属性、`srclang` 属性を持つ `label` 要素を用意してください。必要に応じてこの字幕URLを使っても構いません：`https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. 3番目の `section` 要素の中に、「Transcript」などのセクションタイトル用の `h2` 要素を含めてください。
11. `h2` 要素の下に、ビデオのトランスクリプトを含む `p` 要素を用意してください。

# --hints--

`h1` 要素を用意してください。

```js
assert.isNotNull(document.querySelector("h1"));
```

`h1` にはページのメインタイトルを表すテキストを入れてください。

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

3つの `section` 要素を用意してください。

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

最初の `section` の中に `h2` 要素を含めてください。

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

`h2` には再生中の曲のタイトルを表すテキストを入れてください。

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

最初の `section` の中に `audio` 要素を含めてください。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

`audio` 要素には `controls` 属性を付けてください。

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

`audio` 要素にはオーディオを説明するテキストを持つ `aria-label` 属性を付けてください。

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

`audio` 要素には `source` 要素を含めてください。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

`source` 要素にはオーディオファイルを指す `src` 属性を付けてください。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

`source` 要素にはオーディオの種類を指定する値を持つ `type` 属性を付けてください。

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

2番目の `section` の中に `h2` 要素を含めてください。

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

`h2` には再生中のビデオのタイトルを表すテキストを入れてください。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

2番目の `section` の中に `video` 要素を含めてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

`video` 要素には `controls` 属性を付けてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

`video` 要素にはビデオの幅を指定する値を持つ `width` 属性を付けてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

`video` 要素には `source` 要素を含めてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

`source` 要素にはビデオファイルを指す `src` 属性を付けてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

`source` 要素にはビデオの種類を指定する値を持つ `type` 属性を付けてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

`video` 要素には `track` 要素を含めてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

`track` 要素には字幕ファイルを指す `src` 属性を付けてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

`track` 要素にはファイルの種類を説明するテキストを持つ `kind` 属性を付けてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

`track` 要素には字幕の言語を説明するテキストを持つ `srclang` 属性を付けてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

`track` 要素には字幕の言語を説明するテキストを持つ `label` 属性を付けてください。

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

3番目の `section` の中に `h2` 要素を含めてください。

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

`h2` にはセクションのタイトルを表すテキストを入れてください。

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

3番目の `section` の中にトランスクリプト用の `p` 要素を含めてください。

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
