---
id: 61329b210dac0b08047fd6ab
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`meta` 요소를 계속해서, `viewport` 정의는 브라우저가 페이지를 어떻게 렌더링할지 알려줍니다. 하나를 포함하면 모바일에서 시각적 접근성이 향상되고, _SEO_ (검색 엔진 최적화)도 개선됩니다.

페이지의 `viewport`와 `content`를 자세히 설명하는 `width` 속성을 가진 `initial-scale` 정의를 추가하세요.

# --hints--

`meta` 안에 또 다른 `head` 요소를 만들어야 합니다.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

`meta`에 `name` 속성으로 `viewport`을 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

`meta`에 `content` 속성으로 `width=device-width, initial-scale=1`을 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
