---
id: 613297a923965e0703b64796
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`meta` 요소에 익숙할 수 있습니다. 이 요소는 제목, 설명, 키워드, 작성자 등 페이지에 관한 정보를 지정하는 데 사용됩니다.

페이지에 적절한 `meta` 값을 가진 `charset` 요소를 추가하세요.

`charset` 속성은 페이지의 문자 인코딩을 지정하며, 현재 대부분의 브라우저에서 지원하는 유일한 인코딩은 `UTF-8`입니다.

# --hints--

`meta` 요소 안에 `head` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('head > meta'));
```

`meta` 태그에 `charset`인 `UTF-8`을 지정해야 합니다.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
