---
id: 613e2546d0594208229ada50
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`header` 내에서 `img`, `h1`, `nav` 요소를 중첩하여 페이지에 대한 문맥을 제공하세요.

`img`는 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`를 가리켜야 하며, `id`의 `logo`를 가지고, `alt`의 `freeCodeCamp` 텍스트를 가져야 합니다.

`h1`에는 `HTML/CSS Quiz` 텍스트가 포함되어야 합니다.

# --hints--

`img` 요소에 `header` 요소를 추가하세요.

```js
assert.exists(document.querySelector('header > img'));
```

`h1` 요소에 `header` 요소를 추가하세요.

```js
assert.exists(document.querySelector('header > h1'));
```

`nav` 요소에 `header` 요소를 추가하세요.

```js
assert.exists(document.querySelector('header > nav'));
```

`img`, `h1`, `nav` 요소를 그 순서대로 배치하세요.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

`img` 요소에 `src`의 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` 속성을 부여하세요.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

`img` 요소에 `id`의 `logo` 속성을 부여하세요.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

`img` 요소에 `alt`의 `freeCodeCamp` 속성을 부여하세요.

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

`h1` 요소에 `HTML/CSS Quiz` 텍스트를 부여하세요.

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
