---
id: 613e275749ebd008e74bb62e
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

현재 `img`는 기본 크기를 사용하고 있는데, 이 크기는 너무 큽니다. CSS에는 쉼표로 구분된 값들 중 가장 큰 값을 반환하는 `max` 함수가 있습니다. 예를 들어:

```css
img {
  width: max(250px, 25vw);
}
```

위 예제에서 뷰포트 너비가 1000픽셀 미만이면 이미지 너비는 250px가 됩니다. 뷰포트 너비가 1000픽셀보다 크면 이미지 너비는 25vw가 됩니다. 이는 25vw가 뷰포트 너비의 25%와 같기 때문입니다.

이미지의 `id`를 선택자로 사용하고, `width`를 `10rem`와 `18vw` 중 최대값으로 설정하여 이미지를 크기 조절하세요.

# --hints--

`#logo` 선택자를 사용해 `img` 요소를 선택하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

`img`에 `width`를 `max(10rem, 18vw)`으로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
