---
id: 61408f155e798909b6908712
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

페이지에서 탐색을 가능하게 하려면, 다음 세 개의 목록 항목을 포함하는 순서 없는 목록을 추가하세요:

- `INFO`
- `HTML`
- `CSS`

목록 항목 텍스트는 앵커 태그로 감싸야 합니다.

# --hints--

하나의 `ul` 요소를 `nav` 안에 중첩하세요.

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

세 개의 `li` 요소를 `ul` 요소 안에 중첩하세요.

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

각 `a` 요소 안에 하나의 `li` 요소를 중첩하세요.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

첫 번째 `a` 요소에 `INFO` 텍스트를 넣으세요.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[0]?.textContent, 'INFO');
```

두 번째 `a` 요소에 `HTML` 텍스트를 넣으세요.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[1]?.textContent, 'HTML');
```

세 번째 `a` 요소에 `CSS` 텍스트를 넣으세요.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[2]?.textContent, 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
--fcc-editable-region--
      <nav>

      </nav>
--fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```
