---
id: 6141fabd6f75610664e908fd
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

퀴즈이므로 사용자가 답변을 제출할 수 있는 폼이 필요합니다. `section` 요소를 사용해 폼 내 콘텐츠를 의미론적으로 구분하세요.

`main` 요소 안에 세 개의 중첩된 `section` 요소가 있는 폼을 만드세요. 그런 다음, 올바른 메서드를 사용해 폼이 `https://freecodecamp.org/practice-project/accessibility-quiz`로 제출되도록 하세요.

# --hints--

`form` 요소 안에 `main` 요소를 중첩해야 합니다.

```js
assert.exists(document.querySelector('main > form'));
```

`section` 요소 안에 세 개의 `form` 요소를 중첩해야 합니다.

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

`form` 요소에 `action` 속성을 부여해야 합니다.

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

`action` 속성에 `https://freecodecamp.org/practice-project/accessibility-quiz` 값을 부여해야 합니다.

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

`form` 요소에 `method` 속성을 부여해야 합니다.

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

`form` 요소에 `method` 속성을 부여하고, 그 값은 `post`이어야 합니다.

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
