---
id: 614206033d366c090ca7dd42
title: 17단계
challengeType: 0
dashedName: step-17
---

# --description--

타입페이스는 페이지 접근성에서 중요한 역할을 합니다. 일부 폰트는 다른 폰트보다 읽기 쉽고, 특히 저해상도 화면에서 그렇습니다.

`h1`와 `h2` 요소 모두의 폰트를 `Verdana`로 변경하고, 대체 수단으로 산세리프 계열의 다른 웹 안전 폰트를 사용하세요.

또한, `border-bottom` 요소에 `4px solid #dfdfe2`의 `h2`를 추가하여 섹션을 구분되게 만드세요.

# --hints--

`h1`와 `h2` 요소를 대상으로 여러 요소 선택자를 사용해야 합니다.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
assert.exists(gs('h1, h2') || gs('h2, h1'));
```

`font-family` 속성의 첫 번째 값을 `Verdana`로 설정해야 합니다.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const style = gs('h1, h2') || gs('h2, h1');
assert.include(style?.fontFamily, 'Verdana');
```

`font-family` 속성의 두 번째 값을 다른 산세리프 웹 안전 폰트로 설정해야 합니다. _힌트: Tahoma를 선택하는 것이 좋습니다_.

```js
// Acceptable fonts: Arial, sans-serif, Helvetica, Tahoma, Trebuchet MS.
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const style = gs('h1, h2') || gs('h2, h1');
assert.match(style?.fontFamily, /(Tahoma)|(Arial)|(sans-serif)|(Helvetica)|(Trebuchet MS)/);
```

`h2` 요소를 대상으로 `h2` 요소 선택자를 사용해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h2'));
```

`h2`에 `border-bottom` 값의 `4px solid #dfdfe2` 속성을 부여해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h2')?.borderBottom, '4px solid rgb(223, 223, 226)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
        </section>
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--

```
