---
id: 6144e818fd5ea704fe56081d
title: 28단계
challengeType: 0
dashedName: step-28
---

# --description--

각 `fieldset`에 적절한 `name` 속성을 지정하세요. 그런 다음, 두 개의 순서 없는 목록에 `class`를 `answers-list`로 지정하세요.

마지막으로, `legend`를 사용해 `fieldset`의 내용을 캡션으로 만들고, 텍스트 내용으로 참/거짓 문제를 배치하세요.

# --hints--

첫 번째 `fieldset`에 적절한 `name` 속성을 지정해야 합니다. _힌트: `html-question-one`를 사용하세요_

```js
assert.notEmpty(document.querySelectorAll('fieldset')?.[0]?.name);
```

두 번째 `fieldset`에 적절한 `name` 속성을 지정해야 합니다. _힌트: `html-question-two`를 사용하세요_

```js
assert.notEmpty(document.querySelectorAll('fieldset')?.[1]?.name);
```

첫 번째 `ul` 요소에 `class`를 `answers-list`로 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset > ul')?.[0]?.className, 'answers-list');
```

두 번째 `ul` 요소에 `class`를 `answers-list`로 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset > ul')?.[1]?.className, 'answers-list');
```

첫 번째 `legend` 요소에 텍스트 내용을 지정해야 합니다.

```js
assert.notEmpty(document.querySelectorAll('legend')?.[0]?.textContent);
```

두 번째 `legend` 요소에 텍스트 내용을 지정해야 합니다.

```js
assert.notEmpty(document.querySelectorAll('legend')?.[1]?.textContent);
```

두 `legend` 요소에 같은 텍스트 내용을 사용하면 안 됩니다.

```js
assert.notEqual(document.querySelectorAll('legend')?.[0]?.textContent, document.querySelectorAll('legend')?.[1]?.textContent);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label for="student-name">Name:</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">Email:</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">Date of Birth:</label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
          <div class="question-block">
            <h3><span class="sr-only">Question</span>1</h3>
            <fieldset class="question">
              <legend></legend>
              <ul>
                <li></li>
                <li></li>
              </ul>
            </fieldset>
          </div>
          <div class="question-block">
            <h3><span class="sr-only">Question</span>2</h3>
            <fieldset class="question">
              <legend></legend>
              <ul>
                <li></li>
                <li></li>
              </ul>
            </fieldset>
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

.sr-only {
  position: absolute;
  width: 1px;
  height: 1px;
  overflow: hidden;
  clip: rect(0, 0, 0, 0);
  clip-path: inset(50%);
  white-space: nowrap;
}

```
