---
id: 614ccc21ea91ef1736b9b578
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

접근성 퀴즈 첫 번째 부분에 오신 것을 환영합니다. HTML과 CSS 개발자로서 경험이 쌓여 가는 만큼, 기본 보일러플레이트부터 시작합니다.

이 접근성 여정을 `lang` 요소에 `html` 속성을 추가하는 것으로 시작하세요. 이렇게 하면 화면 읽기 프로그램이 페이지의 언어를 식별하는 데 도움이 됩니다.

# --hints--

`html` 요소에 `lang` 속성을 부여해야 합니다. _힌트: 영어의 경우 `en` 값을 사용할 수 있습니다._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
