---
id: 66fbcdfed6848e48947dbbc9
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, crearás un rastreador de recetas utilizando objetos de JavaScript.

Comienza creando un array vacío llamado `recipes`. Este es el array al que agregarás los objetos de recetas más adelante.

# --hints--

Debe crear un array llamado `recipes`.

```js
assert.isArray(recipes)
```

El array `recipes` debe estar vacío.

```js
assert.isEmpty(recipes)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
