---
id: 670e4f45f7116c0f216a5177
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Antes de continuar, debe practicar cómo acceder a propiedades de un objeto.

Puede utilizar notación de punto (`.`) o corchete (`[]`) para hacer esto. Aquí hay un ejemplo:

```js
const person = {
  name: 'John',
  age: 30,
  job: 'Software Engineer'
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Acceda a las propiedades `name` de `recipe1` y `recipe2`, y asígnelas a las variables `recipe1Name` y `recipe2Name`, respectivamente.

Luego, acceda a las propiedades `cookingTime` de ambas recetas y asígnelas a las variables `recipe1CookingTime` y `recipe2CookingTime`, respectivamente.

Asegúrese de que todas las variables que creó se registren en la consola.

# --hints--

Debe crear una variable `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

Debe asignar el valor de la propiedad `name` de `recipe1` a su variable `recipe1Name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Debe registrar `recipe1Name` en la consola.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Debe crear una variable `recipe2Name`.

```js
assert.isDefined(recipe2Name);
```

Debe asignar el valor de la propiedad `name` de `recipe2` a su variable `recipe2Name`.

```js
assert.strictEqual(recipe2Name, recipe2.name);
```

Debe registrar `recipe2Name` en la consola.

```js
assert.match(code, /console\.log\(recipe2Name\)/);
```

Debe crear una variable `recipe1CookingTime`.

```js
assert.isDefined(recipe1CookingTime);
```

Debe asignar el valor de la propiedad `cookingTime` de `recipe1` a su variable `recipe1CookingTime`.

```js
assert.strictEqual(recipe1CookingTime, recipe1.cookingTime);
```

Debe registrar `recipe1CookingTime` en la consola.

```js
assert.match(code, /console\.log\(recipe1CookingTime\)/);
```

Debe crear una variable `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

Debe asignar el valor de la propiedad `cookingTime` de `recipe2` a su variable `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Debe registrar `recipe2CookingTime` en la consola.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

--fcc-editable-region--

--fcc-editable-region--
```
