---
id: 66c07238b01053abaf812065
title: Construye una Calculadora Factorial
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

Un factorial es el producto de un número entero y todos los enteros por debajo de él. Por ejemplo, el factorial de `5` es `5 * 4 * 3 * 2 * 1 = 120`. En este laboratorio, crearás un calculador de factoriales que toma un número del usuario y calcula el factorial de ese número.

Cumple con las historias de usuario a continuación y pasa todas las pruebas para completar el laboratorio.

**Historia de usuario:**

1. You should declare a variable `num` and assign it a number of your choice. El número asignado debe estar entre `1` y `20` inclusive.
1. Crea una función llamada `factorialCalculator` que toma un número como argumento y devuelve el factorial de ese número.
1. Dentro de la función, declara una variable `result` y asígnale el valor de `1`. Usando un bucle, itera a través de todos los números desde `1` hasta el número de entrada (inclusive) y para cada número, multiplica la variable `result` por el número actual y asigna el resultado a la variable `result`. Puedes elegir usar un bucle `for`, bucle `while`, o bucle `do...while` aquí.
1. Debes llamar a la función `factorialCalculator` con `num` como argumento y asignar el resultado a la variable `factorial`.
1. Debes almacenar la salida final en el formato `Factorial of [num] is [factorial]` y asignarlo a la variable `resultMsg`.
1. Debes imprimir el valor de `resultMsg` en la consola.

# --hints--

Debes tener una variable `num`.

```js
assert.isDefined(num);
```

Debes asignar un valor a la variable `num`

```js
assert.isNotNull(num);
```

El valor de `num` debe estar entre `1` y `20`.

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

El valor de `num` debe ser un número.

```js
assert.isNumber(num);
```

Debes tener una función llamada `factorialCalculator`.

```js
assert.isFunction(factorialCalculator);
```

La función `factorialCalculator` debe tomar un número como argumento.

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

El factorial de `5` debe devolver `120`.

```js
assert.strictEqual(factorialCalculator(5), 120);
```

El factorial de `7` debe devolver `5040`.

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

Debes llamar a tu función `factorialCalculator` con la variable `num` como argumento.

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s?$/m);
```

Debes definir una variable `factorial` y asignarle el resultado de la función `factorialCalculator`.

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

Tu `factorialCalculator` debe producir el resultado correcto.

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

Tu `resultMsg` debe contener una cadena en un formato `Factorial of [num] is [factorial]`.

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

Debes mostrar el valor de `resultMsg` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
