---
id: 65ef181dad3d96b14a21a03e
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Projekt finden Sie die ungefähre Quadratwurzel einer gegebenen Zahl mithilfe der Bisektionsmethode.

Die Bisektionsmethode ist eine Technik, um die Nullstellen einer reellwertigen Funktion zu finden. Sie funktioniert, indem sie ein Intervall eingrenzt, in dem die Quadratwurzel liegt, bis sie auf einen Wert innerhalb einer vorgegebenen Toleranz konvergiert.

Beginnen Sie damit, eine Funktion namens `square_root_bisection` zu erstellen. Lassen Sie die Parameter vorerst leer.

# --hints--

Sie sollten die Funktion `square_root_bisection` ohne Parameter definieren.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
