---
id: 65ef1be2b958cabf712296cd
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Jetzt werden Sie das Intervall wiederholt eingrenzen, indem Sie den Mittelpunkt des aktuellen Intervalls finden und das Quadrat des Mittelpunkts mit dem Zielwert vergleichen.

Erstellen Sie dazu innerhalb des `else`-Blocks eine `for`-Schleife, die bis zu `max_iterations`-mal ausgeführt wird.

Verwenden Sie für Ihre Schleife die `range`-Funktion, die eine Folge von Zahlen erzeugt, die Sie durchlaufen können. Die Syntax lautet `range(start, stop, step)`, wobei `start` die Startzahl (einschließlich), `stop` die Endzahl (ausschließlich) und `step` die Differenz zwischen zwei aufeinanderfolgenden Zahlen in der Folge ist.

Verwenden Sie außerdem `_` als Schleifenvariable. `_` dient als Platzhalter und ist nützlich, wenn Sie eine Variable benötigen, deren Wert Sie aber nicht tatsächlich verwenden müssen.

# --hints--

Sie sollten eine `for`-Schleife erstellen, um `range(max_iterations)` zu durchlaufen. Verwenden Sie `_` als Schleifenvariable.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
