---
id: 65ef198fde24dfb7ff675b42
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Die `raise`-Anweisung ermöglicht es Ihnen, eine bestimmte Ausnahme absichtlich auszulösen. Sie besteht aus dem Schlüsselwort `raise` gefolgt vom Ausnahmetyp und erlaubt es Ihnen, eine benutzerdefinierte Fehlermeldung anzugeben:

```py
raise ValueError("Invalid value")
```

Wenn der obige Code ausgeführt wird, wird eine `ValueError` ausgelöst und die Meldung `"Invalid value"` dem Benutzer angezeigt.

Wenn die `square_target` kleiner als `0` ist, kann keine reelle Quadratwurzel berechnet werden. Daher lösen Sie eine `ValueError` mit der Meldung `'Square root of negative number is not defined in real numbers'` aus. Vergessen Sie nicht, das Schlüsselwort `pass` zu entfernen.

# --hints--

Sie sollten das Schlüsselwort `pass` entfernen.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

Sie sollten eine `ValueError` mit der Meldung `'Square root of negative number is not defined in real numbers'` innerhalb des `if`-Blocks auslösen.

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
