---
id: 65ef19425d1b27b6c930bae6
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Wenn die Zahl, für die Sie die Quadratwurzel finden möchten, negativ ist, sollte der Code einen Fehler auslösen, da die Quadratwurzel einer negativen Zahl in den reellen Zahlen nicht definiert ist.

Entfernen Sie die `pass`-Anweisung und erstellen Sie eine `if`-Anweisung, um zu prüfen, ob `square_target` kleiner als `0` ist.

# --hints--

Sie sollten das Schlüsselwort `pass` entfernen.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

Sie sollten eine `if`-Anweisung erstellen, um zu prüfen, ob `square_target < 0`.

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
