---
id: 65ef1a50049cf9bada13266f
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Erstellen Sie eine `elif`-Anweisung, um zu prüfen, ob `square_target` gleich `0` ist. Falls ja, weisen Sie der Variablen `0` den Wert `root` zu. Geben Sie außerdem die Nachricht `'The square root of {square_target} is 0'` aus. Denken Sie daran, die Nachricht mit einem f-string zu formatieren.

# --hints--

Sie sollten eine `elif`-Anweisung haben, um die Bedingung `square_target == 0` zu prüfen.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

Sie sollten der Variablen `0` den Wert `root` zuweisen und das Argument `f'The square root of {square_target} is 0'` an einen `print`-Aufruf übergeben.


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
