---
id: 65ef190c6b51e9b5a5f7ed29
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Geben Sie der `square_root_bisection`-Funktion die folgenden Parameter:

- `square_target`: Die Zahl, für die Sie die Quadratwurzel finden möchten.
- `tolerance` (optional): Die akzeptable Differenz zwischen dem Quadrat des ungefähren Wurzelwerts und dem tatsächlichen Zielwert (Standard ist `1e-7`). Die Toleranz `1e-7` bedeutet, dass die Lösung bis auf `0.0000001` des tatsächlichen Werts genau ist und ist eine gute Standardwahl, die Genauigkeit und Leistung ausbalanciert.
- `max_iterations` (optional): Die maximale Anzahl der durchzuführenden Iterationen (Standard ist `100`). Wenn die Methode innerhalb dieses Limits nicht konvergiert, gehen Sie davon aus, dass keine Lösung gefunden wurde.
  
# --hints--

Ihre Funktion sollte diese Parameter haben: `square_target`, `tolerance = 1e-7` und `max_iterations = 100`. Die Reihenfolge ist wichtig.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```

