---
id: 65ef1d5e3d2927c5e0f4997b
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

In Python überprüft das Schlüsselwort `is` die Objektidentität. Es wird verwendet, um festzustellen, ob zwei Variablen auf dasselbe Objekt im Speicher zeigen. Im Gegensatz zum `is`, dem Gleichheitsoperator (`==`), bestimmt es, ob die Werte zweier Objekte gleich sind, unabhängig davon, ob es sich um dasselbe Objekt im Speicher handelt.

Außerhalb der for-Schleife erstellen Sie eine `if`-Anweisung, um zu prüfen, ob `root is None` ist. Wenn ja, geben Sie die Nachricht `'Failed to converge within {max_iterations} iterations.'` aus. Denken Sie daran, die Nachricht mit einem f-String zu formatieren.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, die prüft, ob `root is None`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

Sie sollten das Argument `f'Failed to converge within {max_iterations} iterations.'` an die `print`-Funktion übergeben.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
