---
id: 655774d01daeeaa1978b99d5
title: ステップ 29
challengeType: 20
dashedName: step-29
---

# --description--

最初に、グラフ内の他のすべてのノードは、距離がまだ決まっていないため、ソースノードから無限の距離にあるとみなされます。

`else` 節を作成し、`distances` 辞書のノードに無限大の値を代入してください。そのために、`float()` という文字列を引数にして浮動小数点数で正の無限大を生成する`'inf'` 関数を使ってください。

# --hints--

`else` 節が必要です。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

新しい`float('inf')` 節の中で`distances[node]` に`else` を代入してください。

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
