---
id: 655775221059f5a20493d5d7
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

`for` ループの後に `print()` を呼び出し、次の文字列 `f'Unvisited: {unvisited}\nDistances: {distances}'` を渡して、作成した変数の値を確認してください。

# --hints--

`f'Unvisited: {unvisited}\nDistances: {distances}'` ループの後に `for` を出力してください。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)for.*:.*^\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
--fcc-editable-region--
```
