---
id: 65578d0f6c78a0b868a43b9c
title: ステップ 40
challengeType: 20
dashedName: step-40
---

# --description--

`while` ループの中で、最初に行うことは現在訪問するノードを定義することです。そのために `min()` 関数を使うことができます。この関数は引数として渡されたイテラブルの中で最小のアイテムを返します。

`pass` を削除し、`current` という変数を作成して `min(unvisited)` を代入してください。

# --hints--

`current` ループの中で `while` 変数を作成する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

`min(unvisited)` 変数に `current` を代入してください。`pass` は削除することを忘れないでください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
