---
id: 65577a17564ce8a8e06c1460
title: ステップ 35
challengeType: 20
dashedName: step-35
---

# --description--

辞書内包表記は条件付き`if`/`else`構文もサポートしています：

```py
{key: val_1 if condition else val_2 for key in dict}
```

上の例では、`dict`が既存の辞書です。`condition`が`True`と評価されるとき、`key`は`val_1`の値を持ち、それ以外の場合は`val_2`になります。

辞書内包表記を使って`graph`に基づく辞書を作成し、`distances`変数に代入してください。ノードが開始ノードと等しい場合はキーにゼロの値を与え、それ以外の場合は無限大を与えてください。後者を実現するには`float('inf')`を使ってください。

# --hints--

`distances`変数に値を与えるには辞書内包表記の構文を使うべきです。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
