---
id: 655776db1eeae0a620e42a0d
title: ステップ 31
challengeType: 20
dashedName: step-31
---

# --description--

今、`my_graph` と `'A'` を引数として渡して関数を呼び出してください。

# --hints--

`shortest_path` を呼び出す際に、`my_graph` と `'A'` を引数として渡す必要があります。

```js
({ test: () => assert.match(code, /^shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
--fcc-editable-region--
```
