---
id: 65578cb031cd93b77a285db2
title: ステップ 38
challengeType: 20
dashedName: step-38
---

# --description--

あなたの関数は、開始ノードに接続されているすべてのノードを探索します。すべてのノードに対して最短経路を計算します。その後、開始ノードを未訪問ノードから削除します。

次に、最も近い隣接ノードを訪問し、このプロセスをすべてのノードが訪問されるまで繰り返します。

ここからは、グラフ内のノードを探索するメインループの作成に取り組みます。アルゴリズムの開発中に無限ループが発生する問題を避けるため、関数の呼び出しをコメントにしてください。

# --hints--

関数の呼び出しをコメントにしてください。

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
