---
id: 65577791ad8c26a7705e2919
title: ステップ 34
challengeType: 20
dashedName: step-34
---

# --description--

辞書内包表記を使うと、既存の辞書から新しい辞書を作成できます。

```py
{key: val for key in dict}
```

上の例では、`val` が新しい辞書で `key` に対応する値で、`dict` は既存の辞書です。

開始ノードと他の各ノード間の経路を追跡したいです。

`distances` 変数の後に、`paths` 変数を作成し、`graph` のすべてのキーを持つ辞書を代入してください。各キーには空のリストを代入し、辞書内包表記を使って辞書を作成してください。

# --hints--

`paths` という名前の変数があるはずです。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

`paths` 変数は辞書内包表記の構文を使い、graph の各ノードに空のリストを代入するようにしてください。

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
