---
id: 655791e6cf5e03be3de73451
title: ステップ 50
challengeType: 20
dashedName: step-50
---

# --description--

関数の呼び出しのコメントを外そうとすると、動作しません。修正すべきバグがいくつかあります。最初のバグは、ネストされた`if`の中で、`paths[node]`リストに存在しないかもしれない要素にアクセスしようとしていることが原因です。ですので、`paths[node]`にアクセスする前に`paths[node][-1]`が空でないことを確認する必要があります。

ネストされた`if`文に追加の条件を加えて、`paths[node]`にアクセスする前に`paths[node][-1]`が空でないことを確かめてください。

# --hints--

ネストされた`paths[node]`文の最初の条件として`if`を追加してください。条件を組み合わせるには`and`演算子を使いましょう。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                   
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
