---
id: 655774955b097ea14897db12
title: ステップ 28
challengeType: 20
dashedName: step-28
---

# --description--

開始ノードからの距離はゼロです。なぜならアルゴリズムはそこから評価を始めるからです。

ループ内で`node`を`unvisited`に追加した後、ノードが開始ノードと等しい場合に発動する`if`文を作成してください。そして`0`辞書内のそのノードに`distances`を代入してください。

# --hints--

`if`が`node`と等しいときに実行される`start`文を作成してください。

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

新しい`if`文の中で、`0`辞書のノードに`distances`を代入してください。

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
