---
id: 6557712d77ce2d9bd7e63afd
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

`B` に接続された別のノードをグラフに追加し、`C` と名付けてください。

既存の辞書をこの配置を表すように変更してください：`'C'` にもう一つのキー `my_graph` を追加し、その値に文字列 `'B'` を設定してください。

また、既存の `'B'` キーの値をリスト `['A', 'C']` に変更して、`'B'` ノードの複数の接続を表してください。

# --hints--

辞書には `'A'`、`'B'`、`'C'` の3つのキーが必要です。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` の値は文字列 `'B'` であるべきです。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` はリストであるべきです。

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']` の値は `'A'` と `'C'` を含むリストであるべきです。

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']` の値は文字列 `'B'` であるべきです。

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
