---
id: 6557924d47c325bf27afbe51
title: ステップ 53
challengeType: 20
dashedName: step-53
---

# --description--

アルゴリズムは完成していますが、出力を改善できます。また、関数に追加の引数を渡して、2つのノード間のパスのみを返すこともできます。

関数宣言に`target`を3番目のパラメータとして追加し、デフォルト値を空の文字列にしてください。

# --hints--

関数は3つのパラメータを受け取る必要があります：`graph`、`start`、および`target`です。順番が重要です。

```js
({ test: () => assert.match(code, /^def\s+shortest_path\s*\(\s*graph\s*,\s*start\s*,\s*target\s*=?\s*.*\s*\)\s*:/m) })
```

`target`パラメータはデフォルト値として空の文字列を持つ必要があります。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == "(graph, start, target='')"
  `))
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
--fcc-editable-region--
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
