---
id: 65577739f57ecca6c39bb4e9
title: ステップ 33
challengeType: 20
dashedName: step-33
---

# --description--

`list()` 型コンストラクタを使うと、イテラブルからリストを作成できます。

`unvisited` 変数の代入を `list()` を使うように変更し、イテラブルとして `graph` を渡してください。

# --hints--

`list()` 辞書からリストを生成するには `graph` を使うべきです。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/list\s*\(\s*graph\s*\)/));
  }
})
```

`list(graph)` に `unvisited` を代入してください。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\s*=\s*list\s*\(\s*graph\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
