---
id: 65578d4fc3afc3b8f554c882
title: ステップ 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` はキーワード専用の引数も受け取ります。`min()` に関数を追加の引数として渡すことで、リストの項目の比較方法を変更できます。

前のステップで書いた行の結果は、アルファベット順で最初に来るノードでした。代わりに、開始ノードからの距離が最も小さい未訪問のノードを選びたいです。

`key=distances.get` を `min()` の呼び出しの2番目の引数として渡してください。こうすることで、比較は `unvisited` 辞書内の各 `distances` リスト項目の値に基づいて行われます。

# --hints--

`key=distances.get` を `min()` の呼び出しの2番目の引数として渡す必要があります。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
