---
id: 65577236b056379d5dbc7000
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

これから、新しいグラフ内の各ノード間の最短経路を計算するアルゴリズムの開発を始めます。

空の関数を `shortest_path` という名前で宣言してください。関数本体は `pass` キーワードを使って埋めてください。

# --hints--

`shortest_path` という名前の関数を定義してください。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
