---
id: 655771d889132f9ccd341060
title: ステップ 20
challengeType: 20
dashedName: step-20
---

# --description--

グラフの辺に距離や時間、その他の定量的な値を表す重みが付いている場合、そのグラフは*重み付き*グラフと呼ばれます。

今回の場合、これらの重みは各ノード、つまり空間上の点間の距離になります。重み付きグラフを表現するには、辞書の値をタプルのリストに変更できます。

タプルの最初の要素は接続されているノードで、2番目の要素は距離を示す整数になります。

`my_graph['A']`をタプルのリストに変更してください。`A-B`の距離は`3`で、`A-D`の距離は`1`であることを考慮してください。

# --hints--

`my_graph["A"]`は`('B', 3)`と`('D', 1)`というタプルを含むリストである必要があります。

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph`は`'A'`、`'B'`、`'C'`、`'D'`という4つのキーを持つ必要があります。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
