---
id: 655777060d8ddea6741be1b1
title: ステップ 32
challengeType: 20
dashedName: step-32
---

# --description--

`distances` のすべての距離は、開始ノードを除いて無限大に設定されています。`unvisited` リストにはグラフ内のすべてのノードが含まれています。しかし実際には、この結果を得るために `for` ループは必要ありません。

`for` ループとその本体全体を削除してください。

# --hints--

`for` ループとそのネストされたコードすべてを削除してください。

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)distances\s*=\s*\{\s*\}\s*\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
