---
id: 6557716aadbd2d9c42c0e69a
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

最後にもう1つのノード、`'D'`を追加してください。これは`'A'`と`'C'`に接続されています。

この構造を表すように辞書を修正してください。複数の接続を表すには、再びリストを使ってください。

# --hints--

辞書には`'A'`、`'B'`、`'C'`、`'D'`という4つのキーが必要です。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']`はリストであるべきです。

```js
({ test: () => assert(runPython(`
    type(my_graph["A"]) is list
  `))
})
```

`my_graph['A']`は`'B'`と`'D'`を含むリストであるべきです。

```js
({ test: () => assert(runPython(`
    len(my_graph["A"]) == 2 and "B" in my_graph["A"] and "D" in my_graph["A"]
  `))
})
```

`my_graph['B']`はリストであるべきです。

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']`は`'A'`と`'C'`を含むリストであるべきです。

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']`はリストであるべきです。

```js
({ test: () => assert(runPython(`
    type(my_graph["C"]) is list
  `))
})
```

`my_graph['C']`は`'B'`と`'D'`を含むリストであるべきです。

```js
({ test: () => assert(runPython(`
    len(my_graph["C"]) == 2 and "B" in my_graph["C"] and "D" in my_graph["C"]
  `))
})
```

`my_graph['D']`はリストであるべきです。

```js
({ test: () => assert(runPython(`
    type(my_graph["D"]) is list
  `))
})
```

`my_graph['D']`は`'A'`と`'C'`を含むリストであるべきです。

```js
({ test: () => assert(runPython(`
    len(my_graph["D"]) == 2 and "A" in my_graph["D"] and "C" in my_graph["D"]
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': ['A', 'C'],
    'C': 'B'
}
--fcc-editable-region--
```
