---
id: 6557709b0aee699a6a00528c
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

例えば、グラフは空間内の2つの点、`A` と `B` をパスでつなぐ様子を表すために使えます。このようなグラフは、エッジでつながれた2つのノードで構成されます。

既存の `'species'` キーを文字列 `'A'` に置き換えてください。次に、対応する値を文字列 `'B'` に置き換えて、`'A'` と `'B'` のノード間の接続を表してください。

# --hints--

辞書には `'A'` キーが必要です。

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

`'A'` キーの値は `'B'` にしてください。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

辞書には2つのキーが必要です。

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
