---
id: 655773b0591c5f9f4045883e
title: ステップ 24
challengeType: 20
dashedName: step-24
---

# --description--

アルゴリズムは指定されたノードから開始します。次に、グラフを探索して、開始ノード（または*source*）と他のすべてのノード間の最短経路を見つけます。

そのために、関数には`graph`と`start`という2つのパラメータが必要です。これらを関数の宣言に追加してください。

# --hints--

関数は`graph`と`start`をこの順番でパラメータとして受け取る必要があります。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
