---
id: 655773f8b8b5db9fc6d0ae76
title: ステップ 25
challengeType: 20
dashedName: step-25
---

# --description--

訪問済みのノードを追跡するには、グラフ内のすべてのノードのリストが必要です。ノードが訪問されると、そのリストから削除されます。

さて、`pass` キーワードを `unvisited` という名前の変数に置き換え、空のリストを代入してください。

# --hints--

`unvisited` 関数の中に `shortest_path` という変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

`unvisited` 変数に空のリストを代入してください。`pass` は削除することを忘れないでください。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
