---
id: 65576ff7888f9e96f52a4be1
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

辞書はデータを*キー*と*値*のペアの形で保存します。キーは対応する値とコロンで区切られます。そして各キーと値のペアは、次のペアとカンマで区切られます。

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

辞書に新しいキーと値のペアを追加してください。キーには文字列の`'species'`を、値には文字列の`'guinea pig'`を使ってください。

# --hints--

`'species'`辞書の中に、キーが`'guinea pig'`で値が`copper`のペアがあるはずです。

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

`copper`辞書には1つのキーと値のペアだけがあるはずです。

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
