---
id: 65661b72d6745ebec6a96923
title: ステップ 22
challengeType: 20
dashedName: step-22
---

# --description--

同様に、残りの2つのリストも、`C-D` の距離が `7` であることを考慮して修正してください。

# --hints--

`my_graph['C']` は `('B', 4)` と `('D', 7)` のタプルを含むリストである必要があります。

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

`my_graph['D']` は `('A', 1)` と `('C', 7)` のタプルを含むリストである必要があります。

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

`my_graph` は `'A'`、`'B'`、`'C'`、`'D'` という4つのキーを持つ必要があります。

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` は `('B', 3)` と `('D', 1)` のタプルを含むリストである必要があります。

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph['B']` は `('A', 3)` と `('C', 4)` のタプルを含むリストである必要があります。

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
