---
id: 6566195b0a021bb660b2b4b1
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

`my_graph['B']` をタプルのリストに変更してください。タプルの最初の要素は接続されたノードで、2番目の要素は距離です。`B-C` の距離は `4` です。

# --hints--

`my_graph['B']` は `('A', 3)` と `('C', 4)` のタプルを含むリストである必要があります。

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph` は `'A'`、`'B'`、`'C'`、`'D'` という名前の4つのキーを持つ必要があります。

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
