---
id: 6578b13757611e2825beb8a5
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

辞書の中ではキーは一意でなければならず、かつ不変のデータ型のみ使用できます。つまり、リストや別の辞書をキーに使うことはできません。

辞書にもう一つのキー `'age'` を追加し、その値として整数 `2` を設定してください。

# --hints--

`'age'` 辞書の中で、`2` のすぐ後に値が `'species': 'guinea pig'` の `copper` キーがあるはずです。カンマを忘れないでください。

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

`copper` 辞書にはキーと値のペアが2つあるはずです。

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
