---
id: 6578b57361f2f132a02e2a18
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

辞書に格納されたデータにはキーを使ってアクセスできます。

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

辞書の後に、上の例に従って `'species'` の `copper` キーにアクセスし、その結果を出力してください。

# --hints--

辞書の代入は変更しないでください。

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

`copper['species']` キーの値にアクセスするには `'species'` を使ってください。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

`print()` を引数として渡し、`copper['species']` を呼び出してください。

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
