---
id: 657891ab9c1903f4e55433ba
title: ステップ 44
challengeType: 20
dashedName: step-44
---

# --description--

新しい `if` の条件が true のとき、隣接ノードへのより短い経路が見つかっています。

新しい `if` ブロック内で、`pass` を削除し、隣接ノードの距離を隣接ノードの距離と `current` の距離の合計に代入してください。

# --hints--

新しい `distance + distances[current]` 内で `distances[node]` に `if` を代入する必要があります。

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+distance\s*\+\s*distances\s*\[\s*current\s*\]\s*<\s*distances\s*\[\s*node\s*\]\s*/);
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*distance\s*\+\s*distances\s*\[\s*current\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)

    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
--fcc-editable-region--
            if distance + distances[current] < distances[node]:
                pass
--fcc-editable-region--
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
