---
id: 65789506b30453080f77470c
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

これまでに、さまざまなデータ型に触れてきました：

- 整数、文字列、タプル、ブール値のような不変のデータ型
- リストや辞書のような可変のデータ型

辞書は波括弧、中括弧のペア `{}` で識別されます。

まず、`copper` という変数を作成し、空の辞書を波括弧、中括弧のペアで代入してください。これは、空のリストを大括弧のペアで作成するのと同じ方法です。

# --hints--

`copper` という変数を用意してください。

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

`copper` 変数には空の辞書の値を代入してください。そのために波括弧、中括弧のペアを使います。

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
