---
id: 6579ca0923cfa7162089d2f0
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

同じ構文を使って、既存のキーの値を変更できます。

`print()` を呼び出す直前に、`'species'` キーにアクセスして、その値を `'Cavia porcellus'` に再代入してください。

# --hints--

`copper['species']` を呼び出す前に、`'Cavia porcellus'` に文字列 `print()` を再代入する必要があります。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
