---
id: 6579dd49fa8a8e1fd06b85a9
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

グラフは、要素のペア間の関係を表すデータ構造です。
これらの要素は*ノード*と呼ばれ、実際の物体、エンティティ、空間上の点などを表すことができます。ノード間の接続は*エッジ*と呼ばれます。

以下はグラフの視覚的な表現です：

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


`copper` 辞書の名前を `my_graph` に変更してください。これがアルゴリズムをテストするグラフを表します。

# --hints--

`copper` 辞書の名前を `my_graph` に変更してください。

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

`my_graph` 変数は辞書である必要があります。

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
