---
id: 6579cbab9825b8170974c69a
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

辞書からキーと値のペアを削除するには、`del` キーワードを使います。

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

`for` ループの直前に、`del` キーワードを使って `'age'` キーとその値を `copper` から削除してください。

# --hints--

`del` ループの前に `copper['age']` を削除するには、`for` キーワードを使う必要があります。

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
