---
id: 6579cd5f6dd62c189e53ddbb
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

辞書の基本的な側面をレビューしたので、最短経路アルゴリズムの構築に進みましょう。

`copper` 辞書の宣言の後のすべてのコード行を削除してください。

# --hints--

辞書の宣言の後のすべてのコード行を削除する必要があります。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
