---
id: 65796fac81f983127558f3f4
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

今の`print()`の呼び出しを修正して、`age`キーの値を出力してください。


# --hints--

`copper['age']`キーの値にアクセスするには`age`を使うべきです。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

`print()`を引数として渡し、`copper['age']`を呼び出すべきです。

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
