---
id: 6579717f0920131304286804
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

辞書を宣言した後に新しいキーと値のペアを追加するには、既存のキーにアクセスするのと同じ方法でキーを指定し、代入演算子を使って新しいキーの値を設定できます。

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

`print()` の呼び出しを削除してください。その後、`copper` の宣言の後に、辞書にキー `'food'` を追加し、その値を `'hay'` に設定してください。

# --hints--

コードに `print(copper['age'])` が含まれていてはいけません。

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

辞書を宣言した後に、キー `'food'` を `copper` に追加してください。

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

辞書を宣言した後に、`copper['food']` を `'hay'` に設定してください。

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
