---
id: 65affe602a2a828b0382a058
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

`'age'` キーを文字列 `'B'` に置き換え、その値を文字列 `'A'` に設定して、両方向のノード間の接続を表してください。

# --hints--

辞書には `'B'` キーが必要です。

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

`'B'` キーの値は `'A'` であるべきです。

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

辞書には `'A'` キーが必要です。

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

`'A'` キーの値は `'B'` であるべきです。

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

辞書には2つのキーが必要です。

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
