---
id: 65b7cd2b7bd9a684ccf1dc16
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

辞書のキーを繰り返すには、辞書をそのまま`for`ループに入れればいいです。以下のコードは、辞書`dict`の各キーを出力します。

```py
for i in dict:
   print(i)
```

`print()`の呼び出しを、`for`ループに置き換えて、`copper`を繰り返し処理し、各キーを出力してください。

# --hints--

コードに`print(copper)`があってはいけません。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

`for`ループを作成して、`copper`辞書を繰り返してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

`copper`ループの中で、`for`の各キーを出力してください。

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
