---
id: 65b7d3319ebec69b983fb91b
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

出力からわかるように、`.items()` は各キーと値のペアを別々のタプルにストアするデータ構造を作成します。これらのタプルの要素を繰り返すには、2つ目のループ変数を追加できます。

```py
for i, j in dict.items():
    print(i, j)
```

`for` のループを修正して、2つのループ変数を受け取り、ループ内で両方を出力してください。

# --hints--

`for` のループを2つの変数を使うように修正してください。

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

`for` のループ内で両方のループ変数を出力してください。

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
