---
id: 65b7cf140d34058d7ea8935f
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

辞書のキーの値を繰り返したい場合、1つの方法は`.values()`メソッドを使うことです。

`for`ループを`copper.values()`ではなく`copper`を繰り返すように変更して、出力を確認してください。

# --hints--

`for`ループを`copper.values()`ではなく`copper`を繰り返すように変更してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
