---
id: 65b7d25992879698180e6a71
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

最後に、キーと値のペアを繰り返すことができるようにしたい場合は、`.items()` メソッドを使えます。

`for` のループを修正して、`copper.items()` の代わりに `copper.values()` を繰り返すようにしてください。

# --hints--

`for` のループを修正して、`copper.items()` の代わりに `copper.values()` を繰り返すようにしてください。

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
