---
id: ab306dbdcc907c7ddfc30830
title: Crear una herramienta de aplanamiento profundo
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

En este laboratorio, estarás implementando un algoritmo para aplanar un array.

Aplanar un array significa convertir un array anidado de cualquier profundidad en un solo array unidimensional. El proceso extrae todos los elementos en orden, desempaquetando solo arrays. Otros tipos permanecen sin cambios.

Por ejemplo:

| Original | Flattened |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función llamada `steamrollArray`.
1. La función `steamrollArray` debe aceptar un argumento: un arreglo anidado.
1. La función debe aplanar el array anidado, teniendo en cuenta los diferentes niveles de anidación.
1. Tu solución no debe utilizar los métodos `Array.prototype.flat()` o `Array.prototype.flatMap()`.
1. Las variables globales no deben ser usadas.

# --hints--

Deberías tener una función `steamrollArray`.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` debe devolver `["a", "b"]`.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` debe devolver `[1, 2, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` debe devolver `[1, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` debe devolver `[1, {}, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

Tu solución no debe utilizar los métodos `Array.prototype.flat()` o `Array.prototype.flatMap()`.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

No debes usar variables globales.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
