---
id: 66bba6fff611169359d9d36a
title: Jenga ukurasa wa barua pepe wa siri
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kipengele cha `main` chenye `id` ya `email`.
1. Kipengele chako cha `#email` kinapaswa kuwa na `padding` ya `50px`, ukingo wa juu wa `50px`, `width` ya `500px`, na `border` ambayo ni `2px`.
1. Upana wa jumla wa kipengele chako cha `#email`, ikiwa ni pamoja na nafasi za ndani na mipaka, unapaswa kuwa `500px`.
1. Unapaswa kuwa na vipengele viwili vya `div`, kimoja chenye kitambulisho cha `confidential` na kingine chenye kitambulisho cha `top-secret`, ndani ya kipengele chako cha `#email`.
1. Vipengele vyako vya `#confidential` na `#top-secret` vinapaswa kuwa na `display` ya `inline-block`.
1. Vipengele vyako vya `#confidential` na `#top-secret` vinapaswa kuwa na `padding`, ukingo wa kushoto, na `border`.
1. Kipengele cha `#confidential` kinapaswa kuwa na maandishi ya `CONFIDENTIAL`.
1. Kipengele cha `#top-secret` kinapaswa kuwa na maandishi ya `TOP SECRET`.
1. Vipengele vyako vya `#confidential` na `#top-secret` vinapaswa kuzungushwa kwa kutumia mabadiliko ya CSS.
1. Unapaswa kuwa na angalau vipengele vitatu vya aya ndani ya kipengele chako cha `#email`.
1. Unapaswa kuwa na angalau vipengele vitatu vya `span` vyenye darasa la vitu la `blurred`, ndani ya vipengele vyako vya aya.
1. Unapaswa kuwa na kichaguzi cha `blurred` kinachofanya kipengele cha `3px` kuwa na mng'ao mdogo kwa kutumia kichujio cha CSS.

# --hints--

Unapaswa kuwa na kipengele cha `main` chenye `id` ya `email`.

```js
assert.exists(document.querySelector('main#email'));
```

Unapaswa kuwa na kichaguzi cha `#email` kinachoweka `padding` ya vipengele vyake pande zote kuwa `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

Unapaswa kuwa na kichaguzi cha `#email` kinachoweka `margin-top` ya vipengele vyake kuwa `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

Unapaswa kuwa na kichaguzi cha `#email` kinachoweka `width` ya vipengele vyake kuwa `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

Kipengele chako cha `#email` kinapaswa kuwa na mpaka wa `2px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

Kipengele chako cha `#email` kinapaswa kuwa na `box-sizing` ya `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

Unapaswa kuwa na kipengele cha `div` chenye `id` ya `confidential` ndani ya kipengele chako cha `#email`.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

Unapaswa kuwa na kipengele cha `div` chenye `id` ya `top-secret` ndani ya kipengele chako cha `#email`.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

Kipengele chako cha `#confidential` kinapaswa kuwa na `display` iliyowekwa kuwa `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

Kipengele chako cha `#top-secret` kinapaswa kuwa na `display` iliyowekwa kuwa `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

Kipengele chako cha `#confidential` kinapaswa kuwa na `padding` pande zote.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Kipengele chako cha `#top-secret` kinapaswa kuwa na `padding` pande zote.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Kipengele chako cha `#confidential` kinapaswa kuwa na ukingo wa kushoto.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

Kipengele chako cha `#top-secret` kinapaswa kuwa na ukingo wa kushoto.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

Kipengele chako cha `#confidential` kinapaswa kuwa na `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

Kipengele chako cha `#top-secret` kinapaswa kuwa na `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

Maandishi ya kipengele chako cha `#confidential` yanapaswa kuwa `CONFIDENTIAL`. Umeacha maandishi au kuna makosa ya tahajia.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

Maandishi ya kipengele chako cha `#top-secret` yanapaswa kuwa `TOP SECRET`. Umeacha maandishi au kuna makosa ya tahajia.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

Kipengele chako cha `#confidential` kinapaswa kutumia mabadiliko ya CSS ya `transform` kuzungusha kipengele.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

Kipengele chako cha `#top-secret` kinapaswa kutumia mabadiliko ya CSS ya `transform` kuzungusha kipengele.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

Unapaswa kuwa na angalau vipengele vitatu vya aya ndani ya kipengele chako cha `#email`.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

Unapaswa kuwa na angalau vipengele vitatu vya `span` vyenye darasa la vitu la `blurred` ndani ya aya zako.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

Vipengele vyako vya `.blurred` havipaswi kuwa tupu.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

Unapaswa kuwa na kichaguzi cha `.blurred` kinachoweka `filter` ya vipengele vyake kuwa `blur(3px)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
