---
id: 68ce56d54d3cad1149a1c4cf
title: Eine Seite zur Funktionsauswahl gestalten
challengeType: 25
dashedName: design-a-feature-selection-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Ihre Seite sollte ein `h1`-Element mit dem Text `Feature Selection` enthalten.
2. Ihre Seite sollte ein `div`-Element mit der Klasse `feature-card-container` enthalten.
3. Ihre Seite sollte mindestens zwei `label`-Elemente enthalten, die jeweils die Klasse `feature-card` innerhalb des `div` haben.
4. Jedes Ihrer `label`-Elemente sollte Bezeichnungstext und ein `input`-Element vom Typ `checkbox` enthalten.
5. Erstellen Sie einen Selektor, der die Kontrollkästchen anspricht, und wenden Sie die folgende Gestaltung an:
    - Alle Ihre Kontrollkästchen-Elemente sollten in Ihrem CSS auf `appearance: none;` gesetzt sein.
    - Alle Ihre Kontrollkästchen-Elemente sollten eine Rahmenbreite, -farbe und -stil Ihrer Wahl haben.
6. Wenn das Kontrollkästchen aktiviert ist, sollte in der Mitte des Kontrollkästchens ein Häkchen `✓` angezeigt werden.
7. Wenn das Kontrollkästchen aktiviert ist, sollte sich die Hintergrundfarbe des Kontrollkästchens in eine Farbe Ihrer Wahl ändern.
8. Wenn das Kontrollkästchen aktiviert ist, sollte sich die Rahmenfarbe des Kontrollkästchens in eine Farbe Ihrer Wahl ändern.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihre CSS-Datei in Ihrem HTML verlinken.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
assert.exists(document.querySelector('h1'));
```

Ihr `h1` sollte den Text `Feature Selection` enthalten.

```js
assert.equal(document.querySelector('h1').innerText.trim(), `Feature Selection`);
```

Sie sollten ein `div`-Element mit der Klasse `feature-card-container` haben.

```js
assert.exists(document.querySelector('div.feature-card-container'));
```

Sie sollten mindestens zwei `label`-Elemente innerhalb von `.feature-card-container` haben.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isAtLeast(labels.length, 2);
```

Jedes `label`-Element innerhalb von `.feature-card-container` sollte die Klasse `feature-card` haben.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isTrue(label.classList.contains('feature-card')));
```

Innerhalb jedes `label`-Elements sollten Sie Bezeichnungstext haben.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isNotEmpty(label.innerText.trim()));
```

Innerhalb jedes label-Elements sollte ein `input`-Element vom Typ `checkbox` sein.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.exists(label.querySelector('input[type="checkbox"]')));
```

Sie sollten einen Selektor für die Kontrollkästchen haben.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Ihre Kontrollkästchen sollten auf `appearance: none` gesetzt sein.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.equal(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('appearance'), 'none');
```

Ihre Kontrollkästchen sollten eine Rahmenbreite Ihrer Wahl haben.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-width'));
```

Ihre Kontrollkästchen sollten eine Rahmenfarbe Ihrer Wahl haben.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-color'));
```

Ihre Kontrollkästchen sollten einen Rahmenstil Ihrer Wahl haben.

```js
const selectors = [  
  '.feature-card [type="checkbox"]',
  '.feature-card input[type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input', // there aren't other input elements, so this is a valid approach
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-style'));
```

Wenn die Kontrollkästchen ausgewählt sind, sollte in der Mitte des Kontrollkästchens ein Häkchen `✓` sein.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;

    checkbox.checked = true;
    const after = window.getComputedStyle(checkbox, ':after');
    assert.equal(after.content, '"✓"');
});
```

Wenn das Kontrollkästchen aktiviert ist, sollte sich die Hintergrundfarbe des Kontrollkästchens in eine Farbe Ihrer Wahl ändern.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBG = window.getComputedStyle(checkbox).backgroundColor;
    checkbox.checked = true;
    const checkedBG = window.getComputedStyle(checkbox).backgroundColor;

    assert.notEqual(checkedBG, uncheckedBG);
});
```

Wenn das Kontrollkästchen aktiviert ist, sollte sich die Rahmenfarbe des Kontrollkästchens in eine Farbe Ihrer Wahl ändern.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBC = window.getComputedStyle(checkbox).borderColor;
    checkbox.checked = true;
    const checkedBC = window.getComputedStyle(checkbox).borderColor;

    assert.notEqual(checkedBC, uncheckedBC);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Selection Feature Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Feature Selection</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Feature Selection</h1>
    <div class="feature-card-container">
      <label class="feature-card">
        <input type="checkbox" />
        <h4>Cloud Storage</h4>
        <p>100 Gigabyte secure storage</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Dedicated Support</h4>
        <p>24/7 customer help</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Advanced Analytics</h4>
        <p>Insights & reports</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Custom User Themes</h4>
        <p>Personalized dashboard design</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Multi-User Collab</h4>
        <p>Team access and sharing</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>API Access</h4>
        <p>Integrate with your custom built tools</p>
      </label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  flex-direction: column;
  height: 100vh;
}

.feature-card-container {
  display: flex;
  flex-wrap: wrap;
  gap: 20px;
  max-width: 600px;
  justify-content: center;
}

.feature-card-container > * {
  flex: 1 1 180px;
}

.feature-card {
  border: 2px solid #ccc;
  border-radius: 12px;
  padding: 20px;
  text-align: center;
  position: relative;
  cursor: pointer;
}

.feature-card:hover {
  border-color: #f1be32;
  box-shadow: 0 4px 12px rgba(255, 201, 53, 0.6);
}

.feature-card input[type='checkbox'] {
  position: absolute;
  top: 5px;
  right: 5px;
  width: 24px;
  height: 24px;
  cursor: pointer;
  appearance: none;
  -webkit-appearance: none;
  border: 2px solid #ccc;
  border-radius: 4px;
  background-color: white;
}

.feature-card input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #f1be32;
}

.feature-card input[type='checkbox']:checked::after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  line-height: 20px;
}

.feature-card h4 {
  margin: 10px 0 5px;
  font-size: 1.2rem;
}

@media (max-width: 600px) {
  body {
    height: auto;
  }
}
```
