---
id: 587d824a367417b2b2512c43
title: 개인 도서관
challengeType: 4
forumTopicId: 301571
dashedName: personal-library
---

# --description--

다음과 기능적으로 유사한 풀스택 JavaScript 앱을 만드세요: <a href="https://personal-library.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://personal-library.freecodecamp.rocks/</a>. 이 프로젝트를 진행하려면 다음 방법 중 하나를 사용해 코드를 작성해야 합니다:

-   <a href="https://github.com/freeCodeCamp/boilerplate-project-library" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 프로젝트를 완성하세요.
-   원하는 사이트 빌더를 사용해 프로젝트를 완성하세요. 반드시 GitHub 저장소의 모든 파일을 포함해야 합니다.

# --instructions--

1.  MongoDB 연결 문자열을 `.env`에 따옴표 없이 `DB`로 추가하세요  
    예시: `DB=mongodb://admin:pass@1234.mlab.com:1234/fccpersonallib`
2.  `.env` 파일에서 `NODE_ENV`를 따옴표 없이 `test`로 설정하세요
3.  모든 라우트를 `routes/api.js` 내에 생성해야 합니다
4.  모든 기능 테스트는 `tests/2_functional-tests.js`에서 작성하세요

# --hints--

예시 URL 대신 자신만의 프로젝트를 제공할 수 있습니다.

```js
  assert(
    !/.*\/personal-library\.freecodecamp\.rocks/.test(code)
  );
```

`/api/books`에 `title`를 폼 데이터로 포함해 <b>POST</b> 요청을 보내 책을 추가할 수 있습니다. 반환되는 응답은 `title`과 고유한 `_id`를 키로 가진 객체입니다. 요청에 `title`가 포함되지 않으면 반환 응답은 문자열 `missing required field title`이어야 합니다.

```js
  try {
    const response1 = await fetch(code + '/api/books', {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book 1' })
    });
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    let data1 = await response1.json();
    assert.isObject(data1);
    assert.property(data1, 'title');
    assert.equal(data1.title, 'Faux Book 1');
    assert.property(data1, '_id');
    const response2 = await fetch(code + '/api/books', {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({})
    });
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    let data2 = await response2.text();
    assert.isString(data2);
    assert.equal(data2, 'missing required field title');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`/api/books`에 <b>GET</b> 요청을 보내면 모든 책을 나타내는 JSON 응답을 받을 수 있습니다. JSON 응답은 각 객체(책)가 `title`, `_id`, `commentcount` 속성을 가진 배열입니다.

```js
  try {
    let url = code + '/api/books';
    let a = fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book A' })
    }).then(response => {
      if (!response.ok) throw Error(response.text());
      return response.json();
    });
    let b = fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book B' })
    }).then(response => {
      if (!response.ok) throw Error(response.text());
      return response.json();
    });
    let c = fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book C' })
    }).then(response => {
      if (!response.ok) throw Error(response.text());
      return response.json();
    });
    await Promise.all([a, b, c]).then(async () => {
      const response = await fetch(url);
      if (!response.ok) {
        throw Error(await response.text());
      }
      let data = await response.json();
      assert.isArray(data);
      assert.isAtLeast(data.length, 3);
      data.forEach((book) => {
        assert.isObject(book);
        assert.property(book, 'title');
        assert.isString(book.title);
        assert.property(book, '_id');
        assert.property(book, 'commentcount');
        assert.isNumber(book.commentcount);
      });
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`/api/books/{_id}`에 <b>GET</b> 요청을 보내면 `title`, `_id`, 그리고 `comments` 배열(주석이 없으면 빈 배열)을 포함하는 단일 책 객체를 가져올 수 있습니다. 책을 찾지 못하면 문자열 `no book exists`를 반환하세요.

```js
  try {
    let url = code + '/api/books';
    const noBookResponse = await fetch(url + '/5f665eb46e296f6b9b6a504d');
    if (!noBookResponse.ok) {
      throw Error(await noBookResponse.text());
    }
    let noBook = await noBookResponse.text();
    assert.isString(noBook);
    assert.equal(noBook, 'no book exists');
    const createResponse = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Faux Book Alpha' })
    });
    if (!createResponse.ok) {
      throw Error(await createResponse.text());
    }
    let sampleBook = await createResponse.json();
    assert.isObject(sampleBook);
    let bookId = sampleBook._id;
    const queryResponse = await fetch(url + '/' + bookId);
    if (!queryResponse.ok) {
      throw Error(await queryResponse.text());
    }
    let bookQuery = await queryResponse.json();
    assert.isObject(bookQuery);
    assert.property(bookQuery, 'title');
    assert.equal(bookQuery.title, 'Faux Book Alpha');
    assert.property(bookQuery, 'comments');
    assert.isArray(bookQuery.comments);
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`comment`에 `/api/books/{_id}`를 폼 본문 데이터로 포함한 <b>POST</b> 요청을 보내 책에 주석을 추가할 수 있습니다. 반환 응답은 앞선 테스트의 <b>GET</b> `/api/books/{_id}` 요청과 유사한 책 객체입니다. 요청에 `comment`가 포함되지 않으면 문자열 `missing required field comment`를 반환하세요. 책을 찾지 못하면 문자열 `no book exists`를 반환하세요.

```js
  try {
    let url = code + '/api/books';
    const createResponse = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Notable Book' })
    });
    if (!createResponse.ok) {
      throw Error(await createResponse.text());
    }
    let commentTarget = await createResponse.json();
    assert.isObject(commentTarget);
    let bookId = commentTarget._id;
    const comment1Response = await fetch(url + '/' + bookId, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ comment: 'This book is fab!' })
    });
    if (!comment1Response.ok) {
      throw Error(await comment1Response.text());
    }
    let bookCom1 = await comment1Response.json();
    const comment2Response = await fetch(url + '/' + bookId, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ comment: 'I did not care for it' })
    });
    if (!comment2Response.ok) {
      throw Error(await comment2Response.text());
    }
    let bookCom2 = await comment2Response.json();
    assert.isObject(bookCom2);
    assert.property(bookCom2, '_id');
    assert.property(bookCom2, 'title');
    assert.property(bookCom2, 'comments');
    assert.lengthOf(bookCom2.comments, 2);
    bookCom2.comments.forEach((comment) => {
      assert.isString(comment);
      assert.oneOf(comment, ['This book is fab!', 'I did not care for it']);
    });
    const commentErrResponse = await fetch(url + '/' + bookId, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({})
    });
    if (!commentErrResponse.ok) {
      throw Error(await commentErrResponse.text());
    }
    let commentErr = await commentErrResponse.text();
    assert.isString(commentErr);
    assert.equal(commentErr, 'missing required field comment');
    const failingCommentResponse = await fetch(url + '/5f665eb46e296f6b9b6a504d', {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ comment: 'Never Seen Comment' })
    });
    if (!failingCommentResponse.ok) {
      throw Error(await failingCommentResponse.text());
    }
    let failingComment = await failingCommentResponse.text();
    assert.isString(failingComment);
    assert.equal(failingComment, 'no book exists');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`/api/books/{_id}`에 <b>DELETE</b> 요청을 보내 컬렉션에서 책을 삭제할 수 있습니다. 성공하면 문자열 `delete successful`를 반환합니다. 책을 찾지 못하면 문자열 `no book exists`를 반환하세요.

```js
  try {
    let url = code + '/api/books';
    const createResponse = await fetch(url, {
      method: 'POST',
      headers: { 'Content-Type': 'application/json' },
      body: JSON.stringify({ title: 'Deletable Book' })
    });
    if (!createResponse.ok) {
      throw Error(await createResponse.text());
    }
    let deleteTarget = await createResponse.json();
    assert.isObject(deleteTarget);
    let bookId = deleteTarget._id;
    const deleteResponse = await fetch(url + '/' + bookId, { method: 'DELETE' });
    if (!deleteResponse.ok) {
      throw Error(await deleteResponse.text());
    }
    let doDelete = await deleteResponse.text();
    assert.isString(doDelete);
    assert.equal(doDelete, 'delete successful');
    const failingDeleteResponse = await fetch(url + '/5f665eb46e296f6b9b6a504d', { method: 'DELETE' });
    if (!failingDeleteResponse.ok) {
      throw Error(await failingDeleteResponse.text());
    }
    let failingDelete = await failingDeleteResponse.text();
    assert.isString(failingDelete);
    assert.equal(failingDelete, 'no book exists');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

`/api/books`에 <b>DELETE</b> 요청을 보내 데이터베이스의 모든 책을 삭제할 수 있습니다. 성공하면 문자열 `complete delete successful`를 반환합니다.

```js
  try {
    const response = await fetch(code + '/api/books', { method: 'DELETE' });
    if (!response.ok) {
      throw Error(await response.text());
    }
    const deleteAll = await response.text();
    assert.isString(deleteAll);
    assert.equal(deleteAll, 'complete delete successful');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

필요한 10개의 기능 테스트가 모두 완료되어 통과했습니다.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    assert.isAtLeast(getTests.length, 10, 'At least 10 tests passed');
    getTests.forEach((test) => {
      assert.equal(test.state, 'passed', 'Test in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```
