---
id: 587d8249367417b2b2512c41
title: 미터-야드 변환기
challengeType: 4
forumTopicId: 301570
dashedName: metric-imperial-converter
---

# --description--

다음과 기능적으로 유사한 풀스택 JavaScript 앱을 만드세요: <a href="https://metric-imperial-converter.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://metric-imperial-converter.freecodecamp.rocks/</a>. 이 프로젝트를 진행하려면 다음 방법 중 하나를 사용해 코드를 작성해야 합니다:

- <a href="https://github.com/freeCodeCamp/boilerplate-project-metricimpconverter/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 프로젝트를 완성하세요.
- 원하는 사이트 빌더를 사용해 프로젝트를 완성하세요. GitHub 저장소의 모든 파일을 반드시 포함해야 합니다.

**참고:** 이 프로젝트의 테스트는 `glitch.com`를 사용할 때 작동하지 않습니다.

# --instructions--

- `/controllers/convertHandler.js`에서 필요한 변환 로직을 완성하세요.
- `/routes/api.js`에서 필요한 라우트를 완성하세요.
- `sample.env` 파일을 `.env`로 복사하고 변수를 적절히 설정하세요.
- 테스트를 자동으로 실행하려면 `NODE_ENV=test` 파일에 `.env`를 추가하세요.
- 콘솔에서 테스트를 실행하려면 `npm run test` 명령어를 사용하세요.

`tests/1_unit-tests.js`에 다음 테스트를 작성하세요:

- `convertHandler`는 정수 입력을 올바르게 읽어야 합니다.
- `convertHandler`는 소수 입력을 올바르게 읽어야 합니다.
- `convertHandler`는 분수 입력을 올바르게 읽어야 합니다.
- `convertHandler`는 소수가 포함된 분수 입력을 올바르게 읽어야 합니다.
- `convertHandler`는 이중 분수(예: `3/2/3`)에 대해 올바르게 오류를 반환해야 합니다.
- `convertHandler`는 숫자 입력이 없을 때 `1`의 숫자 입력을 기본값으로 올바르게 처리해야 합니다.
- `convertHandler`는 각 유효한 입력 단위를 올바르게 읽어야 합니다.
- `convertHandler`는 유효하지 않은 입력 단위에 대해 올바르게 오류를 반환해야 합니다.
- `convertHandler`는 각 유효한 입력 단위에 대해 올바른 반환 단위를 반환해야 합니다.
- `convertHandler`는 각 유효한 입력 단위에 대해 철자화된 문자열 단위를 올바르게 반환해야 합니다.
- `convertHandler`는 `gal`를 `L`으로 올바르게 변환해야 합니다.
- `convertHandler`는 `L`를 `gal`으로 올바르게 변환해야 합니다.
- `convertHandler`는 `mi`를 `km`으로 올바르게 변환해야 합니다.
- `convertHandler`는 `km`를 `mi`으로 올바르게 변환해야 합니다.
- `convertHandler`는 `lbs`를 `kg`으로 올바르게 변환해야 합니다.
- `convertHandler`는 `kg`를 `lbs`으로 올바르게 변환해야 합니다.

`tests/2_functional-tests.js`에 다음 테스트를 작성하세요:

- `10L` 같은 유효한 입력을 변환하세요: `GET` 요청을 `/api/convert`로 변환합니다.
- `32g` 같은 유효하지 않은 입력을 변환하세요: `GET` 요청을 `/api/convert`로 변환합니다.
- `3/7.2/4kg` 같은 유효하지 않은 숫자를 변환하세요: `GET` 요청을 `/api/convert`로 변환합니다.
- `3/7.2/4kilomegagram` 같은 유효하지 않은 숫자와 단위를 변환하세요: `GET` 요청을 `/api/convert`로 변환합니다.
- `kg` 같은 숫자 없이 변환하세요: `GET` 요청을 `/api/convert`로 변환합니다.

# --hints--

자신만의 프로젝트를 제공할 수 있으며, 예시 URL을 사용할 필요는 없습니다.

```js
  assert(
    !/.*\/metric-imperial-converter\.freecodecamp\.rocks/.test(
      code
    )
  );
```

하나의 매개변수로 허용된 숫자와 단위를 포함하는 `GET`를 `/api/convert`할 수 있습니다. (힌트: 단위 시작을 나타내는 첫 번째 문자의 인덱스를 찾아 입력을 분리하세요)

```js

```

`'gal'`를 `'L'`로, 그 반대로 변환할 수 있습니다. (1 gal은 3.78541 L)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 3.78541);
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10gal');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 37.8541);
    assert.equal(data2.returnUnit, 'L');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.26417);
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10l');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 2.64172);
    assert.equal(data4.returnUnit, 'gal');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

`'lbs'`를 `'kg'`로, 그 반대로 변환할 수 있습니다. (1 lbs는 0.453592 kg)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1lbs');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 0.45359);
    assert.equal(data1.returnUnit, 'kg');
    const response2 = await fetch(code + '/api/convert?input=10lbs');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 4.53592);
    assert.equal(data2.returnUnit, 'kg');
    const response3 = await fetch(code + '/api/convert?input=1kg');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 2.20462);
    assert.equal(data3.returnUnit, 'lbs');
    const response4 = await fetch(code + '/api/convert?input=10kg');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 22.04624);
    assert.equal(data4.returnUnit, 'lbs');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

`'mi'`를 `'km'`로, 그 반대로 변환할 수 있습니다. (1 mi는 1.60934 km)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 1.60934);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=10mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 16.0934);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(code + '/api/convert?input=1km');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.62137);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(code + '/api/convert?input=10km');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 6.21373);
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

모든 입력 단위는 대소문자 구분 없이 허용되어야 하며, 반환 시에는 `initUnit`와 `returnUnit` 모두 소문자로 반환되어야 합니다. 단, liter는 대문자 `'L'`로 나타내야 합니다.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.initUnit, 'gal');
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10L');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.initUnit, 'L');
    assert.equal(data2.returnUnit, 'gal');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.initUnit, 'L');
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10KM');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.initUnit, 'km');
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

측정 단위가 유효하지 않으면 `'invalid unit'`가 반환됩니다.

```js
  try {
    const response = await fetch(code + '/api/convert?input=1min');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid unit' || data === 'invalid unit');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

숫자가 유효하지 않으면 `'invalid number'`가 반환됩니다.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2gal'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid number' || data === 'invalid number');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

단위와 숫자 모두 유효하지 않으면 `'invalid number and unit'`가 반환됩니다.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2min'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(
      data.error === 'invalid number and unit' ||
        data === 'invalid number and unit'
    );
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

매개변수에 분수, 소수 또는 둘 다 사용할 수 있습니다(예: 5, 1/2, 2.5/6). 아무것도 제공되지 않으면 기본값은 1입니다.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.approximately(data1.initNum, 1, 0.001);
    assert.approximately(data1.returnNum, 1.60934, 0.001);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=1/5mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.approximately(data2.initNum, 1 / 5, 0.1);
    assert.approximately(data2.returnNum, 0.32187, 0.001);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(
      code + '/api/convert?input=1.5/7km'
    );
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.approximately(data3.initNum, 1.5 / 7, 0.001);
    assert.approximately(data3.returnNum, 0.13315, 0.001);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(
      code + '/api/convert?input=3/2.7km'
    );
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.approximately(data4.initNum, 3 / 2.7, 0.001);
    assert.approximately(data4.returnNum, 0.69041, 0.001);
    assert.equal(data4.returnUnit, 'mi');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

반환값은 `initNum`, `initUnit`, `returnNum`, `returnUnit`, `string`로 구성되며, 결과는 `'{initNum} {initUnitString} converts to {returnNum} {returnUnitString}'` 형식으로 소수점 다섯 자리까지 반올림됩니다.

```js
  try {
    const response = await fetch(code + '/api/convert?input=2mi');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.json();
    assert.equal(data.initNum, 2);
    assert.equal(data.initUnit, 'mi');
    assert.approximately(data.returnNum, 3.21868, 0.001);
    assert.equal(data.returnUnit, 'km', 'returnUnit did not match');
    assert.equal(data.string, '2 miles converts to 3.21868 kilometers');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

모든 16개 단위 테스트가 완료되어 통과했습니다.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const unitTests = getTests.filter(test => {
      return !!test.context.match(/Unit Tests/gi);
    });
    assert.isAtLeast(unitTests.length, 16, 'At least 16 tests passed');
    unitTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

모든 5개 기능 테스트가 완료되어 통과했습니다.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const functTests = getTests.filter(test => {
      return !!test.context.match(/Functional Tests/gi);
    });
    assert.isAtLeast(functTests.length, 5, 'At least 5 tests passed');
    functTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```
