---
id: 6720b3191bbafd269bae0ae8
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Crea un elemento `div` che abbia un `id` di `greeting-card` e un `class` di `card`.

All'interno dell'elemento `div`, aggiungi un `h1` con il testo `Happy Birthday!`. Poi aggiungi un elemento paragrafo con un `class` chiamato `message` e il testo `Wishing you all the happiness and joy on your special day!`.

# --hints--

Dovresti avere un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

L'elemento `div` dovrebbe avere un `id` di `greeting-card`.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

L'elemento `div` dovrebbe avere un `class` di `card`.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

Dovresti avere un elemento `h1` all'interno dell'elemento `div`.

```js
assert.exists(document.querySelector("div > h1"));
```

L'elemento `h1` dovrebbe contenere il testo `Happy Birthday!`.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

Dovresti avere un elemento `p` all'interno di `div`.

```js
assert.exists(document.querySelector("div > p"));
```

L'elemento `p` dovrebbe avere un `class` di `message`.

```js
assert.exists(document.querySelector("div > p.message"));
```

L'elemento `p` dovrebbe contenere il testo `Wishing you all the happiness and joy on your special day!`.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
