---
id: 6720b7f3155a9278d6b6a8b8
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Ora è il momento di definire lo stile del tuo biglietto di auguri.

Aggiungi un selettore per l'elemento `body`, poi: 

- cambia il `font-family` in `Arial` seguito dal generico `sans-serif`,
- dai una spaziatura interna su tutti i lati di `40px`,
- imposta la proprietà `text-align` su `center`.

# --hints--

Dovresti avere un selettore `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Dovresti impostare `font-family` su `"Arial", sans-serif` nel selettore `body`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

Dovresti impostare `padding` su `40px` nel selettore `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

Dovresti impostare `text-align` su `center` nel selettore `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
