---
id: 6720b1f2680ec706ed534acf
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop, eserciterai l'uso di diverse pseudo-classi e pseudo-elementi progettando un biglietto di auguri. Il codice boilerplate HTML ti è stato fornito.

Inizia il workshop collegando il file `styles.css`.

# --hints--

Il tuo elemento `link` dovrebbe essere all'interno del tuo elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con valore `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con valore `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
