---
id: 672b4434dd6d6a1edc405419
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Aggiungi due elementi `a` all'interno dell'elemento `.card-links`.

Il primo deve avere un testo `Send Card`, un `class` con valore `send-link` e un `href` con valore `#send`.

Il secondo deve avere un testo `Share on Social Media`, un `class` con valore `share-link` e un `href` con valore `#share`.

# --hints--

Il primo elemento `a` all'interno di `.card-links` deve avere un testo `Send Card`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

Il primo elemento `a` deve avere un `class` con valore `send-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

Il primo elemento `a` deve avere un `href` con valore `#send`.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

Il secondo elemento `a` all'interno di `.card-links` deve avere un testo `Share on Social Media`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

Il secondo elemento `a` deve avere un `class` con valore `share-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

Il secondo elemento `a` deve avere un `href` con valore `#share`.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
