---
id: 67332db0f793bf08f2a18528
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

È ora di riempire il secondo `section`!

Aggiungi un elemento `h2` all'elemento `#share` che contiene il testo `Sharing your card...`, poi aggiungi un elemento `p` con il testo `Your card was shared on social media!`.

# --hints--

L'elemento `#share` dovrebbe contenere un elemento `h2`.

```js
assert.exists(document.querySelector("#share > h2"));
```

L'elemento `h2` dovrebbe contenere il testo `Sharing your card...`.

```js
assert.strictEqual(document.querySelector("#share > h2").innerText.trim(), "Sharing your card...");
```

Il secondo elemento all'interno di `#share` dovrebbe essere un elemento `p`.

```js
assert.strictEqual(document.querySelector("#share").children[1].tagName, "P");
```

L'elemento `p` dovrebbe avere il testo `Your card was shared on social media!`.

```js
assert.strictEqual(document.querySelector("#share").children[1].innerText.trim(), "Your card was shared on social media!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
