---
id: 66db57ad34c7089b9b41bfd6
title: Build a Lightbox Viewer
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

라이트박스는 웹사이트에서 페이지의 주요 콘텐츠 위에 팝업이나 모달 창을 띄워 멀티미디어 콘텐츠를 더 흥미롭게 보여줄 때 사용합니다.

이 실습에서는 썸네일을 클릭하면 전체 크기 이미지를 보여주는 라이트박스 갤러리를 만듭니다. 각 이미지에는 썸네일과 전체 크기 이미지 두 가지 버전이 제공됩니다. 전체 크기 이미지는 썸네일과 같지만 `-thumbnail` 접미사가 없습니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**User Stories:**

1. `div` 안에 `gallery` 클래스를 가진 `body`이 있어야 합니다.
1. `.gallery` 요소 안에 `gallery-item` 클래스를 가진 세 개의 이미지 썸네일이 있어야 합니다. 썸네일 이미지로는 다음 링크를 사용하세요:
   
    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. `div` 안에 `lightbox` 클래스를 가진 `body`이 있어야 합니다.
1. `button` 요소 안에 `id`라는 `close-btn`을 가진 `.lightbox`가 있어야 합니다. 원한다면 `&times;`을 텍스트로 사용할 수 있습니다.
1. `img` 요소 안에 `id`라는 `lightbox-image`을 가진 `.lightbox`이 있어야 합니다.
1. `.lightbox` 요소는 미리보기가 현재 이미지 위에 열리도록 고정 위치여야 합니다.
1. `.lightbox` 요소는 컨테이너의 높이와 너비를 100%로 설정해 뷰포트를 완전히 덮어야 합니다. `.lightbox` 요소가 컨테이너의 좌상단에서 시작하도록 해야 합니다.
1. `.lightbox`은 배경색이 있어야 합니다. 처음에는 `display` 속성이 `none`로 설정되어 숨겨져 있어야 합니다.
1. `.gallery-item` 요소 중 하나를 클릭하면 `.lightbox` 요소의 `display` 속성이 `flex`으로 설정되어 `.lightbox` 요소와 그 안의 두 요소가 보이게 해야 합니다.
1. `.gallery-item` 요소 중 하나를 클릭하면 `#lightbox-image` 요소의 `src`가 클릭한 이미지의 `-thumbnail` 속성에서 `src`를 제거한 전체 크기 이미지로 설정되어야 합니다. 전체 크기 이미지는 다음 링크에 있습니다:

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`
  
1. `.lightbox` 요소가 보이는 상태에서 `#close-btn` 또는 `.lightbox` 요소를 클릭하면 `.lightbox` 요소의 `display`가 다시 `none`로 설정되어야 합니다.

**참고:** HTML에 스타일시트와 JavaScript 파일을 반드시 연결하세요.

# --hints--

`div` 요소 안에 `gallery` 클래스를 가진 `body`이 있어야 합니다.

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

`.gallery` 요소 안에 `img` 클래스를 가진 세 개의 `gallery-item` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

`.gallery` 요소 안에 `img`로 설정된 `src`를 가진 `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg` 요소가 있어야 합니다.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

`.gallery` 요소 안에 `img`로 설정된 `src`를 가진 `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg` 요소가 있어야 합니다.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

`.gallery` 요소 안에 `img`로 설정된 `src`를 가진 `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg` 요소가 있어야 합니다.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

`div` 요소 안에 `lightbox` 클래스를 가진 `body` 요소가 있어야 합니다.

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

`.lightbox` 요소 안에 `button`로 설정된 `id`를 가진 `close-btn` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

`.lightbox` 요소 안에 `img`로 설정된 `id`를 가진 `lightbox-image` 요소가 있어야 합니다.

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

`.lightbox` 요소는 고정 위치여야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

`.lightbox` 요소는 뷰포트를 완전히 덮어야 합니다.

```js
const style = new __helpers.CSSHelp(document).getStyle(".lightbox");

const width = style?.width;
const height = style?.height;

assert.oneOf(
  width,
  ["100%", "100vw"]
);

assert.oneOf(
  height,
  ["100%", "100vh"]
);
```

`.lightbox` 요소는 컨테이너의 좌상단에 정렬되어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

`.lightbox` 요소는 배경색이 있어야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

`.lightbox` 요소는 처음에 숨겨져 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

`.gallery-item` 요소 중 하나를 클릭하면 `.lightbox` 요소의 `display` 속성이 `flex`으로 설정되어 `.lightbox`와 그 안의 두 요소가 보이게 해야 합니다.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

`.gallery-item` 요소 중 하나를 클릭하면 `#lightbox-image` 요소의 `src`가 클릭한 이미지의 `-thumbnail` 속성에서 `src`를 제거한 전체 크기 이미지로 설정되어야 합니다.

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

`.lightbox` 요소가 보이는 상태에서 `#close-btn` 버튼을 클릭하면 `.lightbox` 요소의 `display`가 다시 `none`로 설정되어야 합니다.

```js
const lightbox = document.querySelector(".lightbox");
const closeBtn = document.getElementById("close-btn"); 
const galleryItem = document.querySelector(".gallery-item");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

closeBtn.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

`.lightbox` 요소가 보이는 상태에서 `.lightbox` 요소를 클릭하면 `.lightbox` 요소의 `display`가 다시 `none`로 설정되어야 합니다.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

lightbox.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```
