---
id: 614385513d91ae5c251c2052
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Agregue un elemento `title` con el texto `Magazine`, un elemento `link` para la hoja de estilo de fuente`https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`, una hoja de estilos `link` para la hoja de estilos `https://use.fontawesome.com/releases/v5.8.2/css/all.css` FontAwesome, y un `link` para su hoja de estilos `./styles.css`.

La hoja de estilo de fuente cargará tres fuentes separadas: `Anton`, `Baskervville` y `Raleway`.

# --hints--

Your code should have three `link` elements.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Your `link` elements should be inside your `head` element.

```js
const headContentRegex = /(?<=<head\s*>)(?:.|\s*)*?(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

Los tres elementos `link` deben tener un atributo `rel` con el valor `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Uno de los elementos del enlace debe tener el `href` establecido en`https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Uno de los elementos del enlace debe tener el `href` establecido en`https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

One of your `link` elements should have the `href` set to `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

El código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.isNotNull(title);
```

Your code should have one `title` element.

```js
const title = document.querySelectorAll('title');
assert.strictEqual(title?.length, 1);
```

Your `title` element should be inside your `head` element.

```js
const headContentRegex = /(?<=<head\s*>)(?:.|\s*)*?(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

El proyecto debe tener el título `Magazine`. Recuerde, la mayúscula y la ortografía importan para el título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
