---
id: 681cb1b0dab50c87ddb2e51b
title: "Challenge 10: 3 Strikes"
challengeType: 29
dashedName: challenge-10
---

# --description--

1から10,000までの整数が与えられたとき、その整数までの数字のうち、平方の中に少なくとも1つの数字の3が含まれている数の個数を返してください。

# --hints--

`squares_with_three(1)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1), 0)`)
}})
```

`squares_with_three(10)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10), 1)`)
}})
```

`squares_with_three(100)` should return `19`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(100), 19)`)
}})
```

`squares_with_three(1000)` should return `326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1000), 326)`)
}})
```

`squares_with_three(10000)` should return `4531`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10000), 4531)`)
}})
```

# --seed--

## --seed-contents--

```py
def squares_with_three(n):

    return n
```

# --solutions--

```py
def squares_with_three(n):
    count = 0
    for i in range(1, n + 1):
        square = i * i
        if '3' in str(square):
            count += 1
    return count
```
