---
id: 681cb1afdab50c87ddb2e516
title: "Challenge 5: Jbelmud Text"
challengeType: 29
dashedName: challenge-5
---

# --description--

文字列が与えられたとき、各単語を次の条件で変換した文字列のシャッフル版を返してください。

- 単語の最初と最後の文字はそのままの位置に残す
- 最初と最後の文字の間のすべての文字はアルファベット順に並べ替える
- 入力される文字列には句読点が含まれず、すべて小文字です

# --hints--

`jbelmu("hello world")` should return `"hello wlord"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("hello world"), "hello wlord")`)
}})
```

`jbelmu("i love jumbled text")` should return `"i love jbelmud text"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("i love jumbled text"), "i love jbelmud text")`)
}})
```

`jbelmu("freecodecamp is my favorite place to learn to code")` should return `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe")`)
}})
```

`jbelmu("the quick brown fox jumps over the lazy dog")` should return `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def jbelmu(text):

    return text
```

# --solutions--

```py
def jbelmu(text):
    words = text.split()
    jumbled = []

    for word in words:
        if len(word) <= 3:
            jumbled.append(word)
        else:
            first = word[0]
            last = word[-1]
            middle = ''.join(sorted(word[1:-1]))
            jumbled.append(first + middle + last)
    
    return ' '.join(jumbled)
```
