---
id: 681cb1a2dab50c87ddb2e514
title: "Challenge 3: Fibonacci Sequence"
challengeType: 29
dashedName: challenge-3
---

# --description--

フィボナッチ数列は、各数字が直前の2つの数字の合計になる数の列です。`0` と `1` から始めると、数列の最初の10個の数字は `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34` です。

フィボナッチ数列の最初の2つの数字を含む配列と、数列の長さを表す整数が与えられたとき、その長さの数列を含む配列を返してください。

- 関数は長さが0以上の任意の数列に対応する必要があります。
- 長さが0の場合は、空の配列を返してください。
- 最初の数字は数列の一部であることに注意してください。

# --hints--


`fibonacci_sequence([0, 1], 20)` should return `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181])`)
}})
```

`fibonacci_sequence([21, 32], 1)` should return `[21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([21, 32], 1), [21])`)
}})
```

`fibonacci_sequence([0, 1], 0)` should return `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 0), [])`)
}})
```

`fibonacci_sequence([10, 20], 2)` should return `[10, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([10, 20], 2), [10, 20])`)
}})
```

`fibonacci_sequence([123456789, 987654321], 5)` should return `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541])`)
}})
```

# --seed--

## --seed-contents--

```py
def fibonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def fibonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2]
        sequence.append(next_value)

    return sequence
```
