---
id: 681cb1afdab50c87ddb2e515
title: "チャレンジ 4: S  P  A  C  E  J  A  M"
challengeType: 29
dashedName: challenge-4
---

# --description--

文字列が与えられたら、その文字列からすべてのスペースを取り除き、各文字の間に2つのスペースを挿入し、すべてのアルファベットを大文字に変換して結果を返してください。

- アルファベット以外の文字は変更しないでください（スペースを除く）。

# --hints--

`space_jam("freeCodeCamp")` は `"F  R  E  E  C  O  D  E  C  A  M  P"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("   free   Code   Camp   ")` は `"F  R  E  E  C  O  D  E  C  A  M  P"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("Hello World?!")` は `"H  E  L  L  O  W  O  R  L  D  ?  !"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !")`)
}})
```

`space_jam("C@t$ & D0g$")` は `"C  @  T  $  &  D  0  G  $"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $")`)
}})
```

`space_jam("all your base")` は `"A  L  L  Y  O  U  R  B  A  S  E"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E")`)
}})
```

# --seed--

## --seed-contents--

```py
def space_jam(s):

    return s
```

# --solutions--

```py
def space_jam(s):
    s = s.replace(" ", "")
    s = "  ".join(s)
    return s.upper()
```
