---
id: 6821ebea237de8297eaee795
title: "Challenge 19: Candlelight"
challengeType: 29
dashedName: challenge-19
---

# --description--

最初に持っているキャンドルの数を表す整数と、新しいキャンドルを作るのに必要な燃えたキャンドルの数を表す整数が与えられたとき、できるだけ多く作って燃やした後に使ったキャンドルの合計数を返してください。

例えば、7本のキャンドルがあり、新しいキャンドルを作るのに燃えたキャンドルが2本必要な場合：

1. 7本のキャンドルを燃やして7本の残りを得る、
2. 6本の残りをリサイクルして3本の新しいキャンドルを作る（1本の残りが残る）、
3. 3本のキャンドルを燃やしてさらに3本の残りを得る（合計4本）、
4. 4本の残りをリサイクルして2本の新しいキャンドルを作る、
5. 2本のキャンドルを燃やして2本の残りを得る、
6. 2本の残りをリサイクルして1本の新しいキャンドルを作る、
7. 1本のキャンドルを燃やす。

この例では合計で13本のキャンドルを燃やしたことになります。

# --hints--

`burn_candles(7, 2)` should return `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

`burn_candles(10, 5)` should return `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

`burn_candles(20, 3)` should return `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

`burn_candles(17, 4)` should return `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

`burn_candles(2345, 3)` should return `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```
