---
id: 6821ebda237de8297eaee792
title: "Challenge 16: Reverse Parenthesis"
challengeType: 29
dashedName: challenge-16
---

# --description--

正しくネストされた丸括弧を含む文字列が与えられたとき、次のルールに従って文字列のデコード版を返してください。

- 丸括弧の各ペア内のすべての文字は反転されるべきです。
- 丸括弧は最終結果から取り除かれます。
- 丸括弧がネストされている場合、最も内側のペアから先に反転し、その結果を外側のペアの反転に含めます。
- すべての丸括弧は均等にバランスが取れて正しくネストされていると仮定してください。

# --hints--

`decode("(f(b(dc)e)a)")` should return `"abcdef"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` should return `"Can you read this?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` should return `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
