---
id: 6821ebe4237de8297eaee794
title: "チャレンジ18: 2番目に高い価格"
challengeType: 29
dashedName: challenge-18
---

# --description--

異なるノートパソコンの価格を表す整数の配列と、予算を表す整数が与えられたとき、次のいずれかを返してください。

1. 予算内であればリストの中で2番目に高価なノートパソコン、または
2. そうでなければ購入可能な中で最も高価なノートパソコン、または
3. 予算内にノートパソコンがない場合は `0` を返してください。

- 重複する価格は無視してください。

# --hints--

`get_laptop_cost([1500, 2000, 1800, 1400], 1900)` は `1800` を返すべきです

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900)` は `1800` を返すべきです

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 2200)` は `1899` を返すべきです

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 2200), 1899)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 1000)` は `0` を返すべきです

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 1000), 0)`)
}})
```

`get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` は `1400` を返すべきです

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_laptop_cost(laptops, budget):

    return laptops
```

# --solutions--

```py
def get_laptop_cost(laptops, budget):
    laptops = sorted(set(laptops), reverse=True)

    if len(laptops) >= 2 and budget >= laptops[1]:
        return laptops[1]
    if not laptops or budget < laptops[-1]:
        return 0

    for price in laptops[2:]:
        if budget >= price:
            return price

    return 0
```
