---
id: 6821ebf8237de8297eaee798
title: "Challenge 22: Tribonacci Sequence"
challengeType: 29
dashedName: challenge-22
---

# --description--

トリボナッチ数列は、各数字が直前の3つの数字の合計になる数列です。`0`、`0`、`1`から始めると、数列の最初の10個の数字は`0`、`0`、`1`、`1`、`2`、`4`、`7`、`13`、`24`、`44`です。

トリボナッチ数列の最初の3つの数字を含む配列と、数列の長さを表す整数が与えられたとき、その長さの数列を含む配列を返してください。

- 関数は長さが0以上の任意の数列に対応する必要があります。
- 長さが0の場合は、空の配列を返してください。
- 最初の数字は数列の一部であることに注意してください。

# --hints--

`tribonacci_sequence([0, 0, 1], 20)` should return `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513])`)
}})
```

`tribonacci_sequence([21, 32, 43], 1)` should return `[21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([21, 32, 43], 1), [21])`)
}})
```

`tribonacci_sequence([0, 0, 1], 0)` should return `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 0), [])`)
}})
```

`tribonacci_sequence([10, 20, 30], 2)` should return `[10, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 2), [10, 20])`)
}})
```

`tribonacci_sequence([10, 20, 30], 3)` should return `[10, 20, 30]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 3), [10, 20, 30])`)
}})
```

`tribonacci_sequence([123, 456, 789], 8)` should return `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134])`)
}})
```

# --seed--

## --seed-contents--

```py
def tribonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def tribonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return [start_sequence[0], start_sequence[1]]
    if length == 3:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2] + sequence[-3]
        sequence.append(next_value)

    return sequence
```
