---
id: 6821ebf3237de8297eaee797
title: "チャレンジ 21: 16進数ジェネレーター"
challengeType: 29
dashedName: challenge-21
---

# --description--

名前付きのCSSカラー文字列が与えられたとき、その色が支配的になるランダムな16進数（hex）カラーコードを生成してください。

- 関数は`"red"`、`"green"`、または`"blue"`を入力引数として処理する必要があります。
- 入力がこれらのいずれでもない場合、関数は`"Invalid color"`を返す必要があります。
- 関数は、入力色の値が他の色よりも大きいランダムな6文字の16進数カラーコードを返す必要があります。
- 入力に対する有効な出力の例：

| 入力   | 出力   |
|---------|----------|
| `"red"`   | `"FF0000"` |
| `"green"` | `"00FF00"` |
| `"blue"`  | `"0000FF"` |

# --hints--

`generate_hex("yellow")`は`"Invalid color"`を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_hex("yellow"), "Invalid color")`);
}})
```

`generate_hex("red")`は6文字の文字列を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(len(generate_hex("red")), 6)`);
}})
```

`generate_hex("red")`は有効な6文字の16進数カラーコードを返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase

hex = generate_hex("red").upper()
is_valid_hex = len(hex) == 6 and all(c in "0123456789ABCDEF" for c in hex)
TestCase().assertTrue(is_valid_hex)`)
}})
```

`generate_hex("red")`は他の色よりも赤の値が高い有効な16進数カラーを返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase

hex = generate_hex("red").upper()
is_valid_hex = len(hex) == 6 and all(c in "0123456789ABCDEF" for c in hex)
TestCase().assertTrue(is_valid_hex)

r = int(hex[:2], 16)
g = int(hex[2:4], 16)
b = int(hex[4:], 16)

TestCase().assertGreater(r, g)
TestCase().assertGreater(r, b)`)
}})
```

`generate_hex("red")`を2回呼び出すと、赤が支配的な異なる2つの16進数カラー値が返される必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase

hex1 = generate_hex("red").upper()
is_valid_hex1 = len(hex1) == 6 and all(c in "0123456789ABCDEF" for c in hex1)
TestCase().assertTrue(is_valid_hex1)

r1 = int(hex1[:2], 16)
g1 = int(hex1[2:4], 16)
b1 = int(hex1[4:], 16)

TestCase().assertGreater(r1, g1)
TestCase().assertGreater(r1, b1)

hex2 = generate_hex("red").upper()
is_valid_hex2 = len(hex2) == 6 and all(c in "0123456789ABCDEF" for c in hex2)
TestCase().assertTrue(is_valid_hex2)

r2 = int(hex2[:2], 16)
g2 = int(hex2[2:4], 16)
b2 = int(hex2[4:], 16)

TestCase().assertGreater(r2, g2)
TestCase().assertGreater(r2, b2)
TestCase().assertNotEqual(hex1, hex2)`);
}})
```

`generate_hex("green")`を2回呼び出すと、緑が支配的な異なる2つの16進数カラー値が返される必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase

hex1 = generate_hex("green").upper()
is_valid_hex1 = len(hex1) == 6 and all(c in "0123456789ABCDEF" for c in hex1)
TestCase().assertTrue(is_valid_hex1)

r1 = int(hex1[:2], 16)
g1 = int(hex1[2:4], 16)
b1 = int(hex1[4:], 16)

TestCase().assertGreater(g1, r1)
TestCase().assertGreater(g1, b1)

hex2 = generate_hex("green").upper()
is_valid_hex2 = len(hex2) == 6 and all(c in "0123456789ABCDEF" for c in hex2)
TestCase().assertTrue(is_valid_hex2)

r2 = int(hex2[:2], 16)
g2 = int(hex2[2:4], 16)
b2 = int(hex2[4:], 16)

TestCase().assertGreater(g2, r2)
TestCase().assertGreater(g2, b2)
TestCase().assertNotEqual(hex1, hex2)`);
}})
```

`generate_hex("blue")`を2回呼び出すと、青が支配的な異なる2つの16進数カラー値が返される必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase

hex1 = generate_hex("blue").upper()
is_valid_hex1 = len(hex1) == 6 and all(c in "0123456789ABCDEF" for c in hex1)
TestCase().assertTrue(is_valid_hex1)

r1 = int(hex1[:2], 16)
g1 = int(hex1[2:4], 16)
b1 = int(hex1[4:], 16)

TestCase().assertGreater(b1, r1)
TestCase().assertGreater(b1, g1)

hex2 = generate_hex("blue").upper()
is_valid_hex2 = len(hex2) == 6 and all(c in "0123456789ABCDEF" for c in hex2)
TestCase().assertTrue(is_valid_hex2)

r2 = int(hex2[:2], 16)
g2 = int(hex2[2:4], 16)
b2 = int(hex2[4:], 16)

TestCase().assertGreater(b2, r2)
TestCase().assertGreater(b2, g2)
TestCase().assertNotEqual(hex1, hex2)`);
}})
```

# --seed--

## --seed-contents--

```py
def generate_hex(color):

    return color
```

# --solutions--

```py
import random
def generate_hex(color):
    def to_hex(n):
        return hex(n)[2:].upper().zfill(2)

    dominant = random.randint(170, 255)
    weak1 = random.randint(0, 169)
    weak2 = random.randint(0, 169)

    if color.lower() == "red":
        r = dominant
        g = weak1
        b = weak2
    elif color.lower() == "green":
        r = weak1
        g = dominant
        b = weak2
    elif color.lower() == "blue":
        r = weak1
        g = weak2
        b = dominant
    else:
        return "Invalid color"

    return f'{to_hex(r)}{to_hex(g)}{to_hex(b)}'
```
