---
id: 68adce01c0e1144d0a902956
title: "チャレンジ 25: 母音リピーター"
challengeType: 29
dashedName: challenge-25
---

# --description--

文字列が与えられたとき、各母音が前に出てきた母音よりも1回多く繰り返される新しい文字列を返してください。たとえば、文中の最初の母音はそのままです。2番目の母音は連続して2回、3番目の母音は連続して3回、というように続きます。

- 大文字または小文字の `a`、`e`、`i`、`o`、`u` の文字は母音とみなします。
- 元の母音はそのままの大文字・小文字を保持してください。
- 繰り返す母音は小文字にしてください。
- 母音以外の文字は元の大文字・小文字を保持してください。

# --hints--

`repeat_vowels("hello world")` は `"helloo wooorld"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` は `"freeeCooodeeeeCaaaaamp"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` は `"AEeIiiOoooUuuuu"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` は `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
