---
id: 68b06e589bf2273243814773
title: "Challenge 37: Sentence Capitalizer"
challengeType: 29
dashedName: challenge-37
---

# --description--

段落が与えられたとき、各文の最初の文字を大文字にした新しい段落を返してください。

- その他の文字はすべてそのまま保持してください。
- 文の終わりはピリオド (`.`)、1つ以上の疑問符 (`?`)、または1つ以上の感嘆符 (`!`) で終わる場合があります。

# --hints--

`capitalize("this is a simple sentence.")` should return `"This is a simple sentence."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("this is a simple sentence."), "This is a simple sentence.")`)
}})
```

`capitalize("hello world. how are you?")` should return `"Hello world. How are you?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("hello world. how are you?"), "Hello world. How are you?")`)
}})
```

`capitalize("i did today's coding challenge... it was fun!!")` should return `"I did today's coding challenge... It was fun!!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!")`)
}})
```

`capitalize("crazy!!!strange???unconventional...sentences.")` should return `"Crazy!!!Strange???Unconventional...Sentences."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.")`)
}})
```

`capitalize("there's a space before this period . why is there a space before that period ?")` should return `"There's a space before this period . Why is there a space before that period ?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?")`)
}})
```

# --seed--

## --seed-contents--

```py
def capitalize(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def capitalize(paragraph):
    def repl(match):
        return match.group(1) + match.group(2) + match.group(3).upper()

    result = re.sub(r'([.!?]+)(\s*)([a-z])', repl, paragraph)

    if result and result[0].islower():
        result = result[0].upper() + result[1:]

    return result
```
