---
id: 68b06e589bf227324381476f
title: "Challenge 35: Word Frequency"
challengeType: 29
dashedName: challenge-35
---

# --description--

段落が与えられたとき、最も頻繁に出現する単語を3つ返す配列を作成してください。

- 段落内の単語はスペースで区切られています。
- 段落内の大文字・小文字は区別しません。例えば、`Hello` と `hello` は同じ単語として扱います。
- 段落内の句読点は無視してください。句読点にはカンマ（`,`）、ピリオド（`.`）、感嘆符（`!`）が含まれます。
- 返す配列の単語はすべて小文字にしてください。
- 返す配列は、最も頻繁に出現する単語が最初に来るように降順で並べてください。

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` should return `["coding", "python", "in"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` should return `["i", "like", "coding"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` should return `["debug", "test", "deploy"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
