---
id: 68b06e589bf2273243814777
title: "チャレンジ 39: タンクを満たす"
challengeType: 29
dashedName: challenge-39
---

# --description--

燃料タンクの容量、現在の燃料量、ガロンあたりの価格が与えられたとき、タンクを満タンにするための費用を返してください。

- 最初の引数はタンクの総容量（ガロン単位）です。
- 2番目の引数はタンク内の現在の燃料量（ガロン単位）です。
- 3番目の引数は1ガロンあたりの燃料の価格です。
- 戻り値は小数点以下2桁で四捨五入し、`"$d.dd"` の形式で返してください。

# --hints--

`cost_to_fill(20, 0, 4.00)` は `"$80.00"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` は `"$17.50"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` は `"$29.25"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` は `"$0.00"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` は `"$21.89"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
