---
id: 68b7cadffed0e75a517da675
title: "Challenge 53: Decimal to Binary"
challengeType: 29
dashedName: challenge-53
---

# --description--

0以上の整数が与えられたら、その2進数表現を文字列として返してください。

2進数は数字の`0`と`1`だけを使ってどんな数でも表します。10進数を2進数に変換するには、その数を`2`で繰り返し割り、余りを記録します。数が0になるまで繰り返してください。最後に記録した余りから最初の余りまでを読みます。例えば、`12`を2進数に変換すると：

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12`の2進数表現は`1100`です。

# --hints--

`to_binary(5)` should return `"101"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(5), "101")`)
}})
```

`to_binary(12)` should return `"1100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(12), "1100")`)
}})
```

`to_binary(50)` should return `"110010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(50), "110010")`)
}})
```

`to_binary(99)` should return `"1100011"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(99), "1100011")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_binary(decimal):

    return decimal
```

# --solutions--

```py
def to_binary(decimal):
    if decimal == 0:
        return "0"
    binary = ""
    while decimal > 0:
        binary = str(decimal % 2) + binary
        decimal //= 2
    return binary
```
