---
id: 68b7cadffed0e75a517da677
title: "Challenge 54: P@ssw0rd Str3ngth!"
challengeType: 29
dashedName: challenge-54
---

# --description--

パスワードの文字列が与えられたら、その強度に応じて `"weak"`、`"medium"`、または `"strong"` を返してください。

パスワードは次のルールに基づいて評価されます：

- 8文字以上であること。
- 大文字と小文字の両方を含むこと。
- 少なくとも1つの数字を含むこと。
- 次のセットの中から少なくとも1つの特殊文字を含むこと：`!`、`@`、`#`、`$`、`%`、`^`、`&`、または `*`。

ルールを2つ未満しか満たさない場合は `"weak"` を返してください。
ルールを2つまたは3つ満たす場合は `"medium"` を返してください。
ルールをすべて4つ満たす場合は `"strong"` を返してください。

# --hints--

`check_strength("123456")` should return `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("123456"), "weak")`)
}})
```

`check_strength("pass!!!")` should return `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("pass!!!"), "weak")`)
}})
```

`check_strength("Qwerty")` should return `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("Qwerty"), "weak")`)
}})
```

`check_strength("PASSWORD")` should return `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD"), "weak")`)
}})
```

`check_strength("PASSWORD!")` should return `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("PassWord%^!")` should return `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PassWord%^!"), "medium")`)
}})
```

`check_strength("qwerty12345")` should return `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("qwerty12345"), "medium")`)
}})
```

`check_strength("S3cur3P@ssw0rd")` should return `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("S3cur3P@ssw0rd"), "strong")`)
}})
```

`check_strength("C0d3&Fun!")` should return `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("C0d3&Fun!"), "strong")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_strength(password):

    return password
```

# --solutions--

```py
import re
def check_strength(password):
    rules_met = 0

    if len(password) >= 8:
        rules_met += 1
    if re.search(r'[a-z]', password) and re.search(r'[A-Z]', password):
        rules_met += 1
    if re.search(r'\d', password):
        rules_met += 1
    if re.search(r'[!@#$%^&*]', password):
        rules_met += 1

    if rules_met < 2:
        return "weak"
    elif rules_met < 4:
        return "medium"
    else:
        return "strong"
```
