---
id: 68b7cadffed0e75a517da673
title: "Challenge 52: Binary to Decimal"
challengeType: 29
dashedName: challenge-52
---

# --description--

2進数を表す文字列が与えられたら、その10進数の値を数値として返してください。

2進数は数字の`0`と`1`だけを使ってどんな数でも表します。2進数を10進数に変換するには、各桁を`2`のべき乗で乗算して、それらをすべて足し合わせます。右端の桁を`2^0`で乗算し、その左の桁を`2^1`で乗算し、以下同様に続けます。すべての桁を`2`のべき乗で乗算したら、その結果を合計してください。

たとえば、2進数の`101`は10進数で`5`に等しいです。なぜなら：

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` should return `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```
