---
id: 68c1a929005bf54d342aa8d4
title: "Challenge 57: Space Week Day 3: Phone Home"
challengeType: 29
dashedName: challenge-57
---

# --description--

スペースウィークの3日目では、通信ルート上の自分、衛星、そして母星の間の距離（キロメートル単位）を表す配列が与えられます。次の条件を使って、メッセージがルートを通って目的の惑星に届くまでにかかる時間を求めてください。

- 配列の最初の値は、自分の位置から最初の衛星までの距離です。
- 最後の値を除く各値は、次の衛星までの距離です。
- 配列の最後の値は、前の衛星から母星までの距離です。
- メッセージは秒速300,000キロメートルで移動します。
- メッセージが通過する各衛星は0.5秒の伝送遅延を追加します。
- 小数点以下4桁に丸め、末尾のゼロは削除して数値を返してください。

# --hints--

`send_message([300000, 300000])` should return `2.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` should return `3.0627`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` should return `364.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` should return `1674.3333`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` should return `2.4086`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` should return `21.1597`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
