---
id: 68c1a929005bf54d342aa8d6
title: "Challenge 59: Space Week Day 5: Goldilocks Zone"
challengeType: 29
dashedName: challenge-59
---

# --description--

宇宙週間の5日目では、恒星の「ゴールディロックスゾーン」を計算します。これは、液体の水が存在するのに「ちょうど良い」条件が整っている恒星の周囲の領域です。

恒星の質量が与えられたら、そのゴールディロックスゾーンの開始距離と終了距離を天文単位で示す配列を返してください。

ゴールディロックスゾーンを計算するには：

1. 恒星の光度を、質量を3.5乗した値で求めます。
2. ゾーンの開始距離は、その光度の平方根に0.95を掛けた値です。
3. ゾーンの終了距離は、その光度の平方根に1.37を掛けた値です。

- 距離は小数点以下2桁に丸めて返してください。

例えば、質量が`1`の場合、`[0.95, 1.37]`を返します。

# --hints--

`goldilocks_zone(1)` should return `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` should return `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` should return `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` should return `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` should return `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
