---
id: 68c1a929005bf54d342aa8d3
title: "Challenge 56: Space Week Day 2: Exoplanet Search"
challengeType: 29
dashedName: challenge-56
---

# --description--

スペースウィークの2日目では、各文字が星の光度の読み取り値を表す文字列が与えられます。トランジット法を使って太陽系外惑星が検出されたかどうかを判定してください。トランジット法とは、惑星が星の前を通過して観測される光度が減少する現象のことです。

- 光度の読み取り値は `0-9` と `A-Z` の文字のみで構成され、それぞれの読み取り値は次の数値に対応します：
- 文字 `0-9` は光度レベル `0-9` に対応します。
- 文字 `A-Z` は光度レベル `10-35` に対応します。

星に太陽系外惑星があるとみなす条件は、任意の単一の読み取り値が全読み取り値の平均の80％以下であることです。例えば、星の平均光度が10の場合、任意の読み取り値が8以下であれば太陽系外惑星があるとみなされます。

# --hints--

`has_exoplanet("665544554")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```
