---
id: 68c497f3aaefc9fd9f1b0e26
title: "Challenge 62: Hex to Decimal"
challengeType: 29
dashedName: challenge-62
---

# --description--

16進数（基数16）を表す文字列が与えられたら、その10進数（基数10）の値を整数として返してください。

16進数は16個の数字を使う数値システムです：

- `0-9` は `0` から `9` までの値を表します。
- `A-F` は `10` から `15` までの値を表します。

部分的な変換テーブルは以下の通りです：

| 16進数 | 10進数 |
|:-------:|:-------:|
| 0       | 0       |
| 1       | 1       |
| ...     | ...     |
| 9       | 9       |
| A       | 10      |
| ...     | ...     |
| F       | 15      |
| 10      | 16      |
| ...     | ...     |
| 9F      | 159     |
| A0      | 160     |
| ...     | ...     |
| FF      | 255     |
| 100     | 256     |

- 文字列には `0–9` と `A–F` の文字のみが含まれます。

# --hints--

`hex_to_decimal("A")` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A"), 10)`)
}})
```

`hex_to_decimal("15")` should return `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("15"), 21)`)
}})
```

`hex_to_decimal("2E")` should return `46`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("2E"), 46)`)
}})
```

`hex_to_decimal("FF")` should return `255`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("FF"), 255)`)
}})
```

`hex_to_decimal("A3F")` should return `2623`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A3F"), 2623)`)
}})
```

# --seed--

## --seed-contents--

```py
def hex_to_decimal(hex):

    return hex
```

# --solutions--

```py
def hex_to_decimal(hex):

    return int(hex, 16)
```
