---
id: 68c497f3aaefc9fd9f1b0e24
title: "Challenge 60: Space Week Day 6: Moon Phase"
challengeType: 29
dashedName: challenge-60
---

# --description--

宇宙週間の6日目では、`"YYYY-MM-DD"` の形式で日付が与えられ、その日の月の満ち欠けを次のルールに従って判定してください。

28日間の簡略化した月の周期を使い、4つの等しいフェーズに分けます。

- `"New"`: 1日目から7日目まで
- `"Waxing"`: 8日目から14日目まで
- `"Full"`: 15日目から21日目まで
- `"Waning"`: 22日目から28日目まで

28日目の後は周期が繰り返され、1日目は新月となります。

- `"2000-01-06"` を基準の新月（周期の1日目）として使い、与えられた日のフェーズを判定してください。
- 基準日より前の日付は与えられません。
- 正しいフェーズを文字列で返してください。

**注意:** 1日目は新月の日を表し、前回の新月から0日経過していることを意味します。

# --hints--

`moon_phase("2000-01-12")` should return `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` should return `"Waxing"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` should return `"Full"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` should return `"Waning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` should return `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
