---
id: 68cae5b538ff798bbd4da001
title: "Challenge 63: Battle of Words"
challengeType: 29
dashedName: challenge-63
---

# --description--

あなたのチームと対戦チームを表す2つの文が与えられます。各単語が対応する相手チームの単語と戦うとき、次のルールに従ってどちらのチームが勝つか判定してください。

- 与えられる文は常に同じ単語数です。
- 単語は1つのスペースで区切られ、文字のみで構成されます。
- 各単語の値は、その単語の文字の合計値です。
- 文字`a`から`z`は値`1`から`26`に対応します。例えば、`a`は`1`であり、`z`は`26`です。
- 大文字は文字の値を2倍にします。例えば、`A`は`2`であり、`Z`は`52`です。
- 単語は順番に戦います：あなたのチームの最初の単語が対戦チームの最初の単語と戦い、以下同様です。
- 単語の値が相手の単語の値より大きい場合、その単語が勝ちます。
- 勝った単語の数が多いチームが勝者です。

あなたのチームが勝者の場合は`"We win"`を、負けた場合は`"We lose"`を、両チームの勝ち数が同じ場合は`"Draw"`を返してください。

# --hints--

`battle("hello world", "hello word")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` should return `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` should return `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
