---
id: 68cae5b538ff798bbd4da004
title: "チャレンジ 66: HTML タグストリッパー"
challengeType: 29
dashedName: challenge-66
---

# --description--

HTMLコードの文字列が与えられたら、タグを取り除いてプレーンテキストの内容を返してください。

- 入力の文字列は有効なHTMLのみを含みます。
- HTMLタグはネストされている場合があります。
- タグとその属性はすべて取り除いてください。

例えば、`'<a href="#">Click here</a>'` は `"Click here"` を返すべきです。

# --hints--

`strip_tags('<a href="#">Click here</a>')` は `"Click here"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` は `"Hello World!"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` は空文字列（`""`）を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` は `sectionsection` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
