---
id: 68cae5b538ff798bbd4da007
title: "Challenge 69: Missing Socks"
challengeType: 29
dashedName: challenge-69
---

# --description--

最初に持っていた靴下のペア数を表す整数と、洗濯を何回行ったかを表す別の整数が与えられたとき、以下の条件に従って現在持っている完全な靴下のペア数を返してください。

- 洗濯を2回行うごとに、靴下が1足失われます。
- 洗濯を3回行うごとに、失くしていた靴下が1足見つかります。
- 洗濯を5回行うごとに、靴下が1足擦り切れて捨てなければなりません。
- 洗濯を10回行うごとに、靴下のペアを1組買います。
- 靴下の合計数が0未満になることはありません。
- ルールは重複する場合があります。例えば、洗濯10回目では、靴下を1足失い、1足捨て、新しいペアを1組買います。
- 完全な靴下のペア数を返してください。

# --hints--

`sock_pairs(2, 5)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(2, 5), 1)`)
}})
```

`sock_pairs(1, 2)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 2), 0)`)
}})
```

`sock_pairs(5, 11)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(5, 11), 4)`)
}})
```

`sock_pairs(6, 25)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(6, 25), 3)`)
}})
```

`sock_pairs(1, 8)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 8), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sock_pairs(pairs, cycles):

    return pairs
```

# --solutions--

```py
def sock_pairs(pairs, cycles):
    socks = pairs * 2

    for i in range(1, cycles + 1):
        if i % 2 == 0:
            socks -= 1
        if i % 3 == 0:
            socks += 1
        if i % 5 == 0:
            socks -= 1
        if i % 10 == 0:
            socks += 2

        if socks < 0:
            socks = 0

    return socks // 2
```
