---
id: 68cae5b538ff798bbd4da008
title: "Challenge 70: HTML Attribute Extractor"
challengeType: 29
dashedName: challenge-70
---

# --description--

有効なHTML要素の文字列が与えられたとき、次の条件に従って要素の属性を返してください。

- 与えられるのは1つの要素だけです。
- 属性は `attribute="value"` の形式で与えられます。
- 各属性のプロパティと値をカンマで区切り、`["attribute1, value1", "attribute2, value2"]` の形式で文字列の配列として返してください。
- 属性は与えられた順番で返してください。
- 属性が見つからなければ、空の配列を返してください。

# --hints--

`extract_attributes('<span class="red"></span>')` should return `["class, red"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<span class="red"></span>'), ["class, red"])`)
}})
```

`extract_attributes('<meta charset="UTF-8" />')` should return `["charset, UTF-8"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<meta charset="UTF-8" />'), ["charset, UTF-8"])`)
}})
```

`extract_attributes("<p>Lorem ipsum dolor sit amet</p>")` should return `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes("<p>Lorem ipsum dolor sit amet</p>"), [])`)
}})
```

`extract_attributes('<input name="email" type="email" required="true" />')` should return `["name, email", "type, email", "required, true"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"])`)
}})
```

`extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>')` should return `["id, submit", "class, btn btn-primary"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"])`)
}})
```

# --seed--

## --seed-contents--

```py
def extract_attributes(element):

    return element
```

# --solutions--

```py
import re
def extract_attributes(element):
    pattern = r'([\w-]+)="([^"]*)"'
    matches = re.findall(pattern, element)
    return [f"{attr}, {val}" for attr, val in matches]
```
