---
id: 68cae5b538ff798bbd4da006
title: "Challenge 68: Credit Card Masker"
challengeType: 29
dashedName: challenge-68
---

# --description--

クレジットカード番号の文字列が与えられたとき、以下の条件に従ってマスクしたバージョンを返してください。

- 文字列は4桁ずつの数字が4セット含まれており（`0-9`）、すべてのセットは単一のスペースまたは単一のハイフン（`-`）で区切られています。
- 最後の4桁を除くすべての数字をアスタリスク（`*`）に置き換えます。
- 残りの文字は変更しません。

例えば、`"4012-8888-8888-1881"` が与えられた場合、`"****-****-****-1881"` を返してください。

# --hints--

`mask("4012-8888-8888-1881")` should return `"****-****-****-1881"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("4012-8888-8888-1881"), "****-****-****-1881")`)
}})
```

`mask("5105 1051 0510 5100")` should return `"**** **** **** 5100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("5105 1051 0510 5100"), "**** **** **** 5100")`)
}})
```

`mask("6011 1111 1111 1117")` should return `"**** **** **** 1117"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("6011 1111 1111 1117"), "**** **** **** 1117")`)
}})
```

`mask("2223-0000-4845-0010")` should return `"****-****-****-0010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("2223-0000-4845-0010"), "****-****-****-0010")`)
}})
```

# --seed--

## --seed-contents--

```py
def mask(card):

    return card
```

# --solutions--

```py
def mask(card):
    if '-' in card:
        split = card.split('-')
        return f"****-****-****-{split[3]}"
    else:
        split = card.split(' ')
        return f"**** **** **** {split[3]}"
```
