---
id: 68cae5b538ff798bbd4da00a
title: "Challenge 72: Thermostat Adjuster 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

部屋の現在の温度（華氏）と目標温度（摂氏）が与えられたとき、以下の条件に基づいて部屋の温度を調整する方法を示す文字列を返してください。

- 現在の温度が目標より低い場合は `"Heat: X degrees Fahrenheit"` を返してください。`X` は目標に達するために部屋を暖める必要がある華氏の度数で、小数点以下1桁に丸めます。
- 現在の温度が目標より高い場合は `"Cool: X degrees Fahrenheit"` を返してください。`X` は目標に達するために部屋を冷やす必要がある華氏の度数で、小数点以下1桁に丸めます。
- 現在の温度が目標と等しい場合は `"Hold"` を返してください。

摂氏を華氏に変換するには、摂氏温度に1.8を乗算し、結果に32を加えます（`F = (C * 1.8) + 32`）。

# --hints--

`adjust_thermostat(32, 0)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` should return `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` should return `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` should return `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
