---
id: 68d2ba1468508398389487ce
title: "Challenge 73: Speak Wisely, You Must"
challengeType: 29
dashedName: challenge-73
---

# --description--

文が与えられたら、賢い教師からの助言のように聞こえるバージョンを次のルールに従って返してください。

- 単語は1つのスペースで区切られています。
- 文中で次の単語のいずれかが最初に現れる箇所を見つけます：`"have"`、`"must"`、`"are"`、`"will"`、`"can"`。
- その単語を含め、それより前のすべての単語を文の最後に移動し、以下を行います。
  - 移動する単語の順序はそのままにします。
  - すべて小文字にします。
  - その前にカンマとスペースを追加します。
- 新しい文の最初の単語の最初の文字を大文字にします。
- 与えられた文はすべて1つの句読点で終わります。元の句読点は保持し、新しい文の最後に移動します。
- 新しい文を返し、単語間は1つのスペースにし、文の先頭や末尾にスペースがないことを確認してください。

例えば、`"You must speak wisely."`が与えられたら`"Speak wisely, you must."`を返します。

# --hints--

`wise_speak("You must speak wisely.")` should return `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` should return `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` should return `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` should return `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` should return `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
