---
id: 68e39ed6106dac2f0a98fd63
title: "Challenge 81: Nth Prime"
challengeType: 29
dashedName: challenge-81
---

# --description--

素数とは、1より大きい正の整数で、1とその数自身以外に割り切れる数がないものです。最初の5つの素数は `2`、`3`、`5`、`7`、および `11` です。

正の整数 `n` が与えられたとき、`n` 番目の素数を返してください。例えば、`5` が与えられた場合、5番目の素数である `11` を返します。

# --hints--

`nth_prime(5)` should return `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(5), 11)`)
}})
```

`nth_prime(10)` should return `29`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(10), 29)`)
}})
```

`nth_prime(16)` should return `53`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(16), 53)`)
}})
```

`nth_prime(99)` should return `523`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(99), 523)`)
}})
```

`nth_prime(1000)` should return `7919`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(1000), 7919)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_prime(n):

    return n
```

# --solutions--

```py
def nth_prime(n):
    primes = []
    num = 2

    while len(primes) < n:
        is_prime = True
        for i in range(2, int(num**0.5) + 1):
            if num % i == 0:
                is_prime = False
                break
        if is_prime:
            primes.append(num)
        num += 1

    return primes[n - 1]
```
