---
id: 68e39ed6106dac2f0a98fd62
title: "Challenge 80: Email Sorter"
challengeType: 29
dashedName: challenge-80
---

# --description--

1971年10月29日に、最初のメールが送信され、現在も使われている`username@domain`形式が導入されました。現在では、数十億のメールアドレスがあります。

このチャレンジでは、メールアドレスのリストが与えられ、ドメイン名（`@`の後の部分）でアルファベット順に、次にユーザー名（`@`の前の部分）でソートする必要があります。

- ソートは大文字・小文字を区別しません。
- 同じドメインのメールが複数ある場合は、ユーザー名でソートしてください。
- ソートされたアドレスの配列を返してください。
- 返すアドレスは元の大文字・小文字を保持してください。

例えば、`["jill@mail.com", "john@example.com", "jane@example.com"]`が与えられた場合、`["jane@example.com", "john@example.com", "jill@mail.com"]`を返してください。

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` should return `["jane@example.com", "john@example.com", "jill@mail.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"])`)
}})
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` should return `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"])`)
}})
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` should return `["user@x.com", "user@y.com", "user@z.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"])`)
}})
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` should return `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"])`)
}})
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` should return `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"])`)
}})
```

# --seed--

## --seed-contents--

```py
def sort(emails):

    return emails
```

# --solutions--

```py
def sort(emails):
    return sorted(
        emails,
        key=lambda email: (email.split('@')[1].lower(), email.split('@')[0].lower())
    )
```
