---
id: 68e39ed6106dac2f0a98fd64
title: "Challenge 82: SpOoKy~CaSe"
challengeType: 29
dashedName: challenge-82
---

# --description--

変数名を表す文字列が与えられたとき、次の条件に従って「spooky case」に変換してください。

- すべてのアンダースコア（`_`）とハイフン（`-`）をチルダ（`~`）に置き換えます。
- 文字列の最初の文字と、その後の文字のうちチルダを除いた偶数番目の文字を大文字にします。その他の文字はすべて小文字にします。

例えば、`hello_world`が与えられたら、`HeLlO~wOrLd`を返してください。

# --hints--

`spookify("hello_world")` should return `"HeLlO~wOrLd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("hello_world"), "HeLlO~wOrLd")`)
}})
```

`spookify("Spooky_Case")` should return `"SpOoKy~CaSe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("Spooky_Case"), "SpOoKy~CaSe")`)
}})
```

`spookify("TRICK-or-TREAT")` should return `"TrIcK~oR~tReAt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt")`)
}})
```

`spookify("c_a-n_d-y_-b-o_w_l")` should return `"C~a~N~d~Y~~b~O~w~L"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L")`)
}})
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` should return `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS")`)
}})
```

# --seed--

## --seed-contents--

```py
def spookify(boo):

    return boo
```

# --solutions--

```py
def spookify(boo):
    replaced = boo.replace("_", "~").replace("-", "~")
    
    result = []
    capitalize = True

    for char in replaced:
        if char == "~":
            result.append(char)
        else:
            result.append(char.upper() if capitalize else char.lower())
            capitalize = not capitalize

    return "".join(result)
```
