---
id: 68e39ed6106dac2f0a98fd65
title: "Challenge 83: Signature Validation"
challengeType: 29
dashedName: challenge-83
---

# --description--

メッセージの文字列、秘密鍵の文字列、署名の数値が与えられたとき、このエンコード方法を使って署名が有効かどうかを判定してください。

- メッセージと秘密鍵の文字は次の値を持ちます：
  - `a` から `z` はそれぞれ `1` から `26` の値を持ちます。
  - `A` から `Z` はそれぞれ `27` から `52` の値を持ちます。
- その他の文字は値を持ちません。
- 署名はメッセージの合計値と秘密鍵の合計値を足したものとして計算します。

例えば、メッセージが `"foo"`、秘密鍵が `"bar"` の場合、署名は `57` になります：

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

計算した署名が与えられた署名と一致するかどうかを確認してください。

# --hints--

`verify("foo", "bar", 57)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("foo", "bar", 57), True)`)
}})
```

`verify("foo", "bar", 54)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("foo", "bar", 54), False)`)
}})
```

`verify("freeCodeCamp", "Rocks", 238)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("freeCodeCamp", "Rocks", 238), True)`)
}})
```

`verify("Is this valid?", "No", 210)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Is this valid?", "No", 210), False)`)
}})
```

`verify("Is this valid?", "Yes", 233)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Is this valid?", "Yes", 233), True)`)
}})
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def verify(message, key, signature):

    return message
```

# --solutions--

```py
def verify(message, key, signature):
    def charValue(ch) -> int:
        if 'a' <= ch <= 'z':
            return ord(ch) - ord('a') + 1
        elif 'A' <= ch <= 'Z':
            return ord(ch) - ord('A') + 27
        else:
            return 0

    def compute_sum(s):
        return sum(charValue(ch) for ch in s)

    total = compute_sum(message) + compute_sum(key)
    return total == signature
```
