---
id: 68ee9e3066cfd4eb2328e8a8
title: "Challenge 89: Counting Cards"
challengeType: 29
dashedName: challenge-89
---

# --description--

標準的なトランプのデッキには、各スートに13枚のユニークなカードがあります。デッキから選ぶカードの枚数を表す整数が与えられたとき、選べるユニークなカードの組み合わせの数を返してください。

- 順序は関係ありません。カードAを選んでからカードBを選ぶのは、カードBを選んでからカードAを選ぶのと同じです。

例えば、`52`が与えられたら`1`を返してください。52枚のデッキから52枚選ぶ組み合わせは1通りだけです。また、`2`が与えられたら`1326`を返してください。デッキから2枚選ぶと1326通りのカードの組み合わせが考えられます。

# --hints--

`combinations(52)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` should return `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` should return `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` should return `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
