---
id: 68f6587287ad1f4ad39b0c7f
title: "チャレンジ 93: 母音と子音"
challengeType: 29
dashedName: challenge-93
---

# --description--

文字列が与えられたら、その文字列に含まれる母音の数と子音の数を配列で返してください。

- 母音は大文字・小文字を問わず `a`、`e`、`i`、`o`、`u` です。
- 子音は大文字・小文字を問わずその他のすべての文字です。
- 文字以外の文字は無視してください。

たとえば、`"Hello World"` が与えられたら、`[3, 7]` を返します。

# --hints--

`count("Hello World")` は `[3, 7]` を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("JavaScript")` は `[3, 7]` を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("JavaScript"), [3, 7])`)
}})
```

`count("Python")` は `[1, 5]` を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Python"), [1, 5])`)
}})
```

`count("freeCodeCamp")` は `[5, 7]` を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("freeCodeCamp"), [5, 7])`)
}})
```

`count("Hello, World!")` は `[3, 7]` を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello, World!"), [3, 7])`)
}})
```

`count("The quick brown fox jumps over the lazy dog.")` は `[11, 24]` を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24])`)
}})
```

# --seed--

## --seed-contents--

```py
def count(s):

    return s
```

# --solutions--

```py
def count(s):
    vowels = "aeiou"
    consonants = "bcdfghjklmnpqrstvwxyz"
    v = 0
    c = 0

    for char in s:
        ch = char.lower()
        if ch in vowels:
            v += 1
        if ch in consonants:
            c += 1

    return [v, c]
```
