---
id: 68f6587287ad1f4ad39b0c80
title: "Challenge 94: Email Signature Generator"
challengeType: 29
dashedName: challenge-94
---

# --description--

名前、役職、会社名の文字列が与えられたとき、以下のルールに従ってメール署名を1つの文字列として返してください。

- 名前は最初に表示し、名前の最初の文字に応じた接頭辞を付けます（大文字・小文字は区別しません）。
  - `A-I`で始まる名前には`>>`を接頭辞として使います。
  - `J-R`で始まる名前には`--`を接頭辞として使います。
  - `S-Z`で始まる名前には`::`を接頭辞として使います。
- 名前の後にはカンマとスペース（`, `）を付けます。
- 役職と会社名はカンマとスペースの後に続き、`" at "`（前後にスペースあり）で区切ります。

例えば、`"Quinn Waverly"`、`"Founder and CEO"`、`"TechCo"`が与えられた場合、`"--Quinn Waverly, Founder and CEO at TechCo"`を返します。

# --hints--

`generate_signature("Quinn Waverly", "Founder and CEO", "TechCo")` should return `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo")`)
}})
```

`generate_signature("Alice Reed", "Engineer", "TechCo")` should return `">>Alice Reed, Engineer at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo")`)
}})
```

`generate_signature("Tina Vaughn", "Developer", "example.com")` should return `"::Tina Vaughn, Developer at example.com"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com")`)
}})
```

`generate_signature("B. B.", "Product Tester", "AcmeCorp")` should return `">>B. B., Product Tester at AcmeCorp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp")`)
}})
```

`generate_signature("windstorm", "Cloud Architect", "Atmospheronics")` should return `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_signature(name, title, company):

    return name
```

# --solutions--

```py
def generate_signature(name, title, company):
    first_letter = name[0].upper()
    if first_letter in "ABCDEFGHI":
        prefix = ">>"
    elif first_letter in "JKLMNOPQR":
        prefix = "--"
    else:
        prefix = "::"

    return f"{prefix}{name}, {title} at {company}"
```
