---
id: 68ffb91507a5b645769328c5
title: "Challenge 102: Longest Word"
challengeType: 29
dashedName: challenge-102
---

# --description--

文章の文字列が与えられたら、その文章の中で最も長い単語を返してください。

- 単語はスペース1つで区切られています。
- 単語の長さにカウントされるのは文字（`a-z`、大文字小文字は区別しません）のみです。
- 同じ長さの単語が複数ある場合は、最初に出てくる単語を返してください。
- 単語は与えられた文字列に現れる通りに返し、句読点は取り除いてください。

# --hints--

`longest_word("The quick red fox")` should return `"quick"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("The quick red fox"), "quick")`)
}})
```

`longest_word("Hello coding challenge.")` should return `"challenge"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Hello coding challenge."), "challenge")`)
}})
```

`longest_word("Do Try This At Home.")` should return `"This"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Do Try This At Home."), "This")`)
}})
```

`longest_word("This sentence... has commas, ellipses, and an exclamation point!")` should return `"exclamation"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation")`)
}})
```

`longest_word("A tie? No way!")` should return `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("A tie? No way!"), "tie")`)
}})
```

`longest_word("Wouldn't you like to know.")` should return `"Wouldnt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Wouldn't you like to know."), "Wouldnt")`)
}})
```

# --seed--

## --seed-contents--

```py
def longest_word(sentence):

    return sentence
```

# --solutions--

```py
import re
def longest_word(sentence):
    words = sentence.split(" ")
    longest = ""

    for word in words:
        cleaned = re.sub(r'[^a-zA-Z]', '', word)
        if len(cleaned) > len(longest):
            longest = cleaned

    return longest
```
