---
id: 68ffb91507a5b645769328c9
title: "Challenge 106: Message Validator"
challengeType: 29
dashedName: challenge-106
---

# --description--

メッセージの文字列とバリデーションの文字列が与えられたとき、メッセージが有効かどうかを判定してください。

- メッセージが有効なのは、メッセージ内の各単語がバリデーション文字列の対応する文字で始まっている場合です。順番も一致している必要があります。
- 文字の大文字・小文字は区別しません。
- メッセージ内の単語は単一のスペースで区切られています。

# --hints--

`is_valid_message("hello world", "hw")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("hello world", "hw"), True)`)
}})
```

`is_valid_message("ALL CAPITAL LETTERS", "acl")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("ALL CAPITAL LETTERS", "acl"), True)`)
}})
```

`is_valid_message("Coding challenge are boring.", "cca")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("Coding challenge are boring.", "cca"), False)`)
}})
```

`is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"), True)`)
}})
```

`is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_message(message, validation):

    return message
```

# --solutions--

```py
def is_valid_message(message, validation):
    words = message.split()
    
    if len(words) != len(validation):
        return False

    for word, val_letter in zip(words, validation):
        if word[0].lower() != val_letter.lower():
            return False

    return True
```
