---
id: 69162d64f96574d9bb629f00
title: "Challenge 115: Markdown Ordered List Item Converter"
challengeType: 29
dashedName: challenge-115
---

# --description--

Markdownで表された番号付きリストの項目を表す文字列が与えられたとき、それに対応するHTMLの文字列を返してください。

Markdownの有効な番号付きリスト項目は以下の条件を満たす必要があります：

- 0個以上のスペースで始まり、
- 1以上の数字とピリオド（`.`）が続き、
- その後に少なくとも1つのスペースがあり、
- リスト項目のテキストが続く。

文字列が上記の正確な形式でない場合は`"Invalid format"`を返してください。そうでなければ、リスト項目のテキストを`li`タグで囲んだ文字列を返してください。

例えば、`"1. My item"`が与えられたら、`"<li>My item</li>"`を返します。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

`convert_list_item("1. My item")` should return `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` should return `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` should return `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
