---
id: 69162d64f96574d9bb629efd
title: "Challenge 112: AI Detector"
challengeType: 29
dashedName: challenge-112
---

# --description--

本日のチャレンジは、2022年11月30日にリリースされたChatGPTに触発されています。

1つ以上の文からなる文字列が与えられたとき、以下のルールに基づいてAIによって生成された可能性が高いかどうかを判定してください。

- 2つ以上のダッシュ（`-`）が含まれている。
- 2つ以上の括弧のセット（`()`）が含まれている。括弧内にテキストがあってもよい。
- 7文字以上の単語が3つ以上含まれている。

- 単語は1つのスペースで区切られ、文字（`A-Z`）のみで構成されます。句読点やその他の文字は単語の一部として含めないでください。

与えられた文が上記のいずれかのルールを満たす場合は`"AI"`を返し、そうでなければ`"Human"`を返してください。

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
