---
id: 69162d64f96574d9bb629eff
title: "Challenge 114: Camel to Snake"
challengeType: 29
dashedName: challenge-114
---

# --description--

キャメルケースの文字列が与えられたら、次のルールに従ってスネークケースのバージョンを返してください。

- 入力の文字列は文字（`A-Z` と `a-z`）のみを含み、必ず小文字で始まります。
- キャメルケースの文字列の大文字はすべて新しい単語の始まりです。
- すべての文字を小文字に変換してください。
- 単語はアンダースコア（`_`）で区切ってください。

# --hints--

`to_snake("helloWorld")` should return `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` should return `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` should return `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
