---
id: 69162d64f96574d9bb629efb
title: "Challenge 110: Word Guesser"
challengeType: 29
dashedName: challenge-110
---

# --description--

同じ長さの2つの文字列、秘密の単語と推測した単語が与えられたとき、以下のルールで推測した単語を秘密の単語と比較してください。

- 秘密の単語と推測は大文字のアルファベット（`"A"`から`"Z"`）のみで構成されます。
- 推測の各文字について、秘密の単語との一致の仕方に応じて数字に置き換えます：
  - `"2"`：文字が秘密の単語にあり、かつ正しい位置にある場合。
  - `"1"`：文字が秘密の単語にあるが、位置が違う場合。
  - `"0"`：文字が秘密の単語にない場合。
- 秘密の単語の各文字は最大で1回だけ使えます。
- 完全一致（`"2"`）は最初に割り当てられ、その後、部分一致（`"1"`）が左から右へ残りの文字に割り当てられます。
- 推測の中で同じ文字が複数回出現する場合、その文字は秘密の単語に出現する回数までしか一致できません。

例えば、秘密の単語が`"APPLE"`で推測が`"POPPA"`の場合、`"10201"`を返します：

最初の`"P"`は正しい位置ではない（`"1"`）、`"O"`は秘密の単語にない（`"0"`）、2番目の`"P"`は正しい位置にある（`"2"`）、3番目の`"P"`はゼロ（`"0"`）で、秘密の単語の`"P"`が2回使われているため、`"A"`は正しい位置ではない（`"1"`）です。

# --hints--

`compare("APPLE", "POPPA")` should return `"10201"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("APPLE", "POPPA"), "10201")`)
}})
```

`compare("REACT", "TRACE")` should return `"11221"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("REACT", "TRACE"), "11221")`)
}})
```

`compare("DEBUGS", "PYTHON")` should return `"000000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("DEBUGS", "PYTHON"), "000000")`)
}})
```

`compare("JAVASCRIPT", "TYPESCRIPT")` should return `"0000222222"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222")`)
}})
```

`compare("ORANGE", "ROUNDS")` should return `"110200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("ORANGE", "ROUNDS"), "110200")`)
}})
```

`compare("WIRELESS", "ETHERNET")` should return `"10021000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("WIRELESS", "ETHERNET"), "10021000")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare(word, guess):

    return word
```

# --solutions--

```py
def compare(word, guess):
    n = len(word)
    result = ["0"] * n
    used_secret = [False] * n

    for i in range(n):
        if guess[i] == word[i]:
            result[i] = "2"
            used_secret[i] = True

    for i in range(n):
        if result[i] == "0":
            for j in range(n):
                if not used_secret[j] and guess[i] == word[j]:
                    result[i] = "1"
                    used_secret[j] = True
                    break

    return "".join(result)
```
