---
id: 69162d64f96574d9bb629f01
title: "Challenge 116: Permutation Count"
challengeType: 29
dashedName: challenge-116
---

# --description--

文字列が与えられたとき、その文字から作成できる異なる順列の数を返してください。

- 順列とは、文字列の文字を並べ替えたものです。
- 重複する順列は数えないでください。
- 文字列に繰り返しの文字が含まれている場合、同じ並びは一度だけ数えます。
- 文字列は文字（`A-Z`、`a-z`）のみを含みます。

たとえば、`"abb"`が与えられた場合、文字の並べ方が3通りあるので`3`を返してください：`"abb"`、`"bab"`、および`"bba"`。

# --hints--

`count_permutations("abb")` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` should return `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` should return `39916800`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
