---
id: 691b559495c5cb5a37b9b483
title: "Challenge 123: Roman Numeral Builder"
challengeType: 29
dashedName: challenge-123
---

# --description--

整数が与えられたら、その整数に対応するローマ数字の値を返してください。

ローマ数字は以下の記号を使います：

| Symbol | Value |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

ローマ数字は左から右へ、最大の値から最小の値へ向かって次のルールに従って書かれます：

- 記号の後に同じかそれより小さい値の記号が続く場合は加算されます。例えば、`18`は`XVIII`と書かれます（10 + 5 + 1 + 1 + 1 = 18）。
- 小さい記号が大きい記号の前に現れる場合は減算を表し、4や9をどの位の値でも表現します。例えば、19は`XIX`と書かれます（10 + (10 - 1)）。
- 同じ記号を3回以上連続して使うことはできません。3回以上連続して書く必要がある場合は減算を使います。したがって、書ける最大の数字は3999です。

もう一つ例を挙げると、`1464`が与えられたら、`"MCDLXIV"`を返します（1000 + (500 - 100) + 50 + 10 + (5 - 1)）。

# --hints--

`to_roman(18)` should return `"XVIII"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` should return `"XIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` should return `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` should return `"MMXXV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` should return `"MMMCMXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```
