---
id: 691b559495c5cb5a37b9b482
title: "Challenge 122: Markdown Bold Parser"
challengeType: 29
dashedName: challenge-122
---

# --description--

Markdownで太字のテキストが含まれている可能性がある文字列が与えられたとき、対応するHTML文字列を返してください。

- Markdownの太字テキストは、2つのアスタリスク（`**`）または2つのアンダースコア（`__`）で始まり終わるテキストです。
- テキストとアスタリスクまたはアンダースコアの間にスペースは入れられませんが、テキスト内にはスペースがあってもかまいません。
- すべての太字の部分をHTMLの`b`タグに変換して文字列を返してください。

例えば、`"**This is bold**"`が与えられたら、`"<b>This is bold</b>"`を返します。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

`parse_bold("**This is bold**")` should return `"<b>This is bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is bold**"), "<b>This is bold</b>")`)
}})
```

`parse_bold("__This is also bold__")` should return `"<b>This is also bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__This is also bold__"), "<b>This is also bold</b>")`)
}})
```

`parse_bold("**This is not bold **")` should return `"**This is not bold **"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is not bold **"), "**This is not bold **")`)
}})
```

`parse_bold("__ This is also not bold__")` should return `"__ This is also not bold__"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__ This is also not bold__"), "__ This is also not bold__")`)
}})
```

`parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog.")` should return `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_bold(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_bold(markdown):
    markdown = re.sub(r'\*\*(\S(?:.*?\S)?)\*\*', r'<b>\1</b>', markdown)
    markdown = re.sub(r'__(\S(?:.*?\S)?)__', r'<b>\1</b>', markdown)

    return markdown
```
