---
id: 691b559495c5cb5a37b9b489
title: "Challenge 129: Markdown Blockquote Parser"
challengeType: 29
dashedName: challenge-129
---

# --description--

Markdown のブロッククオートを含む文字列が与えられたとき、対応する HTML 文字列を返してください。

Markdown のブロッククオートとは、次の条件を満たす任意の行です。

- 0 個以上のスペースで始まる
- 続いて大なり記号 (`>`)
- その後に 1 個以上のスペース
- 最後にブロッククオートのテキストが続く

ブロッククオートのテキストを、開始および終了の HTML `blockquote` タグで囲んで返してください。

例えば、`"> This is a quote"` が与えられたら、`<blockquote>This is a quote</blockquote>` を返します。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

`parse_blockquote("> This is a quote")` should return `"<blockquote>This is a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` should return `"<blockquote>This is also a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` should return `"<blockquote>So  Is  This</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
