---
id: 691b559495c5cb5a37b9b480
title: "Challenge 120: Pounds to Kilograms"
challengeType: 29
dashedName: challenge-120
---

# --description--

ポンドで表された重さの数値が与えられたら、文字列`"(lbs) pounds equals (kgs) kilograms."`を返してください。

- `"(lbs)"`を入力された数値に置き換えてください。
- `"(kgs)"`をキログラムに換算した値に置き換えてください。小数点以下2桁に丸め、常に2桁の小数を含めてください。
- 1ポンドは0.453592キログラムです。
- 入力が`1`の場合は、`"pound"`の代わりに`"pounds"`を使ってください。
- 換算後の値が`1`の場合は、`"kilogram"`の代わりに`"kilograms"`を使ってください。

# --hints--

`convert_to_kgs(1)` should return `"1 pound equals 0.45 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` should return `"0 pounds equals 0.00 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` should return `"100 pounds equals 45.36 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` should return `"2.5 pounds equals 1.13 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` should return `"2.20462 pounds equals 1.00 kilogram."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
