---
id: 691b559495c5cb5a37b9b487
title: "Challenge 127: Speed Check"
challengeType: 29
dashedName: challenge-127
---

# --description--

時速マイル（MPH）での走行速度と、時速キロメートル（KPH）での速度制限が与えられたとき、速度超過しているかどうか、警告を受けるか罰金を受けるかを判定してください。

- 1マイルは1.60934キロメートルです。
- 速度制限以下で走行している場合は `"Not Speeding"` を返してください。
- 速度制限を5KPH以下で超過している場合は `"Warning"` を返してください。
- 速度制限を5KPH以上超過している場合は `"Ticket"` を返してください。

# --hints--

`speed_check(30, 70)` should return `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` should return `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` should return `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` should return `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` should return `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` should return `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
