---
id: 6925e2068081f40f549ced1c
title: "Challenge 138: Sum of Divisors"
challengeType: 29
dashedName: challenge-138
---

# --description--

正の整数が与えられたとき、その約数のすべての合計を返してください。

- 約数とは、その数を割り切る整数のことです（余りは`0`です）。
- 各約数は一度だけ数えます。

例えば、`6`が与えられた場合、`12`を返してください。なぜなら`6`の約数は`1`、`2`、`3`、`6`であり、それらの合計が`12`だからです。

# --hints--

`sum_divisors(6)` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` should return `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` should return `56`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` should return `224`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` should return `806`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` should return `23520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
