---
id: 6925e2068081f40f549ced1b
title: "Challenge 137: Snowflake Generator"
challengeType: 29
dashedName: challenge-137
---

# --description--

改行文字 (`\n`) を使って改行を表現している複数行の文字列が与えられたとき、各行を水平方向に鏡像反転し、元の行の末尾に付け加えた新しい文字列を返してください。

- 行を鏡像反転するとは、スペースを含むすべての文字を逆順に並べることです。

例えば、`"* \n *\n* "` が与えられた場合、コンソールには次のように表示されます。

```sh
* 
 *
* 
```

返すべき文字列は `"*  *\n ** \n*  *"` で、コンソールには次のように表示されます。

```sh
*  *
 ** 
*  *
```

与えられた文字列と返す文字列の空白に注意してください。空白を削除しないようにしてください。

# --hints--

`generate_snowflake("* \n *\n* ")` should return `"*  *\n ** \n*  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` should return `"X=~~=X"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` should return `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` should return `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` should return `"*  --  *\n * -- * \n*  --  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
