---
id: 6925e2068081f40f549ced1a
title: "Challenge 136: Markdown Image Parser"
challengeType: 29
dashedName: challenge-136
---

# --description--

Markdown の画像を表す文字列が与えられたとき、それに対応する HTML の文字列を返してください。

Markdown の画像は次の形式です：`"![alt text](image_url)"`。ここで：

- `alt text` は画像の説明（`alt` 属性の値）です。
- `image_url` は画像のソース URL（`src` 属性の値）です。

HTML の `img` タグの文字列を返し、`src` を画像の URL に、`alt` を代替テキストに設定してください。

例えば、`"![Cute cat](cat.png)"` が与えられたら `'<img src="cat.png" alt="Cute cat">'` を返してください。

- タグ、属性の順序、スペース、引用符の使い方が上記と同じであることを確認してください。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを見るにはブラウザのコンソールを確認してください。

# --hints--

`parse_image("![Cute cat](cat.png)")` should return `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` should return `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` should return `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
