---
id: 6925e2068081f40f549ced1d
title: "Challenge 139: Rock, Paper, Scissors"
challengeType: 29
dashedName: challenge-139
---

# --description--

2つの文字列が与えられます。最初の文字列はプレイヤー1を、2番目の文字列はプレイヤー2を表します。じゃんけんの勝者を判定してください。

- 入力される文字列は常に`"Rock"`、`"Paper"`、または`"Scissors"`のいずれかです。
- `"Rock"`は`"Scissors"`に勝ちます。
- `"Paper"`は`"Rock"`に勝ちます。
- `"Scissors"`は`"Paper"`に勝ちます。

Return:

- プレイヤー1が勝った場合は`"Player 1 wins"`を返してください。
- プレイヤー2が勝った場合は`"Player 2 wins"`を返してください。
- 両方のプレイヤーが同じ選択をした場合は`"Tie"`を返してください。

# --hints--

`rock_paper_scissors("Rock", "Rock")` should return `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Rock"), "Tie")`)
}})
```

`rock_paper_scissors("Rock", "Paper")` should return `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Paper"), "Player 2 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Paper")` should return `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Paper"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Rock", "Scissors")` should return `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Scissors"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Scissors")` should return `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Scissors"), "Tie")`)
}})
```

`rock_paper_scissors("Scissors", "Rock")` should return `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Rock"), "Player 2 wins")`)
}})
```

# --seed--

## --seed-contents--

```py
def rock_paper_scissors(player1, player2):

    return player1
```

# --solutions--

```py
def rock_paper_scissors(player1, player2):
    if player1 == player2:
        return "Tie"
    
    if (player1 == "Rock" and player2 == "Scissors") or \
       (player1 == "Paper" and player2 == "Rock") or \
       (player1 == "Scissors" and player2 == "Paper"):
        return "Player 1 wins"
    else:
        return "Player 2 wins"
```
