---
id: 69272dcf1c24b44fd79137c6
title: "Challenge 143: Markdown Italic Parser"
challengeType: 29
dashedName: challenge-143
---

# --description--

マークダウンでイタリック体のテキストが含まれている可能性がある文字列が与えられたとき、対応するHTML文字列を返してください。

- マークダウンのイタリック体のテキストは、単一のアスタリスク（`*`）または単一のアンダースコア（`_`）で始まり終わる任意のテキストです。
- テキストとアスタリスクまたはアンダースコアの間にスペースは入れられませんが、テキスト内にはスペースがあってもかまいません。
- すべてのイタリック体の部分をHTMLの`i`タグに変換して文字列を返してください。

たとえば、`"*This is italic*"`が与えられたら、`"<i>This is italic</i>"`を返します。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

`parse_italics("*This is italic*")` should return `"<i>This is italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is italic*"), "<i>This is italic</i>")`)
}})
```

`parse_italics("_This is also italic_")` should return `"<i>This is also italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_This is also italic_"), "<i>This is also italic</i>")`)
}})
```

`parse_italics("*This is not italic *")` should return `"*This is not italic *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is not italic *"), "*This is not italic *")`)
}})
```

`parse_italics("_ This is also not italic_")` should return `"_ This is also not italic_"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_ This is also not italic_"), "_ This is also not italic_")`)
}})
```

`parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog.")` should return `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_italics(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_italics(markdown):
    pattern = r'(\*|_)([^\s][^]*?[^\s])\1'
    pattern = r'(\*|_)([^\s].*?[^\s])\1'
    return re.sub(pattern, r'<i>\2</i>', markdown)
```
