---
id: 69272dcf1c24b44fd79137c4
title: "Challenge 141: Takeoff Fuel"
challengeType: 29
dashedName: challenge-141
---

# --description--

現在の飛行機の燃料ガロン数と、目的地に到達するために必要な燃料リットル数が与えられたとき、追加で何ガロンの燃料を補給すべきかを求めてください。

- 1ガロンは3.78541リットルです。
- 飛行機にすでに十分な燃料がある場合は、`0`を返してください。
- 補給できるのは整数のガロンのみです。
- 返す数値に小数点は含めないでください。

# --hints--

`fuel_to_add(0, 1)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(0, 1), 1)`)
}})
```

`fuel_to_add(5, 40)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(5, 40), 6)`)
}})
```

`fuel_to_add(10, 30)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(10, 30), 0)`)
}})
```

`fuel_to_add(896, 20500)` should return `4520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(896, 20500), 4520)`)
}})
```

`fuel_to_add(1000, 50000)` should return `12209`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(1000, 50000), 12209)`)
}})
```

# --seed--

## --seed-contents--

```py
def fuel_to_add(current_gallons, required_liters):

    return current_gallons
```

# --solutions--

```py
import math
def fuel_to_add(current_gallons, required_liters):
    liters_per_gallon = 3.78541
    current_liters = current_gallons * liters_per_gallon
    if current_liters >= required_liters:
        return 0
    liters_needed = required_liters - current_liters
    return math.ceil(liters_needed / liters_per_gallon)
```
