---
id: 69306364df283fcaff2e1ad9
title: "Challenge 148: Tire Pressure"
challengeType: 29
dashedName: challenge-148
---

# --description--

車両の4つのタイヤの空気圧をpsi単位で表す4つの数字の配列と、タイヤの最小および最大空気圧をbar単位で表す2つの数字の配列が与えられたとき、各タイヤの状態を表す4つの文字列の配列を返してください。

- 1 bar は14.5038 psiに相当します。

各タイヤについて、次の値を含む配列を返してください。

- 空気圧が許容される最小値を下回っている場合は `"Low"`。
- 空気圧が最小値と最大値の間にある場合は `"Good"`。
- 空気圧が最大値を超えている場合は `"High"`。

# --hints--

`tire_status([32, 28, 35, 29], [2, 3])` should return `["Good", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"])`)
}})
```

`tire_status([32, 28, 35, 30], [2, 2.3])` should return `["Good", "Low", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"])`)
}})
```

`tire_status([29, 26, 31, 28], [2.1, 2.5])` should return `["Low", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"])`)
}})
```

`tire_status([31, 31, 30, 29], [1.5, 2])` should return `["High", "High", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"])`)
}})
```

`tire_status([30, 28, 30, 29], [1.9, 2.1])` should return `["Good", "Good", "Good", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"])`)
}})
```

# --seed--

## --seed-contents--

```py
def tire_status(pressures_psi, range_bar):

    return pressures_psi
```

# --solutions--

```py
def tire_status(pressures_psi, range_bar):
    psi_to_bar = 1 / 14.5038
    min_bar, max_bar = range_bar
    result = []

    for psi in pressures_psi:
        bar = psi * psi_to_bar
        if bar < min_bar:
            result.append("Low")
        elif bar > max_bar:
            result.append("High")
        else:
            result.append("Good")
    return result
```
