---
id: 69373793f5a867f769cde135
title: "Challenge 150: Markdown Unordered List Parser"
challengeType: 29
dashedName: challenge-150
---

# --description--

有効なMarkdownの番号なしリストの文字列が与えられたとき、それに対応するHTMLの文字列を返してください。

番号なしリストは1つ以上のリスト項目で構成されます。有効なリスト項目は独立した行に現れ、以下の条件を満たします。

- ダッシュ（`"-"`）で始まり、
- その後に少なくとも1つのスペースがあり、
- その後にリスト項目のテキストが続きます。

リストは改行文字（`"\n"`）で区切られた単一の文字列として与えられます。項目のテキストに改行文字は含めないでください。

各リスト項目をHTMLの`li`タグで囲み、リスト全体を`ul`タグで囲んでください。

例えば、`"- Item A\n- Item B"`が与えられたら、`"<ul><li>Item A</li><li>Item B</li></ul>"`を返します。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

`parse_unordered_list("- Item A\n- Item B")` should return `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- Item A\\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>")`)
}})
```

`parse_unordered_list("-  JavaScript\n-  Python")` should return `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("-  JavaScript\\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>")`)
}})
```

`parse_unordered_list("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` should return `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- 2 C Flour\\n- 1/2 C Sugar\\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>")`)
}})
```

`parse_unordered_list("- A-1\n- A-2\n- B-1")` should return `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- A-1\\n- A-2\\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_unordered_list(markdown):

    return markdown
```

# --solutions--

```py
def parse_unordered_list(markdown):
    lines = markdown.split("\n")
    list_items = []

    for line in lines:
        list_items.append(f"<li>{line[1:].strip()}</li>")

    return f"<ul>{''.join(list_items)}</ul>"
```
