---
id: 69373793f5a867f769cde138
title: "Challenge 153: Tic-Tac-Toe"
challengeType: 29
dashedName: challenge-153
---

# --description--

3×3の行列（配列の配列）で表された完成した三目並べのゲームが与えられたとき、勝者を判定してください。

- 与えられた行列の各要素は`"X"`か`"O"`のいずれかです。

プレイヤーは、横、縦、または斜めに自分の文字が3つ並んでいれば勝ちです。

Return:

- プレイヤーXが3つ並んでいれば`"X wins"`を返してください。
- プレイヤーOが3つ並んでいれば`"O wins"`を返してください。
- どちらのプレイヤーも3つ並んでいなければ`"Draw"`を返してください。

# --hints--

`tic_tac_toe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` should return `"X wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins")`)
}})
```

`tic_tac_toe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])` should return `"X wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins")`)
}})
```

`tic_tac_toe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw")`)
}})
```

`tic_tac_toe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]])` should return `"O wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]]), "O wins")`)
}})
```

`tic_tac_toe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` should return `"X wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins")`)
}})
```

`tic_tac_toe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def tic_tac_toe(board):

    return board
```

# --solutions--

```py
def tic_tac_toe(board):
    lines = []

    for row in board:
        lines.append(row)

    for c in range(3):
        lines.append([board[0][c], board[1][c], board[2][c]])

    lines.append([board[0][0], board[1][1], board[2][2]])
    lines.append([board[0][2], board[1][1], board[2][0]])

    for line in lines:
        if all(cell == "X" for cell in line):
            return "X wins"
        if all(cell == "O" for cell in line):
            return "O wins"

    return "Draw"
```
