---
id: 6939b873185d8e00d453563c
title: "チャレンジ157：Markdownリンクパーサー"
challengeType: 29
dashedName: challenge-157
---

# --description--

Markdownのリンクの文字列が与えられたとき、それに対応するHTMLの文字列を返してください。

マークダウンのリンクは次の形式です：`"[link_text](link_url)"`。`href`に`link_url`を設定し、タグの内容に`link_text`を使ったHTMLの`a`タグの文字列を返してください。

例えば、`"[freeCodeCamp](https://freecodecamp.org/)"`が与えられたら`'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`を返します。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを見るにはブラウザのコンソールを確認してください。

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` は `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` は `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` は `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
