---
id: 6939b873185d8e00d453563f
title: "Challenge 160: Knight Moves"
challengeType: 29
dashedName: challenge-160
---

# --description--

チェス盤上のナイトの位置が与えられたとき、ナイトが移動できる有効なマスの数を返してください。

標準的なチェス盤は8×8で、列は左から右へ`A`から`H`まで、行は下から上へ`1`から`8`までラベル付けされています。見た目は次のようになります：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

ナイトは「L字型」に動きます：一方向（横または縦）に2マス、そして直角方向に1マス動きます。

つまり、ナイトは最大で8つの位置に移動できますが、盤の端に近い場合は移動可能なマスが少なくなります。例えば、ナイトが`A1`にいる場合、移動できるのは`B3`か`C2`だけです。

# --hints--

`knight_moves("A1")` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("A1"), 2)`)
}})
```

`knight_moves("D4")` should return `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("D4"), 8)`)
}})
```

`knight_moves("G6")` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("G6"), 6)`)
}})
```

`knight_moves("B8")` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("B8"), 3)`)
}})
```

`knight_moves("H3")` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("H3"), 4)`)
}})
```

# --seed--

## --seed-contents--

```py
def knight_moves(position):

    return position
```

# --solutions--

```py
def knight_moves(position):
    col = ord(position[0]) - ord('A')
    row = int(position[1]) - 1

    moves = [
        (2, 1), (2, -1),
        (-2, 1), (-2, -1),
        (1, 2), (1, -2),
        (-1, 2), (-1, -2)
    ]

    valid = 0

    for dx, dy in moves:
        new_col = col + dx
        new_row = row + dy

        if 0 <= new_col < 8 and 0 <= new_row < 8:
            valid += 1

    return valid
```
