---
id: 6939b873185d8e00d453563e
title: "Challenge 159: Integer Hypotenuse"
challengeType: 29
dashedName: challenge-159
---

# --description--

2つの正の整数が直角三角形の2つの脚（短い2辺）の長さを表すとき、斜辺の長さが整数かどうかを判定してください。

斜辺の長さは、2つの脚の長さの2乗を足し合わせ、その合計の平方根を取ることで計算されます（a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>）。

# --hints--

`is_integer_hypotenuse(3, 4)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(3, 4), True)`)
}})
```

`is_integer_hypotenuse(2, 3)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(2, 3), False)`)
}})
```

`is_integer_hypotenuse(5, 12)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(5, 12), True)`)
}})
```

`is_integer_hypotenuse(10, 10)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(10, 10), False)`)
}})
```

`is_integer_hypotenuse(780, 1040)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(780, 1040), True)`)
}})
```

`is_integer_hypotenuse(250, 333)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(250, 333), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_integer_hypotenuse(a, b):

    return a
```

# --solutions--

```py
import math
def is_integer_hypotenuse(a, b):
    total = a*a + b*b
    c = math.isqrt(total)
    return c*c == total
```
