---
id: 694596b0585c11170ac7c7fa
title: "Challenge 162: Energy Consumption"
challengeType: 29
dashedName: challenge-162
---

# --description--

運動中に消費したカロリー数と、その運動中に電子機器が使用したワット時数が与えられたとき、どちらがより多くのエネルギーを使ったかを判定してください。

比較するために、以下の換算を使って両方の値をジュールに変換してください。

- 1カロリーは4184ジュールに相当します。
- 1ワット時は3600ジュールに相当します。

Return:

- 運動の方が多くエネルギーを使った場合は `"Workout"` を返してください。
- デバイスの方が多くエネルギーを使った場合は `"Devices"` を返してください。
- 両方が同じ量のエネルギーを使った場合は `"Equal"` を返してください。

# --hints--

`compare_energy(250, 50)` should return `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` should return `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` should return `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` should return `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` should return `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` should return `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
