---
id: 694596b0585c11170ac7c7fc
title: "Challenge 164: Markdown Inline Code Parser"
challengeType: 29
dashedName: challenge-164
---

# --description--

1つ以上のインラインコードブロックを含むMarkdownの文字列が与えられたとき、それに対応するHTML文字列を返してください。

Markdownのインラインコードブロックは、コードブロックのテキストの開始と終了に単一のバックティック（``` ` ```）を使います。

与えられた文字列のすべてのコードブロックをHTMLの`code`タグに変換して返してください。

例えば、文字列```"Use `let` to declare the variable."```, return `"Use <code>let</code> to declare the variable."`が与えられた場合です。

注意：コンソールは文字列内のHTMLタグをログに表示しない場合があります。タグを含むログを確認するにはブラウザのコンソールをチェックしてください。

# --hints--

```parse_inline_code("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` to declare the variable."), "Use <code>let</code> to declare the variable.")`)
}})
```

```parse_inline_code("Use `let` or `const` to declare a variable.")``` should return `"Use <code>let</code> or <code>const</code> to declare a variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` or \`const\` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.")`)
}})
```

```parse_inline_code("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Run \`npm install\` then \`npm start\`."), "Run <code>npm install</code> then <code>npm start</code>.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_inline_code(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_inline_code(markdown):
    return re.sub(r"`([^`]+)`", r"<code>\1</code>", markdown)
```
