---
id: 696655d24b614176d4c9b78b
title: "Challenge 168: Scaled Image"
challengeType: 29
dashedName: challenge-168
---

# --description--

画像の幅と高さを表す文字列と、画像をスケールする数値が与えられたとき、スケールされた幅と高さを返してください。

- 入力の文字列は`"WxH"`の形式です。例えば、`"800x600"`です。
- スケールは幅と高さに乗算する数値です。

スケールされた寸法を同じ`"WxH"`形式で返してください。

# --hints--

`scale_image("800x600", 2)` should return `"1600x1200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("800x600", 2), "1600x1200")`)
}})
```

`scale_image("100x100", 10)` should return `"1000x1000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("100x100", 10), "1000x1000")`)
}})
```

`scale_image("1024x768", 0.5)` should return `"512x384"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("1024x768", 0.5), "512x384")`)
}})
```

`scale_image("300x200", 1.5)` should return `"450x300"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("300x200", 1.5), "450x300")`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_image(size, scale):

    return size
```

# --solutions--

```py
def scale_image(size, scale):
    width, height = map(int, size.split("x"))

    new_width = round(width * scale)
    new_height = round(height * scale)

    return f"{new_width}x{new_height}"
```
