---
id: 696655d24b614176d4c9b78d
title: "Challenge 170: Odd or Even Day"
challengeType: 29
dashedName: challenge-170
---

# --description--

Unixエポックからのミリ秒数で表されたタイムスタンプが与えられたとき、次を返してください：

- そのタイムスタンプの日付が奇数日なら `"odd"` を返す。
- そのタイムスタンプの日付が偶数日なら `"even"` を返す。

例えば、`1769472000000`（2026年1月27日のタイムスタンプ）が与えられた場合、日付の数字（`"odd"`）が奇数なので `27` を返します。

注意：タイムスタンプはミリ秒単位で、ローカル時間ではなくUTCタイムゾーンの日付を使う必要があります。

# --hints--

`odd_or_even_day(1769472000000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
