---
id: 69738771fb5a7b8b24cca2a5
title: "Challenge 179: Pocket Change"
challengeType: 29
dashedName: challenge-179
---

# --description--

ポケットの中のコインを表す整数の配列が与えられます。各整数はコインの価値をセント単位で表しています。合計金額を `"$D.CC"` の形式で返してください。

- 100セントは1ドルに相当します。
- 戻り値では、1ドル未満の金額には先頭にゼロを付け、セントは常に2桁で表示してください。

# --hints--

`count_change([25, 10, 5, 1])` should return `"$0.41"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1]), "$0.41")`)
}})
```

`count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` should return `"$1.43"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43")`)
}})
```

`count_change([100, 25, 100, 1000, 5, 500, 2000, 25])` should return `"$37.55"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55")`)
}})
```

`count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` should return `"$0.70"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70")`)
}})
```

`count_change([1])` should return `"$0.01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([1]), "$0.01")`)
}})
```

`count_change([25, 25, 25, 25])` should return `"$1.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 25, 25, 25]), "$1.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_change(change):

    return change
```

# --solutions--

```py
def count_change(change):
    total_cents = sum(change)
    dollars = total_cents // 100
    cents = total_cents % 100
    return f"${dollars}.{cents:02d}"
```
