---
id: 697a49e9860d24853adef67f
title: "チャレンジ 194: 冬季大会 第15日目: フリースタイルスキー"
challengeType: 29
dashedName: challenge-194
---

# --description--

2語からなるトリック名が、フリースタイルスキーの有効なトリック名かどうか判定してください。

トリックが有効なのは、最初の単語が有効な最初の単語のリストにあり、2番目の単語が有効な2番目の単語のリストにある場合です。

- 2語は1つのスペースで区切られます。

有効な最初の単語:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

有効な2番目の単語:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`is_valid_trick("Polar Vortex")` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Polar Vortex"), True)`)
}})
```

`is_valid_trick("Solar Icequake")` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Solar Icequake"), True)`)
}})
```

`is_valid_trick("Thunder Blizzard")` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Thunder Blizzard"), True)`)
}})
```

`is_valid_trick("Phantom Frostbite")` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Phantom Frostbite"), True)`)
}})
```

`is_valid_trick("Ghost Avalanche")` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Ghost Avalanche"), True)`)
}})
```

`is_valid_trick("Snowstorm Shadow")` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Snowstorm Shadow"), False)`)
}})
```

`is_valid_trick("Solar Sky")` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Solar Sky"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_trick(trick_name):

    return trick_name
```

# --solutions--

```py
def is_valid_trick(trick_name):
    valid_first = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"]
    valid_second = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"]

    words = trick_name.split(" ")
    first, second = words
    return first in valid_first and second in valid_second
```
