---
id: 697a49e9860d24853adef67c
title: "Challenge 191: 2026 Winter Games Day 12: Bobsled"
challengeType: 29
dashedName: challenge-191
---

# --description--

ボブスレーチームの選手の体重を表す配列とボブスレー自体の重さを表す数値が与えられたとき、チームがレースに出場できるかどうかを判定してください。

- 配列の長さはチームの人数を決めます：1人、2人、または4人チームです。
- すべての体重の値はキログラム（kg）単位です。

ボブスレー（そり単体）の最小重量は以下の通りです：

- 1人チームの場合は162kg
- 2人チームの場合は170kg
- 4人チームの場合は210kg

ボブスレーの総重量（選手の体重＋そりの重さ）は以下を超えてはいけません：

- 1人チームの場合は247kg
- 2人チームの場合は390kg
- 4人チームの場合は630kg

チームがすべての条件を満たす場合は`"Eligible"`を、1つ以上の条件を満たさない場合は`"Not Eligible"`を返してください。

# --hints--

`check_eligibility([78], 165)` should return `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([78], 165), "Eligible")`)
}})
```

`check_eligibility([80], 160)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 160), "Not Eligible")`)
}})
```

`check_eligibility([80], 170)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 170), "Not Eligible")`)
}})
```

`check_eligibility([85, 90], 170)` should return `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 90], 170), "Eligible")`)
}})
```

`check_eligibility([85, 95], 168)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 95], 168), "Not Eligible")`)
}})
```

`check_eligibility([112, 97], 185)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([112, 97], 185), "Not Eligible")`)
}})
```

`check_eligibility([110, 102, 90, 106], 222)` should return `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([110, 102, 90, 106], 222), "Eligible")`)
}})
```

`check_eligibility([106, 99, 90, 88], 205)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 90, 88], 205), "Not Eligible")`)
}})
```

`check_eligibility([106, 99, 103, 96], 227)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 103, 96], 227), "Not Eligible")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_eligibility(athlete_weights, sled_weight):

    return athlete_weights
```

# --solutions--

```py
def check_eligibility(athlete_weights, sled_weight):
    team_size = len(athlete_weights)

    rules = {
        1: {"min_sled": 162, "max_total": 247},
        2: {"min_sled": 170, "max_total": 390},
        4: {"min_sled": 210, "max_total": 630}
    }

    min_sled = rules[team_size]["min_sled"]
    max_total = rules[team_size]["max_total"]

    total_weight = sum(athlete_weights) + sled_weight

    if sled_weight < min_sled:
        return "Not Eligible"

    if total_weight > max_total:
        return "Not Eligible"

    return "Eligible"
```
