---
id: 697a49e6ff50d756c9b69365
title: "Challenge 188: 2026 Winter Games Day 9: Skeleton"
challengeType: 29
dashedName: challenge-188
---

# --description--

スケルトントラックのカーブを表す文字列が与えられたとき、そのトラックの難易度を判定してください。

- 与えられる文字列は以下の文字のみで構成されます：
  - `"L"` は左カーブを表します
  - `"R"` は右カーブを表します
  - `"S"` は直線区間を表します

- 方向転換がある場合は15点加算されます（`"L"`の後に`"R"`が続く、またはその逆の場合）。
- その他のカーブはそれぞれ5点加算されます（その他の`"L"`または`"R"`の文字）。
- 直線区間は0点です。

トラックの難易度は合計得点に基づきます。以下を返してください：

- 合計が0～100の場合は`"Easy"`
- 合計が101～200の場合は`"Medium"`
- 合計が200を超える場合は`"Hard"`

# --hints--

`get_difficulty("SLSLLSRRLSRLRL")` should return `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLSLLSRRLSRLRL"), "Easy")`)
}})
```

`get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` should return `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard")`)
}})
```

`get_difficulty("SRRRRLSLLRLRSSRLSRL")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRRRRLSLLRLRSSRLSRL"), "Medium")`)
}})
```

`get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` should return `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard")`)
}})
```

`get_difficulty("SLLSSLRLSLSLRSLSSLRL")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium")`)
}})
```

`get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` should return `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_difficulty(track):

    return track
```

# --solutions--

```py
def get_difficulty(track):
  score = 0

  for i in range(len(track)):
    current = track[i]
    previous = track[i - 1] if i > 0 else None

    if current == 'S':
      score += 0
    elif current == 'L' or current == 'R':
      if previous and previous != 'S' and previous != current:
        score += 15
      else:
        score += 5

  if score <= 100:
    return 'Easy'
  elif score <= 200:
    return 'Medium'
  else:
    return 'Hard'
```
