---
id: 697a49e6ff50d756c9b69366
title: "Challenge 189: 2026 Winter Games Day 10: Alpine Skiing"
challengeType: 29
dashedName: challenge-189
---

# --description--

スキー場の垂直落差、水平距離、タイプが与えられたとき、その斜面の難易度評価を判定してください。

評価を決めるには：

- 落差を距離で割って斜面の傾斜を計算します。
- 次に、斜面のタイプに基づいて調整された傾斜を計算します：
  - `"Downhill"`: 傾斜に1.2を乗算する
  - `"Slalom"`: 傾斜に0.9を乗算する
  - `"Giant Slalom"`: 傾斜に1.0を乗算する

Return:

- 調整された傾斜が0.1以下の場合は`"Green"`
- 調整された傾斜が0.1より大きく0.25以下の場合は`"Blue"`
- 調整された傾斜が0.25より大きい場合は`"Black"`


# --hints--

`get_hill_rating(95, 900, "Slalom")` should return `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(95, 900, "Slalom"), "Green")`)
}})
```

`get_hill_rating(620, 2800, "Downhill")` should return `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(620, 2800, "Downhill"), "Black")`)
}})
```

`get_hill_rating(420, 1680, "Giant Slalom")` should return `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(420, 1680, "Giant Slalom"), "Blue")`)
}})
```

`get_hill_rating(250, 3000, "Downhill")` should return `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(250, 3000, "Downhill"), "Green")`)
}})
```

`get_hill_rating(110, 900, "Slalom")` should return `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(110, 900, "Slalom"), "Blue")`)
}})
```

`get_hill_rating(380, 1500, "Giant Slalom")` should return `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(380, 1500, "Giant Slalom"), "Black")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_hill_rating(drop, distance, hill_type):

    return drop
```

# --solutions--

```py
def get_hill_rating(drop, distance, hill_type):
    steepness = drop / distance

    if hill_type == "Downhill":
        steepness *= 1.2
    elif hill_type == "Slalom":
        steepness *= 0.9
    elif hill_type == "Giant Slalom":
        steepness *= 1.0

    if steepness <= 0.1:
        return "Green"
    elif steepness <= 0.25:
        return "Blue"
    else:
        return "Black"
```
