---
id: 697a49e9860d24853adef67d
title: "Challenge 192: 2026 Winter Games Day 13: Nordic Combined"
challengeType: 29
dashedName: challenge-192
---

# --description--

選手のジャンプスコアの配列が与えられたとき、ノルディック複合のクロスカントリー部分のスタート遅延時間を計算してください。

ジャンプスコアが最も高い選手が最初にスタートします（遅延時間は0秒）。他の選手は、ジャンプスコアが最も高い選手との差に応じて遅れてスタートします。

各選手の遅延時間を計算するには、最も高いジャンプスコアからその選手のジャンプスコアを引き、その結果に1.5を乗算してください。遅延時間は最も近い整数に切り上げます。

# --hints--

`calculate_start_delays([120, 110, 125])` should return `[8, 23, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([120, 110, 125]), [8, 23, 0])`)
}})
```

`calculate_start_delays([118, 125, 122, 120])` should return `[11, 0, 5, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([118, 125, 122, 120]), [11, 0, 5, 8])`)
}})
```

`calculate_start_delays([100, 105, 95, 110, 120, 115, 108])` should return `[30, 23, 38, 15, 0, 8, 18]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18])`)
}})
```

`calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` should return `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_start_delays(jump_scores):

    return jump_scores
```

# --solutions--

```py
import math
def calculate_start_delays(jump_scores):
    best_jump = max(jump_scores)
    return [math.ceil((best_jump - score) * 1.5) for score in jump_scores]
```
