---
id: 698a1a863194f1f4e63f6460
title: "Challenge 209: Element Size"
challengeType: 29
dashedName: challenge-209
---

# --description--

ウィンドウサイズ、ビューポート幅の `"vw"` 単位での要素の幅、およびビューポート高さの `"vh"` 単位での要素の高さが与えられたとき、要素のサイズをピクセル単位で求めてください。

- 与えられるウィンドウサイズと返される要素サイズは、例えば `"width x height"`、`"1200 x 800"` の形式の文字列です。

- `"vw"` 単位はウィンドウ幅のパーセントです。例えば `"50vw"` はウィンドウ幅の 50% です。
- `"vh"` 単位はウィンドウ高さのパーセントです。例えば `"50vh"` はウィンドウ高さの 50% です。

# --hints--

`get_element_size("1200 x 800", "50vw", "50vh")` should return `"600 x 400"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "50vw", "50vh"), "600 x 400")`)
}})
```

`get_element_size("320 x 480", "25vw", "50vh")` should return `"80 x 240"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("320 x 480", "25vw", "50vh"), "80 x 240")`)
}})
```

`get_element_size("1000 x 500", "7vw", "3vh")` should return `"70 x 15"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1000 x 500", "7vw", "3vh"), "70 x 15")`)
}})
```

`get_element_size("1920 x 1080", "95vw", "100vh")` should return `"1824 x 1080"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1920 x 1080", "95vw", "100vh"), "1824 x 1080")`)
}})
```

`get_element_size("1200 x 800", "0vw", "0vh")` should return `"0 x 0"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "0vw", "0vh"), "0 x 0")`)
}})
```

`get_element_size("1440 x 900", "100vw", "114vh")` should return `"1440 x 1026"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1440 x 900", "100vw", "114vh"), "1440 x 1026")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_element_size(window_size, element_vw, element_vh):

    return window_size
```

# --solutions--

```py
def get_element_size(window_size, element_vw, element_vh):
    window_width, window_height = map(int, window_size.split(" x "))

    vw = int(element_vw.replace("vw", ""))
    vh = int(element_vh.replace("vh", ""))

    width_px = round((vw / 100) * window_width)
    height_px = round((vh / 100) * window_height)

    return f"{width_px} x {height_px}"
```
