---
id: 698a1a863194f1f4e63f645f
title: "Challenge 208: Trail Traversal"
challengeType: 29
dashedName: challenge-208
---

# --description--

トレイルマップを表す文字列の配列が与えられたとき、目的地に到達するために必要な移動の文字列を返してください。

与えられた文字列には以下の値が含まれます：

- `"C"`: 現在の位置
- `"G"`: 目的地
- `"T"`: 通行可能なトレイルの部分
- `"-"`: 通行不可能なマップの部分

現在地から目的地までトレイルに沿って進むために必要な移動を表す文字列を返してください。移動は以下の通りです：

- `"R"`: 右への移動
- `"D"`: 下への移動
- `"L"`: 左への移動
- `"U"`: 上への移動

- 現在地から目的地まで、分岐のない連続したトレイルが必ず1つだけ存在します。
- 各トレイルの位置は、最大で2つの通行可能な場所と接しています。

例えば、次のような場合：

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

Return `"RDDRDD"`.

# --hints--

`navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` should return `"RDDRDD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD")`)
}})
```

`navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` should return `"RRUUURR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR")`)
}})
```

`navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` should return `"DLDDRRRRD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD")`)
}})
```

`navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` should return `"RRRDDL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL")`)
}})
```

`navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` should return `"ULLLUUURRRRRRDDDR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate_trail(map):

    return map
```

# --solutions--

```py
def navigate_trail(map):
    rows = len(map)
    cols = len(map[0])

    directions = [
        (0, 1, "R"),
        (1, 0, "D"),
        (0, -1, "L"),
        (-1, 0, "U"),
    ]

    for r in range(rows):
        for c in range(cols):
            if map[r][c] == "C":
                row, col = r, c

    result = ""
    prev = None

    while map[row][col] != "G":
        for dr, dc, move in directions:
            new_row = row + dr
            new_col = col + dc

            if (
                0 <= new_row < rows and
                0 <= new_col < cols and
                (map[new_row][new_col] == "T" or
                 map[new_row][new_col] == "G") and
                (not prev or (new_row, new_col) != prev)
            ):
                result += move
                prev = (row, col)
                row, col = new_row, new_col
                break

    return result
```
