---
id: 698a1a73ade5ac0e19180fa0
title: "Challenge 197: Blood Type Compatibility"
challengeType: 29
dashedName: challenge-197
---

# --description--

ドナーの血液型とレシピエントの血液型が与えられたとき、ドナーがレシピエントに血液を提供できるかどうかを判定してください。

各血液型は以下で構成されます：

- 1文字のアルファベット：`"A"`、`"B"`、`"AB"`、または `"O"`
- そして Rh因子：`"+"` または `"-"`

血液型は有効なアルファベットのいずれかにRh因子が続く形になります。例えば、`"AB+"` と `"O-"` は有効な血液型です。

Letter Rules:

- `"O"` は他のアルファベット型に提供できます。
- `"A"` は `"A"` と `"AB"` に提供できます。
- `"B"` は `"B"` と `"AB"` に提供できます。
- `"AB"` は `"AB"` のみへ提供できます。

Rh Rules:

- マイナス（`"-"`）は `"-"` と `"+"` の両方に提供できます。
- プラス（`"+"`）は `"+"` のみへ提供できます。

ドナーがレシピエントに血液を提供できるには、アルファベットのルールとRhのルールの両方を満たす必要があります。

# --hints--

`can_donate("B+", "B+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B+", "B+"), True)`)
}})
```

`can_donate("O-", "AB-")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "AB-"), True)`)
}})
```

`can_donate("O+", "A-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "A-"), False)`)
}})
```

`can_donate("A+", "AB+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A+", "AB+"), True)`)
}})
```

`can_donate("A-", "B-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A-", "B-"), False)`)
}})
```

`can_donate("B-", "AB+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "AB+"), True)`)
}})
```

`can_donate("B-", "A+")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "A+"), False)`)
}})
```

`can_donate("O-", "O+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "O+"), True)`)
}})
```

`can_donate("O+", "O-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "O-"), False)`)
}})
```

`can_donate("AB+", "AB-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("AB+", "AB-"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_donate(donor, recipient):

    return donor
```

# --solutions--

```py
def can_donate(donor, recipient):
    donor_type = donor[:-1]
    donor_rh = donor[-1]

    recipient_type = recipient[:-1]
    recipient_rh = recipient[-1]

    abo_compatibility = {
        "O": ["A", "B", "AB", "O"],
        "A": ["A", "AB"],
        "B": ["B", "AB"],
        "AB": ["AB"]
    }

    abo_match = recipient_type in abo_compatibility[donor_type]
    rh_match = (
        donor_rh == "-" or
        (donor_rh == "+" and recipient_rh == "+")
    )

    return abo_match and rh_match
```
