---
id: 698a1a73ade5ac0e19180fa4
title: "Challenge 201: Matrix Shift"
challengeType: 29
dashedName: challenge-201
---

# --description--

数値のマトリックス（配列の配列）と整数が与えられたとき、マトリックス内のすべての値を指定された量だけシフトしてください。

- 正のシフトは値を右に移動させます。
- 負のシフトは値を左に移動させます。

値はラップ（循環）します：

- マトリックスを1つの連続したシーケンスとして扱います。
- 値が行の終わりを超えて移動すると、次の行の先頭から続きます。
- 値がマトリックスの最後の位置を超えると、最初の位置に戻ります。
- 左にシフトするときも同様に逆方向で適用されます。

例えば、次のような場合：

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

`1` のシフトで、すべての数字を右に1つ移動させます：

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shift_matrix([[1, 2, 3], [4, 5, 6]], 1)` should return `[[6, 1, 2], [3, 4, 5]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6]], -1)` should return `[[2, 3, 4], [5, 6, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` should return `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` should return `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` should return `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` should return `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_matrix(matrix, shift):

    return matrix
```

# --solutions--

```py
def shift_matrix(matrix, shift):
    rows = len(matrix)
    cols = len(matrix[0])
    flat = [num for row in matrix for num in row]
    n = len(flat)
    s = shift % n
    shifted = flat[-s:] + flat[:-s] if s != 0 else flat[:]
    result = []
    for i in range(rows):
        result.append(shifted[i*cols:(i+1)*cols])

    return result
```
