---
id: 698a1a863194f1f4e63f645e
title: "Challenge 207: Smallest Gap"
challengeType: 29
dashedName: challenge-207
---

# --description--

文字列が与えられたとき、最も文字数が少ない（最小のギャップ）同じ文字の間にある部分文字列を返してください。

- 常に少なくとも1組の一致する文字があります。
- 返す部分文字列には一致する文字自体は含めません。
- もし2つ以上のギャップが同じ長さの場合は、最初のものの文字を返してください。

例えば、`"ABCDAC"`が与えられた場合、`"DA"`を返します。理由は：

- 文字列内で繰り返されるのは`"A"`と`"C"`だけです。
- 2つの`"A"`文字の間の文字数は3で、`"C"`文字の間は2です。
- したがって、2つの`"C"`文字の間の文字列を返します。

# --hints--

`smallest_gap("ABCDAC")` should return `"DA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` should return `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` should return `"#q6e&rkf(p"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` should return `""`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` should return `"fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
