---
id: 6994cff2290543b3aec9f50f
title: "Challenge 210: HSL Validator"
challengeType: 29
dashedName: challenge-210
---

# --description--

文字列が有効な CSS `hsl()` カラー値かどうかを判定してください。

- 有効な HSL 値は `"hsl(h, s%, l%)"` の形式である必要があります。具体的には：
  - `h`（色相）は 0 から 360（両端含む）の数値でなければなりません。
  - `s`（彩度）は 0% から 100% のパーセンテージでなければなりません。
  - `l`（明度）は 0% から 100% のパーセンテージでなければなりません。

- スペースは以下の場所にのみ許可されます：
  - 開き括弧の直後
  - カンマの前後いずれかまたは両方
  - 閉じ括弧の前後いずれかまたは両方

- オプションで、値の末尾にセミコロン（`";"`）を付けることができます。

例えば、`"hsl(240, 50%, 50%)"` は有効な HSL 値です。

# --hints--

`is_valid_hsl("hsl(240, 50%, 50%)")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(240, 50%, 50%)"), True)`)
}})
```

`is_valid_hsl("hsl( 200 , 50% , 75% )")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl( 200 , 50% , 75% )"), True)`)
}})
```

`is_valid_hsl("hsl(99,60%,80%);")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(99,60%,80%);"), True)`)
}})
```

`is_valid_hsl("hsl(0, 0%, 0%) ;")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(0, 0%, 0%) ;"), True)`)
}})
```

`is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;"), True)`)
}})
```

`is_valid_hsl("hsl(361, 50%, 80%)")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(361, 50%, 80%)"), False)`)
}})
```

`is_valid_hsl("hsl(300, 101%, 70%)")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(300, 101%, 70%)"), False)`)
}})
```

`is_valid_hsl("hsl(200, 55%, 75)")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(200, 55%, 75)"), False)`)
}})
```

`is_valid_hsl("hsl (80, 20%, 10%)")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl (80, 20%, 10%)"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_hsl(hsl):

    return hsl
```

# --solutions--

```py
import re
def is_valid_hsl(hsl):
    pattern = r"^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$"
    
    match = re.match(pattern, hsl)
    if not match:
        return False

    hue = int(match.group(1))
    sat = int(match.group(2))
    light = int(match.group(3))

    if not (0 <= hue <= 360):
        return False
    if not (0 <= sat <= 100):
        return False
    if not (0 <= light <= 100):
        return False

    return True
```
