---
id: 699c8e045ee7cb94ed2322d6
title: "Challenge 215: Parking Fee Calculator"
challengeType: 29
dashedName: challenge-215
---

# --description--

車を駐車した時間と引き取った時間を表す2つの文字列が与えられたとき、駐車料金を計算してください。

- 与えられる文字列は24時間制の`"HH:MM"`形式です。例えば`"14:00"`は午後2時を表します。
- 最初の文字列は車を駐車した時間、2つ目は引き取った時間です。
- 引き取り時間が駐車時間よりも早い場合は、駐車が深夜をまたいで翌日にまたがっていることを意味します。

Fee rules:

- 駐車1時間ごとに3ドルかかります。
- 端数の時間は切り上げて次の1時間として計算します。
- 駐車が深夜（午前0時）をまたぐ場合は、追加で10ドルの深夜料金がかかります。
- 最小料金は5ドルです（合計が5ドル未満の場合にのみ適用されます）。

合計料金を`"$cost"`、`"$5"`の形式で返してください。

# --hints--

`calculate_parking_fee("09:00", "11:00")` should return `"$6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("09:00", "11:00"), "$6")`)
}})
```

`calculate_parking_fee("10:00", "10:30")` should return `"$5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("10:00", "10:30"), "$5")`)
}})
```

`calculate_parking_fee("08:10", "10:45")` should return `"$9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("08:10", "10:45"), "$9")`)
}})
```

`calculate_parking_fee("14:40", "23:10")` should return `"$27"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("14:40", "23:10"), "$27")`)
}})
```

`calculate_parking_fee("18:15", "01:30")` should return `"$34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("18:15", "01:30"), "$34")`)
}})
```

`calculate_parking_fee("11:11", "11:10")` should return `"$82"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("11:11", "11:10"), "$82")`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_parking_fee(park_time, pickup_time):

    return park_time
```

# --solutions--

```py
import math

def calculate_parking_fee(park_time, pickup_time):
    def to_minutes(time_str):
        hours, minutes = map(int, time_str.split(":"))
        return hours * 60 + minutes

    entry_minutes = to_minutes(park_time)
    exit_minutes = to_minutes(pickup_time)

    overnight = False

    if exit_minutes < entry_minutes:
        overnight = True
        total_minutes = (24 * 60 - entry_minutes) + exit_minutes
    else:
        total_minutes = exit_minutes - entry_minutes

    hours = math.ceil(total_minutes / 60)

    cost = hours * 3

    if overnight:
        cost += 10

    if cost < 5:
        cost = 5

    return f"${cost}"
```
