---
id: 699c8e045ee7cb94ed2322dd
title: "Challenge 222: Equinox Shadows"
challengeType: 29
dashedName: challenge-222
---

# --description--

今日は春分の日で、太陽が赤道の真上にあり、正午には真上にあります。ある時刻が与えられたとき、4フィートの垂直な棒が作る影を求めてください。

- 時刻は`"HH:MM"`の24時間表記の文字列です（例えば、`"15:00"`は午後3時です）。
- 時刻は30分刻みでのみ与えられます。

Rules:

- 太陽は午前6時に`"east"`から昇り、午後6時に`"west"`に沈みます。
- 影は常に太陽の反対方向を指します。
- 影の長さ（フィート単位）は正午からの時間差の3乗です。
- 日の出前（午前6時前）、日没後（午後6時以降）、および正午には影はありません。

Return:

- 影が存在する場合は`"(length)ft (direction)"`を返してください。例えば、`"8ft west"`です。
- それ以外の場合は`"No shadow"`を返してください。

例えば、`"10:00"`が与えられた場合、午前10時は正午から2時間離れているので2<sup>3</sup> = 8フィートで、太陽が東にあるため影は西を指すので`"8ft west"`を返します。

# --hints--

`get_shadow("10:00")` should return `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` should return `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` should return `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` should return `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` should return `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
