---
id: 699c8e045ee7cb94ed2322d8
title: "Challenge 217: Captured Chess Pieces"
challengeType: 29
dashedName: challenge-217
---

# --description--

チェスの盤上にまだある駒を表す文字列の配列が与えられたとき、相手に捕獲された駒の合計価値を計算してください。

チェスでは、最初に16個の駒を持っています：

| 駒     | 略称           | 数量     | 価値  |
| ------ | -------------- | -------- | ----- |
| ポーン | `"P"`          | 8        | 1     |
| ルーク | `"R"`          | 2        | 5     |
| ナイト | `"N"`          | 2        | 3     |
| ビショップ | `"B"`          | 2        | 3     |
| クイーン | `"Q"`          | 1        | 9     |
| キング | `"K"`          | 1        | 0     |

- 与えられた配列には上記の略称のみが含まれます。
- 16個の駒のうち、配列に含まれていないものはすべて捕獲されています。
- 捕獲されたすべての駒の合計価値を返してください。ただし…
- キングが捕獲されている場合は、`"Checkmate"`を返してください。

# --hints--

`get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` should return `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8)`)
}})
```

`get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"])` should return `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"]), 26)`)
}})
```

`get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` should return `16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16)`)
}})
```

`get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4)`)
}})
```

`get_captured_value(["P", "K"])` should return `38`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "K"]), 38)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` should return `"Checkmate"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_captured_value(pieces):

    return pieces
```

# --solutions--

```py
def get_captured_value(pieces):
    if "K" not in pieces:
        return "Checkmate"

    values = {"P":1, "R":5, "N":3, "B":3, "Q":9, "K":0}
    remaining = "PPPPPPPPRRNNBBQK"

    for p in pieces:
        remaining = remaining.replace(p, "", 1)

    return sum(values[p] for p in remaining)
```
