---
id: 69a890af247de743333bd4d1
title: "Challenge 228: Movie Night"
challengeType: 29
dashedName: challenge-228
---

# --description--

曜日を表す文字列、上映時間を表す別の文字列、そしてチケットの枚数を表す整数が与えられたとき、その上映の映画チケットの合計金額を返してください。

与えられる曜日は次のいずれかです：

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

上映時間は `"H:MMam"` または `"H:MMpm"` の形式で与えられます。例えば `"10:00am"` や `"10:00pm"` のようにです。

次のルールに従い、合計金額を `"$D.CC"` の形式で返してください：

- 週末（金曜日から日曜日）：チケット1枚あたり12.00ドル。
- 平日（月曜日から木曜日）：チケット1枚あたり10.00ドル。
- マチネ（午後5時前）：チケット1枚あたり2.00ドル割引（火曜日を除く）。
- 火曜日：すべてのチケットが1枚5.00ドル。

# --hints--

`get_movie_night_cost("Saturday", "10:00pm", 1)` should return `"$12.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Saturday", "10:00pm", 1), "$12.00")`)
}})
```

`get_movie_night_cost("Sunday", "10:00am", 1)` should return `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Sunday", "10:00am", 1), "$10.00")`)
}})
```

`get_movie_night_cost("Tuesday", "7:20pm", 2)` should return `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "7:20pm", 2), "$10.00")`)
}})
```

`get_movie_night_cost("Wednesday", "5:40pm", 3)` should return `"$30.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Wednesday", "5:40pm", 3), "$30.00")`)
}})
```

`get_movie_night_cost("Monday", "11:50am", 4)` should return `"$32.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Monday", "11:50am", 4), "$32.00")`)
}})
```

`get_movie_night_cost("Friday", "4:30pm", 5)` should return `"$50.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Friday", "4:30pm", 5), "$50.00")`)
}})
```

`get_movie_night_cost("Tuesday", "11:30am", 1)` should return `"$5.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "11:30am", 1), "$5.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):

    return day
```

# --solutions--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):
    if day == "Tuesday":
        price = 5
    else:
        if day in ["Friday", "Saturday", "Sunday"]:
            price = 12
        else:
            price = 10

        hour = int(showtime.split(":")[0])
        period = showtime[-2:]

        hour24 = hour

        if period == "pm" and hour != 12:
            hour24 += 12
        if period == "am" and hour == 12:
            hour24 = 0

        if hour24 < 17:
            price -= 2

    total = price * number_of_tickets

    return f"${total:.2f}"
```
