---
id: 69b1028d6e265413d0198a2f
title: "Challenge 236: Browser History"
challengeType: 29
dashedName: challenge-236
---

# --description--

ブラウザのコマンドの配列が与えられたとき、履歴をURLの配列として、現在のページのインデックスとともに2つの値を持つ配列を返してください。

有効なコマンドは以下の通りです：

- `"URL"` - URLはウェブアドレスです（例えば`"freecodecamp.org"`）。指定されたURLに移動し、履歴の次の位置に追加し、先の履歴は破棄します。
- `"Back"` - 履歴の前のページに移動します。前のページがなければ現在のページに留まります。
- `"Forward"` - 履歴の次のページに移動します。次のページがなければ現在のページに留まります。

例えば、`["freecodecamp.org", "freecodecamp.org/learn", "Back"]`が与えられたら、`[["freecodecamp.org", "freecodecamp.org/learn"], 0]`を返してください。

# --hints--

`get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` should return `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` should return `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"])` should return `[["example.com", "example.com/contact", "example.com/blog"], 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` should return `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "Back"])` should return `[["example.com", "example.com/about"], 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Forward"])` should return `[["example.com", "example.com/about"], 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_browser_history(commands):

    return commands
```

# --solutions--

```py
def get_browser_history(commands):
    stack = []
    index = -1

    for command in commands:
        if command == "Back":
            if index > 0:
                index -= 1
        elif command == "Forward":
            if index < len(stack) - 1:
                index += 1
        else:
            del stack[index + 1:]
            stack.append(command)
            index += 1

    return [stack, index]
```
