---
id: 69b5b2be76ec8135a7fbe973
title: "Challenge 244: Rook and Bishop Attack"
challengeType: 29
dashedName: challenge-244
---

# --description--

チェス盤上のルークの位置を示す文字列と、ビショップの位置を示す別の文字列が与えられたとき、どちらかの駒がもう一方を攻撃できるかどうかを判定してください。

標準的なチェス盤は8×8で、列は左から右へ`A`から`H`まで、行は下から上へ`1`から`8`までラベル付けされています。見た目は次のようになります：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- ルークは水平または垂直方向に好きなだけマスを移動できます。
- ビショップは任意の斜め方向に好きなだけマスを移動できます。
- ある駒がもう一方の駒の位置に移動できる場合、その駒はもう一方を攻撃できます。

Return:

- ルークがビショップを攻撃できる場合は`"rook"`を返してください。
- ビショップがルークを攻撃できる場合は`"bishop"`を返してください。
- どちらの駒も互いを攻撃できない場合は`"neither"`を返してください。

# --hints--

`rook_bishop_attack("A1", "A5")` should return `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("A1", "A5"), "rook")`)
}})
```

`rook_bishop_attack("C3", "F6")` should return `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("C3", "F6"), "bishop")`)
}})
```

`rook_bishop_attack("D4", "D7")` should return `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("D4", "D7"), "rook")`)
}})
```

`rook_bishop_attack("B7", "H1")` should return `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B7", "H1"), "bishop")`)
}})
```

`rook_bishop_attack("B3", "C5")` should return `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B3", "C5"), "neither")`)
}})
```

`rook_bishop_attack("G3", "E8")` should return `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("G3", "E8"), "neither")`)
}})
```

# --seed--

## --seed-contents--

```py
def rook_bishop_attack(rook, bishop):

    return rook
```

# --solutions--

```py
def rook_bishop_attack(rook, bishop):
    col_diff = abs(ord(rook[0]) - ord(bishop[0]))
    row_diff = abs(int(rook[1]) - int(bishop[1]))

    if rook[0] == bishop[0] or rook[1] == bishop[1]:
        return "rook"
    if col_diff == row_diff:
        return "bishop"
    return "neither"
```
