---
id: 69b559d2903b9e4afe9075f8
title: "Challenge 238: Digit Rotation Escape"
challengeType: 29
dashedName: challenge-238
---

# --description--

正の整数が与えられたとき、その数自身またはその回転のいずれかが、その桁数で割り切れるかどうかを判定してください。

回転とは、最初の桁を末尾に移動させることを意味します。例えば、1回回転させると、123は231になります。

- まず回転なしの`0`（与えられた数）をチェックしてください。
- 与えられる数にはゼロは含まれません。
- 割り切れる回転数が見つかった場合はその数を返し、なければ`"none"`を返してください。

# --hints--

`get_rotation(123)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(123), 0)`)
}})
```

`get_rotation(13579)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(13579), 3)`)
}})
```

`get_rotation(24681)` should return `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(24681), "none")`)
}})
```

`get_rotation(84138789345)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(84138789345), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rotation(n):

    return n
```

# --solutions--

```py
def get_rotation(n):
    s = str(n)
    digit_count = len(s)
    current = s

    for i in range(digit_count):
        if int(current) % digit_count == 0:
            return i
        current = current[1:] + current[0]

    return "none"
```
