---
id: 69b58ce40693f140c84c855a
title: "Challenge 241: FizzBuzz Validator"
challengeType: 29
dashedName: challenge-241
---

# --description--

連続した整数の配列が与えられ、3と5の倍数が置き換えられている場合、それが有効なFizzBuzzのシーケンスかどうかを判定してください。

有効なFizzBuzzのシーケンスでは：

- 3の倍数は "Fizz" に置き換えられます。
- 5の倍数は "Buzz" に置き換えられます。
- 3と5の両方の倍数は "FizzBuzz" に置き換えられます。
- その他の数字は整数のままです。

# --hints--

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([13, 14, "FizzBuzz", 16, 17])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([13, 14, "FizzBuzz", 16, 17]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, 5])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, 5]), False)`)
}})
```

`is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", "Buzz", 5])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", "Buzz", 5]), False)`)
}})
```

`is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]), False)`)
}})
```

`is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_fizz_buzz(arr):

    return arr
```

# --solutions--

```py
def is_fizz_buzz(arr):
    start = None
    for i, val in enumerate(arr):
        if isinstance(val, int):
            start = val - i
            break

    if start is None:
        return False

    for i in range(len(arr)):
        n = start + i
        if n % 15 == 0:
            expected = "FizzBuzz"
        elif n % 3 == 0:
            expected = "Fizz"
        elif n % 5 == 0:
            expected = "Buzz"
        else:
            expected = n
        if arr[i] != expected:
            return False

    return True
```
