---
id: 69bc6cb30c1d112a2e110a07
title: "チャレンジ250：隠されたキー"
challengeType: 29
dashedName: challenge-250
---

# --description--

250回目のデイリーチャレンジへようこそ！

エンコードされた文字列が与えられたら、暗号キーを使ってデコードし、その結果を返してください。

キーの見つけ方：

- チャレンジ番号が1から今日までのすべてのデイリーチャレンジのうち、番号が25の倍数であるもの（このチャレンジも含む）を見てください。
- それらのチャレンジタイトルのそれぞれの最初の文字を取り出し、文字列に結合してください。タイトルが文字以外で始まる場合は、最初の文字を見つけてください。

メッセージをデコードするには、エンコードされたメッセージの各文字について次の操作を行います：

  1. キーの対応する文字を見ます（メッセージがキーより長い場合はキーを繰り返します）。
  2. キーの文字を対応する数字に変換します：`"A"` = 1、`"B"` = 2、…、`"Z"` = 26。
  3. エンコードされた文字をアルファベットでその数字分だけ後ろにシフトします。
  4. シフトが`"A"`より前に行く場合は、`"Z"`に戻ります。

例えば、エンコードされたメッセージが`"Y"`で始まり、最初のキー文字が`"V"`（22）なら、`"Y"`を22文字分後ろにシフトして`"C"`を得ます。この処理を各文字に繰り返してメッセージ全体をデコードしてください。

- シフトされるのは文字だけで、スペースはそのまま返します。
- 与えられる文字も返す文字もすべて大文字です。

# --hints--

`decode("YAVJYNXE")`は`"CONGRATS"`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")`は`"CODING LEGEND"`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")`は`"YOU ARE AWESOME"`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")`は`"KEEP GOING"`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")`は`"A WINNER IS A DREAMER WHO NEVER GIVES UP"`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
