---
id: 69c5f3d787b1725d5f00c8b9
title: "Challenge 255: Earth Day Cleanup Crew"
challengeType: 29
dashedName: challenge-255
---

# --description--

今日はアースデイです。清掃したアイテムの配列が与えられたとき、以下のルールに基づいて合計のクリーンアップスコアを返してください。

与えられるアイテムは以下のいずれかです：

| アイテム | 基本値 |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

レアアイテムは `["rare", value]` で表されます。例えば `["rare", 80]` です。レアアイテムには連続ボーナスは付きません。

- 連続ボーナス：同じアイテムが連続して現れると、ボーナスポイントが増加します。
  - 1回目の連続出現：基本値
  - 2回目：基本値 + 1
  - 3回目：基本値 + 2
  - 以降同様
- 5個目のアイテム乗数：5個目ごとに乗数が付きます。
  - 5個目のアイテム：×2
  - 10個目のアイテム：×3
  - 以降同様

- ボーナス計算後に乗数を適用してください。

# --hints--

`get_cleanup_score(["bottle", "straw", "shoe", "battery"])` should return `44`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["bottle", "straw", "shoe", "battery"]), 44)`)
}})
```

`get_cleanup_score(["electronics", "straw", "newspaper", "bottle", "bag"])` should return `58`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["electronics", "straw", "newspaper", "bottle", "bag"]), 58)`)
}})
```

`get_cleanup_score(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` should return `79`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79)`)
}})
```

`get_cleanup_score(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` should return `358`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358)`)
}})
```

`get_cleanup_score(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` should return `383`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_cleanup_score(items):

    return items
```

# --solutions--

```py
def get_cleanup_score(items):
    values = {
        'bottle': 10, 'can': 6, 'bag': 8, 'tire': 35,
        'straw': 4, 'cardboard': 3, 'newspaper': 3,
        'shoe': 12, 'electronics': 25, 'battery': 18, 'mattress': 38
    }
    total = 0
    prev_item = None
    streak = 0
    for i, item in enumerate(items):
        is_rare = isinstance(item, list) and item[0] == 'rare'
        if is_rare:
            base = item[1]
            prev_item = None
            streak = 0
        else:
            base = values[item]
            streak = streak + 1 if item == prev_item else 0
            prev_item = item
        points = base if is_rare else base + streak
        if (i + 1) % 5 == 0:
            multiplier = (i + 1) // 5 + 1
            points *= multiplier
        total += points
    return total
```
