---
id: 69c5f3d787b1725d5f00c8bd
title: "チャレンジ 259: FizzBuzz エクスプロージョン"
challengeType: 29
dashedName: challenge-259
---

# --description--

整数が与えられたとき、次のルールを使って単語 `"fizzbuzz"` を少なくとも指定された数の `"z"` を含む文字列に変えるまでにかかるステップ数を返してください。

- 文字列 `"fizzbuzz"` から始めます。
- 各ステップで、文字列内の文字の位置（最初の `"f"` は位置1）を使って標準のFizzBuzzルールを適用します。
  - 文字の位置が3で割り切れる場合、その文字を `"fizz"` に置き換えます。
  - 文字の位置が5で割り切れる場合、その文字を `"buzz"` に置き換えます。
  - 文字の位置が3と5の両方で割り切れる場合、その文字を `"fizzbuzz"` に置き換えます。

つまり、1ステップ後に `"fizzbuzz"` は `"fifizzzbuzzfizzzz"` に変わり、9つの `"z"` を含みます。

# --hints--

`explode_fizzbuzz(9)` は `1` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` は `2` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` は `3` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` は `4` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` は `5` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` は `6` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` は `11` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` は `14` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
