---
id: 69c6ff713a52713463aa7927
title: "チャレンジ 261: 数字の単語"
challengeType: 29
dashedName: challenge-261
---

# --description--

0から99までの整数が与えられたら、その英語の単語表現を返してください。

- `0` は `"zero"` を返します。
- 1から19までの数字はそれぞれ固有の名前があります（`"one"`、`"two"`、...、`"ten"`、`"eleven"`、...、`"eighteen"`、`"nineteen"`）。
- 20から90までの10の倍数はそれぞれ固有の名前があります（`"twenty"`、`"thirty"`、...、`"eighty"`、`"ninety"`）。
- 21から99までの10の倍数でない数字は、ハイフンでつながれた2つの単語で書かれます。例えば `"forty-two"` と `"fifty-three"` のようにです。

# --hints--

`get_number_words(0)` は `"zero"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(0), "zero")`)
}})
```

`get_number_words(10)` は `"ten"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(10), "ten")`)
}})
```

`get_number_words(19)` は `"nineteen"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(19), "nineteen")`)
}})
```

`get_number_words(30)` は `"thirty"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(30), "thirty")`)
}})
```

`get_number_words(53)` は `"fifty-three"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(53), "fifty-three")`)
}})
```

`get_number_words(7)` は `"seven"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(7), "seven")`)
}})
```

`get_number_words(12)` は `"twelve"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(12), "twelve")`)
}})
```

`get_number_words(60)` は `"sixty"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(60), "sixty")`)
}})
```

`get_number_words(67)` は `"sixty-seven"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(67), "sixty-seven")`)
}})
```

`get_number_words(98)` は `"ninety-eight"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(98), "ninety-eight")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_words(n):

    return n
```

# --solutions--

```py
def get_number_words(n):
    ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
            "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
            "seventeen","eighteen","nineteen"]
    tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"]

    if n < 20:
        return ones[n]
    if n % 10 == 0:
        return tens[n // 10]
    return tens[n // 10] + "-" + ones[n % 10]
```
