---
id: 69cfca90e8a0a6d4d6871c51
title: "チャレンジ 267: パーセク変換器"
challengeType: 29
dashedName: challenge-267
---

# --description--

遠い銀河系では、パーセクは時間と距離の両方を測る単位として使われています。整数のパーセク数が与えられたら、それに対応する時間または距離を返してください。

- 与えられた整数が奇数なら時間を表します。偶数なら距離を表します。

以下の変換率を使ってください：

| パーセク | 時間／距離 |
| - | - |
| 1 | 2時間 |
| 2 | 6光年 |

変換した値は整数で返してください。

# --hints--

`convert_parsecs(1)` は `2` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(1), 2)`)
}})
```

`convert_parsecs(2)` は `6` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(2), 6)`)
}})
```

`convert_parsecs(31)` は `62` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(31), 62)`)
}})
```

`convert_parsecs(88)` は `264` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(88), 264)`)
}})
```

`convert_parsecs(17)` は `34` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(17), 34)`)
}})
```

`convert_parsecs(14)` は `42` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(14), 42)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_parsecs(parsecs):

    return parsecs
```

# --solutions--

```py
def convert_parsecs(parsecs):
    if parsecs % 2 != 0:
        return parsecs * 2
    else:
        return parsecs * 3
```
