---
id: 69e2383af7832c8032603b94
title: "チャレンジ 278: コーヒー注文パーサー"
challengeType: 29
dashedName: challenge-278
---

# --description--

コーヒーの注文を表す文字列が与えられたとき、メニューのアイテムを特定し、フォーマットされた注文を返してください。

以下のメニューアイテムと価格を使用します：

| アイテム | 価格 |
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

一致したアイテムを `" + "` でつなぎ、その後にコロンとスペース（`": "`）を付けて、合計金額を返してください。

例えば、`"I'd like an oat latte with vanilla syrup and an extra shot please."` が与えられたら、`"oat latte + vanilla syrup + extra shot: $6.25"` を返します。

アイテムはメニューに出てくる順番で表示し、合計金額は常に小数点以下2桁で表示してください。

# --hints--

`format_coffee_order("I'd like an oat latte with vanilla syrup and an extra shot please.")` は `"oat latte + vanilla syrup + extra shot: $6.25"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25")`)
}})
```

`format_coffee_order("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` は `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85")`)
}})
```

`format_coffee_order("Can I get a cold brew with some cream and an extra shot.")` は `"cold brew + extra shot + cream: $5.75"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75")`)
}})
```

`format_coffee_order("Just an espresso please.")` は `"espresso: $3.00"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Just an espresso please."), "espresso: $3.00")`)
}})
```

`format_coffee_order("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` は `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_coffee_order(order):

    return order
```

# --solutions--

```py
def format_coffee_order(order):
    menu = [
        ("cold brew", 450),
        ("oat latte", 500),
        ("cappuccino", 475),
        ("espresso", 300),
        ("vanilla syrup", 75),
        ("caramel drizzle", 60),
        ("extra shot", 50),
        ("oat milk", 75),
        ("cream", 75),
    ]
    found = [(name, price) for name, price in menu if name in order.lower()]
    total = sum(price for _, price in found)
    items = " + ".join(name for name, _ in found)
    return f"{items}: ${total / 100:.2f}"
```
