---
id: 69e2383af7832c8032603b90
title: "チャレンジ274：最年長者"
challengeType: 29
dashedName: challenge-274
---

# --description--

`"name"`と`"age"`プロパティを持つオブジェクトの配列が与えられたとき、最年長者の名前を含む配列を返してください。

複数の人が最年長の場合は、入力に現れる順番で全員の名前を返してください。

# --hints--

`get_oldest([{ "name": "Brenda", "age": 40 }])`は`["Brenda"]`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "Brenda", "age": 40}]), ["Brenda"])`)
}})
```

`get_oldest([{ "name": "Alice", "age": 30 }, { "name": "Bob", "age": 25 }])`は`["Alice"]`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "Alice", "age": 30}, {"name": "Bob", "age": 25}]), ["Alice"])`)
}})
```

`get_oldest([{ "name": "Allison", "age": 25 }, { "name": "Bill", "age": 30 }, { "name": "Carol", "age": 30 }])`は`["Bill", "Carol"]`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "Allison", "age": 25}, {"name": "Bill", "age": 30}, {"name": "Carol", "age": 30}]), ["Bill", "Carol"])`)
}})
```

`get_oldest([{ "name": "George", "age": 50 }, { "name": "Shirley", "age": 42 }, { "name": "Beth", "age": 48 }, { "name": "Holly", "age": 50 }, { "name": "Kevin", "age": 44 }, { "name": "Frank", "age": 47 }, { "name": "Zach", "age": 50 }, { "name": "Jennifer", "age": 43 }])`は`["George", "Holly", "Zach"]`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_oldest([{"name": "George", "age": 50}, {"name": "Shirley", "age": 42}, {"name": "Beth", "age": 48}, {"name": "Holly", "age": 50}, {"name": "Kevin", "age": 44}, {"name": "Frank", "age": 47}, {"name": "Zach", "age": 50}, {"name": "Jennifer", "age": 43}]), ["George", "Holly", "Zach"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_oldest(people):

    return people
```

# --solutions--

```py
def get_oldest(people):
    max_age = max(p["age"] for p in people)
    return [p["name"] for p in people if p["age"] == max_age]
```
