---
id: 69f35a5bb823ed620fcb7cbb
title: "Challenge 282: Sleep Debt"
challengeType: 29
dashedName: challenge-282
---

# --description--

今日までの各夜に眠った時間の配列と、1晩あたりの目標時間が与えられたとき、睡眠負債を解消するために今夜何時間眠る必要があるかを返してください。

- 今夜の睡眠時間も、追いつくために必要な合計時間に含めてください。
- 今夜の目標時間以上にすでに眠っている場合は、`0`を返してください。

# --hints--

`sleep_debt([6, 6, 6, 6, 6, 6], 8)` should return `20`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 6, 6, 6, 6, 6], 8), 20)`)
}})
```

`sleep_debt([6, 7, 8, 4, 8, 6], 7)` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 7, 8, 4, 8, 6], 7), 10)`)
}})
```

`sleep_debt([10, 10, 9, 10, 9, 11], 9)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([10, 10, 9, 10, 9, 11], 9), 4)`)
}})
```

`sleep_debt([8, 7, 6, 7, 6, 8], 6)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 7, 6, 7, 6, 8], 6), 0)`)
}})
```

`sleep_debt([8, 9, 10, 9, 10, 7], 7)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 9, 10, 9, 10, 7], 7), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sleep_debt(hours_slept, target_hours):

    return hours_slept
```

# --solutions--

```py
def sleep_debt(hours_slept, target_hours):
    debt = (len(hours_slept) + 1) * target_hours - sum(hours_slept)
    return max(0, debt)
```
