---
id: 69f8c998d78ad3171a0713b9
title: "チャレンジ 286: 未解決のIssue"
challengeType: 29
dashedName: challenge-286
---

# --description--

Issue番号の配列とプルリクエスト（PR）番号の配列が与えられたとき、すべてのPRがマージされた後に未解決のまま残っているIssueの配列を返してください。

- PRは、その数字が回転した関係にある場合にIssueをクローズします。たとえば、Issue 123はPR 231や312によってクローズされます。
- PRは、まったく同じ番号のIssueをクローズしません。たとえば、PR 123はIssue 123をクローズしません。つまり、すべて同じ数字のIssueはクローズされません。
- どちらの番号も先頭のゼロが削除されている場合があります。たとえば、PR 201はIssue 12（012、201の回転）をクローズします。同様に、Issue 201はPR 12によってクローズされます。

未解決のIssueは、与えられた順番で返してください。

# --hints--

`get_open_issues([123, 234], [231])` は `[234]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` は `[345, 16]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` は `[456, 332, 15]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` は `[115, 296, 24]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` は `[95, 395, 754, 296, 709, 237, 1802]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```
