---
id: 69f8c998d78ad3171a0713be
title: "チャレンジ 291: FizzBuzz カウント"
challengeType: 29
dashedName: challenge-291
---

# --description--

開始番号と終了番号が与えられたとき、その範囲内（両端含む）での fizz と buzz の出現回数を数えてください。

- 3で割り切れる数字は fizz としてカウントします。
- 5で割り切れる数字は buzz としてカウントします。
- 3と5の両方で割り切れる数字は fizz と buzz の両方としてカウントします。

カウント結果を `{ fizz, buzz }` の形式でオブジェクトまたは辞書として返してください。

# --hints--

`fizz_buzz_count(1, 11)` は `{"fizz": 3, "buzz": 2}` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(1, 11), {"fizz": 3, "buzz": 2})`)
}})
```

`fizz_buzz_count(14, 41)` は `{"fizz": 9, "buzz": 6}` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(14, 41), {"fizz": 9, "buzz": 6})`)
}})
```

`fizz_buzz_count(24, 100)` は `{"fizz": 26, "buzz": 16}` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(24, 100), {"fizz": 26, "buzz": 16})`)
}})
```

`fizz_buzz_count(-635, -14)` は `{"fizz": 207, "buzz": 125}` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-635, -14), {"fizz": 207, "buzz": 125})`)
}})
```

`fizz_buzz_count(-5432, 6789)` は `{"fizz": 4074, "buzz": 2444}` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-5432, 6789), {"fizz": 4074, "buzz": 2444})`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_count(start, end):

    return start
```

# --solutions--

```py
def fizz_buzz_count(start, end):
    fizz = sum(1 for i in range(start, end + 1) if i % 3 == 0)
    buzz = sum(1 for i in range(start, end + 1) if i % 5 == 0)
    return {"fizz": fizz, "buzz": buzz}
```
