---
id: 69f8c998d78ad3171a0713bc
title: "チャレンジ 289: 差の合計"
challengeType: 29
dashedName: challenge-289
---

# --description--

数値の配列が与えられたとき、各数値とその次の数値の差の合計を返してください。

例えば、`[1, 3, 4]`が与えられた場合、`3`（2 + 1）を返します。

# --hints--

`sum_of_differences([1, 3, 4])`は`3`を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([1, 3, 4]), 3)`)
}})
```

`sum_of_differences([5, -3, 3, 9, 10])`は`5`を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([5, -3, 3, 9, 10]), 5)`)
}})
```

`sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7])`は`-16`を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16)`)
}})
```

`sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])`は`25`を返すはずです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_of_differences(arr):

    return arr
```

# --solutions--

```py
def sum_of_differences(arr):
    return sum(arr[i + 1] - arr[i] for i in range(len(arr) - 1))
```
