---
id: 69f90c1329a94b37e2a2086d
title: "Challenge 294: Parentheses Combinations"
challengeType: 29
dashedName: challenge-294
---

# --description--

整数`n`が与えられたとき、`n`組の丸括弧の有効な組み合わせの数を返してください。

- 有効な組み合わせとは、すべての開き丸括弧に対応する閉じ丸括弧があり、対応する開き丸括弧より前に閉じ丸括弧が現れない文字列のことです。

例えば、`2`が与えられた場合、有効な組み合わせは2つあります。

```json
(())
()()
```

# --hints--

`get_combinations(2)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(2), 2)`)
}})
```

`get_combinations(3)` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(3), 5)`)
}})
```

`get_combinations(5)` should return `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(5), 42)`)
}})
```

`get_combinations(8)` should return `1430`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(8), 1430)`)
}})
```

`get_combinations(13)` should return `742900`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(13), 742900)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_combinations(n):

    return n
```

# --solutions--

```py
def get_combinations(n):
    count = 0

    def generate(op, cl):
        nonlocal count
        if op == 0 and cl == 0:
            count += 1
            return
        if op > 0:
            generate(op - 1, cl)
        if cl > op:
            generate(op, cl - 1)

    generate(n, n)
    return count
```
