---
id: 6a0dcc730cb92a616f86f0c2
title: "チャレンジ 298: スキーマバリデーター パート4"
challengeType: 29
dashedName: challenge-298
---

# --description--

オブジェクト（JavaScriptの場合）または辞書（Pythonの場合）が、以下のスキーマに合致するかどうかを判定してください。

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean
}
```

- パイプ記号（`|`）は「または」を意味します。`role`はリストにある`Roles`のいずれかでなければなりません。
- `?`の後の疑問符（`supporter`）は、そのフィールドがオプションであることを示しますが、存在する場合は指定された型でなければなりません。
- 追加のキーは許可されます

# --hints--

`is_valid_schema({"username": "vivian", "posts": 1, "verified": False, "role": "user", "supporter": True})`は`True`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "vivian", "posts": 1, "verified": False, "role": "user", "supporter": True}), True)`)
}})
```

`is_valid_schema({"username": "rudolph", "posts": 15, "verified": True, "role": "creator"})`は`True`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "rudolph", "posts": 15, "verified": True, "role": "creator"}), True)`)
}})
```

`is_valid_schema({"username": "hernandez", "posts": 35, "verified": True, "role": "moderator", "supporter": False, "followers": 55})`は`True`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "hernandez", "posts": 35, "verified": True, "role": "moderator", "supporter": False, "followers": 55}), True)`)
}})
```

`is_valid_schema({"username": "julia", "posts": 50, "verified": True, "role": "admin", "supporter": "true"})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "julia", "posts": 50, "verified": True, "role": "admin", "supporter": "true"}), False)`)
}})
```

`is_valid_schema({"username": "bernard", "posts": 0, "verified": True, "role": "friend", "supporter": True})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bernard", "posts": 0, "verified": True, "role": "friend", "supporter": True}), False)`)
}})
```

`is_valid_schema({"username": "felix", "posts": 40, "verified": "yes", "role": "staff", "supporter": False})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "felix", "posts": 40, "verified": "yes", "role": "staff", "supporter": False}), False)`)
}})
```

`is_valid_schema({"username": "jimmy", "posts": True, "verified": False, "role": "creator", "supporter": True})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "jimmy", "posts": True, "verified": False, "role": "creator", "supporter": True}), False)`)
}})
```

`is_valid_schema({"username": True, "posts": 30, "verified": True, "role": "moderator", "supporter": False})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": True, "posts": 30, "verified": True, "role": "moderator", "supporter": False}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    roles = ["user", "creator", "moderator", "staff", "admin"]
    return (
        isinstance(obj.get("username"), str) and
        isinstance(obj.get("posts"), int) and not isinstance(obj.get("posts"), bool) and
        isinstance(obj.get("verified"), bool) and
        obj.get("role") in roles and
        ("supporter" not in obj or isinstance(obj["supporter"], bool))
    )
```
