---
id: 6a0dcd03ee4e68698080ef6d
title: "チャレンジ 309: 数字の並べ替え"
challengeType: 29
dashedName: challenge-309
---

# --description--

カンマで区切られた数字の文字列が与えられたとき、数字を小さい順に並べ替えた配列を返してください。

# --hints--

`sort_numbers("3,1,2")` は `[1, 2, 3]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort_numbers("3,1,2"), [1, 2, 3])`)
}})
```

`sort_numbers("5,3,8,1,9,2")` は `[1, 2, 3, 5, 8, 9]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort_numbers("5,3,8,1,9,2"), [1, 2, 3, 5, 8, 9])`)
}})
```

`sort_numbers("12,61,49,80,19,50,77,38")` は `[12, 19, 38, 49, 50, 61, 77, 80]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort_numbers("12,61,49,80,19,50,77,38"), [12, 19, 38, 49, 50, 61, 77, 80])`)
}})
```

`sort_numbers("0,6,-19,44,-2,7,0")` は `[-19, -2, 0, 0, 6, 7, 44]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort_numbers("0,6,-19,44,-2,7,0"), [-19, -2, 0, 0, 6, 7, 44])`)
}})
```

# --seed--

## --seed-contents--

```py
def sort_numbers(s):

    return s
```

# --solutions--

```py
def sort_numbers(s):
    return sorted(int(x) for x in s.split(","))
```
