---
id: 6a0dcd03ee4e68698080ef68
title: "チャレンジ 304: 旅程の手配"
challengeType: 29
dashedName: challenge-304
---

# --description--

1日旅行のための少なくとも2つのオプションのストップを含む配列が与えられたとき、有効な旅程の手配の数を返してください。

旅程には常に `"breakfast"`、`"lunch"`、および `"dinner"` が含まれ、これらは引数として渡されません。オプションのストップは旅程のどこにでも配置できますが、以下のルールに従う必要があります：

- `"breakfast"` は常に最初で、`"lunch"` の前に少なくとも1つのストップがあります。
- `"lunch"` は `"dinner"` の前に現れ、間に少なくとも1つのストップがあります。
- `"dinner"` の後に現れるオプションのストップは最大で1つまでです。

有効な手配の数を返してください。

# --hints--

`get_itinerary_count(["library", "park"])` は `2` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` は `18` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` は `120` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` は `840` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` は `55440` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
