---
id: 6a0dcc730cb92a616f86f0bf
title: "チャレンジ 295: スキーマバリデーター パート1"
challengeType: 29
dashedName: challenge-295
---

# --description--

オブジェクト（JavaScript）または辞書（Python）が、以下のスキーマに一致するかどうかを判定してください。

```json
{
  username: string
}
```

- 追加のキーは許可されています

# --hints--

`is_valid_schema({"username": "bob"})` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bob"}), True)`)
}})
```

`is_valid_schema({"username": "jen", "posts": 30})` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "jen", "posts": 30}), True)`)
}})
```

`is_valid_schema({"username": ""})` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": ""}), True)`)
}})
```

`is_valid_schema({"username": 7})` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 7}), False)`)
}})
```

`is_valid_schema({"posts": 25})` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"posts": 25}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return isinstance(obj.get("username"), str)
```
