---
id: 6a0dcc730cb92a616f86f0c1
title: "チャレンジ 297: スキーマバリデーター パート3"
challengeType: 29
dashedName: challenge-297
---

# --description--

オブジェクト（JavaScript）または辞書（Python）が、以下のスキーマに一致するかどうかを判定してください。

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles
}
```

- パイプ記号（`|`）は「または」を意味します。`role`はリストにある`Roles`のいずれかでなければなりません。
- 追加のキーは許可されます

# --hints--

`is_valid_schema({"username": "henry", "posts": 0, "verified": True, "role": "staff"})`は`True`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "henry", "posts": 0, "verified": True, "role": "staff"}), True)`)
}})
```

`is_valid_schema({"username": "sara", "posts": 45, "verified": False, "role": "creator", "followers": 70})`は`True`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "sara", "posts": 45, "verified": False, "role": "creator", "followers": 70}), True)`)
}})
```

`is_valid_schema({"username": "penelope", "posts": 20, "verified": True, "role": "admin"})`は`True`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "penelope", "posts": 20, "verified": True, "role": "admin"}), True)`)
}})
```

`is_valid_schema({"username": "kevin", "posts": 0, "verified": False, "role": "user"})`は`True`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "kevin", "posts": 0, "verified": False, "role": "user"}), True)`)
}})
```

`is_valid_schema({"username": "george", "posts": 15, "verified": True, "role": "moderator"})`は`True`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "george", "posts": 15, "verified": True, "role": "moderator"}), True)`)
}})
```

`is_valid_schema({"username": "david", "posts": 0, "verified": False, "role": "guest"})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "david", "posts": 0, "verified": False, "role": "guest"}), False)`)
}})
```

`is_valid_schema({"username": "wendy", "posts": 10, "verified": True})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "wendy", "posts": 10, "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "fabian", "posts": 1, "verified": True, "role": True})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "fabian", "posts": 1, "verified": True, "role": True}), False)`)
}})
```

`is_valid_schema({"username": 8, "posts": 1, "verified": True, "role": "user"})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 8, "posts": 1, "verified": True, "role": "user"}), False)`)
}})
```

`is_valid_schema({"username": "penny", "posts": "10", "verified": True, "role": "staff"})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "penny", "posts": "10", "verified": True, "role": "staff"}), False)`)
}})
```

`is_valid_schema({"username": "john", "posts": "1", "verified": "true", "role": "admin"})`は`False`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "john", "posts": "1", "verified": "true", "role": "admin"}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    roles = ["user", "creator", "moderator", "staff", "admin"]
    return (
        isinstance(obj.get("username"), str) and
        isinstance(obj.get("posts"), int) and not isinstance(obj.get("posts"), bool) and
        isinstance(obj.get("verified"), bool) and
        obj.get("role") in roles
    )
```
