---
id: 6a0dcd03ee4e68698080ef6c
title: "チャレンジ 308: クレジットカードバリデーター"
challengeType: 29
dashedName: challenge-308
---

# --description--

クレジットカード番号の数字の文字列が与えられたとき、次の方法を使って有効なカード番号かどうかを判定してください。

- 後ろから2番目の数字から始めて、左に向かって1つおきに数字を2倍にします。
- 2倍にした数字が9より大きい場合は、9を引きます。
- 2倍にした数字とそのままの数字をすべて合計します。
- 合計が10で割り切れれば、その番号は有効です。

# --hints--

`is_valid_card("4532015112830366")` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830366"), True)`)
}})
```

`is_valid_card("5425233430109903")` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("5425233430109903"), True)`)
}})
```

`is_valid_card("371449635398431")` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("371449635398431"), True)`)
}})
```

`is_valid_card("6011111111111117")` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("6011111111111117"), True)`)
}})
```

`is_valid_card("4532015112830367")` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830367"), False)`)
}})
```

`is_valid_card("1234567890123456")` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("1234567890123456"), False)`)
}})
```

`is_valid_card("4532015112830368")` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830368"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_card(number):

    return number
```

# --solutions--

```py
def is_valid_card(number):
    total = 0
    for i, digit in enumerate(reversed(number)):
        d = int(digit)
        if i % 2 == 1:
            d *= 2
            if d > 9:
                d -= 9
        total += d
    return total % 10 == 0
```
