---
id: 6a0dcd03ee4e68698080ef6b
title: "チャレンジ 307: ゾーニング規制"
challengeType: 29
dashedName: challenge-307
---

# --description--

都市の建物配置を表す2次元グリッド（配列の配列）が与えられたとき、ゾーニング規則に違反しているすべての建物の座標を返してください。

グリッドの各セルには、以下のテーブルのラベルのいずれかが含まれています。建物は、その（最大4つの）隣接セルが水平または垂直方向に、隣接できないタイプである場合に違反とみなされます。

| ラベル | 種類 | 隣接できないもの |
| - | - | - |
| `"i"` | 工業用 | `"R"`, `"I"` |
| `"A"` | 農業用 | `"C"` |
| `"R"` | 住宅用 | `"i"`, `"C"` |
| `"I"` | 公共施設用 | `"i"` |
| `"C"` | 商業用 | `"R"`, `"A"` |
| `""`（空文字） | 未開発 | 制限なし |

違反しているすべてのセルの座標を`[row, col]`ペアの配列として、順不同で返してください。違反がなければ空の配列を返してください。

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])`は`[[0, 0], [0, 1]]`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])`は`[[0, 1], [1, 1]]`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])`は`[]`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])`は`[[0, 1], [0, 2], [0, 3]]`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])`は`[[2, 3], [2, 4], [3, 1], [3, 2]]`を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
