---
id: 6a15cadf5f240d05a264955f
title: "チャレンジ 321: 周期表のスペル"
challengeType: 29
dashedName: challenge-321
---

# --description--

単語が与えられたとき、その単語が周期表の元素記号を使って綴れるかどうかを判定してください。

- 単語を綴るときは大文字・小文字を無視してください。`"neon"` は記号 `"Ne"`、`"O"`、`"N"` を使って綴れます。

元素記号の完全なリストは以下の通りです：

```json
["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];
```

単語を綴るために使われた元素を、元の大文字・小文字のまま、単語を綴る順番で配列として返してください。綴れない場合は空の配列を返してください。

# --hints--

`get_periodic_spelling("neon")` は `["Ne", "O", "N"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("neon"), ["Ne", "O", "N"])`)
}})
```

`get_periodic_spelling("rational")` は `["Ra", "Ti", "O", "N", "Al"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("rational"), ["Ra", "Ti", "O", "N", "Al"])`)
}})
```

`get_periodic_spelling("yarn")` は `["Y", "Ar", "N"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("yarn"), ["Y", "Ar", "N"])`)
}})
```

`get_periodic_spelling("carbon")` は `["C", "Ar", "B", "O", "N"]` または `["Ca", "Rb", "O", "N"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("carbon"), [["C", "Ar", "B", "O", "N"], ["Ca", "Rb", "O", "N"]])`)
}})
```

`get_periodic_spelling("noisy")` は `["N", "O", "I", "S", "Y"]` または `["No", "I", "S", "Y"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("noisy"), [["N", "O", "I", "S", "Y"], ["No", "I", "S", "Y"]])`)
}})
```

`get_periodic_spelling("bicycles")` は `["B", "I", "C", "Y", "Cl", "Es"]` または `["Bi", "C", "Y", "Cl", "Es"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("bicycles"), [["B", "I", "C", "Y", "Cl", "Es"], ["Bi", "C", "Y", "Cl", "Es"]])`)
}})
```

`get_periodic_spelling("optics")` は `["O", "P", "Ti", "C", "S"]`、`["O", "P", "Ti", "Cs"]`、`["O", "Pt", "I", "C", "S"]`、または `["O", "Pt", "I", "Cs"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("optics"), [["O", "P", "Ti", "C", "S"], ["O", "P", "Ti", "Cs"], ["O", "Pt", "I", "C", "S"], ["O", "Pt", "I", "Cs"]])`)
}})
```

`get_periodic_spelling("value")` は `[]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("value"), [])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_periodic_spelling(word):

    return word
```

# --solutions--

```py
def get_periodic_spelling(word):
    elements = ["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"]
    lower = {e.lower(): e for e in elements}

    def spell(w, i):
        if i == len(w):
            return []
        two = w[i:i + 2]
        one = w[i:i + 1]
        if two in lower:
            rest = spell(w, i + 2)
            if rest is not None:
                return [lower[two]] + rest
        if one in lower:
            rest = spell(w, i + 1)
            if rest is not None:
                return [lower[one]] + rest
        return None

    result = spell(word.lower(), 0)
    return result if result is not None else []
```
