---
id: 6a15cadf5f240d05a2649556
title: "チャレンジ 312: ストリーミング料金"
challengeType: 29
dashedName: challenge-312
---

# --description--

ストリーミングサービスのカート内の映画を表す配列と、サブスクリプションの階層を表す文字列が与えられたとき、映画の合計料金を返してください。

カート内の各アイテムは、`"format"`（`"HD"` または `"4K"`）と `"type"`（`"rent"` または `"buy"`）を持つオブジェクトです。それぞれの料金は以下の通りです：

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

次のサブスクリプション階層の割引を適用してください：

- `"none"`: 定価
- `"basic"`: 10%割引
- `"premium"`: 25%割引

合計料金を小数点以下2桁に丸めて、`"$D.CC"` の形式で返してください。

# --hints--

`get_streaming_bill([{ "format": "HD", "type": "rent" }], "none")` は `"$3.99"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}], "none"), "$3.99")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "HD", "type": "buy" }], "premium")` は `"$12.73"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "HD", "type": "buy"}], "premium"), "$12.73")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "HD", "type": "rent" }, { "format": "HD", "type": "buy" }], "basic")` は `"$18.87"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "HD", "type": "rent"}, {"format": "HD", "type": "buy"}], "basic"), "$18.87")`)
}})
```

`get_streaming_bill([{ "format": "4K", "type": "buy" }, { "format": "4K", "type": "buy" }], "premium")` は `"$29.98"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "4K", "type": "buy"}, {"format": "4K", "type": "buy"}], "premium"), "$29.98")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "4K", "type": "rent" }, { "format": "HD", "type": "buy" }, { "format": "4K", "type": "buy" }], "none")` は `"$42.96"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}], "none"), "$42.96")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "4K", "type": "rent" }, { "format": "HD", "type": "buy" }, { "format": "4K", "type": "buy" }, { "format": "HD", "type": "buy" }], "basic")` は `"$50.36"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}, {"format": "HD", "type": "buy"}], "basic"), "$50.36")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_streaming_bill(cart, subscription):

    return cart
```

# --solutions--

```py
def get_streaming_bill(cart, subscription):
    pricing = {
        "HD": {"rent": 3.99, "buy": 12.99},
        "4K": {"rent": 5.99, "buy": 19.99},
    }
    discounts = {"none": 1, "basic": 0.9, "premium": 0.75}
    multiplier = discounts[subscription]
    total = sum(pricing[item["format"]][item["type"]] * multiplier for item in cart)
    return f"${total:.2f}"
```
