---
id: 6a15cadf5f240d05a264955a
title: "チャレンジ 316: 1337 スピーク"
challengeType: 29
dashedName: challenge-316
---

# --description--

小文字の文字列が与えられたら、以下の文字をそれぞれのリート置換に変換してリートスピークに翻訳したものを返してください。

| 文字 | リート |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- 置換がない文字はそのままにしてください。

# --hints--

`make_leet("cool")` は `"c001"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("cool"), "c001")`)
}})
```

`make_leet("leet")` は `"1337"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("leet"), "1337")`)
}})
```

`make_leet("hacker")` は `"h4ck3r"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("hacker"), "h4ck3r")`)
}})
```

`make_leet("satellite")` は `"547311173"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("satellite"), "547311173")`)
}})
```

`make_leet("abcdefghijklmnopqrstuvwxyz")` は `"4bcd3f9h1jk1mn0pqr57uvwxyz"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz")`)
}})
```

# --seed--

## --seed-contents--

```py
def make_leet(s):

    return s
```

# --solutions--

```py
def make_leet(s):
    leet_map = {"a": "4", "e": "3", "g": "9", "i": "1", "l": "1", "o": "0", "t": "7", "s": "5"}
    return "".join(leet_map.get(c, c) for c in s)
```
