---
id: 6a151d32d772271dd2448c2d
title: "チャレンジ 311: 呪文使い"
challengeType: 29
dashedName: challenge-311
---

# --description--

あなたが唱えている呪文コードの文字列が与えられたとき、合計スコアを計算してください。

文字列の各文字は呪文を表します：

| Code | 呪文 | カテゴリー | 基本スコア |
| - | - | - | - |
| `"f"` | 火 | 破壊 | 3 |
| `"l"` | 雷 | 破壊 | 3 |
| `"i"` | 氷 | 制御 | 2 |
| `"w"` | 風 | 制御 | 2 |
| `"h"` | 回復 | 回復 | 1 |
| `"s"` | 盾 | 回復 | 1 |

異なるカテゴリーの呪文が連続して唱えられた回数に応じてコンボ倍率が適用されます：

- 最初の呪文は常に基本スコアで計算されます。
- 前の呪文と異なるカテゴリーの呪文が連続するたびに倍率が1ずつ増加します。
- 前の呪文と同じカテゴリーの呪文を唱えると倍率は1にリセットされます。
- 各呪文のスコアは基本スコアに現在の倍率を掛けたものです。

呪文のシーケンスから合計スコアを返してください。

# --hints--

`cast("fihwl")` は `33` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` は `45` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` は `37` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` は `50` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` は `101` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
