---
id: 6a15cadf5f240d05a264955b
title: "チャレンジ 317: BMI計算機"
challengeType: 29
dashedName: challenge-317
---

# --description--

ポンドでの体重とインチでの身長が与えられたら、小数点以下1桁に丸めたBMI（ボディマス指数）を返してください。

BMIを求めるには、体重を身長の二乗で割り、その結果に703を乗算してください。

# --hints--

`calculate_bmi(180, 70)` は `25.8` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(180, 70), 25.8)`)
}})
```

`calculate_bmi(140, 64)` は `24.0` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(140, 64), 24.0)`)
}})
```

`calculate_bmi(160, 76)` は `19.5` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(160, 76), 19.5)`)
}})
```

`calculate_bmi(200, 60)` は `39.1` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(200, 60), 39.1)`)
}})
```

`calculate_bmi(150, 68)` は `22.8` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(150, 68), 22.8)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_bmi(weight, height):

    return weight
```

# --solutions--

```py
def calculate_bmi(weight, height):
    return round((weight / (height * height)) * 703, 1)
```
