---
id: 6a19b1062d1b153d8ac76d73
title: "チャレンジ 325: ラッキーナンバー"
challengeType: 29
dashedName: challenge-325
---

# --description--

人の名前のファーストネームとラストネームの文字列が与えられたとき、次のルールを使ってラッキーナンバーを計算してください。

- ファーストネームとラストネームはスペースで区切られています
- それぞれの名前の母音と子音の数を数えます
- 小さい方の母音と子音の数を掛け合わせ、それに小さい方の名前の長さを掛けます
- 大きい方の母音と子音の数と大きい方の名前の長さでも同じ計算をします
- 大きい方の値から小さい方の値を引いてラッキーナンバーを求めます

最終的な値がゼロ（`0`）の場合は、`13`を返してください。

# --hints--

`get_lucky_number("John Doe")` は `21` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("John Doe"), 21)`)
}})
```

`get_lucky_number("Olivia Lewis")` は `52` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Olivia Lewis"), 52)`)
}})
```

`get_lucky_number("James Wilson")` は `18` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("James Wilson"), 18)`)
}})
```

`get_lucky_number("Elizabeth Hernandez")` は `81` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Elizabeth Hernandez"), 81)`)
}})
```

`get_lucky_number("Mike Walker")` は `32` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Mike Walker"), 32)`)
}})
```

`get_lucky_number("Chloe Perez")` は `13` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Chloe Perez"), 13)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_lucky_number(name):

    return name
```

# --solutions--

```py
def get_lucky_number(name):
    vowels = "aeiouAEIOU"

    def name_stats(s):
        vowel_count = sum(1 for c in s if c in vowels)
        consonant_count = sum(1 for c in s if c.isalpha() and c not in vowels)
        return vowel_count, consonant_count, len(s)

    first, last = name.split(" ")
    first_vowels, first_consonants, first_len = name_stats(first)
    last_vowels, last_consonants, last_len = name_stats(last)

    larger_vowels = max(first_vowels, last_vowels)
    smaller_vowels = min(first_vowels, last_vowels)
    larger_consonants = max(first_consonants, last_consonants)
    smaller_consonants = min(first_consonants, last_consonants)
    larger_length = max(first_len, last_len)
    smaller_length = min(first_len, last_len)

    larger_value = larger_vowels * larger_consonants * larger_length
    smaller_value = smaller_vowels * smaller_consonants * smaller_length
    result = larger_value - smaller_value

    return 13 if result == 0 else result
```
