---
id: 6a19b1062d1b153d8ac76d72
title: "チャレンジ 324: 重複文字のカウント"
challengeType: 29
dashedName: challenge-324
---

# --description--

2つの文字列が与えられたとき、2番目の文字列にある文字のうち、1番目の文字列に見つかる文字の数を返してください。

- 2番目の文字列に重複している文字は、それぞれ別々にカウントされます。

# --hints--

`duplicate_character_count("aloha", "hei")` は `1` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("aloha", "hei"), 1)`)
}})
```

`duplicate_character_count("jambo", "bonjour")` は `4` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("jambo", "bonjour"), 4)`)
}})
```

`duplicate_character_count("hello", "hola")` は `3` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello", "hola"), 3)`)
}})
```

`duplicate_character_count("ola", "hej")` は `0` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ola", "hej"), 0)`)
}})
```

`duplicate_character_count("ciao", "konnichiwa")` は `5` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ciao", "konnichiwa"), 5)`)
}})
```

`duplicate_character_count("merhaba", "xin chao")` は `2` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("merhaba", "xin chao"), 2)`)
}})
```

`duplicate_character_count("hello world", "hello to everyone around the world")` は `26` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello world", "hello to everyone around the world"), 26)`)
}})
```

# --seed--

## --seed-contents--

```py
def duplicate_character_count(str1, str2):

    return str1
```

# --solutions--

```py
def duplicate_character_count(str1, str2):
    return sum(1 for char in str2 if char in str1)
```
