---
id: 6a19dd69388e5a6e59e4f2ef
title: "チャレンジ 326: 最大利益"
challengeType: 29
dashedName: challenge-326
---

# --description--

日々の株価の配列と予算（ドル単位）が与えられたとき、その期間内に株を売買して得られる最大の利益を計算してください。

- 売るのは買った後だけにしてください。
- 買いと売りの取引は最大で1回ずつまで行えます。売った後は再び買うことはできません。
- 株は丸ごと1株単位でしか買えません。
- 最大の利益を文字列で返してください。利益は1セント単位で切り捨て、小数点以下2桁でフォーマットしてください。

# --hints--

`get_max_profit([5, 6], 50)` は `"10.00"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([5, 6], 50), "10.00")`)
}})
```

`get_max_profit([8, 2, 5, 10], 20)` は `"80.00"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([8, 2, 5, 10], 20), "80.00")`)
}})
```

`get_max_profit([4, 5, 3, 6], 20)` は `"18.00"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([4, 5, 3, 6], 20), "18.00")`)
}})
```

`get_max_profit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200)` は `"8.31"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200), "8.31")`)
}})
```

`get_max_profit([15.38, 15.01, 14.99, 14.62, 14.28], 80)` は `"0.00"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([15.38, 15.01, 14.99, 14.62, 14.28], 80), "0.00")`)
}})
```

`get_max_profit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25)` は `"73.80"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25), "73.80")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_max_profit(prices, budget):

    return prices
```

# --solutions--

```py
def get_max_profit(prices, budget):
    max_profit = 0

    for buy in range(len(prices) - 1):
        for sell in range(buy + 1, len(prices)):
            shares = int(budget / prices[buy])
            profit = shares * (prices[sell] - prices[buy])
            if profit > max_profit:
                max_profit = profit

    return f"{int(max_profit * 100) / 100:.2f}"
```
