---
id: 6a1d9f98e819ed70a0e994e1
title: "チャレンジ 335: 5つのサイコロ"
challengeType: 29
dashedName: challenge-335
---

# --description--

1～6の値を持つ5つのサイコロの配列が与えられたら、可能な限り最良の役を返してください。

役は最も低いものから高いものへ次のようにランク付けされています：

| 役 | 説明 |
|------|-------------|
| `"no pair"` | ペアなしまたはそれ以上の役なし |
| `"pair"` | 同じ値のサイコロが2つ |
| `"two pair"` | 2つの異なるペア |
| `"three of a kind"` | 同じ値のサイコロが3つ |
| `"small straight"` | 4つの連続した値 |
| `"large straight"` | 5つの連続した値 |
| `"full house"` | 3つの同じ値とペア |
| `"four of a kind"` | 同じ値のサイコロが4つ |
| `"five of a kind"` | 5つすべてのサイコロが同じ値 |

# --hints--

`five_dice([1, 1, 1, 1, 1])` は `"five of a kind"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 1, 1, 1, 1]), "five of a kind")`)
}})
```

`five_dice([5, 5, 5, 6, 5])` は `"four of a kind"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([5, 5, 5, 6, 5]), "four of a kind")`)
}})
```

`five_dice([2, 5, 6, 4, 3])` は `"large straight"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 5, 6, 4, 3]), "large straight")`)
}})
```

`five_dice([4, 3, 3, 3, 1])` は `"three of a kind"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 3, 3, 3, 1]), "three of a kind")`)
}})
```

`five_dice([4, 6, 2, 6, 5])` は `"pair"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 6, 2, 6, 5]), "pair")`)
}})
```

`five_dice([1, 4, 5, 6, 2])` は `"no pair"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 4, 5, 6, 2]), "no pair")`)
}})
```

`five_dice([1, 3, 4, 6, 2])` は `"small straight"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 3, 4, 6, 2]), "small straight")`)
}})
```

`five_dice([2, 2, 5, 2, 5])` は `"full house"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 2, 5, 2, 5]), "full house")`)
}})
```

`five_dice([6, 4, 5, 6, 4])` は `"two pair"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([6, 4, 5, 6, 4]), "two pair")`)
}})
```

# --seed--

## --seed-contents--

```py
def five_dice(dice):

    return dice
```

# --solutions--

```py
def five_dice(dice):
    from collections import Counter
    counts = Counter(dice)
    values = sorted(counts.values(), reverse=True)
    keys = sorted(counts.keys())

    is_large_straight = len(keys) == 5 and keys[-1] - keys[0] == 4
    is_small_straight = not is_large_straight and any(
        keys[i + 3] - keys[i] == 3 and len(set(keys[i:i + 4])) == 4
        for i in range(len(keys) - 3)
    )

    if values[0] == 5: return "five of a kind"
    if values[0] == 4: return "four of a kind"
    if values[0] == 3 and values[1] == 2: return "full house"
    if is_large_straight: return "large straight"
    if is_small_straight: return "small straight"
    if values[0] == 3: return "three of a kind"
    if values[0] == 2 and values[1] == 2: return "two pair"
    if values[0] == 2: return "pair"
    return "no pair"
```
