---
id: 6a1d9f98e819ed70a0e994e0
title: "チャレンジ 334: 正確なおつり"
challengeType: 29
dashedName: challenge-334
---

# --description--

整数のセント単位の金額が与えられたとき、ペニー（1セント）、ニッケル（5セント）、ダイム（10セント）、クォーター（25セント）を使って正確なおつりを作る異なる方法の数を返してください。

# --hints--

`exact_change(3)` は `1` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(3), 1)`)
}})
```

`exact_change(9)` は `2` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(9), 2)`)
}})
```

`exact_change(17)` は `6` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(17), 6)`)
}})
```

`exact_change(39)` は `24` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(39), 24)`)
}})
```

`exact_change(61)` は `73` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(61), 73)`)
}})
```

`exact_change(99)` は `213` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(99), 213)`)
}})
```

# --seed--

## --seed-contents--

```py
def exact_change(amount):

    return amount
```

# --solutions--

```py
def exact_change(amount):
    coins = [1, 5, 10, 25]
    dp = [0] * (amount + 1)
    dp[0] = 1
    for coin in coins:
        for i in range(coin, amount + 1):
            dp[i] += dp[i - coin]
    return dp[amount]
```
