---
id: 6a1d9f98e819ed70a0e994de
title: "チャレンジ 332: Issue トリアージ"
challengeType: 29
dashedName: challenge-332
---

# --description--

Issue に最後に投稿されてからのミリ秒数と、Issue に最後に投稿されたメッセージが与えられたとき、次のルールに従って Issue に対して何をすべきか判断してください。

- 最後のメッセージが7日未満の場合は、`"leave it"` を返してください。
- 最後のメッセージが7日以上前で、その内容に大文字小文字を区別せずに `"bump"` が含まれている場合は、`"close it"` を返してください。
- それ以外の場合は、`"bump it"` を返してください。

# --hints--

`triage_issue(86400000, "Lets fix it")` は `"leave it"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(86400000, "Lets fix it"), "leave it")`)
}})
```

`triage_issue(1209600000, "still waiting")` は `"bump it"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(1209600000, "still waiting"), "bump it")`)
}})
```

`triage_issue(864000000, "bump")` は `"close it"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(864000000, "bump"), "close it")`)
}})
```

`triage_issue(604800000, "Do we still want this?")` は `"bump it"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Do we still want this?"), "bump it")`)
}})
```

`triage_issue(604800000, "Bumping this")` は `"close it"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Bumping this"), "close it")`)
}})
```

`triage_issue(345600000, "I'll make a PR")` は `"leave it"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(345600000, "I'll make a PR"), "leave it")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(ms, message):

    return ms
```

# --solutions--

```py
def triage_issue(ms, message):
    seven_days = 7 * 24 * 60 * 60 * 1000
    if ms < seven_days:
        return "leave it"
    if "bump" in message.lower():
        return "close it"
    return "bump it"
```
