---
id: 6a1d9f98e819ed70a0e994df
title: "チャレンジ 333: イシュー トリアージ 2"
challengeType: 29
dashedName: challenge-333
---

# --description--

イシューのタイトルと現在のラベルの配列が与えられたら、以下のルールに基づいて更新されたラベルの配列を返してください。

イシューにラベルがない場合は、次を追加してください。

- タイトルに`"error"`または`"bug"`が含まれている場合は`"bug"`と`"needs triage"`を追加
- タイトルに`"feature"`または`"add"`が含まれている場合は`"enhancement"`と`"discussing"`を追加

それ以外の場合、与えられたラベルに次が含まれている場合は、

- `"needs triage"`がありタイトルに`"simple"`または`"easy"`が含まれている場合は`"needs triage"`を削除し`"good first issue"`を追加
- `"discussing"`がありタイトルに`"planned"`または`"next"`が含まれている場合は`"discussing"`を削除し`"on the roadmap"`を追加
- それ以外で`"needs triage"`または`"discussing"`があれば、それを削除して`"help wanted"`を追加

タイトルに次が含まれている場合は、

- `"security"`があれば`"critical"`ラベルを追加してください

# --hints--

`triage_issue("app crashes with error", [])` は `["bug", "needs triage"]` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", [])), sorted(["bug", "needs triage"]))`)
}})
```

`triage_issue("app crashes with error", ["bug", "needs triage"])` は `["bug", "help wanted"]` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", ["bug", "needs triage"])), sorted(["bug", "help wanted"]))`)
}})
```

`triage_issue("add dark mode", [])` は `["enhancement", "discussing"]` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", [])), sorted(["enhancement", "discussing"]))`)
}})
```

`triage_issue("add dark mode", ["enhancement", "discussing"])` は `["enhancement", "help wanted"]` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted"]))`)
}})
```

`triage_issue("xss security bug", [])` は `["bug", "needs triage", "critical"]` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("xss security bug", [])), sorted(["bug", "needs triage", "critical"]))`)
}})
```

`triage_issue("security vulnerability in auth", [])` は `["critical"]` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("security vulnerability in auth", [])), sorted(["critical"]))`)
}})
```

`triage_issue("easy a11y fix", ["bug", "needs triage"])` は `["bug", "good first issue"]` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("easy a11y fix", ["bug", "needs triage"])), sorted(["bug", "good first issue"]))`)
}})
```

`triage_issue("planned api migration", ["enhancement", "discussing"])` は `["enhancement", "on the roadmap"]` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("planned api migration", ["enhancement", "discussing"])), sorted(["enhancement", "on the roadmap"]))`)
}})
```

`triage_issue("improve security", ["enhancement", "discussing"])` は `["enhancement", "help wanted", "critical"]` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("improve security", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted", "critical"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(title, labels):

    return title
```

# --solutions--

```py
def triage_issue(title, labels):
    result = list(labels)
    t = title.lower()

    def has(label):
        return label in result

    def add(label):
        if not has(label):
            result.append(label)

    def remove(label):
        if label in result:
            result.remove(label)

    if len(result) == 0:
        if "error" in t or "bug" in t:
            add("bug")
            add("needs triage")
        elif "feature" in t or "add" in t:
            add("enhancement")
            add("discussing")
    else:
        if has("needs triage") and ("simple" in t or "easy" in t):
            remove("needs triage")
            add("good first issue")
        elif has("discussing") and ("planned" in t or "next" in t):
            remove("discussing")
            add("on the roadmap")
        elif has("needs triage") or has("discussing"):
            remove("needs triage")
            remove("discussing")
            add("help wanted")

    if "security" in t:
        add("critical")

    return result
```
