---
id: 6a1d9f98e819ed70a0e994dc
title: "チャレンジ 330: 小文字の単語"
challengeType: 29
dashedName: challenge-330
---

# --description--

文字列が与えられたら、完全に小文字の単語だけを元の順序で返し、各単語の間にスペースを入れてください。

# --hints--

`get_lowercase_words("hello GOOD world")` は `"hello world"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("hello GOOD world"), "hello world")`)
}})
```

`get_lowercase_words("these are all lowercase")` は `"these are all lowercase"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("these are all lowercase"), "these are all lowercase")`)
}})
```

`get_lowercase_words("less is NoT more")` は `"less is more"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("less is NoT more"), "less is more")`)
}})
```

`get_lowercase_words("DonT eat pizza every OTHER day")` は `"eat pizza every day"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("DonT eat pizza every OTHER day"), "eat pizza every day")`)
}})
```

`get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` は `"the quick brown fox jumped over the lazy dog"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_lowercase_words(s):

    return s
```

# --solutions--

```py
def get_lowercase_words(s):
    return " ".join(word for word in s.split(" ") if word == word.lower())
```
