---
id: 6a2037a68a0bc2aef0006008
title: "チャレンジ 346: 貯金箱"
challengeType: 29
dashedName: challenge-346
---

# --description--

貯金箱を表すオブジェクトが与えられたら、合計金額を`"$D.CC"`の形式の文字列で返してください。

オブジェクトには以下のいずれかが含まれている場合があります：

| コイン | 価値 |
|------|-------|
| `pennies` | $0.01 |
| `nickels` | $0.05 |
| `dimes` | $0.10 |
| `quarters` | $0.25 |

# --hints--

`piggy_bank({"pennies": 3, "nickels": 5, "dimes": 2, "quarters": 6})` は `"$1.98"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"pennies": 3, "nickels": 5, "dimes": 2, "quarters": 6}), "$1.98")`)
}})
```

`piggy_bank({"pennies": 1, "nickels": 1, "dimes": 1, "quarters": 1})` は `"$0.41"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"pennies": 1, "nickels": 1, "dimes": 1, "quarters": 1}), "$0.41")`)
}})
```

`piggy_bank({"nickels": 8, "dimes": 6, "quarters": 5})` は `"$2.25"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"nickels": 8, "dimes": 6, "quarters": 5}), "$2.25")`)
}})
```

`piggy_bank({})` は `"$0.00"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({}), "$0.00")`)
}})
```

`piggy_bank({"pennies": 146, "nickels": 11, "dimes": 0, "quarters": 19})` は `"$6.76"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"pennies": 146, "nickels": 11, "dimes": 0, "quarters": 19}), "$6.76")`)
}})
```

# --seed--

## --seed-contents--

```py
def piggy_bank(coins):

    return coins
```

# --solutions--

```py
def piggy_bank(coins):
    values = {"pennies": 0.01, "nickels": 0.05, "dimes": 0.10, "quarters": 0.25}
    total = sum(values[coin] * count for coin, count in coins.items())
    return f"${total:.2f}"
```
