---
id: 6a2037a68a0bc2aef0006003
title: "チャレンジ 341: 誕生日カウントダウン"
challengeType: 29
dashedName: challenge-341
---

# --description--

今日の日付と誕生日が与えられたとき、その人の次の誕生日までの日数を返してください。

- 今日の日付は先頭にゼロを付けた `"YYYY-MM-DD"` 形式の文字列で与えられます。例：`"2026-07-16"`。
- 誕生日は先頭にゼロを付けない `"M/D"` 形式の文字列で与えられます。例：`"9/7"`。
- 今日が誕生日の場合は、次の誕生日までの日数を返してください（`0` ではありません）。
- 閏年も考慮してください。

# --hints--

`days_until_birthday("2026-07-16", "9/7")` は `53` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "9/7"), 53)`)
}})
```

`days_until_birthday("2026-07-16", "3/22")` は `249` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "3/22"), 249)`)
}})
```

`days_until_birthday("2026-07-16", "7/16")` は `365` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "7/16"), 365)`)
}})
```

`days_until_birthday("2024-02-28", "3/1")` は `2` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-02-28", "3/1"), 2)`)
}})
```

`days_until_birthday("2023-04-24", "12/30")` は `250` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2023-04-24", "12/30"), 250)`)
}})
```

`days_until_birthday("2024-03-01", "2/29")` は `1460` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-03-01", "2/29"), 1460)`)
}})
```

`days_until_birthday("2096-03-01", "2/29")` は `2920` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2096-03-01", "2/29"), 2920)`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_birthday(today, birthday):

    return today
```

# --solutions--

```py
def days_until_birthday(today, birthday):
    from datetime import date

    year, month, day = [int(x) for x in today.split("-")]
    b_month, b_day = [int(x) for x in birthday.split("/")]

    today_date = date(year, month, day)

    def is_leap_year(y):
        return (y % 4 == 0 and y % 100 != 0) or y % 400 == 0

    def next_valid_birthday(from_year):
        for y in range(from_year, from_year + 9):
            if b_month == 2 and b_day == 29 and not is_leap_year(y):
                continue
            d = date(y, b_month, b_day)
            if d > today_date:
                return d

    return (next_valid_birthday(year) - today_date).days
```
