---
id: 6a2037a68a0bc2aef0006007
title: "チャレンジ 345: ワードブレンダー"
challengeType: 29
dashedName: challenge-345
---

# --description--

2つの単語が与えられたら、最初の単語の前半と2番目の単語の後半を組み合わせて新しい単語を返してください。

- 奇数文字の単語の場合、前半は短い方の半分になります。

# --hints--

`blend_words("turtle", "toucan")` は `"turcan"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("turtle", "toucan"), "turcan")`)
}})
```

`blend_words("chipmunk", "flamingo")` は `"chipingo"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("chipmunk", "flamingo"), "chipingo")`)
}})
```

`blend_words("falcon", "pelican")` は `"falican"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("falcon", "pelican"), "falican")`)
}})
```

`blend_words("hyena", "iguana")` は `"hyana"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("hyena", "iguana"), "hyana")`)
}})
```

`blend_words("scorpion", "gorilla")` は `"scorilla"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("scorpion", "gorilla"), "scorilla")`)
}})
```

`blend_words("platypus", "wolverine")` は `"platerine"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("platypus", "wolverine"), "platerine")`)
}})
```

# --seed--

## --seed-contents--

```py
def blend_words(word1, word2):

    return word1
```

# --solutions--

```py
def blend_words(word1, word2):
    first_half = word1[:len(word1) // 2]
    second_half = word2[len(word2) // 2:]
    return first_half + second_half
```
