---
id: 6a2037a68a0bc2aef0006000
title: "チャレンジ 338: ペット年齢計算機"
challengeType: 29
dashedName: challenge-338
---

# --description--

ペットの種類と人間の年齢が与えられたら、次の換算テーブルを使ってペットの年齢に換算した値を返してください。

| ペット | 乗数 |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`pet_years("dog", 5)` は `35` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("dog", 5), 35)`)
}})
```

`pet_years("cat", 9)` は `54` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("cat", 9), 54)`)
}})
```

`pet_years("rabbit", 3)` は `24` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("rabbit", 3), 24)`)
}})
```

`pet_years("hamster", 4)` は `120` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("hamster", 4), 120)`)
}})
```

`pet_years("guinea pig", 5)` は `60` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("guinea pig", 5), 60)`)
}})
```

`pet_years("goldfish", 2)` は `12` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("goldfish", 2), 12)`)
}})
```

`pet_years("bird", 1)` は `5` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("bird", 1), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def pet_years(pet, age):

    return pet
```

# --solutions--

```py
def pet_years(pet, age):
    table = {
        "dog": 7,
        "cat": 6,
        "rabbit": 8,
        "hamster": 30,
        "guinea pig": 12,
        "goldfish": 6,
        "bird": 5,
    }
    return age * table[pet]
```
