---
id: 6a2037a68a0bc2aef0005fff
title: "チャレンジ 337: タリーカウンター"
challengeType: 29
dashedName: challenge-337
---

# --description--

タリーマークの文字列が与えられたら、表されている合計数を返してください。

- 各パイプ `"|"` は1つのカウントを表します。
- 5つ目のマークはスラッシュ `"/"` で表され、5つのグループを完成させます（`"||||/"`）。
- グループはスペースで区切られています。

# --hints--

`get_tally_count("||||")` は `4` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||"), 4)`)
}})
```

`get_tally_count("||||/")` は `5` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/"), 5)`)
}})
```

`get_tally_count("||||/ |||")` は `8` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ |||"), 8)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||")` は `17` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||"), 17)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` は `41` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_tally_count(s):

    return s
```

# --solutions--

```py
def get_tally_count(s):
    return len(s.replace(" ", ""))
```
