---
id: 6a2037a68a0bc2aef0006002
title: "チャレンジ 340: ピッグラテン変換"
challengeType: 29
dashedName: challenge-340
---

# --description--

文字列が与えられたら、次のルールに従ってピッグラテンに変換してください。

- 単語が母音（`"a"`、`"e"`、`"i"`、`"o"`、または`"u"`）で始まる場合は、末尾に`"way"`を追加します。例えば、`"universe"`は`"universeway"`に変換されます。
- 単語が1つ以上の子音で始まる場合は、それらを末尾に移動し、`"ay"`を追加します。例えば、`"hello"`は`"ellohay"`に変換されます。
- 最初の文字の大文字・小文字は保持してください。例えば、`"Hello"`は`"Ellohay"`に変換されます。

# --hints--

`pig_latin("universe")` は `"universeway"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("universe"), "universeway")`)
}})
```

`pig_latin("hello")` は `"ellohay"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello"), "ellohay")`)
}})
```

`pig_latin("hello universe")` は `"ellohay universeway"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello universe"), "ellohay universeway")`)
}})
```

`pig_latin("Hello universe")` は `"Ellohay universeway"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Hello universe"), "Ellohay universeway")`)
}})
```

`pig_latin("Pig Latin is fun")` は `"Igpay Atinlay isway unfay"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Pig Latin is fun"), "Igpay Atinlay isway unfay")`)
}})
```

`pig_latin("The quick brown fox jumped over the lazy dog")` は `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday")`)
}})
```

# --seed--

## --seed-contents--

```py
def pig_latin(s):

    return s
```

# --solutions--

```py
def pig_latin(s):
    vowels = "aeiou"

    def convert_word(word):
        lower_word = word.lower()
        if lower_word[0] in vowels:
            return word + "way"
        i = 0
        while i < len(word) and lower_word[i] not in vowels:
            i += 1
        consonants = word[:i]
        rest = word[i:]
        result = rest + consonants + "ay"
        if word[0].isupper():
            return result[0].upper() + result[1:].lower()
        return result

    return " ".join(convert_word(word) for word in s.split(" "))
```
