---
id: 6a2037a68a0bc2aef0006006
title: "チャレンジ 344: 黄金比"
challengeType: 29
dashedName: challenge-344
---

# --description--

2つの数字が与えられたとき、その比率が黄金比に近いかどうかを判定してください。

- 黄金比は `1.618` を使ってください
- 許容誤差は `0.01` としてください

# --hints--

`is_golden_ratio(21, 34)` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(21, 34), True)`)
}})
```

`is_golden_ratio(15, 20)` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(15, 20), False)`)
}})
```

`is_golden_ratio(8, 13)` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(8, 13), True)`)
}})
```

`is_golden_ratio(10, 16)` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(10, 16), False)`)
}})
```

`is_golden_ratio(1618, 1000)` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(1618, 1000), True)`)
}})
```

`is_golden_ratio(88, 55)` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(88, 55), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_golden_ratio(a, b):

    return a
```

# --solutions--

```py
def is_golden_ratio(a, b):
    golden_ratio = 1.618
    ratio = max(a, b) / min(a, b)
    return abs(ratio - golden_ratio) <= 0.01
```
