---
id: 6a22d77ddf034bc4e35b1d5b
title: "チャレンジ 354: コントラスト評価 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

RGB値を表す2つの配列とテキストが大きいかどうかを示すブール値が与えられたら、次の方法でWCAGのコントラスト評価を返してください。

まず、各RGB値を相対輝度に変換します。

- 各チャンネル`[R, G, B]`を255で割り、`0`から`1`の範囲の値にします
- 各チャンネルにガンマ補正の式を適用します:
  - チャンネル値が`0.04045`以下の場合: `channel / 12.92`
  - それ以外の場合: `((channel + 0.055) / 1.055) ^ 2.4`
- 輝度を計算します: `0.2126 * R + 0.7152 * G + 0.0722 * B`

次に、各輝度値に0.05を加え、明るい方の値を暗い方の値で割ってコントラスト比を計算します。明るい方は常に最初の引数になります。

次のテーブルを使ってコントラスト比に基づく評価を返してください。

| 評価 | 通常テキスト | 大きなテキスト |
|--------|-------------|------------|
| `"AAA"` | 7.0以上 | 4.5以上 |
| `"AA"` | 4.5以上 | 3.0以上 |
| `"Fail"` | 4.5未満 | 3.0未満 |

# --hints--

`get_contrast_rating([255, 255, 255], [0, 0, 0], False)` は `"AAA"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

`get_contrast_rating([215, 188, 188], [55, 55, 55], False)` は `"AA"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

`get_contrast_rating([143, 144, 210], [46, 47, 61], False)` は `"Fail"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

`get_contrast_rating([167, 167, 210], [53, 10, 53], True)` は `"AAA"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

`get_contrast_rating([135, 147, 155], [60, 70, 90], True)` は `"AA"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

`get_contrast_rating([125, 210, 195], [105, 130, 90], True)` は `"Fail"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
