---
id: 6a22d77ddf034bc4e35b1d58
title: "チャレンジ 351: プロニック数"
challengeType: 29
dashedName: challenge-351
---

# --description--

数値がプロニック数かどうかを判定してください。

プロニック数とは、連続する2つの整数の積です。例えば、6は2 * 3 = 6なのでプロニック数です。

# --hints--

`is_pronic(6)` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(6), True)`)
}})
```

`is_pronic(15)` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(15), False)`)
}})
```

`is_pronic(12)` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(12), True)`)
}})
```

`is_pronic(132)` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(132), True)`)
}})
```

`is_pronic(80)` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(80), False)`)
}})
```

`is_pronic(0)` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(0), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_pronic(n):

    return n
```

# --solutions--

```py
def is_pronic(n):
    k = int(n ** 0.5)
    return k * (k + 1) == n
```
