---
id: 6a22d77ddf034bc4e35b1d54
title: "チャレンジ 347: ゲーム理論"
challengeType: 29
dashedName: challenge-347
---

# --description--

2人のプレイヤーの戦略を表す同じ長さの2つの文字列が与えられたとき、スコアを配列 `[player1, player2]` として返してください。

- 与えられた文字列は、`"C"`（協力）か`"D"`（裏切り）のいずれかの文字のみを含みます。
- 各文字は1ラウンドを表し、スコアは次のように計算されます：
  - 両方のプレイヤーが協力した場合、各自3点。
  - 両方のプレイヤーが裏切った場合、各自1点。
  - 片方が裏切り、もう片方が協力した場合、裏切ったプレイヤーは5点、協力したプレイヤーは0点。

# --hints--

`play_game("CCCC", "CCCC")` は `[12, 12]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` は `[4, 4]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` は `[5, 10]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` は `[24, 34]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` は `[66, 21]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
