---
id: 6a22d77ddf034bc4e35b1d5a
title: "チャレンジ 353: コントラスト評価 2"
challengeType: 29
dashedName: challenge-353
---

# --description--

2つの相対輝度値とテキストが大きいかどうかを示すブール値が与えられたら、次の方法でWCAGのコントラスト評価を返してください。

各輝度値に0.05を加え、明るい方の値を暗い方の値で割ってコントラスト比を計算します。明るい方の値は常に最初の引数になります。

次のテーブルを使ってコントラスト比に基づく評価を返してください。

| 評価 | 通常テキスト | 大きなテキスト |
|--------|-------------|------------|
| `"AAA"` | 7.0以上 | 4.5以上 |
| `"AA"` | 4.5以上 | 3.0以上 |
| `"Fail"` | 4.5未満 | 3.0未満 |

# --hints--

`get_contrast_rating(1.0, 0.0, False)` は `"AAA"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(1.0, 0.0, False), "AAA")`)
}})
```

`get_contrast_rating(0.9015, 0.1364, False)` は `"AA"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.9015, 0.1364, False), "AA")`)
}})
```

`get_contrast_rating(0.8965, 0.1628, False)` は `"Fail"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.8965, 0.1628, False), "Fail")`)
}})
```

`get_contrast_rating(0.7469, 0.0957, True)` は `"AAA"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7469, 0.0957, True), "AAA")`)
}})
```

`get_contrast_rating(0.7489, 0.2018, True)` は `"AA"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7489, 0.2018, True), "AA")`)
}})
```

`get_contrast_rating(0.6571, 0.1974, True)` は `"Fail"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.6571, 0.1974, True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(l1, l2, is_large_text):

    return l1
```

# --solutions--

```py
def get_contrast_rating(l1, l2, is_large_text):
    ratio = (l1 + 0.05) / (l2 + 0.05)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
