---
id: 6a22d77ddf034bc4e35b1d59
title: "チャレンジ 352: コントラスト評価 1"
challengeType: 29
dashedName: challenge-352
---

# --description--

コントラスト比とテキストが大きいかどうかを示すブール値が与えられたら、次のテーブルを使ってWCAG評価を返してください。

| 評価 | 通常テキスト | 大きなテキスト |
|--------|-------------|------------|
| `"AAA"` | 7.0以上 | 4.5以上 |
| `"AA"` | 4.5以上 | 3.0以上 |
| `"Fail"` | 4.5未満 | 3.0未満 |

# --hints--

`get_contrast_rating("7.5", False)` は `"AAA"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("7.5", False), "AAA")`)
}})
```

`get_contrast_rating("4.8", False)` は `"AA"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.8", False), "AA")`)
}})
```

`get_contrast_rating("4.2", False)` は `"Fail"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.2", False), "Fail")`)
}})
```

`get_contrast_rating("4.5", True)` は `"AAA"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.5", True), "AAA")`)
}})
```

`get_contrast_rating("3.0", True)` は `"AA"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("3.0", True), "AA")`)
}})
```

`get_contrast_rating("2.7", False)` は `"Fail"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("2.7", False), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(ratio, is_large_text):

    return ratio
```

# --solutions--

```py
def get_contrast_rating(ratio, is_large_text):
    ratio = float(ratio)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
