---
id: 6a26df95efa55a2524399744
title: "チャレンジ 363: バケット塗りつぶし 2"
challengeType: 29
dashedName: challenge-363
---

# --description--

1文字の色の文字列で構成された2Dグリッドとターゲットの色が与えられたとき、グリッド全体をターゲットの色にするために必要な最小のフラッドフィルの「クリック」回数を返してください。

- クリックするたびに、クリックしたセルの色と同じ色でつながっているセルの領域全体がターゲットの色に変わります。
- セルは水平方向および垂直方向に接続されており（斜めは含みません）。

# --hints--

`bucket_fill([["R", "R"], ["R", "R"]], "G")` は `1` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["R", "R"], ["R", "R"]], "G"), 1)`)
}})
```

`bucket_fill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B")` は `0` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B"), 0)`)
}})
```

`bucket_fill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R")` は `3` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R"), 3)`)
}})
```

`bucket_fill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P")` は `5` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P"), 5)`)
}})
```

`bucket_fill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y")` は `12` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y"), 12)`)
}})
```

# --seed--

## --seed-contents--

```py
def bucket_fill(grid, target_color):

    return grid
```

# --solutions--

```py
def bucket_fill(grid, target_color):
    rows = len(grid)
    cols = len(grid[0])
    visited = set()

    def dfs(r, c, color):
        if r < 0 or r >= rows or c < 0 or c >= cols:
            return
        if (r, c) in visited:
            return
        if grid[r][c] != color:
            return
        visited.add((r, c))
        dfs(r + 1, c, color)
        dfs(r - 1, c, color)
        dfs(r, c + 1, color)
        dfs(r, c - 1, color)

    clicks = 0
    for r in range(rows):
        for c in range(cols):
            if (r, c) not in visited and grid[r][c] != target_color:
                dfs(r, c, grid[r][c])
                clicks += 1

    return clicks
```
