---
id: 6a26df3e2988bcdded204895
title: "チャレンジ 359: ゴルフ ハンディキャップ計算機"
challengeType: 29
dashedName: challenge-359
---

# --description--

ゴルフのスコアの配列と対応するコースのパーの配列が与えられたとき、次の方法でゴルファーのハンディキャップ指数を返してください。

- 各ラウンドの差分をスコアからパーを引いて計算し、すべての差分の平均を小数点以下1桁に丸めて返してください。

# --hints--

`calculate_handicap([72, 72, 72], [72, 72, 72])` は `0` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([72, 72, 72], [72, 72, 72]), 0)`)
}})
```

`calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72])` は `6` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72]), 6)`)
}})
```

`calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36])` は `8.3` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3)`)
}})
```

`calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` は `8.8` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8)`)
}})
```

`calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` は `11.7` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_handicap(scores, pars):

    return scores
```

# --solutions--

```py
def calculate_handicap(scores, pars):
    differentials = [s - p for s, p in zip(scores, pars)]
    average = sum(differentials) / len(differentials)
    return int(average * 10 + 0.5) / 10
```
