---
id: 6a26df95efa55a2524399743
title: "チャレンジ 362: ノノグラムバリデーター"
challengeType: 29
dashedName: challenge-362
---

# --description--

手がかりの数字の配列とセルの配列が与えられたとき、そのセルがノノグラムの手がかりを満たしているかどうかを判定してください。

- 手がかりは、連続した塗りつぶされたセルの長さを順番に表す数字の配列です。例えば、手がかりが`[3, 2]`の場合、3つの連続した塗りつぶしセルの後に少なくとも1つの空セルを挟んで2つの連続した塗りつぶしセルがあることを意味します。
- 行は1（塗りつぶし）と0（空）の配列です。

# --hints--

`is_valid_nonogram([3, 2], [1, 1, 1, 0, 1, 1])` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2], [1, 1, 1, 0, 1, 1]), True)`)
}})
```

`is_valid_nonogram([3, 2], [0, 1, 1, 1, 1, 1])` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2], [0, 1, 1, 1, 1, 1]), False)`)
}})
```

`is_valid_nonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]), False)`)
}})
```

`is_valid_nonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]), True)`)
}})
```

`is_valid_nonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])` は `True` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]), True)`)
}})
```

`is_valid_nonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])` は `False` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_nonogram(clue, cells):

    return clue
```

# --solutions--

```py
def is_valid_nonogram(clue, cells):
    runs = []
    count = 0

    for i in range(len(cells) + 1):
        if i < len(cells) and cells[i] == 1:
            count += 1
        elif count > 0:
            runs.append(count)
            count = 0

    if len(runs) != len(clue):
        return False
    return all(runs[i] == clue[i] for i in range(len(runs)))
```
