---
id: 6a26df3e2988bcdded204894
title: "チャレンジ 358: 絵文字翻訳"
challengeType: 29
dashedName: challenge-358
---

# --description--

絵文字の文字列が与えられたら、次のテーブルを使ってフレーズを返してください。

| Emoji | Word |
|-------|------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

単語はスペースで区切って返してください。

# --hints--

`get_emoji_phrase("🪨⭐")` は `"rock star"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🪨⭐"), "rock star")`)
}})
```

`get_emoji_phrase("🥵🐕")` は `"hot dog"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🥵🐕"), "hot dog")`)
}})
```

`get_emoji_phrase("👶🦈")` は `"baby shark"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("👶🦈"), "baby shark")`)
}})
```

`get_emoji_phrase("⭐🐟")` は `"star fish"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("⭐🐟"), "star fish")`)
}})
```

`get_emoji_phrase("🧊🧊👶")` は `"ice ice baby"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🧊🧊👶"), "ice ice baby")`)
}})
```

`get_emoji_phrase("🐱🐟🍲")` は `"cat fish soup"` を返す必要があります。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🐱🐟🍲"), "cat fish soup")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_emoji_phrase(s):

    return s
```

# --solutions--

```py
def get_emoji_phrase(s):
    table = {
        '👶': 'baby',
        '🐱': 'cat',
        '🐕': 'dog',
        '🐟': 'fish',
        '🥵': 'hot',
        '🧊': 'ice',
        '🪨': 'rock',
        '🦈': 'shark',
        '🍲': 'soup',
        '⭐': 'star',
    }
    return ' '.join(table[emoji] for emoji in s)
```
