---
id: 6a26df3e2988bcdded204893
title: "チャレンジ 357: 食物連鎖"
challengeType: 29
dashedName: challenge-357
---

# --description--

`[predator, prey]` のペアの配列が与えられたら、頂点捕食者から最下層までの食物連鎖を返してください。

- 頂点捕食者とは、他の動物に捕食されることがない動物です。
- 連鎖は文字列の配列として返してください。

# --hints--

`get_food_chain([["cat", "mouse"]])` は `["cat", "mouse"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["cat", "mouse"]]), ["cat", "mouse"])`)
}})
```

`get_food_chain([["wolf", "deer"], ["deer", "grass"]])` は `["wolf", "deer", "grass"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"])`)
}})
```

`get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` は `["hawk", "snake", "frog", "fly"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"])`)
}})
```

`get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` は `["eagle", "fox", "rabbit", "grass"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"])`)
}})
```

`get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` は `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_food_chain(pairs):

    return pairs
```

# --solutions--

```py
def get_food_chain(pairs):
    prey_set = {p for _, p in pairs}
    chain_map = dict(pairs)
    apex = next(predator for predator, _ in pairs if predator not in prey_set)
    chain = [apex]
    while chain[-1] in chain_map:
        chain.append(chain_map[chain[-1]])
    return chain
```
