---
id: 6a26df3e2988bcdded204897
title: "チャレンジ 361: 話し言葉の時間"
challengeType: 29
dashedName: challenge-361
---

# --description--

アナログ時計の時針と分針の角度（12時から時計回りの度数）が与えられたら、話し言葉の英語で時間を返してください。

分針の角度を分に変換します（360°＝60分）。その後、以下のルールを使ってください。

| 分 | 話し言葉 |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29（15を除く） | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59（45を除く） | `"X minutes to Z"`（ここで`X`は60 - 分）|

`Y`は現在の時、`Z`は次の時で、どちらも時針の角度（360°＝12時間）から導き出されます。

注意：針の角度は正確に数字に合わない場合があるので、何らかの方法で丸めることを検討してください。

# --hints--

`get_spoken_time(90, 0)` は `"3 o'clock"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(90, 0), "3 o'clock")`)
}})
```

`get_spoken_time(160, 120)` は `"20 minutes past 5"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(160, 120), "20 minutes past 5")`)
}})
```

`get_spoken_time(255, 180)` は `"half past 8"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(255, 180), "half past 8")`)
}})
```

`get_spoken_time(67.5, 92)` は `"quarter past 2"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(67.5, 92), "quarter past 2")`)
}})
```

`get_spoken_time(200, 240)` は `"20 minutes to 7"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(200, 240), "20 minutes to 7")`)
}})
```

`get_spoken_time(322.5, 273)` は `"quarter to 11"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(322.5, 273), "quarter to 11")`)
}})
```

`get_spoken_time(117.5, 335)` は `"5 minutes to 4"` を返すべきです。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(117.5, 335), "5 minutes to 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_time(hour_angle, minute_angle):

    return hour_angle
```

# --solutions--

```py
def get_spoken_time(hour_angle, minute_angle):
    minutes = int(minute_angle / 6)
    hour = int(hour_angle / 30) or 12
    next_hour = (hour % 12) + 1

    if minutes == 0:
        return f"{hour} o'clock"
    if minutes == 15:
        return f"quarter past {hour}"
    if minutes == 30:
        return f"half past {hour}"
    if minutes == 45:
        return f"quarter to {next_hour}"
    if minutes < 30:
        return f"{minutes} minutes past {hour}"
    return f"{60 - minutes} minutes to {next_hour}"
```
