---
id: 587d7b8a367417b2b2512b4f
title: Scrivere dichiarazioni concise di oggetti letterali usando la sintassi abbreviata delle proprietà
challengeType: 1
forumTopicId: 301225
dashedName: write-concise-object-literal-declarations-using-object-property-shorthand
---

# --description--

ES6 aggiunge un ottimo supporto per definire facilmente oggetti letterali.

Considera il seguente codice:

```js
const getMousePosition = (x, y) => ({
  x: x,
  y: y
});
```

`getMousePosition` è una funzione semplice che restituisce un oggetto contenente due proprietà. ES6 fornisce uno zucchero sintattico per eliminare la ridondanza di dover scrivere `x: x`. Puoi semplicemente scrivere `x` una volta, e verrà convertito in `x: x` (o qualcosa di equivalente) sotto il cofano. Ecco la stessa funzione di prima riscritta per usare questa nuova sintassi:

```js
const getMousePosition = (x, y) => ({ x, y });
```

# --instructions--

Usa la sintassi abbreviata delle proprietà con oggetti letterali per creare e restituire un oggetto con le proprietà `name`, `age` e `gender`.

# --hints--

`createPerson("Zodiac Hasbro", 56, "male")` dovrebbe restituire `{name: "Zodiac Hasbro", age: 56, gender: "male"}`.

```js
assert.deepEqual(
  { name: 'Zodiac Hasbro', age: 56, gender: 'male' },
  createPerson('Zodiac Hasbro', 56, 'male')
);
```

Il tuo codice non dovrebbe usare `key:value`.

```js
assert(!__helpers.removeJSComments(code).match(/:/g))
```

# --seed--

## --seed-contents--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  return {
    name: name,
    age: age,
    gender: gender
  };
  // Only change code above this line
};
```

# --solutions--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  /*return {
    name: name,
    age: age,
    gender: gender
  };*/
  return {
    name,
    age,
    gender
  };
  // Only change code above this line
};
```
