---
id: 587d7b8a367417b2b2512b4e
title: Creare stringhe usando template literal
challengeType: 1
forumTopicId: 301200
dashedName: create-strings-using-template-literals
---

# --description--

Una nuova funzionalità di ES6 è il <dfn>template literal</dfn>. Si tratta di un tipo speciale di stringa che rende più semplice creare stringhe complesse.

I template literal ti permettono di creare stringhe su più righe e di usare funzionalità di interpolazione delle stringhe per costruirle.

Considera il codice qui sotto:

```js
const person = {
  name: "Zodiac Hasbro",
  age: 56
};

const greeting = `Hello, my name is ${person.name}!
I am ${person.age} years old.`;

console.log(greeting);
```

La console mostrerà le stringhe `Hello, my name is Zodiac Hasbro!` e `I am 56 years old.`.

Sono successe molte cose. Innanzitutto, l'esempio usa i backtick (`` ` ``), not quotes (`'` or `"`), to wrap the string. Secondly, notice that the string is multi-line, both in the code and the output. This saves inserting `\n` within strings. The `${variable}` syntax used above is a placeholder. Basically, you won't have to use concatenation with the `+` operator anymore. To add variables to strings, you just drop the variable in a template string and wrap it with `${` and `}`. Similarly, you can include other expressions in your string literal, for example `${a + b}`). Questo nuovo modo di creare stringhe ti dà più flessibilità per costruire stringhe robuste.

# --instructions--

Usa la sintassi dei template literal con i backtick per creare un array di stringhe di elementi lista (`li`). Il testo di ogni elemento lista dovrebbe essere uno degli elementi dell'array nella proprietà `failure` dell'oggetto `result` e avere un attributo `class` con valore `text-warning`. La funzione `makeList` dovrebbe restituire l'array di stringhe degli elementi lista.

Usa un metodo iteratore (qualsiasi tipo di loop) per ottenere l'output desiderato (mostrato sotto).

```js
[
  '<li class="text-warning">no-var</li>',
  '<li class="text-warning">var-on-top</li>',
  '<li class="text-warning">linebreak</li>'
]
```

# --hints--

`failuresList` dovrebbe essere un array contenente `result failure` messaggi.

```js
assert(
  typeof makeList(result.failure) === 'object' && failuresList.length === 3
);
```

`failuresList` dovrebbe essere uguale all'output specificato.

```js
assert(
  makeList(result.failure).every(
    (v, i) =>
      v === `<li class="text-warning">${result.failure[i]}</li>` ||
      v === `<li class='text-warning'>${result.failure[i]}</li>`
  )
);
```

Dovresti usare stringhe template e interpolazione di espressioni.

```js
assert.match(code, /(`.*\${.*}.*`)/);
```

Dovresti usare un iteratore.

```js
assert(__helpers.removeJSComments(code).match(/for|map|reduce|forEach|while/));
```

# --seed--

## --seed-contents--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  // Only change code below this line
  const failureItems = [];
  // Only change code above this line

  return failureItems;
}

const failuresList = makeList(result.failure);
```

# --solutions--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  return arr.map(val => `<li class="text-warning">${val}</li>`);
}

const failuresList = makeList(result.failure);
```
