---
id: 587d7b8c367417b2b2512b54
title: "Usa getter e setter per controllare l'accesso a un oggetto"
challengeType: 1
forumTopicId: 301220
dashedName: use-getters-and-setters-to-control-access-to-an-object
---

# --description--

Puoi ottenere valori da un oggetto e impostare il valore di una proprietà all'interno di un oggetto.

Questi sono classicamente chiamati <dfn>getter</dfn> e <dfn>setter</dfn>.

Le funzioni getter servono semplicemente a restituire (ottenere) il valore di una variabile privata di un oggetto all'utente senza che l'utente consulti direttamente la variabile privata.

Le funzioni setter servono a modificare (impostare) il valore di una variabile privata di un oggetto in base al valore passato alla funzione setter. Questa modifica può includere calcoli o anche la completa sovrascrittura del valore precedente.

```js
class Book {
  constructor(author) {
    this._author = author;
  }
  // getter
  get writer() {
    return this._author;
  }
  // setter
  set writer(updatedAuthor) {
    this._author = updatedAuthor;
  }
}
const novel = new Book('anonymous');
console.log(novel.writer);
novel.writer = 'newAuthor';
console.log(novel.writer);
```

La console mostrerebbe le stringhe `anonymous` e `newAuthor`.

Nota la sintassi usata per invocare getter e setter. Non sembrano nemmeno funzioni. Getter e setter sono importanti perché nascondono i dettagli di implementazione interni.

**Nota:** È convenzione anteporre un underscore (`_`) al nome di una variabile privata. Tuttavia, questa pratica di per sé non rende una variabile privata.

# --instructions--

Usa la parola chiave `class` per creare una classe `Thermostat`. La `constructor` accetta una temperatura in Fahrenheit.

Nella classe, crea un `getter` per ottenere la temperatura in Celsius e un `setter` che accetta una temperatura in Celsius.

Ricorda che `C = 5/9 * (F - 32)` e `F = C * 9.0 / 5 + 32`, dove `F` è il valore della temperatura in Fahrenheit e `C` è il valore della stessa temperatura in Celsius.

**Nota:** Quando implementi questo, terrai traccia della temperatura all'interno della classe in una sola scala, Fahrenheit o Celsius.

Questa è la potenza di un getter e di un setter. Stai creando un'API per un altro utente, che può ottenere il risultato corretto indipendentemente da quale scala tieni traccia.

In altre parole, stai astrarre i dettagli di implementazione dall'utente.

# --hints--

`Thermostat` dovrebbe essere una `class` con un metodo `constructor` definito.

```js
assert.isFunction(Thermostat);
assert.isFunction(Thermostat?.constructor);
```

La parola chiave `class` dovrebbe essere usata.

```js
assert.match(code, /class/);
```

`Thermostat` dovrebbe poter essere istanziato.

```js
const _t = new Thermostat(122);
assert.isObject(_t);
```

Quando istanziato con un valore in Fahrenheit, `Thermostat` dovrebbe impostare la `temperature` corretta.

```js
const _t = new Thermostat(122);
assert.strictEqual(_t?.temperature, 50);
```

Un `getter` dovrebbe essere definito.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.get);
```

Un `setter` dovrebbe essere definito.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.set);
```

Chiamare il `setter` con un valore in Celsius dovrebbe impostare il `temperature`.

```js
const _t = new Thermostat(32);
_t.temperature = 26;
const _u = new Thermostat(32);
_u.temperature = 50;
assert.approximately(_t.temperature, 26, 0.1);
assert.approximately(_u.temperature, 50, 0.1);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```

# --solutions--

```js
class Thermostat {
  constructor(fahrenheit) {
    this._tempInCelsius = 5/9 * (fahrenheit - 32);
  }
  get temperature(){
    return this._tempInCelsius;
  }
  set temperature(newTemp){
    this._tempInCelsius = newTemp;
  }
}

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```
