---
id: 587d7b8c367417b2b2512b55
title: Riutilizzare il codice JavaScript usando import
challengeType: 1
forumTopicId: 301208
dashedName: reuse-javascript-code-using-import
---

# --description--

`import` ti permette di scegliere quali parti di un file o modulo caricare. Nella lezione precedente, gli esempi esportavano `add` dal file `math_functions.js`. Ecco come puoi importarlo per usarlo in un altro file:

```js
import { add } from './math_functions.js';
```

Qui, `import` troverà `add` in `math_functions.js`, importerà solo quella funzione per te da usare e ignorerà il resto. `./` indica all'import di cercare il file `math_functions.js` nella stessa cartella del file corrente. Il percorso relativo del file (`./`) e l'estensione del file (`.js`) sono necessari quando si usa import in questo modo.

Puoi importare più di un elemento dal file aggiungendoli nell'istruzione `import` così:

```js
import { add, subtract } from './math_functions.js';
```

# --instructions--

Aggiungi l'istruzione `import` appropriata che permetta al file corrente di usare le funzioni `uppercaseString` e `lowercaseString` che hai esportato nella lezione precedente. Queste funzioni si trovano in un file chiamato `string_functions.js`, che è nella stessa directory del file corrente.

# --hints--

Dovresti importare correttamente `uppercaseString`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

Dovresti importare correttamente `lowercaseString`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

uppercaseString("hello");
lowercaseString("WORLD!");
```

# --solutions--

```js
import { uppercaseString, lowercaseString } from './string_functions.js';

uppercaseString("hello");
lowercaseString("WORLD!");
```
