---
id: 587d7b8c367417b2b2512b57
title: Usa * per importare tutto da un file
challengeType: 1
forumTopicId: 301210
dashedName: use--to-import-everything-from-a-file
---

# --description--

Supponiamo che tu abbia un file e voglia importare tutto il suo contenuto nel file corrente. Questo si può fare con la sintassi `import * as`. Ecco un esempio in cui il contenuto di un file chiamato `math_functions.js` viene importato in un file nella stessa directory:

```js
import * as myMathModule from "./math_functions.js";
```

L'istruzione `import` sopra creerà un oggetto chiamato `myMathModule`. Questo è solo un nome di variabile, puoi chiamarlo come vuoi. L'oggetto conterrà tutte le esportazioni da `math_functions.js`, così potrai consultare le funzioni come faresti con qualsiasi altra proprietà di un oggetto. Ecco come puoi usare le funzioni `add` e `subtract` che sono state importate:

```js
myMathModule.add(2,3);
myMathModule.subtract(5,3);
```

# --instructions--

Il codice in questo file richiede il contenuto del file: `string_functions.js`, che si trova nella stessa directory del file corrente. Usa la sintassi `import * as` per importare tutto dal file in un oggetto chiamato `stringFunctions`.

# --hints--

Il tuo codice dovrebbe usare correttamente la sintassi `import * as`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*\*\s*as\s+stringFunctions\s+from\s*('|")\.\/string_functions\.js\1/g
  )
);
```

# --seed--

## --seed-contents--

```js

// Only change code above this line

stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```

# --solutions--

```js
import * as stringFunctions from "./string_functions.js";

// add code above this line
stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```
