---
id: 587d7b88367417b2b2512b47
title: Usa il parametro rest con i parametri di una funzione
challengeType: 1
forumTopicId: 301221
dashedName: use-the-rest-parameter-with-function-parameters
---

# --description--

Per aiutarci a creare funzioni più flessibili, ES6 introduce il <dfn>parametro rest</dfn> per i parametri delle funzioni. Con il parametro rest, puoi creare funzioni che accettano un numero variabile di argomenti. Questi argomenti vengono archiviati in un array che può essere consultato successivamente all’interno della funzione.

Dai un’occhiata a questo codice:

```js
function howMany(...args) {
  return "You have passed " + args.length + " arguments.";
}
console.log(howMany(0, 1, 2));
console.log(howMany("string", null, [1, 2, 3], { }));
```

La console mostrerebbe le stringhe `You have passed 3 arguments.` e `You have passed 4 arguments.`.

Il parametro rest elimina la necessità di usare l’oggetto `arguments` e ci permette di usare i metodi degli array sull’array di parametri passati alla funzione `howMany`.

# --instructions--

Modifica la funzione `sum` usando il parametro rest in modo che la funzione `sum` possa accettare un numero qualsiasi di argomenti e restituire la loro somma.

# --hints--

Il risultato di `sum(0,1,2)` dovrebbe essere 3

```js
assert(sum(0, 1, 2) === 3);
```

Il risultato di `sum(1,2,3,4)` dovrebbe essere 10

```js
assert(sum(1, 2, 3, 4) === 10);
```

Il risultato di `sum(5)` dovrebbe essere 5

```js
assert(sum(5) === 5);
```

Il risultato di `sum()` dovrebbe essere 0

```js
assert(sum() === 0);
```

`sum` dovrebbe essere una funzione freccia che usa la sintassi del parametro rest (`...`) sul parametro `args`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/sum=\(\.\.\.args\)=>/));
```

# --seed--

## --seed-contents--

```js
const sum = (x, y, z) => {
  const args = [x, y, z];
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```

# --solutions--

```js
const sum = (...args) => {
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```
