---
id: 587d7b89367417b2b2512b49
title: "Usa l'assegnazione destrutturante per assegnare variabili dagli oggetti"
challengeType: 1
forumTopicId: 301215
dashedName: use-destructuring-assignment-to-assign-variables-from-objects
---

# --description--

La destrutturazione ti permette di assegnare un nuovo nome di variabile quando estrai i valori. Puoi farlo mettendo il nuovo nome dopo i due punti durante l'assegnazione del valore.

Usando lo stesso oggetto dell'ultimo esempio:

```js
const user = { name: 'John Doe', age: 34 };
```

Ecco come puoi dare nuovi nomi di variabili nell'assegnazione:

```js
const { name: userName, age: userAge } = user;
```

Puoi leggerlo come "prendi il valore di `user.name` e assegnalo a una nuova variabile chiamata `userName`" e così via. Il valore di `userName` sarebbe la stringa `John Doe`, e il valore di `userAge` sarebbe il numero `34`.

# --instructions--

Sostituisci le due assegnazioni con un'assegnazione destrutturante equivalente. Dovrebbe comunque assegnare alle variabili `highToday` e `highTomorrow` i valori di `today` e `tomorrow` dall'oggetto `HIGH_TEMPERATURES`.

# --hints--

Dovresti rimuovere la sintassi di assegnazione ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/highToday = HIGH_TEMPERATURES\.today/g) &&
    !__helpers.removeJSComments(code).match(/highTomorrow = HIGH_TEMPERATURES\.tomorrow/g)
);
```

Dovresti usare la destrutturazione per creare la variabile `highToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(today\s*:\s*highToday[^}]*|[^,]*,\s*today\s*:\s*highToday\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

Dovresti usare la destrutturazione per creare la variabile `highTomorrow`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(tomorrow\s*:\s*highTomorrow[^}]*|[^,]*,\s*tomorrow\s*:\s*highTomorrow\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

`highToday` dovrebbe essere uguale a `77` e `highTomorrow` dovrebbe essere uguale a `80`.

```js
assert(highToday === 77 && highTomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line
  
const highToday = HIGH_TEMPERATURES.today;
const highTomorrow = HIGH_TEMPERATURES.tomorrow; 

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today: highToday, tomorrow: highTomorrow } = HIGH_TEMPERATURES;
```
