---
id: 587d7b8c367417b2b2512b56
title: Usa export per condividere un blocco di codice
challengeType: 1
forumTopicId: 301219
dashedName: use-export-to-share-a-code-block
---

# --description--

Immagina un file chiamato `math_functions.js` che contiene diverse funzioni relative a operazioni matematiche. Una di queste è memorizzata in una variabile, `add`, che prende in input due numeri e restituisce la loro somma. Vuoi usare questa funzione in diversi file JavaScript. Per condividerla con questi altri file, devi prima `export`.

```js
export const add = (x, y) => {
  return x + y;
}
```

Il modo sopra è comune per esportare una singola funzione, ma puoi ottenere lo stesso risultato così:

```js
const add = (x, y) => {
  return x + y;
}

export { add };
```

Quando esporti una variabile o una funzione, puoi importarla in un altro file e usarla senza dover riscrivere il codice. Puoi esportare più cose ripetendo il primo esempio per ogni elemento che vuoi esportare, oppure mettendole tutte nell’istruzione export del secondo esempio, così:

```js
export { add, subtract };
```

# --instructions--

Nell’editor ci sono due funzioni relative alle stringhe. Esporta entrambe usando il metodo che preferisci.

# --hints--

Dovresti esportare correttamente `uppercaseString`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+uppercaseString|export\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)})/g
  )
);
```

Dovresti esportare correttamente `lowercaseString`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+lowercaseString|export\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)})/g
  )
);
```

# --seed--

## --seed-contents--

```js
const uppercaseString = (string) => {
  return string.toUpperCase();
}

const lowercaseString = (string) => {
  return string.toLowerCase()
}
```

# --solutions--

```js
export const uppercaseString = (string) => {
  return string.toUpperCase();
}

export const lowercaseString = (string) => {
  return string.toLowerCase()
}
```
