---
id: 587d7b8b367417b2b2512b53
title: Usa la sintassi class per definire una funzione costruttore
challengeType: 1
forumTopicId: 301212
dashedName: use-class-syntax-to-define-a-constructor-function
---

# --description--

ES6 fornisce una nuova sintassi per creare oggetti, usando la parola chiave <dfn>class</dfn>.

In ES5, un oggetto può essere creato definendo una funzione `constructor` e usando la parola chiave `new` per istanziare l'oggetto.

In ES6, una dichiarazione `class` ha un metodo `constructor` che viene invocato con la parola chiave `new`. Se il metodo `constructor` non è definito esplicitamente, allora viene definito implicitamente senza argomenti.

```js
// Explicit constructor
class SpaceShuttle {
  constructor(targetPlanet) {
    this.targetPlanet = targetPlanet;
  }
  takeOff() {
    console.log("To " + this.targetPlanet + "!");
  }
}

// Implicit constructor 
class Rocket {
  launch() {
    console.log("To the moon!");
  }
}

const zeus = new SpaceShuttle('Jupiter');
// prints To Jupiter! in console
zeus.takeOff();

const atlas = new Rocket();
// prints To the moon! in console
atlas.launch();
```

Va notato che la parola chiave `class` dichiara una nuova funzione, a cui viene aggiunto un costruttore. Questo costruttore viene invocato quando `new` viene chiamato per creare un nuovo oggetto.

**Nota:** Per convenzione, i nomi delle classi ES6 dovrebbero usare UpperCamelCase, come `SpaceShuttle` usato sopra.

Il metodo `constructor` è un metodo speciale per creare e inizializzare un oggetto creato con una classe. Ne imparerai di più nella sezione di programmazione orientata agli oggetti della certificazione algoritmi e strutture dati in JavaScript.

# --instructions--

Usa la parola chiave `class` e scrivi una `constructor` per creare la classe `Vegetable`.

La classe `Vegetable` ti permette di creare un oggetto vegetable con una proprietà `name` che viene passata al `constructor`.

# --hints--

`Vegetable` dovrebbe essere una `class` con un metodo `constructor` definito.

```js
assert(
  typeof Vegetable === 'function' && typeof Vegetable.constructor === 'function'
);
```

La parola chiave `class` dovrebbe essere usata.

```js
assert(__helpers.removeJSComments(code).match(/class/g));
```

`Vegetable` dovrebbe poter essere istanziato.

```js
assert(() => {
  const a = new Vegetable('apple');
  return typeof a === 'object';
});
```

`carrot.name` dovrebbe restituire `carrot`.

```js
assert(carrot.name == 'carrot');
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const carrot = new Vegetable('carrot');
console.log(carrot.name); // Should display 'carrot'
```

# --solutions--

```js
class Vegetable {
  constructor(name) {
    this.name = name;
  }
}
const carrot = new Vegetable('carrot');
```
