---
id: 587d7b89367417b2b2512b48
title: "Usa l'operatore di espansione per valutare gli array in loco"
challengeType: 1
forumTopicId: 301222
dashedName: use-the-spread-operator-to-evaluate-arrays-in-place
---

# --description--

ES6 introduce il <dfn>operatore di espansione</dfn>, che ci permette di espandere array e altre espressioni in contesti dove sono attesi più parametri o elementi.

Il codice ES5 qui sotto usa `apply()` per calcolare il valore massimo in un array:

```js
var arr = [6, 89, 3, 45];
var maximus = Math.max.apply(null, arr);
```

`maximus` avrebbe un valore pari a `89`.

Abbiamo dovuto usare `Math.max.apply(null, arr)` perché `Math.max(arr)` restituisce `NaN`. `Math.max()` si aspetta argomenti separati da virgola, non un array. L'operatore di espansione rende questa sintassi molto più leggibile e facile da mantenere.

```js
const arr = [6, 89, 3, 45];
const maximus = Math.max(...arr);
```

`maximus` avrebbe un valore pari a `89`.

`...arr` restituisce un array "spacchettato". In altre parole, espande l'array. Tuttavia, l'operatore di espansione funziona solo in loco, come in un argomento di funzione o in un letterale array. Per esempio:

```js
const spreaded = [...arr];
```

Tuttavia, il codice seguente non funzionerà:

```js
const spreaded = ...arr;
```

# --instructions--

Copia tutto il contenuto di `arr1` in un altro array `arr2` usando l'operatore di espansione.

# --hints--

`arr2` dovrebbe essere una copia corretta di `arr1`.

```js
assert(arr2.every((v, i) => v === arr1[i]) && arr2.length);
```

Dovresti usare l'operatore di espansione `...` per duplicare `arr1`.

```js
assert(__helpers.removeJSComments(code).match(/Array\(\s*\.\.\.arr1\s*\)|\[\s*\.\.\.arr1\s*\]/));
```

`arr2` dovrebbe rimanere invariato quando `arr1` viene modificato.

```js
assert((arr1, arr2) => {
  arr1.push('JUN');
  return arr2.length < arr1.length;
});
```

# --seed--

## --seed-contents--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [];  // Change this line

console.log(arr2);
```

# --solutions--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [...arr1];
```
