---
id: 587d7b8d367417b2b2512b59
title: "Importare un'esportazione predefinita"
challengeType: 1
forumTopicId: 301205
dashedName: import-a-default-export
---

# --description--

Nell'ultima sfida, hai imparato a conoscere `export default` e i suoi usi. Per importare un'esportazione predefinita, devi usare una sintassi `import` diversa. Nel seguente esempio, `add` è l'esportazione predefinita del file `math_functions.js`. Ecco come importarla:

```js
import add from "./math_functions.js";
```

La sintassi differisce in un punto chiave. Il valore importato, `add`, non è racchiuso tra parentesi graffe (`{}`). `add` qui è semplicemente un nome di variabile per qualunque sia l'esportazione predefinita del file `math_functions.js`. Puoi usare qualsiasi nome qui quando importi un predefinito.

# --instructions--

Nel codice seguente, importa l'esportazione predefinita dal file `math_functions.js`, che si trova nella stessa directory di questo file. Dai all'importazione il nome `subtract`.

# --hints--

Dovresti importare correttamente `subtract` da `math_functions.js`.

```js
assert(__helpers.removeJSComments(code).match(/import\s+subtract\s+from\s+('|")\.\/math_functions\.js\1/g));
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

subtract(7,4);
```

# --solutions--

```js
import subtract from "./math_functions.js";

subtract(7,4);
```
