---
id: 587d7b8c367417b2b2512b58
title: Crea un fallback di export con export default
challengeType: 1
forumTopicId: 301199
dashedName: create-an-export-fallback-with-export-default
---

# --description--

Nella lezione `export` hai imparato la sintassi chiamata <dfn>named export</dfn>. Questa ti ha permesso di rendere disponibili più funzioni e variabili per l’uso in altri file.

C’è un’altra sintassi `export` che devi conoscere, nota come <dfn>export default</dfn>. Di solito userai questa sintassi se viene esportato solo uno valore da un file. Viene anche usata per creare un valore fallback per un file o modulo.

Qui sotto ci sono esempi che usano `export default`:

```js
export default function add(x, y) {
  return x + y;
}

export default function(x, y) {
  return x + y;
}
```

Il primo è una funzione nominata, il secondo è una funzione anonima.

Poiché `export default` serve a dichiarare un valore fallback per un modulo o file, puoi avere solo un valore come export predefinito in ogni modulo o file. Inoltre, non puoi usare `export default` insieme a `var`, `let` o `const`

# --instructions--

La funzione seguente dovrebbe essere il valore fallback per il modulo. Aggiungi il codice necessario per farlo.

# --hints--

Il tuo codice dovrebbe usare un fallback `export`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /export\s+default\s+function(\s+subtract\s*|\s*)\(\s*x,\s*y\s*\)\s*{/g
  )
);
```

# --seed--

## --seed-contents--

```js
function subtract(x, y) {
  return x - y;
}
```

# --solutions--

```js
export default function subtract(x, y) {
  return x - y;
}
```
