---
id: 598f48a36c8c40764b4e52b3
title: prevenire la mutazione degli oggetti
challengeType: 1
forumTopicId: 301207
dashedName: prevent-object-mutation
---

# --description--

Come visto nella sfida precedente, la sola dichiarazione `const` non protegge davvero i tuoi dati dalla mutazione. Per assicurarti che i dati non cambino, JavaScript fornisce una funzione `Object.freeze` per impedire la mutazione dei dati.

Qualsiasi tentativo di modificare l'oggetto verrà rifiutato, con un errore lanciato se lo script è in esecuzione in modalità strict.

```js
let obj = {
  name:"FreeCodeCamp",
  review:"Awesome"
};
Object.freeze(obj);
obj.review = "bad";
obj.newProp = "Test";
console.log(obj); 
```

Le assegnazioni `obj.review` e `obj.newProp` genereranno errori, perché il nostro editor funziona in modalità strict per impostazione predefinita, e la console mostrerà il valore `{ name: "FreeCodeCamp", review: "Awesome" }`.

# --instructions--

In questa sfida userai `Object.freeze` per impedire che le costanti matematiche vengano modificate. Devi congelare l'oggetto `MATH_CONSTANTS` in modo che nessuno possa alterare il valore di `PI`, aggiungere o eliminare proprietà.

# --hints--

Non dovresti sostituire la parola chiave `const`.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`MATH_CONSTANTS` dovrebbe essere una variabile costante (usando `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+MATH_CONSTANTS/g));
```

Non dovresti modificare la dichiarazione originale di `MATH_CONSTANTS`.

```js
assert(__helpers.removeJSComments(code).match(
   /const\s+MATH_CONSTANTS\s+=\s+{\s+PI:\s+3.14\s+};/g
));
```

`PI` dovrebbe essere uguale a `3.14`.

```js
assert(PI === 3.14);
```

# --seed--

## --seed-contents--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  // Only change code below this line


  // Only change code above this line
  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```

# --solutions--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  Object.freeze(MATH_CONSTANTS);

  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```
