---
id: 5cdafbe72913098997531682
title: Gestire una promise rifiutata con catch
challengeType: 1
forumTopicId: 301204
dashedName: handle-a-rejected-promise-with-catch
---

# --description--

`catch` è il metodo usato quando la tua promise è stata rifiutata. Viene eseguito subito dopo che è stato chiamato il metodo `reject` di una promise. Ecco la sintassi:

```js
myPromise.catch(error => {
  
});
```

`error` è l'argomento passato al metodo `reject`.

# --instructions--

Aggiungi il metodo `catch` alla tua promise. Usa `error` come parametro della sua funzione di callback e registra `error` nella console.

# --hints--

Dovresti chiamare il metodo `catch` sulla promise.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.catch\(/g)
);
```

Il tuo metodo `catch` dovrebbe avere una funzione di callback con `error` come parametro.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

Dovresti registrare `error` nella console.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.catch\(.*?error.*?console.log\(error\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});

makeServerRequest.catch(error => {
  console.log(error);
});
```
