---
id: 5cdafbd72913098997531681
title: Gestire una promise risolta con then
challengeType: 1
forumTopicId: 301203
dashedName: handle-a-fulfilled-promise-with-then
---

# --description--

Le promise sono più utili quando nel tuo codice hai un processo che richiede un tempo sconosciuto (cioè qualcosa di asincrono), spesso una richiesta a un server. Quando fai una richiesta a un server, questa richiede un certo tempo e, una volta completata, di solito vuoi fare qualcosa con la risposta del server. Questo si può ottenere usando il metodo `then`.

```js
Promise.prototype.then(onFulfilled, onRejected)
```

Il metodo `then` programma funzioni di callback per il completamento futuro di una promise — sia per il suo completamento con successo sia per il suo rifiuto. Uno dei gestori `onFulfilled` e `onRejected` verrà eseguito per gestire il completamento o il rifiuto della promise corrente. Quando la promise viene risolta con `resolve` viene chiamato il gestore `onFulfilled`.

```js
myPromise.then(result => {
  
});
```

`result` proviene dall’argomento passato al metodo `resolve`.

# --instructions--

Aggiungi il metodo `then` alla tua promise. Usa `result` come parametro della sua funzione di callback e stampa `result` nella console.

# --hints--

Dovresti chiamare il metodo `then` sulla promise.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.then\(/g)
);
```

Il tuo metodo `then` dovrebbe avere una funzione di callback con `result` come parametro.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

Dovresti stampare `result` nella console.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.then\(.*?result.*?console.log\(result\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```
