---
id: 5cdafbc32913098997531680
title: Completa una promise con resolve e reject
challengeType: 1
forumTopicId: 301196
dashedName: complete-a-promise-with-resolve-and-reject
---

# --description--

Una promise ha tre stati: `pending`, `fulfilled` e `rejected`. La promise che hai creato nella sfida precedente è bloccata per sempre nello stato `pending` perché non hai aggiunto un modo per completarla. I parametri `resolve` e `reject` dati all'argomento della promise servono proprio a questo. `resolve` si usa quando vuoi che la promise abbia successo, mentre `reject` si usa quando vuoi che fallisca. Sono metodi che prendono un argomento, come vedi qui sotto.

```js
const myPromise = new Promise((resolve, reject) => {
  if(condition here) {
    resolve("Promise was fulfilled");
  } else {
    reject("Promise was rejected");
  }
});
```

L'esempio sopra usa stringhe come argomento di queste funzioni, ma in realtà può essere qualsiasi cosa. Spesso può essere un oggetto, da cui prendere dati da mostrare sul tuo sito o altrove.

# --instructions--

Fai in modo che la promise gestisca successo e fallimento. Se `responseFromServer` è `true`, chiama il metodo `resolve` per completare con successo la promise. Passa a `resolve` una stringa con il valore `We got the data`. Se `responseFromServer` è `false`, usa invece il metodo `reject` e passagli la stringa: `Data not received`.

# --hints--

`resolve` dovrebbe essere chiamato con la stringa prevista quando la condizione `if` è `true`.

```js
assert(
  __helpers.removeJSComments(code).match(/if\s*\(\s*responseFromServer\s*\)\s*{\s*resolve\s*\(\s*('|"|`)We got the data\1\s*\)(\s*|\s*;\s*)}/g)
);
```

`reject` dovrebbe essere chiamato con la stringa prevista quando la condizione `if` è `false`.

```js
assert(
  __helpers.removeJSComments(code).match(/}\s*else\s*{\s*reject\s*\(\s*('|"|`)Data not received\1\s*\)(\s*|\s*;\s*)}/g)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;
    
  if(responseFromServer) {
    // Change this line
  } else {  
    // Change this line
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;

  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```
