---
id: 5cdafbb0291309899753167f
title: Crea una promise in JavaScript
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

Una promise in JavaScript è esattamente quello che sembra: la usi per fare una promessa di fare qualcosa, di solito in modo asincrono. Quando l’attività si completa, o mantieni la promessa o non ci riesci. `Promise` è una funzione costruttore, quindi devi usare la parola chiave `new` per crearne una. Prende come argomento una funzione con due parametri: `resolve` e `reject`. Questi sono metodi usati per determinare l’esito della promise. La sintassi è questa:

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

Crea una nuova promise chiamata `makeServerRequest`. Passa al costruttore una funzione con i parametri `resolve` e `reject`.

# --hints--

Dovresti assegnare una promise a una variabile dichiarata chiamata `makeServerRequest`.

```js
assert(makeServerRequest instanceof Promise);
```

La tua promise dovrebbe ricevere una funzione con `resolve` e `reject` come parametri.

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
