---
id: 5cddbfd622f1a59093ec611d
title: creare uno script modulo
challengeType: 6
forumTopicId: 301198
dashedName: create-a-module-script
---

# --description--

JavaScript ha iniziato con un ruolo piccolo in un web altrimenti per lo più HTML. Oggi è enorme, e alcuni siti web sono costruiti quasi interamente con JavaScript. Per rendere JavaScript più modulare, pulito e manutenibile, ES6 ha introdotto un modo per condividere facilmente il codice tra file JavaScript. Questo comporta esportare parti di un file per usarle in uno o più altri file, e importare le parti di cui hai bisogno, dove ti servono. Per sfruttare questa funzionalità, devi creare uno script nel tuo documento HTML con un `type` di `module`. Ecco un esempio:

```html
<script type="module" src="filename.js"></script>
```

Uno script che usa questo tipo `module` può ora usare le funzionalità `import` e `export` che imparerai nelle prossime sfide.

# --instructions--

Aggiungi uno script al documento HTML di tipo `module` e dagli come file sorgente `index.js`

# --hints--

Dovresti creare un tag `script`.

```js
assert(__helpers.removeJSComments(code).match(/<\s*script[^>]*>\s*<\/\s*script\s*>/g));
```

Il tuo tag `script` dovrebbe avere l'attributo `type` con valore `module`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^t]*type\s*=\s*('|")module\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

Il tuo tag `script` dovrebbe avere un `src` di `index.js`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^s]*src\s*=\s*('|")index\.js\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <!-- Only change code below this line -->

    <!-- Only change code above this line -->
  </body>
</html>
```

# --solutions--

```html
<html>
  <body>
    <script type="module" src="index.js"></script>
  </body>
</html>
```
