---
id: 587d8259367417b2b2512c85
title: 实现选择排序算法
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

In this lab you will implement the selection sort algorithm. 选择排序的工作原理是选择列表中的最小值并与列表中的第一个值进行交换。 然后它从第二个位置开始，选择剩余列表中的最小值，并将其与第二个元素交换。 它继续遍历列表并交换元素，直到它到达列表的末尾。 现在列表已排序。 在所有情况下，选择排序都具有二次方的时间复杂度。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `selectionSort` 函数。
1. `selectionSort` 函数应接受一个数组，并返回一个包含相同元素但按从小到大排序的数组。

# --hints--

`selectionSort` 应该是一个函数。

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` 应该返回一个除了顺序外未改变的数组。

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` 应该返回一个排序的数组（从最小到最大）。

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` 不应使用内置的 `.sort()` 方法。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
