---
id: bad87fee1348bd9aedc08830
title: Usa HTML5 per rendere obbligatorio un campo
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMd4EcQ'
forumTopicId: 18360
dashedName: use-html5-to-require-a-field
---

# --description--

Puoi rendere obbligatori specifici campi del modulo in modo che l'utente non possa inviare il modulo finché non li ha compilati.

Per esempio, se vuoi rendere obbligatorio un campo di input di testo, puoi semplicemente aggiungere l'attributo `required` all'interno del tuo elemento `input`, così: `<input type="text" required>`

# --instructions--

Rendi il tuo campo di testo `input` un campo `required` obbligatorio, così l'utente non potrà inviare il modulo senza completare questo campo.

Poi prova a inviare il modulo senza inserire alcun testo. Vedi come il tuo modulo HTML5 ti avvisa che il campo è obbligatorio?

# --hints--

Il tuo elemento di testo `input` dovrebbe avere l'attributo `required`.

```js
assert.isTrue(document.querySelector('input').required);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" required placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```
