---
id: bad87fee1348bd9aecf08801
title: introduzione agli elementi html5
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBkZGpt7'
forumTopicId: 301097
dashedName: introduction-to-html5-elements
---

# --description--

HTML5 introduce tag HTML più descrittivi. Questi includono `main`, `header`, `footer`, `nav`, `video`, `article`, `section` e altri.

Questi tag danno una struttura descrittiva al tuo HTML, rendono il tuo HTML più facile da leggere e aiutano con l'ottimizzazione per i motori di ricerca (SEO) e l'accessibilità. Il tag `main` di HTML5 aiuta i motori di ricerca e altri sviluppatori a consultare il contenuto principale della tua pagina.

Esempio di utilizzo, un elemento `main` con due elementi figli annidati al suo interno:

```html
<main> 
  <h1>Hello World</h1>
  <p>Hello Paragraph</p>
</main>
```

**Nota:** Molti dei nuovi tag HTML5 e i loro vantaggi sono trattati nella sezione Accessibilità applicata.

# --instructions--

Crea un secondo elemento `p` con il seguente testo kitty ipsum: `Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.`

Poi, crea un elemento `main` e annida solo i due elementi `p` all'interno dell'elemento `main`.

# --hints--

Dovresti avere 2 elementi `p` con testo Kitty Ipsum.

```js
assert.lengthOf(document.querySelectorAll('p'),2);
```

Ognuno dei tuoi elementi `p` dovrebbe avere un tag di chiusura.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

Il tuo elemento `p` dovrebbe contenere le prime parole del testo aggiuntivo `kitty ipsum` fornito.

```js
assert.match(document.querySelectorAll('p')[1].textContent,/Purr\s+jump\s+eat/gi);
```

Il tuo codice dovrebbe avere un solo elemento `main`.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

L'elemento `main` dovrebbe avere due elementi paragrafo come figli.

```js
const main = document.querySelector('main');
const children = main.querySelectorAll("p"); 
assert.lengthOf(children,2);
```

Il tag di apertura `main` dovrebbe venire prima del primo tag paragrafo.

```js
assert.match(code,(/<main>\s*?<p>/g));
```

Il tag di chiusura `main` dovrebbe venire dopo il secondo tag paragrafo di chiusura.

```js
assert.match(code,(/<\/p>\s*?<\/main>/g));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
