---
id: bad87fee1348bd9aedf08802
title: Rimuovere i commenti in HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBmG9T7'
forumTopicId: 18329
dashedName: uncomment-html
---

# --description--

Commentare è un modo per lasciare commenti ad altri sviluppatori all’interno del tuo codice senza influenzare l’output risultante che viene mostrato all’utente finale.

Commentare è anche un modo comodo per rendere il codice inattivo senza doverlo eliminare completamente.

I commenti in HTML iniziano con `<!--` e finiscono con `-->`

# --instructions--

Rimuovi i commenti dai tuoi elementi `h1`, `h2` e `p`.

# --hints--

Il tuo elemento `h1` dovrebbe essere visibile sulla pagina rimuovendo il commento.

```js
assert.notEmpty(document.querySelectorAll('h1'));
```

Il tuo elemento `h2` dovrebbe essere visibile sulla pagina rimuovendo il commento.

```js
assert.notEmpty(document.querySelectorAll('h2'));
```

Il tuo elemento `p` dovrebbe essere visibile sulla pagina rimuovendo il commento.

```js
assert.notEmpty(document.querySelectorAll('p'));
```

Non dovrebbero esserci tag di commento residui visibili sulla pagina (cioè `-->`).

```js
const elements = document.querySelectorAll('*');
const potentialComments = Array.from(elements).filter(el => el.textContent.includes('-->'));
assert.notExists(potentialComments[1])
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
