---
id: bad87fee1348bd9aedf08816
title: Collegare a pagine esterne con elementi àncora
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c8EkncB'
forumTopicId: 18226
dashedName: link-to-external-pages-with-anchor-elements
---

# --description--

Puoi usare gli elementi `a` (*àncora*) per collegare a contenuti esterni alla tua pagina web.

Gli elementi `a` hanno bisogno di un indirizzo web di destinazione chiamato attributo `href`. Hanno anche bisogno di un testo dell'àncora. Ecco un esempio:

```html
<a href="https://www.freecodecamp.org">this links to freecodecamp.org</a>
```

Poi il tuo browser mostrerà il testo `this links to freecodecamp.org` come un collegamento su cui puoi cliccare. E quel collegamento ti porterà all'indirizzo web `https://www.freecodecamp.org`.

# --instructions--

Crea un elemento `a` che colleghi a `https://www.freecatphotoapp.com` e abbia come testo dell'àncora "cat photos".

# --hints--

Il tuo elemento `a` dovrebbe avere come testo dell'àncora `cat photos`.

```js
assert.match(document.querySelector('a').textContent,/cat photos/gi);
```

Ti serve un elemento `a` che colleghi a `https://www.freecatphotoapp.com`

```js
assert.match(document.querySelector('a').getAttribute('href'),/^https?:\/\/(www\.)?freecatphotoapp\.com\/?$/i);
```

Il tuo elemento `a` dovrebbe avere un tag di chiusura.

```js
assert.match(code,/<\/a>/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>



  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <a href="https://www.freecatphotoapp.com">cat photos</a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
