---
id: bad87fee1348bd9aede08817
title: "Annidare un elemento àncora all'interno di un paragrafo"
challengeType: 0
forumTopicId: 18244
dashedName: nest-an-anchor-element-within-a-paragraph
---

# --description--

Puoi annidare collegamenti all'interno di altri elementi di testo.

```html
<p>
  Here's a <a target="_blank" href="https://www.freecodecamp.org"> link to www.freecodecamp.org</a> for you to follow.
</p>
```

Analizziamo l'esempio. Il testo normale è racchiuso nell'elemento `p`:

```html
<p> Here's a ... for you to follow. </p>
```

Segue l'*elemento àncora* `<a>` (che richiede un tag di chiusura `</a>`):  

```html
<a> ... </a>
```

`target` è un attributo del tag àncora che specifica dove aprire il collegamento. Il valore `_blank` indica di aprire il collegamento in una nuova scheda. `href` è un attributo del tag àncora che contiene l'indirizzo URL del collegamento:  

```html
<a href="https://www.freecodecamp.org" target="_blank"> ... </a>
```

Il testo, `link to www.freecodecamp.org`, all'interno dell'elemento `a` si chiama <dfn>testo dell'àncora</dfn> e mostrerà il collegamento su cui cliccare:

```html
<a href=" ... " target="...">link to freecodecamp.org</a>
```

L'output finale dell'esempio sarà così:  

Ecco un <a href="https://www.freecodecamp.org" target="_blank">collegamento a www.freecodecamp.org</a> da seguire.

# --instructions--

Annida l'elemento `a` esistente all'interno di un nuovo elemento `p`. Non creare un nuovo tag àncora. Il nuovo paragrafo deve contenere il testo `View more cat photos`, dove `cat photos` è un collegamento e il resto è testo normale.

# --hints--

Dovresti avere solo un elemento `a`.

```js
assert.lengthOf(document.querySelectorAll('a'), 1 );
```

L'elemento `a` dovrebbe collegare a "`https://www.freecatphotoapp.com`".

```js
assert.lengthOf(document.querySelectorAll('a[href="https://www.freecatphotoapp.com"]'),1);
```

Il tuo elemento `a` dovrebbe avere come testo dell'àncora `cat photos`

```js
assert.match(document.querySelector('a').textContent,/cat\sphotos/gi)
```

Devi creare un nuovo elemento `p`. Nel tuo codice HTML dovrebbero esserci almeno 3 tag `p` in totale.

```js
assert.lengthOf(document.querySelectorAll('p'),3)
```

Il tuo elemento `a` dovrebbe essere annidato all'interno del nuovo elemento `p`.

```js
const anchorParent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode;
assert.strictEqual(anchorParent.tagName,"P")
```

Il tuo elemento `p` dovrebbe contenere il testo `View more ` (con uno spazio dopo).

```js
const textContent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode.textContent;
assert.match(textContent,/View\smore\s/gi);
```

Il tuo elemento `a` non dovrebbe <em>contenere</em> il testo `View more`.

```js
assert.notMatch(document.querySelector('a').textContent,/View\smore/gi); 
```

Ognuno dei tuoi elementi `p` dovrebbe avere un tag di chiusura.

```js
assert.match(code,/<\/p>/g);
assert.match(code,/<p/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

Ognuno dei tuoi elementi `a` dovrebbe avere un tag di chiusura.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>View more <a target="_blank" href="https://www.freecatphotoapp.com">cat photos</a></p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
