---
id: bad87fee1348bd9aedf08833
title: Inserisci il testo segnaposto nel campo vuoto
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cgR7Dc7'
forumTopicId: 18178
dashedName: fill-in-the-blank-with-placeholder-text
---

# --description--

Gli sviluppatori web usano tradizionalmente il <dfn>testo lorem ipsum</dfn> come testo segnaposto. Il testo lorem ipsum è estratto in modo casuale da un famoso passo di Cicerone dell'antica Roma.

Il testo lorem ipsum viene usato come testo segnaposto dai tipografi fin dal XVI secolo, e questa tradizione continua sul web.

Beh, 5 secoli sono abbastanza. Dal momento che stiamo buildando una CatPhotoApp, usiamo qualcosa chiamato testo "kitty ipsum".

# --instructions--

Sostituisci il testo all'interno del tuo elemento `p` con le prime parole di questo testo kitty ipsum: `Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.`

# --hints--

Il tuo elemento `p` dovrebbe contenere le prime parole del testo "kitty ipsum" fornito.

```js
assert.match(document.querySelector('p').textContent,/Kitty(\s)+ipsum/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Hello Paragraph</p>
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff</p>
```
