---
id: bad87fee1348bd9aedf08817
title: creare collegamenti non funzionanti usando il simbolo cancelletto
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMdkytL'
forumTopicId: 18230
dashedName: make-dead-links-using-the-hash-symbol
---

# --description--

A volte vuoi aggiungere `a` elementi al tuo sito web prima di sapere a cosa collegheranno.

Questo è utile anche quando cambi il comportamento di un collegamento usando `JavaScript`, che impareremo più avanti.

# --instructions--

Il valore attuale dell'attributo `href` è un collegamento che punta a "`https://www.freecatphotoapp.com`". Sostituisci il valore dell'attributo `href` con un `#`, noto anche come simbolo cancelletto, per creare un collegamento non funzionante.

Per esempio: `href="#"`

# --hints--

Il tuo elemento `a` dovrebbe essere un collegamento non funzionante con il valore dell'attributo `href` impostato su "#".

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'),'#');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#" target="_blank">cat photos</a>.</p>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
