---
id: bad87fee1348bd9aedf08834
title: Creare un insieme di pulsanti di opzione
challengeType: 0
forumTopicId: 16822
dashedName: create-a-set-of-radio-buttons
---

# --description--

Puoi usare i <dfn>pulsanti di opzione</dfn> per domande in cui vuoi che l'utente dia solo una risposta tra più opzioni.

I pulsanti di opzione sono un tipo di `input`.

Ognuno dei tuoi pulsanti di opzione può essere annidato all'interno del proprio elemento `label`. Avvolgendo un elemento `input` dentro un elemento `label` si associa automaticamente l'input del pulsante di opzione con l'elemento etichetta che lo circonda.

Tutti i pulsanti di opzione correlati dovrebbero avere lo stesso attributo `name` per creare un gruppo di pulsanti di opzione. Creando un gruppo di pulsanti, selezionare un singolo pulsante di opzione deselezionerà automaticamente gli altri pulsanti dello stesso gruppo, assicurando che l'utente fornisca solo una risposta.

Ecco un esempio di pulsante di opzione:

```html
<label> 
  <input type="radio" name="indoor-outdoor">Indoor 
</label>
```

È considerata buona pratica impostare un attributo `for` sull'elemento `label`, con un valore che corrisponde al valore dell'attributo `id` dell'elemento `input`. Questo permette alle tecnologie assistive di creare un collegamento tra l'etichetta e l'elemento `input` correlato. Per esempio:

```html
<input id="indoor" type="radio" name="indoor-outdoor">
<label for="indoor">Indoor</label>
```

Possiamo anche annidare l'elemento `input` all'interno dei tag `label`:

```html
<label for="indoor"> 
  <input id="indoor" type="radio" name="indoor-outdoor">Indoor 
</label>
```

# --instructions--

Aggiungi una coppia di pulsanti di opzione al tuo modulo, ognuno annidato nel proprio elemento `label`. Uno dovrebbe avere l'opzione `indoor` e l'altro dovrebbe avere l'opzione `outdoor`. Entrambi dovrebbero condividere l'attributo `name` con valore `indoor-outdoor` per creare un gruppo di pulsanti di opzione.

# --hints--

La tua pagina dovrebbe avere due elementi pulsante `radio`.

```js
assert.lengthOf(document.querySelectorAll('input[type="radio"]'),2);
```

I tuoi pulsanti di opzione dovrebbero avere l'attributo `name` con valore `indoor-outdoor`.

```js
assert.lengthOf([...document.querySelectorAll('input[type="radio"]')].filter(input => input.name === "indoor-outdoor"),2);
```

Ognuno dei tuoi due elementi pulsante di opzione dovrebbe essere annidato nel proprio elemento `label`.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="radio"]:only-child'),2);
```

Ognuno dei tuoi elementi `label` dovrebbe avere un tag di chiusura.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length);
```

Uno dei tuoi pulsanti di opzione dovrebbe avere l'etichetta `indoor`.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label =>  label.match(/indoor/gi));
assert.isTrue(hasMatch); 
```

Uno dei tuoi pulsanti di opzione dovrebbe avere l'etichetta `outdoor`.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label => label.match(/outdoor/gi));
assert.isTrue(hasMatch); 
```

Ognuno dei tuoi elementi pulsante di opzione dovrebbe essere inserito all'interno del tag `form`.

```js
const labelParent = document.querySelector('label').parentNode;
assert.strictEqual(labelParent.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
   <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
