---
id: bad88fee1348bd9aedf08816
title: Collegare a sezioni interne di una pagina con elementi àncora
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cyrDRUL'
forumTopicId: 301098
dashedName: link-to-internal-sections-of-a-page-with-anchor-elements
---

# --description--

Gli elementi `a` (*àncora*) possono essere usati anche per creare collegamenti interni che saltano a sezioni diverse all'interno di una pagina web.

Per creare un collegamento interno, assegni l'attributo `href` di un collegamento a un simbolo cancelletto `#` più il valore dell'attributo `id` per l'elemento a cui vuoi collegarti internamente, di solito più in basso nella pagina. Poi devi aggiungere lo stesso attributo `id` all'elemento a cui stai collegando. Un `id` è un attributo che descrive in modo univoco un elemento.

Qui sotto c'è un esempio di collegamento àncora interno e del suo elemento obiettivo:

```html
<a href="#contacts-header">Contacts</a>
...
<h2 id="contacts-header">Contacts</h2>
```

Quando gli utenti cliccano sul collegamento `Contacts`, verranno portati alla sezione della pagina web con l'elemento di intestazione **Contacts**.

# --instructions--

Trasforma il tuo collegamento esterno in un collegamento interno cambiando l'attributo `href` in `#footer` e il testo da `cat photos` a `Jump to Bottom`.

Rimuovi l'attributo `target="_blank"` dal tag àncora perché questo fa aprire il documento collegato in una nuova scheda o finestra.

Poi aggiungi un attributo `id` con valore `footer` all'elemento `<footer>` in fondo alla pagina.

# --hints--

Dovrebbe esserci solo un tag àncora nella tua pagina.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Dovrebbe esserci solo un tag `footer` nella tua pagina.

```js
assert.lengthOf(document.querySelectorAll('footer'),1);
```

Il tag `a` dovrebbe avere un attributo `href` impostato su "#footer".

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'), '#footer');
```

Il tag `a` non dovrebbe avere un attributo `target`.

```js
assert.notExists(document.querySelector('a').getAttribute('target'));
```

Il testo `a` dovrebbe essere "Jump to Bottom".

```js
assert.match(document.querySelector('a').textContent,/Jump to Bottom/gi);
```

Il tag `footer` dovrebbe avere un attributo `id` impostato su "footer".

```js
assert.strictEqual(document.querySelector('footer').getAttribute('id'),'footer');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer>Copyright Cat Photo App</footer>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="#footer">Jump to Bottom</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer id="footer">Copyright Cat Photo App</footer>
```
