---
id: bad87fee1348bd9aedf08812
title: Aggiungi immagini al tuo sito web
challengeType: 0
forumTopicId: 16640
dashedName: add-images-to-your-website
---

# --description--

Puoi aggiungere immagini al tuo sito web usando l'elemento `img` e puntare all'URL di un'immagine specifica tramite l'attributo `src`.

Un esempio di questo sarebbe:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg">
```

Nota che `img` è un elemento vuoto.

Tutti gli elementi `img` **devono** avere un attributo `alt`. Il testo all'interno di un attributo `alt` viene usato dai lettori di schermo per migliorare l'accessibilità ed è mostrato se l'immagine non si carica.

**Nota:** Se l'immagine è puramente decorativa, usare un attributo `alt` vuoto è una buona pratica.

Idealmente l'attributo `alt` non dovrebbe contenere caratteri speciali a meno che non siano necessari.

Aggiungiamo un attributo `alt` al nostro esempio `img` sopra:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg" alt="freeCodeCamp logo">
```

# --instructions--

Proviamo ad aggiungere un'immagine al nostro sito web:

All'interno dell'elemento `main` esistente, inserisci un elemento `img` prima degli elementi `p` esistenti.

Ora imposta l'attributo `src` in modo che punti all'URL `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

Infine, non dimenticare di dare al tuo elemento `img` un attributo `alt` con un testo appropriato.

# --hints--

La tua pagina dovrebbe avere un elemento immagine.

```js
assert.exists(document.querySelector('img'));
```

La tua immagine dovrebbe avere un attributo `src` che punti all'immagine del gattino.

```js
const url = document.querySelector('img').getAttribute('src');
assert.match(url,/^https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg$/i);
```

L'attributo `alt` del tuo elemento immagine non dovrebbe essere vuoto.

```js
assert.exists(document.querySelector('img').getAttribute('alt'));
assert.isNotEmpty(document.querySelector('img').getAttribute('alt'));
assert.match(__helpers.removeWhiteSpace(code),/<(?:img|IMG)\S*alt=(['"])(?!\1|>)\S+\1\S*\/?>/)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>


  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
