---
id: bad87fee1348bd9aedf08835
title: Creare un insieme di caselle di spunta
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cqrkJsp'
forumTopicId: 16821
dashedName: create-a-set-of-checkboxes
---

# --description--

I moduli usano comunemente le <dfn>caselle di spunta</dfn> per domande che possono avere più di una risposta.

Le caselle di spunta sono un tipo di `input`.

Ognuna delle tue caselle di spunta può essere annidata all'interno del proprio elemento `label`. Inserendo un elemento `input` dentro un elemento `label` si associa automaticamente l'input della casella di spunta con l'elemento etichetta che la circonda.

Tutti gli input delle caselle di spunta correlate dovrebbero avere lo stesso attributo `name`.

È considerata buona pratica definire esplicitamente la relazione tra un `input` della casella di spunta e il suo corrispondente `label` impostando l'attributo `for` sull'elemento `label` in modo che corrisponda all'attributo `id` dell'elemento `input` associato.

Ecco un esempio di casella di spunta:

```html
<label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
```

# --instructions--

Aggiungi al tuo modulo un insieme di tre caselle di spunta. Ognuna deve essere annidata nel proprio elemento `label`. Tutte e tre devono condividere l'attributo `name` con valore `personality`.

# --hints--

La tua pagina dovrebbe avere tre elementi casella di spunta.

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'),3);
```

Ognuno dei tuoi tre elementi casella di spunta dovrebbe essere annidato nel proprio elemento `label`.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="checkbox"]:only-child'),3);
```

Assicurati che ognuno dei tuoi elementi `label` abbia un tag di chiusura.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length)
```

Le tue caselle di spunta dovrebbero avere l'attributo `name` con valore `personality`.

```js
assert.lengthOf([...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.name === "personality"),3);
```

Ognuna delle tue caselle di spunta dovrebbe essere inserita all'interno del tag `form`.

```js
assert.strictEqual(document.querySelector('label').parentNode.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="playful"><input id="playful" type="checkbox" name="personality">Playful</label>
    <label for="lazy"><input id="lazy" type="checkbox" 
name="personality">Lazy</label>
    <label for="evil"><input id="evil" type="checkbox" 
name="personality">Evil</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
