---
id: bad87fee1348bd9aedf08830
title: aggiungere un testo segnaposto a un campo di testo
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cKdJDhg'
forumTopicId: 16647
dashedName: add-placeholder-text-to-a-text-field
---

# --description--

Il testo segnaposto è ciò che viene mostrato nel tuo elemento `input` prima che l’utente inserisca qualcosa.

Puoi creare un testo segnaposto così:

```html
<input type="text" placeholder="this is placeholder text">
```

**Nota:** Ricorda che `input` è un elemento void.

# --instructions--

Imposta il valore `placeholder` del tuo campo di testo `input` su "cat photo URL".

# --hints--

Dovresti aggiungere un attributo `placeholder` all’elemento di testo `input` esistente.

```js
assert.notEmpty(document.querySelectorAll('input[placeholder]'));
```

Dovresti impostare il valore del tuo attributo `placeholder` su `cat photo URL`.

```js
assert.exists(document.querySelector('input'));
assert.exists(document.querySelector('input').getAttribute('placeholder'));
const placeholder = document.querySelector('input').getAttribute('placeholder');
assert.match(placeholder,/cat\s+photo\s+URL/gi);
```

L’elemento `input` completato non deve avere un tag di chiusura.

```js
assert.notMatch(code,/<input.*\/?>.*<\/input>/gi);
```

L’elemento `input` completato deve avere una sintassi valida.

```js
assert.notEmpty(document.querySelectorAll('input[type=text]'));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```
