---
id: 68ed661e9f463c155bb677a1
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過設計一組彩色方塊來練習使用 CSS flexbox。

首先，在 `body` 元素內，添加一個 `header` 元素。在 `header` 元素內巢狀一個帶有文字 `Colored Boxes Layout` 的 `h1` 元素。

# --hints--

你應該在 `body` 元素內有一個 `header` 元素。

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

你的 `header` 元素應該包含一個 `h1` 元素。

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

你的 `h1` 元素內應該包含文字 `Colored Boxes Layout`。

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
