---
id: 68ef0d6d156f86a6ac1824c6
title: 第 23 步
challengeType: 0
dashedName: step-23
---

# --description--

現在你將使用 `align-content` 屬性。此屬性控制項目沿交叉軸的排列方式。雖然 `align-content` 的預設值是 `stretch`，但你可以用於 `align-content` 的最相關選項是：`space-between`、`start`、`center`、`end`、`space-around` 和 `space-evenly`。

為你的 `.flex-container` 容器添加屬性 `align-content` 並設為 `space-between`，以查看方塊如何對齊。

# --hints--

你的 `.flex-container` 選擇器應該有屬性 `align-content`，其值為 `space-between`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('align-content'), 'space-between');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}

.box {
	max-height: 120px;
	color: #000;
	border: 1px solid #000;
	display: flex;
	flex: 1 1 100px;
	flex-direction: column;
	align-items: center;
	margin: 10px;
	font-weight: bold;
	font-size: 1.125rem;
	border-radius: 5px;
	order: 0; 
}

.flex-container {
	display: flex;
	flex-wrap: wrap;
	width: 70%;
	height: 600px;
	padding: 10px;
	margin: 20px auto;
--fcc-editable-region--
	
--fcc-editable-region--
}
```
