---
id: 68ef0b3237fef991c92d9c19
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

現在在你的 `header` 下方添加一個帶有類別 `flex-container` 的 `div` 元素。然後在你剛剛創建的 `div` 裡面建立六個帶有類別 `box` 的 `div` 元素。

請記住，類別屬性通常用來指向樣式表中的類別名稱。在這種情況下，所有的 `div` 元素將根據 `.flex-container` 和 `.box` 樣式定義被同等地套用樣式。

# --hints--

你應該在 `header` 元素下方有一個 `div` 元素。

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

你在 `header` 元素下方的 `div` 元素應該有 `flex-container` 類別。

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

你應該在 `.flex-container` 元素內巢狀六個 `div` 元素。

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

你在 `.flex-container` 元素內的所有六個 `div` 元素都應該有 `box` 類別。

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
