---
id: 68f46fcc34da869b1005455d
title: 步驟 37
challengeType: 0
dashedName: step-37
---

# --description--

現在要繼續為你的盒子設計樣式，創建一個新的 `.box3` 選擇器，並為它添加 `background` 屬性，值為 `#ffd166`。

然後，為你的新選擇器添加值為 `2` 的 `order` 屬性。你會看到這個方塊移動到 `.box1` 元素之後。

同時為你的元素添加 `flex-shrink` 屬性，值為 `3`。這表示當該行空間不足時，這個盒子會縮小為其他盒子的三倍。

# --hints--

你應該有一個 `.box3` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box3'));
```

你的 `.box3` 選擇器應該有 `background` 屬性，值為 `#ffd166`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box3')?.getPropVal('background'), 'rgb(255, 209, 102)');
```

你的 `.box3` 選擇器應該有 `order` 屬性且值為 `2`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box3')?.getPropVal('order'), '2');
```

你的 `.box3` 選擇器應該有 `flex-shrink` 屬性，值為 `3`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box3')?.getPropVal('flex-shrink'), '3');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}

--fcc-editable-region--

--fcc-editable-region--
```
