---
id: 68f46b758eb8ea7043abd73c
title: 步驟 33
challengeType: 0
dashedName: step-33
---

# --description--

現在你將要整理你的 flexbox 容器內的方塊。首先，創建一個 `.box1` 選擇器，並為它添加 `background` 屬性，值為 `#f16e79`。

然後，為你的 `order` 屬性添加值為 `1`。你會看到第一個方塊移動到最後。這是因為 `box` 類別已經有一個 `order` 值為 `0`，而具有較高 order 值的項目會出現在較後面。

同時為你的盒子添加 `flex-grow` 屬性，值為 `1`。這將使盒子在其所在的列中擴展以填滿額外空間。

# --hints--

你應該有一個 `.box1` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box1'));
```

你的 `.box1` 選擇器應該有 `background` 屬性，值為 `#f16e79`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('background'), 'rgb(241, 110, 121)');
```

你的 `.box1` 選擇器應該有 `order` 屬性，值為 `1`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('order'), '1');
```

你的 `.box1` 選擇器應該有屬性 `flex-grow`，其值為 `1`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--
```
