---
id: 68f93a8b58cbf330b4ba45c6
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

要完成 `flex-container` 類別的樣式，為你的 `flex-container` 類別添加 `padding` 屬性，值為 `10px`，以及 `margin` 屬性，值為 `20px auto`。

請記住：`padding` 是在元素內部（邊框與內容之間）添加空間，而 `margin` 是在元素外部（邊框與其他元素之間）添加空間。

# --hints--

你的 `.flex-container` 選擇器應該有屬性 `padding`，其值為 `10px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('padding'), '10px');
```

你的 `.flex-container` 選擇器應該有屬性 `margin`，值為 `20px auto`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('margin'), '20px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
