---
id: 68f91df5776e693049fb8942
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

現在是時候為 `styles.css` 檔案添加一些 CSS 規則。首先為 `h1` 元素建立一個選擇器。 

將 `h1` 元素的內容置中，方法是將其 `text-align` 屬性設為 `center`。同時創建一個值為 `10px` 的 `margin-bottom` 屬性，以在 `h1` 元素底部與任何位於下方的 HTML 元素之間添加間距。

# --hints--

你應該有一個 `h1` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

你的 `h1` 選擇器應該將 `text-align` 屬性設為 `center`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

你的 `h1` 選擇器應該將 `margin-bottom` 屬性設為 `10px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

