---
id: 692f318859d5e4bf8fd77ba1
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

`order` 屬性指定同一個容器內的彈性項目相對於其他彈性項目的順序。預設情況下，所有彈性項目的 `order` 值為 `0`，表示它們會依照在 HTML 中撰寫的順序出現。`order` 值較低的項目會先出現，而值較高的項目會最後出現。

為你的 `.box` 選擇器添加屬性 `order`，值為 `0`。

# --hints--

你的 `.box` 選擇器應該有 `order` 屬性，值為 `0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('order'), '0');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
