---
id: 692f30e86027a4b253d0a42d
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

為你的 `.box` 選擇器添加屬性 `font-size`，值為 `1.125rem`。

`font-size` 屬性指定文字的大小。它可以使用不同的單位設定，例如像素、em、rem、百分比等等，使其在響應式設計中具有彈性。

在這種情況下，該值使用 `rem` 單位，該單位是相對於根 `html` 元素，而非父元素。這表示字型大小保持一致，不會受到父元素大小的影響，使版面配置更具可預測性。

# --hints--

你的 `.box` 選擇器應該有 `font-size` 屬性，值為 `1.125rem`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('font-size'), '1.125rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
