---
id: 692f2d37f34ed866ed42fcd6
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

`flex-direction` 屬性控制彈性項目在其容器中的排列方式。預設值是 `row`，會將項目從左到右水平排列。此屬性僅適用於具有 `display: flex` 或 `display: inline-flex` 的元素。

使用 `column` 值時，項目會從上到下垂直堆疊。其他可能的值有 `row`、`row-reverse` 和 `column-reverse`。

現在為 `.box` 選擇器添加 `flex-direction` 屬性，值為 `column`。

# --hints--

你的 `.box` 選擇器應該有屬性 `flex-direction`，其值為 `column`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex-direction'), 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
