---
id: 692f22a2ad44ae887a24c8f9
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在每個 `.box` 元素內，添加一個 `h2` 元素，內容為文字 `Box`，後面接一個空格及從 `1` 開始的連續數字。

# --hints--

你應該在第一個具有類別 `box` 的 `div` 內有一個 `h2` 元素，裡面包含文字 `Box 1`。

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

你應該在第二個具有類別 `box` 的 `div` 內有一個 `h2` 元素，裡面包含文字 `Box 2`。

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

你應該在第三個帶有類別 `box` 的 `div` 裡面有一個 `h2` 元素，裡面的文字是 `Box 3`。

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

你應該在第四個具有類別 `box` 的 `div` 中有一個 `h2` 元素，裡面包含文字 `Box 4`。

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

你應該在第五個具有類別 `box` 的 `div` 中有一個 `h2` 元素，裡面包含文字 `Box 5`。

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

你應該在最後一個具有 `box` 類別的 `div` 中有一個 `h2` 元素，裡面包含文字 `Box 6`。

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
