---
id: 692f22933ba1cb86cd75a70c
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

現在是時候為這些方塊創建顏色名稱。為此，請在每個 `h2` 元素下方添加一個 `p` 元素，並依序放入這些顏色名稱：`Red`、`Orange`、`Yellow`、`Green`、`Blue` 和 `Indigo`。每個方塊一個顏色。

# --hints--

你應該在第一個 `.box` 元素內的 `h2` 元素下方添加一個帶有文字 `Red` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

你應該在第二個 `.box` 元素內的 `h2` 元素下方添加一個帶有文字 `Orange` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

你應該在第三個 `.box` 元素內的 `h2` 元素下方添加一個帶有文字 `Yellow` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

你應該在第四個 `.box` 元素內的 `h2` 元素下方添加一個帶有文字 `Green` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

你應該在第五個 `.box` 元素內的 `h2` 元素下方添加一個帶有文字 `Blue` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

你應該在第六個 `.box` 元素內的 `h2` 元素下方添加一個帶有文字 `Indigo` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
