---
id: 692f2b9f2b9aa9458f6fdc97
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

現在你將要為具有 `box` 類別的 `div` 元素設計樣式。首先創建一個 `.box` 類別選擇器，並為該 `box` 類別選擇器添加屬性 `max-height`，值為 `120px`。

# --hints--

你的 `.box` 選擇器應該有 `max-height` 屬性，值為 `120px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('max-height'), '120px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
