---
id: 692f2c20cb7f9e5050b935a7
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

為你的 `.box` 選擇器添加屬性 `border`，值為 `1px solid #000`。`border` 屬性用於創建元素周圍可見的輪廓。它可以在 `border-width`、`border-style` 和 `border-color` 方面進行自訂。 

在這種情況下，`1px` 是 `border-width`，它定義了邊框的厚度。`solid` 是 `border-style`，設定邊框的外觀（實線、虛線、點線等等）。而 `#000` 是 `border-color`，它定義了邊框的顏色。

# --hints--

你的 `.box` 選擇器應該有 `border` 屬性，值為 `1px solid #000`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border'), '1px solid rgb(0, 0, 0)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
