---
id: 692f2c8ea9aa5a593758ccff
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

`display` 屬性控制元素在頁面上的呈現方式。當你將它定義為 `flex` 值時，這會啟用彈性盒模型佈局模式，使你能控制容器內項目的排列和對齊。

現在為 `.box` 選擇器添加屬性 `display`，值為 `flex`。

# --hints--

你的 `.box` 選擇器應該有屬性 `display` 且值為 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('display'), 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
