---
id: bad87fee1348bd9aec908848
title: 부트스트랩 웰 만들기
challengeType: 0
forumTopicId: 16825
dashedName: create-bootstrap-wells
---

# --description--

부트스트랩에는 열에 깊이감을 시각적으로 줄 수 있는 `well`라는 클래스가 있습니다.

각 `col-xs-6` `div` 요소 안에 `well` 클래스를 가진 `div` 요소를 중첩하세요.

# --hints--

각 `div` 요소 안에 `well` 클래스를 가진 `div` 요소를 추가하세요. `div` 요소는 `col-xs-6` 클래스를 가지고 있어야 합니다.

```js
const wells = document.querySelectorAll('div.col-xs-6 > div.well');
assert.lengthOf( wells,2 ); 
```

`col-xs-6` 클래스를 가진 두 `div` 요소 모두 `row` 클래스를 가진 `div` 요소 안에 중첩되어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

모든 `div` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g)?.length , code.match(/<div/g)?.length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

    </div>
    <div class="col-xs-6">

    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
  </div>
</div>
```
