---
id: bad87fee1348bd9aec908853
title: 부트스트랩 요소에 id 속성 추가
challengeType: 0
forumTopicId: 16639
dashedName: add-id-attributes-to-bootstrap-elements
---

# --description--

클래스 속성 외에도 각 요소에 `id` 속성을 줄 수 있다는 점을 기억하세요.

각 id는 특정 요소에 고유해야 하며 페이지당 한 번만 사용해야 합니다.

클래스 `div`인 `well` 요소 각각에 고유한 id를 부여해 봅시다.

요소에 id를 이렇게 줄 수 있다는 점을 기억하세요:

```html
<div class="well" id="center-well">
```

왼쪽 웰에 `left-well`라는 id를 주세요. 오른쪽 웰에는 `right-well`라는 id를 주세요.

# --hints--

왼쪽 `well`에는 `left-well`라는 id가 있어야 합니다.

```js
const column = document.querySelectorAll('.col-xs-6')[0];
const leftWall = column?.querySelectorAll(`:scope ${'#left-well'}`);
assert.lengthOf(leftWall,1)
```

오른쪽 `well`에는 `right-well`라는 id가 있어야 합니다.

```js
const column = document.querySelectorAll('.col-xs-6')[1];
const rightWall = column?.querySelectorAll(`:scope ${'#right-well'}`);
assert.lengthOf(rightWall,1)
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well" id="left-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well" id="right-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```
