---
id: bad87fee1348bd9aec908857
title: 코드를 명확하게 하기 위해 주석 사용하기
challengeType: 0
forumTopicId: 18347
dashedName: use-comments-to-clarify-code
---

# --description--

jQuery를 사용하기 시작하면 HTML 요소를 실제로 HTML에서 변경하지 않고도 수정할 수 있습니다.

이 코드를 직접 수정해서는 안 된다는 점을 모두가 알도록 하세요.

주석은 `<!--`로 시작해서 `-->`로 끝낼 수 있다는 점을 기억하세요.

HTML 맨 위에 `Code below this line should not be changed`라는 주석을 추가하세요.

# --hints--

HTML 맨 위에 `<!--`로 주석을 시작해야 합니다.

```js
assert.match(code,(/^\s*<!--/));
```

주석에 `Code below this line should not be changed`라는 텍스트가 있어야 합니다.

```js
assert.match(code,(/<!--(?!(>|->|.*-->.*this line))\s*.*this line.*\s*-->/gi));
```

주석은 `-->`로 닫아야 합니다.

```js
assert.match(code,(/-->.*\n+.+/g));
```

주석 시작과 끝이 같은 개수여야 합니다.

```js
assert.equal(code.match(/<!--/g).length,code.match(/-->/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<!-- Code below this line should not be changed -->
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
