---
id: bad87fee1348bd9aec908856
title: 부트스트랩 버튼에 레이블 지정하기
challengeType: 0
forumTopicId: 18222
dashedName: label-bootstrap-buttons
---

# --description--

우리가 우물에 레이블을 붙였던 것처럼, 버튼에도 레이블을 붙이고 싶습니다.

각 `button` 요소에 해당 id 선택자에 맞는 텍스트를 넣으세요.

# --hints--

id가 `button`인 `target1` 요소에는 텍스트 `#target1`가 있어야 합니다.

```js
assert.match(document.querySelector('#target1')?.textContent,/#target1/gi)
```

id가 `button`인 `target2` 요소에는 텍스트 `#target2`가 있어야 합니다.

```js
assert.match(document.querySelector('#target2')?.textContent,/#target2/gi)
```

id가 `button`인 `target3` 요소에는 텍스트 `#target3`가 있어야 합니다.

```js
assert.match(document.querySelector('#target3')?.textContent,/#target3/gi)
```

id가 `button`인 `target4` 요소에는 텍스트 `#target4`가 있어야 합니다.

```js
assert.match(document.querySelector('#target4')?.textContent,/#target4/gi)
```

id가 `button`인 `target5` 요소에는 텍스트 `#target5`가 있어야 합니다.

```js
assert.match(document.querySelector('#target5')?.textContent,/#target5/gi)
```

id가 `button`인 `target6` 요소에는 텍스트 `#target6`가 있어야 합니다.

```js
assert.match(document.querySelector('#target6')?.textContent,/#target6/gi)
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1"></button>
        <button class="btn btn-default target" id="target2"></button>
        <button class="btn btn-default target" id="target3"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4"></button>
        <button class="btn btn-default target" id="target5"></button>
        <button class="btn btn-default target" id="target6"></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
