---
id: bad87fee1348bd9aec908849
title: 부트스트랩 웰 안에 요소 추가하기
challengeType: 0
forumTopicId: 16636
dashedName: add-elements-within-your-bootstrap-wells
---

# --description--

이제 각 열의 행에 `div` 요소가 여러 개 중첩되어 있습니다. 이 정도 깊이면 충분합니다. 이제 `button` 요소를 추가할 수 있습니다.

클래스 이름이 `button`인 각 `div` 요소 안에 `well` 요소 세 개를 중첩하세요.

# --hints--

클래스가 `button`인 각 `div` 요소 안에 `well` 요소 세 개가 중첩되어야 합니다.

```js
const buttonOne = document.querySelectorAll('div.well')?.[0];
const buttonOneChildren = buttonOne?.querySelectorAll(`:scope ${'button'}`);
assert.lengthOf(buttonOneChildren,3);

const buttonTwo = document.querySelectorAll('div.well')?.[1]
const buttonTwoChildren = buttonTwo?.querySelectorAll(`:scope ${'button'}`);
assert.lengthOf(buttonTwoChildren,3);
```

총 6개의 `button` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('button'), 6);
```

모든 `button` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/button>/g);
assert.match(code,/<button/g);
assert.equal(code.match(/<\/button>/g).length,code.match(/<button/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">



      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">



      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
  </div>
</div>
```
