---
id: bad87fee1348bd9aec908746
title: 부트스트랩 container-fluid div 안에 페이지를 배치하기
challengeType: 0
forumTopicId: 18198
dashedName: house-our-page-within-a-bootstrap-container-fluid-div
---

# --description--

이제 페이지의 모든 콘텐츠가 모바일 반응형인지 확인해 보겠습니다.

`h3` 요소를 `div` 요소 안에 `container-fluid` 클래스를 가진 상태로 중첩하세요.

# --hints--

`div` 요소는 `container-fluid` 클래스를 가져야 합니다.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container-fluid'));
```

각 `div` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);

assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

`h3` 요소는 `div` 요소 안에 중첩되어야 합니다.

```js
const divElement = document.querySelector('div');
const divChildren = divElement?.querySelectorAll(`:scope ${'h3'}`)
assert.lengthOf(divChildren,1);
```

# --seed--

## --seed-contents--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```

# --solutions--

```html
<div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
</div>
```
