---
id: bad87fee1348bd9aec908846
title: 부트스트랩 헤드라인 만들기
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

이제 HTML, CSS, 부트스트랩 기술을 연습하기 위해 처음부터 무언가를 만들어 보겠습니다.

곧 jQuery 도전 과제에서 사용할 jQuery 놀이터를 만들 것입니다.

먼저, 텍스트가 `h3`인 `jQuery Playground` 요소를 만드세요.

`h3` 요소에 부트스트랩 클래스 `text-primary`를 적용해 색상을 지정하고, 부트스트랩 클래스 `text-center`를 적용해 가운데 정렬하세요.

# --hints--

페이지에 `h3` 요소를 추가해야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

`h3` 요소에는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

`h3` 요소는 클래스 `text-primary`를 적용해 색상을 지정해야 합니다.

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

`h3` 요소는 클래스 `text-center`를 적용해 가운데 정렬해야 합니다.

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

`h3` 요소에는 텍스트 `jQuery Playground`가 있어야 합니다.

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```
