---
id: bad87fee1348bd9bec908846
title: 부트스트랩 행 만들기
challengeType: 0
forumTopicId: 16813
dashedName: create-a-bootstrap-row
---

# --description--

이제 인라인 요소를 위한 부트스트랩 행을 만들겠습니다.

`div` 태그 아래에 `h3` 요소를 만들고, 클래스는 `row`로 지정하세요.

# --hints--

`div` 요소 아래에 `h3` 요소를 추가하세요.

```js
assert.lengthOf(document.querySelectorAll('div'),2);
assert.lengthOf(document.querySelectorAll('div.row h3.text-primary'),0);
assert.lengthOf(document.querySelectorAll('div.row + h3.text-primary'),0)
assert.lengthOf(document.querySelectorAll('h3.text-primary + div.row'),1);
```

`div` 요소에 클래스 `row`가 있어야 합니다.

```js
const newDiv = document.querySelectorAll('div')?.[1]; 
assert.isTrue(newDiv?.classList?.contains('row'));
```

`row div`는 `container-fluid div` 안에 중첩되어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div.container-fluid div.row'),1);
```

`div` 요소에는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>

</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row"></div>
</div>
```
