---
id: bad87fee1348bd9aec908847
title: 부트스트랩 행을 두 개의 열로 나누기
challengeType: 0
forumTopicId: 18306
dashedName: split-your-bootstrap-row
---

# --description--

부트스트랩 행이 준비되었으니, 이제 두 개의 열로 나누어 요소들을 배치해 보세요.

행 안에 클래스가 `div`인 `col-xs-6` 요소 두 개를 만드세요.

# --hints--

`div class="col-xs-6"` 요소 안에 `div class="row"` 요소 두 개가 중첩되어 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

모든 `div` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length , code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">


  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6"></div>
    <div class="col-xs-6"></div>
  </div>
</div>
```
