---
id: bad87fee1348bd9aec908854
title: 레이블 부트스트랩 웰
challengeType: 0
forumTopicId: 18223
dashedName: label-bootstrap-wells
---

# --description--

명확성을 위해 두 웰 모두에 아이디를 레이블로 붙이세요.

왼쪽 웰 위, 그 `col-xs-6` `div` 요소 안에 `h4` 요소를 추가하고 텍스트는 `#left-well`로 하세요.

오른쪽 웰 위, 그 `col-xs-6` `div` 요소 안에 `h4` 요소를 추가하고 텍스트는 `#right-well`로 하세요.

# --hints--

각 `h4` 요소에 `<div class="col-xs-6">` 요소를 추가해야 합니다.

```js
const columnSixes = document.querySelectorAll('.col-xs-6');
const columnSixOneChildren = columnSixes?.[0]?.querySelectorAll(`:scope ${'h4'}`);
assert.lengthOf(columnSixOneChildren,1); 

const columnSixTwoChildren = columnSixes?.[1]?.querySelectorAll(`:scope ${'h4'}`);
assert.lengthOf(columnSixTwoChildren,1); 
```

하나의 `h4` 요소에는 텍스트 `#left-well`가 있어야 합니다.

```js
const firstH4 = document.querySelectorAll('h4')?.[0]; 
assert.match(firstH4?.textContent,/#left-well/gi); 
```

하나의 `h4` 요소에는 텍스트 `#right-well`가 있어야 합니다.

```js
const secondH4 = document.querySelectorAll('h4')?.[1]; 
assert.match(secondH4?.textContent,/#right-well/gi); 
```

모든 `h4` 요소에는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/h4>/g);
assert.match(code,/<h4/g);
assert.equal(code.match(/<\/h4>/g).length , code.match(/<h4/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

      <div class="well" id="left-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">

      <div class="well" id="right-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```
