---
id: a77dbc43c33f39daa4429b4f
title: بناء دالة تحقق من القيمة المنطقية
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

في هذا المختبر، ستبني دالة تتحقق مما إذا كانت القيمة مصنفة كقيمة منطقية أولية.

القيم المنطقية الأولية هي `true` و `false`.

**الهدف:** نفذ قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة تسمى `booWho` تستقبل معلمة واحدة.
1. إذا كانت المعلمة المستلمة قيمة منطقية أولية، يجب أن تُرجع الدالة `true`.
1. إذا كانت المعلمة أي قيمة أخرى، يجب أن تُرجع الدالة `false`.

# --hints--

يجب أن تملك دالة `booWho`.

```js
assert.isFunction(booWho);
```

`booWho(true)` يجب أن تُرجع `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` يجب أن تُرجع `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` يجب أن تُرجع `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` يجب أن تُرجع `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` يجب أن تُرجع `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` يجب أن تُرجع `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` يجب أن تُرجع `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` يجب أن تُرجع `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` يجب أن تُرجع `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` يجب أن تُرجع `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
