---
id: 66bb3e20d3dc5b6d0a21f5dd
title: 실시간 카운터 만들기
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

이 실습에서는 JavaScript를 사용해 실시간 문자 수 카운터를 만드세요. 문자 수 카운터는 `textarea` 요소에 입력된 문자 수를 표시합니다. 사용자가 `textarea`에 입력할 때마다 카운터가 실시간으로 업데이트됩니다.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `textarea`의 `id`를 가진 `text-input` 요소가 있어야 합니다.
2. `p`의 `id`를 가진 `char-count` 요소가 있어야 합니다. 처음에는 `Character Count: 0/50` 텍스트가 포함되어야 합니다. 이 자리 표시자 텍스트는 `textarea` 요소에 있는 현재 문자 수로 교체되어야 합니다.
3. `#text-input` 요소에 `hello coder` 텍스트가 있을 때 `#char-count` 요소에는 `"Character Count: 11/50"` 텍스트가 있어야 합니다.
4. `#text-input` 요소에 `i am a programmer` 텍스트가 있을 때 `#char-count` 요소에는 `"Character Count: 17/50"` 텍스트가 있어야 합니다.
5. `#text-input` 요소에 `hello world` 텍스트가 있을 때 `#char-count` 요소에는 `"Character Count: 11/50"` 텍스트가 있어야 합니다.
6. `#text-input` 요소에 `I am learning a new language and it is called c++.` 텍스트가 있을 때 `#char-count` 요소에는 `"Character Count: 50/50"` 텍스트가 있어야 합니다. **참고:** `maxlength` 속성을 사용하면 올바른 기능을 구현할 수 있지만, 이 실습은 JavaScript만 사용하는 해결책을 요구합니다.  
7. 사용자는 `50` 요소에 `textarea`자를 초과해 입력할 수 없어야 합니다. 문자 수가 `50`에 도달하면 초과 입력은 자동으로 잘리고 `Character Count: 50/50`에 `red` 텍스트가 표시되어야 합니다.

**참고:** HTML에 스타일시트와 JavaScript 파일을 반드시 연결하세요.

# --hints--

`textarea`의 `id`를 가진 `text-input` 요소가 있어야 합니다.

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

`p`의 `id`를 가진 `char-count` 요소가 있어야 합니다.

```js
assert.exists("p#char-count");
```

`#char-count` 요소는 처음에 `Character Count: 0/50` 텍스트를 포함해야 합니다.

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

`#text-input` 요소에 `hello coder` 텍스트가 있을 때 `#char-count` 요소에는 `"Character Count: 11/50"` 텍스트가 있어야 합니다.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

문자 수가 `50`일 때 텍스트가 `red`에 표시되어야 합니다.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

문자 수가 `50` 이상이면 초과 입력이 잘려야 합니다.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
