---
id: 66a8347f97a33e36ffc81b9f
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

To prepare to create some actual content, add a `body` element below the `head` element.

# --hints--

You should have an opening `body` tag.

```js
assert.match(code, /<body>/i);
```

You should have a closing `body` tag.

```js
assert.match(code, /<\/body>/i);
```

لا يجب عليك تغيير عنصر `head`. تأكد من أنك لم تقم بحذف علامة الإغلاق الخاصة بك.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

عنصر `body` الخاص بك يجب أن يأتي بعد عنصر `head` الخاص بك.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>

  --fcc-editable-region--
</html>
```
