---
id: 67ed03ac474c48692f41749e
title: 해시 테이블 만들기
challengeType: 27
dashedName: build-a-hash-table
saveSubmissionToDB: true
---

# --description--

이 실습에서는 해시 테이블을 처음부터 만들어 보세요. 해시 테이블은 키-값 쌍을 저장하는 자료 구조입니다. 해시 테이블은 키를 입력으로 받아 특정 해싱 함수에 따라 이 키를 해싱하는 방식으로 작동합니다.

이 실습에서는 해싱 함수가 간단합니다: 키의 각 문자의 유니코드 값을 모두 더합니다. 그런 다음 해시 값이 실제 키로 사용되어 연관된 값을 저장합니다. 같은 해시 값은 키와 연관된 값을 검색하거나 삭제할 때도 사용됩니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `HashTable`라는 클래스를 정의해야 하며, `collection`의 새 인스턴스가 생성될 때 `HashTable` 속성이 빈 딕셔너리로 초기화되어야 합니다. `collection` 딕셔너리는 키의 해시 값을 기준으로 키-값 쌍을 저장해야 합니다.

2. `HashTable` 클래스는 네 개의 인스턴스 메서드 `hash`, `add`, `remove`, `lookup`를 가져야 합니다.

3. `hash` 메서드는:

    - 문자열을 매개변수로 받아야 합니다.
    - 문자열의 각 문자의 유니코드(ASCII) 값을 모두 더한 해시 값을 반환해야 합니다. 이 계산에는 `ord` 함수를 사용할 수 있습니다.

4. `add` 메서드는:
    - 키-값 쌍을 나타내는 두 인수를 받고, 키의 해시를 계산해야 합니다.
    - 계산된 해시 값을 키로 사용하여 `collection` 딕셔너리 안에 키-값 쌍을 포함하는 딕셔너리를 저장해야 합니다.
    - 여러 키가 같은 해시 값을 생성하면, 해당 해시 값 아래 기존 중첩된 딕셔너리에 키-값 쌍을 저장해야 합니다.

5. `remove` 메서드는:

    - 키를 인수로 받아 해시를 계산해야 합니다.
    - 컬렉션에 키가 존재하는지 확인해야 합니다.
    - 해시 테이블에서 해당 키-값 쌍을 삭제해야 합니다.
    - 키가 컬렉션에 없으면 오류를 발생시키거나 아무것도 삭제하지 않아야 합니다.

6. `lookup` 메서드는:

    - 키를 인수로 받아야 합니다.
    - 키의 해시를 계산하고, 해시 테이블에 저장된 해당 값을 반환해야 합니다.
    - 키가 컬렉션에 없으면 `None`를 반환해야 합니다.

# --hints--

`HashTable` 클래스를 정의해야 합니다.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_class("HashTable")
    `)) 
})
```

`HashTable` 클래스의 새 인스턴스가 생성될 때, `collection` 속성은 빈 딕셔너리로 초기화되어야 합니다.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.collection == {}
    `) 
})
```

`HashTable` 클래스는 `hash` 메서드를 가져야 합니다.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('hash')`))}})
```

`hash` 메서드는 문자열을 매개변수로 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.hash)
    len(sig.parameters) == 1
  `))
})
```

`hash` 메서드는 문자열을 인수로 받아 문자열의 각 문자의 유니코드 값 합계를 반환해야 합니다.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("fcc")
        assert hash_result == 300

        # prevent hardcoding
        assert ht.hash("golf") == 424
        assert ht.hash("read") == 412
`) })
```

`HashTable` 클래스는 `add` 메서드를 가져야 합니다.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('add')`))}})
```

`add` 메서드는 키와 값을 매개변수로 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.add)
    len(sig.parameters) == 2
  `))
})
```

`HashTable` 클래스는 `remove` 메서드를 가져야 합니다.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('remove')`))}})
```

`remove` 메서드는 키를 매개변수로 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.remove)
    len(sig.parameters) == 1
  `))
})
```

컬렉션에 존재하지 않는 키를 삭제하려 할 때, 오류를 발생시키거나 아무것도 삭제하지 않아야 합니다.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("rose", "flower")
        index = ht.hash("rose")
        original = ht.collection.copy()

        ht.remove("tulip")
        ht.remove("sore")

        assert ht.collection == original
        assert "rose" in ht.collection[index]
        assert "tulip" not in ht.collection.get(index, {})
    `) 
})
```

여러 키가 같은 인덱스로 해시될 경우, `remove` 메서드는 해당 인덱스의 전체 딕셔너리를 삭제하지 않고 특정 키-값 쌍만 삭제해야 합니다.

```js
({
  test: () => runPython(`
    ht = HashTable()
    ht.add("rose", "flower")
    ht.add("sore", "pain")  # "rose" and "sore" both hash to the same index

    index = ht.hash("rose")
    ht.remove("rose")

    assert index in ht.collection
    assert "rose" not in ht.collection[index]
    assert "sore" in ht.collection[index]
    assert ht.collection[index]["sore"] == "pain"
  `)
})
```

`HashTable` 클래스는 `lookup` 메서드를 가져야 합니다.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('lookup')`))}})
```

`lookup` 메서드는 키를 매개변수로 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.lookup)
    len(sig.parameters) == 1
  `))
})
```

`HashTable().hash('golf')`는 `424`를 반환해야 합니다.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("golf")
        assert hash_result == 424

        # prevent hardcoding
        assert ht.hash("dear") == 412
        assert ht.hash("cat") == 312
    `) 
})
```

`HashTable().add('golf', 'sport')`는 키-값 쌍을 `424` 키에 컬렉션에 추가해야 합니다.

```js

({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")
    expected_value = {424: {'golf': 'sport'}}

    assert ht.collection == expected_value
    `) 
})
```

`HashTable().add('dear', 'friend')`과 `HashTable().add('read', 'book')`는 `412` 인덱스에 중첩된 딕셔너리로 두 키-값 쌍을 컬렉션에 추가해야 합니다.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("dear", "friend")
    ht.add("read", "book")

    expected_value = {
        "dear": "friend",
        "read": "book"
    }

    assert ht.collection.get(412) == expected_value
    `) 
})
```

키가 해시 테이블에 존재할 때, `remove()` 메서드는 주어진 키와 해당 값을 컬렉션에서 삭제해야 합니다.

```js
({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")

    expected_value_before_removal = {
        "golf": "sport"
    }

    index = ht.hash("golf")

    assert ht.collection.get(index) == expected_value_before_removal

    ht.remove("golf")

    assert "golf" not in ht.collection.get(index, {})
    `) 
})
```

`'golf', 'sport'` 키-값 쌍이 해시 테이블에 존재할 때, `HashTable().lookup('golf')`는 `sport`를 반환해야 합니다.

```js
({ 
    test: () => runPython(`
    ht = HashTable()

    ht.add("golf", "sport")

    expected_value = "sport"
    assert ht.lookup("golf") == expected_value
    `) 
})

```

`'golf', 'sport'` 키-값 쌍이 컬렉션에 존재하지 않을 때, `HashTable().lookup('golf')`는 `None`를 반환해야 합니다.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.lookup("golf") is None
`) })

```

`'fcc'` 키가 컬렉션에 존재할 때, `HashTable().lookup('cfc')`는 `None`를 반환해야 합니다.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("fcc", "coding")
        assert ht.lookup("cfc") is None
`) })
```

`('rose', 'flower')`를 해시 테이블에 추가하면, `collection` 속성은 다음과 같아야 합니다: `{ 441: { 'rose': 'flower' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("rose", "flower")

    expected_value = {
        "rose": "flower"
    }

    assert ht.collection.get(441) == expected_value

    # prevent hardcoding

    ht2 = HashTable()

    ht2.add("kebab", "food")
    expected_value = {
        "kebab": "food"
    }

    assert ht2.collection.get(501) == expected_value
    `)
})
```

기존 키와 같은 해시 값을 갖는 키를 추가하면, 예를 들어 `fcc`와 `cfc`처럼, `collection`는 다음과 같아야 합니다: `{ 300: { 'fcc': 'coding', 'cfc':  'chemical' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("fcc", "coding")
    ht.add("cfc", "chemical")

    expected_value = {
        "fcc": "coding",
        "cfc": "chemical"
    }

    assert ht.collection.get(300) == expected_value

    # prevent hardcoding

    ht2 = HashTable()
    ht2.add("cat", "animal")
    ht2.add("act", "verb")

    expected_value = {
        "cat": "animal",
        "act": "verb"
    }

    assert ht2.collection.get(312) == expected_value
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class HashTable:
    def __init__(self):
        self.collection = {}

    def hash(self, string):
        hashed = 0
        for char in string:
            hashed += ord(char)  
        return hashed

    def add(self, key, val):
        the_hash = self.hash(key)
        if the_hash not in self.collection:
            self.collection[the_hash] = {}
        self.collection[the_hash][key] = val

    def remove(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            del self.collection[the_hash][key]
            if not self.collection[the_hash]: 
                del self.collection[the_hash]

    def lookup(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            return self.collection[the_hash][key]
        return None
```
