---
id: 68e51d8f163d3f8b70f9c4e5
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会中，你将练习使用 `includes()` 和 `slice()` 方法。

首先，创建一个名为 `fccSentence` 的变量，并赋值为字串 `"freeCodeCamp is a great place to learn web development."`。

# --hints--

你应该声明一个名为 `fccSentence` 的变量。

```js
assert.exists(fccSentence);
```

你的 `fccSentence` 变量应是一个字串。

```js
assert.isString(fccSentence);
```

你的 `fccSentence` 变量应该赋值为字串 `"freeCodeCamp is a great place to learn web development."`。

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
