---
id: 68e6172975008fea742a63cd
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

现在你将探索 `slice()` 方法，它用于提取字串的一部分。

创建一个名为 `message` 的变量，并赋值为字串 `"Welcome to freeCodeCamp!"`。

然后使用 `console.log()` 将 `"Here are some examples of the slice() method:"` 记录到控制台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该声明一个名为 `message` 的变量。

```js
assert.exists(message);
```

你的 `message` 变量应是一个字串。

```js
assert.isString(message);
```

你的 `message` 变量应为 `"Welcome to freeCodeCamp!"`。

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

你应该将 `"Here are some examples of the slice() method:"` 记录到控制台日志。

```js
assert.deepInclude(spy.calls, ["Here are some examples of the slice() method:"]);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
