---
id: 68e6172975008fea742a63cb
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

现在是时候看看当未找到子字串时会发生什么。

创建一个名为 `hasJavaScript` 的变量。然后将使用 `includes()` 方法查看 `fccSentence` 是否包含 `"JavaScript"` 的结果赋值给它。

然后使用模板字面或字串串接，将消息 `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` 记录到控制台。将 `<hasJavaScript>` 替换为变量的实际值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该声明一个名为 `hasJavaScript` 的变量。

```js
assert.exists(hasJavaScript);
```

你应该使用 `fccSentence.includes("JavaScript")`。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

你应该将 `fccSentence.includes("JavaScript")` 的结果赋值给你的 `hasJavaScript` 变量。

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

`fccSentence.includes("JavaScript")` 的结果应为 `false`。 

```js
assert.isFalse(hasJavaScript);
```

你应该记录日志消息 `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`，其中 `<hasJavaScript>` 应替换为变量的实际值。

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
