---
id: 68e6172875008fea742a63c9
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

正如你从之前的课程中查全率的，`includes()` 方法查看字串是否包含特定的子串并返回 `true` 或 `false`。

这是一个使用 `includes()` 方法的示例：

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

创建一个名为 `hasFreeCodeCamp` 的变量。然后将使用 `includes()` 方法对 `fccSentence` 进行查看是否包含 `"freeCodeCamp"` 的结果赋值给它。

# --hints--

你应该声明一个名为 `hasFreeCodeCamp` 的变量。

```js
assert.exists(hasFreeCodeCamp);
```

你应该使用 `fccSentence.includes("freeCodeCamp")`。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

你应该将 `fccSentence.includes("freeCodeCamp")` 的结果赋值给你的 `hasFreeCodeCamp` 变量。

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

`fccSentence.includes("freeCodeCamp")` 的结果应为 `true`。 

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
