---
id: 68e6172975008fea742a63cc
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

`includes()` 方法区分大小写，这意味着它区分大写字母和小写字母。

创建一个名为 `hasLowercaseFCC` 的变量。然后赋值为使用 `includes()` 方法查看 `fccSentence` 是否包含 `"freecodecamp"`（全小写）的结果。

然后使用模板字面或字串串接将消息 `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` 记录到控制台。将 `<hasLowercaseFCC>` 替换为变量的实际值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该声明一个名为 `hasLowercaseFCC` 的变量。

```js
assert.exists(hasLowercaseFCC);
```

你应该使用 `fccSentence.includes("freecodecamp")`。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

你应该将 `fccSentence.includes("freecodecamp")` 的结果赋值给你的 `hasLowercaseFCC` 变量。

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

`fccSentence.includes("freecodecamp")` 的结果应为 `false`。 

```js
assert.isFalse(hasLowercaseFCC);
```

你应该记录包含 `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` 的消息，其中 `<hasLowercaseFCC>` 应替换为变量的实际值。

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
