---
id: 68e6172975008fea742a63ce
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

`slice()` 方法提取字串的一部分。它接受两个参数：起始索引（包含）和结束索引（不包含）。

这是一个示例：

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

字串索引从 0 开始。在 `message` 字串中，索引 11 处的字符是 `f`（"freeCodeCamp" 的起始字符）。

创建一个名为 `platform` 的变量。然后将使用 `message.slice(11, 23)` 提取 `"freeCodeCamp"` 的结果赋值给它。

# --hints--

你应该声明一个名为 `platform` 的变量。

```js
assert.exists(platform);
```

你的 `platform` 变量应为字串。

```js
assert.isString(platform);
```

你应该使用 `message.slice(11, 23)`。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

你应该将 `message.slice(11, 23)` 的结果赋值给你的 `platform` 变量。

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

你的 `platform` 变量应与 `"freeCodeCamp"` 相等。

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
