---
id: 68e6172975008fea742a63d1
title: 步骤11
challengeType: 1
dashedName: step-11
---

# --description--

`slice()` 方法可以使用负索引从字串末尾开始计数。`-1` 指的是最后一个字符。

这是一个示例：

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

创建一个名为 `endPunctuation` 的变量。然后将使用 `message.slice(-1)` 提取最后一个字符的结果赋值给它。

然后使用模板字面或字串串接，将消息 `The ending punctuation mark is a "<endPunctuation>"` 记录到控制台。将 `<endPunctuation>` 替换为变量的实际值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该声明一个名为 `endPunctuation` 的变量。

```js
assert.exists(endPunctuation);
```

你的 `endPunctuation` 变量应为字串。

```js
assert.isString(endPunctuation);
```

你应该使用 `message.slice(-1)`。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

你应该将 `message.slice(-1)` 的结果赋值给你的 `endPunctuation` 变量。

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

你的 `endPunctuation` 变量应相等于 `"!"`。

```js
assert.strictEqual(endPunctuation, "!");
```

你应该记录消息 `The ending punctuation mark is a "<endPunctuation>"`，其中 `<endPunctuation>` 应替换为变量的实际值。

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
