---
id: 68ef6e42fa6e8a9a18a401fa
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

使用 `console.log()` 将 `"Here are some examples of the includes() method:"` 记录到控制台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该将 `"Here are some examples of the includes() method:"` 记录到控制台日志。

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
