---
id: 66c057041df6394ca796bf33
title: Einen Satzgenerator erstellen
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

In diesem Labor erstellen Sie zwei verschiedene Geschichten mit einer Satzvorlage. Sie verwenden Variablen, um verschiedene Teile der Geschichte zu speichern, und geben die Geschichten dann in der Konsole aus.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten die folgenden Variablen mit `let` deklarieren:
   
    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. Sie sollten den oben genannten Variablen einige String-Werte Ihrer Wahl zuweisen.

1. Sie sollten eine `firstStory`-Variable deklarieren.
   
1. Sie sollten die folgende Satzvorlage verwenden, um die erste Geschichte zu erstellen und sie der Variablen `firstStory` zuzuweisen: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`

1. Sie sollten Ihre erste Geschichte mit der Nachricht `"First story: [firstStory]"` in der Konsole ausgeben.

1. Sie sollten den Variablen `adjective`, `noun`, `verb`, `place`, `adjective2` und `noun2` neue Werte zuweisen.

1. Sie sollten eine `secondStory`-Variable deklarieren.

1. Sie sollten eine weitere Geschichte mit derselben Vorlage erstellen und sie der Variablen `secondStory` zuweisen.

1. Sie sollten Ihre zweite Geschichte mit der Nachricht `"Second story: [secondStory]"` in der Konsole ausgeben.


# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call.join(' '));
```

# --hints--

Sie sollten eine `adjective`-Variable deklarieren.

```js
assert.isNotNull(adjective);
```

Sie sollten eine `noun`-Variable deklarieren.

```js
assert.isNotNull(noun);
```

Sie sollten eine `verb`-Variable deklarieren.

```js
assert.isNotNull(verb);
```

Sie sollten eine `place`-Variable deklarieren.

```js
assert.isNotNull(place);
```

Sie sollten eine `adjective2`-Variable deklarieren.

```js
assert.isNotNull(adjective2);
```

Sie sollten eine `noun2`-Variable deklarieren.

```js
assert.isNotNull(noun2);
```

Sie sollten der `adjective`-Variable einen String-Wert zuweisen.

```js
assert.isString(adjective);
```

Sie sollten der `noun`-Variable einen String-Wert zuweisen.

```js
assert.isString(noun);
```

Sie sollten der `verb`-Variable einen String-Wert zuweisen.

```js
assert.isString(verb);
```

Sie sollten der `place`-Variable einen String-Wert zuweisen.

```js
assert.isString(place);
```

Sie sollten der `adjective2`-Variable einen String-Wert zuweisen.

```js
assert.isString(adjective2);
```

Sie sollten der `noun2`-Variable einen String-Wert zuweisen.

```js
assert.isString(noun2);
```

Sie sollten eine `firstStory`-Variable deklarieren.

```js
assert.isNotNull(firstStory);
```

Sie sollten für die erste Geschichte das korrekte Geschichtsformat verwenden: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Achten Sie auf Leerzeichen.

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

Sie sollten Ihre erste Geschichte mit den deklarierten Variablen in der richtigen Reihenfolge zusammensetzen.

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Sie sollten Ihre erste Geschichte mit der Nachricht `"First story: [firstStory]"` in der Konsole ausgeben.

```js
assert.include(getLogs(), `First story: ${firstStory}`);
```

Sie sollten eine `secondStory`-Variable deklarieren.

```js
assert.isNotNull(secondStory);
```

Sie sollten die `adjective`-Variable für die zweite Geschichte neu zuweisen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

Sie sollten die `noun`-Variable für die zweite Geschichte neu zuweisen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

Sie sollten die `verb`-Variable für die zweite Geschichte neu zuweisen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

Sie sollten die `place`-Variable für die zweite Geschichte neu zuweisen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

Sie sollten die `adjective2`-Variable für die zweite Geschichte neu zuweisen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

Sie sollten die `noun2`-Variable für die zweite Geschichte neu zuweisen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

Sie sollten für die zweite Geschichte das korrekte Geschichtsformat verwenden: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Achten Sie auf Leerzeichen.

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

Sie sollten Ihre zweite Geschichte mit den deklarierten Variablen in der richtigen Reihenfolge zusammensetzen.

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Sie sollten Ihre zweite Geschichte mit der Nachricht `"Second story: [secondStory]"` in der Konsole ausgeben.

```js
assert.include(getLogs(), `Second story: ${secondStory}`);
```

Ihre `firstStory` sollte nicht dieselbe sein wie Ihre `secondStory`.

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```

