---
id: 60a3e3396c7b40068ad6996b
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Within the `head` element, add a `meta` tag which sets the `charset` to `UTF-8`, and a `title` element with the value `Rothko Painting`.

Füge innerhalb des `body`-Elements ein `img`-Element mit einem `src`-Attribut mit `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` als Wert hinzu.
# --hints--

Your code should have a `meta` tag.

```js
assert(document.querySelectorAll('meta').length === 1);
```

The `meta` tag should set the `charset` to `UTF-8`.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Your code should have a `title` element.

```js
assert(document.querySelectorAll('title').length === 1);
```

The `title` should be `Rothko Painting`.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

Your code should have an `img` element.

```js
assert(document.querySelectorAll('img').length === 1);
```

The `img` element should have a `src` of `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
