---
id: 60a3e3396c7b40068ad69970
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Before you can start styling the `div` you added, you need to link your CSS to your HTML.

Füge ein `link`-Element hinzu, um deine `styles.css`-Datei zu verknüpfen. Setze den `href` auf `styles.css` und denke daran, das `rel`-Attribut auf `stylesheet` zu setzen.

# --hints--

Dein Code sollte ein `link`-Element haben.

```js
assert(/<link/.test(code))
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` enthalten.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` enthalten.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
