---
id: 60a3e3396c7b40068ad6996f
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Add a `div` element in the `body`.

Setze das `class`-Attribut gleich `canvas`.

Dies dient als Canvas für dein Bild.

# --hints--

Dein Code sollte ein `div`-Element enthalten.

```js
assert(document.querySelectorAll('div').length === 1)
```

Das `div`-Element sollte eine `class` mit dem Wert `canvas` enthalten.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
