---
id: 60a3e3396c7b40068ad6996a
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

By now, you should be familiar with the basic elements an HTML page should have.

Richte deinen Code mit einer `DOCTYPE`-Deklaration, einem `html`-Element, dessen Sprache Englisch ist, einem `head`-Element und einem `body`-Element ein.

# --hints--

Dein Code sollte eine `<!DOCTYPE html>`-Deklaration enthalten.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Du solltest ein öffnendes `<html>`-Tag mit einem `lang`-Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein Code sollte innerhalb des `html`-Elements ein `head`-Element haben.

```js
assert(document.querySelectorAll('head').length === 1);
```

Dein Code sollte ein `body`-Element innerhalb des `html`-Elements haben.

```js
assert(document.querySelectorAll('body').length === 1);
```

Dein `head`-Element sollte vor deinem `body`-Element stehen.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
