---
id: 63507bcbfe4ede356e624395
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

입력값을 얻었으니, 이제 그것을 숫자 배열로 나누어야 합니다. `.split()` 메서드를 사용하세요.

`.split()` 메서드는 문자열을 받아 문자열 배열로 나눕니다. 구분자로 사용할 문자열이나 정규식을 전달할 수 있습니다. 예를 들어, `string.split(",")`는 문자열을 각 쉼표에서 나누어 문자열 배열을 반환합니다.

`/,\s*/g` 정규식을 사용해 `value` 문자열을 쉼표로 나누세요. 값 사이에 있는 공백 수에 따라 정규식을 조정할 수 있습니다. 배열을 `array` 변수에 저장하세요.

# --hints--

`calculate` 함수에는 `array` 변수가 있어야 합니다.

```js
assert.match(code.toString(), /array/);
```

`calculate` 함수는 `.split()` 변수에서 `value` 메서드를 호출해야 합니다.

```js
assert.match(code.toString(), /value\.split()/);
```

`calculate` 함수는 `.split()` 메서드의 결과를 `array` 변수에 할당해야 합니다.

```js
assert.match(code.toString(),  /array\s*=\s*value\.split()/);
```

`/,\s*/g`를 `split()` 메서드의 구분자로 사용하세요.

```js
assert.match(code.toString(), /value\.split\(\s*\/,\s*\\s*\*\s*\/g\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
}
--fcc-editable-region--
```
