---
id: 63508c898d753754757bd5e3
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

기본적으로 `.sort()` 메서드는 배열의 요소를 문자열로 변환한 다음 알파벳순으로 정렬합니다. `.sort()` 메서드는 원본 배열을 변경합니다. 이 방법은 문자열에는 잘 작동하지만 숫자에는 잘 맞지 않습니다. 예를 들어, 문자열로 정렬할 때는 `10`가 `2`보다 앞서지만, 숫자로 정렬할 때는 `2`가 `10`보다 앞섭니다.

이 문제를 해결하려면 `.sort()` 메서드에 콜백 함수를 전달할 수 있습니다. 이 함수는 비교되는 두 요소를 나타내는 두 개의 인수를 받습니다. 첫 번째 요소가 두 번째 요소보다 앞서야 하면 `0`보다 작은 값을 반환하고, 첫 번째 요소가 두 번째 요소보다 뒤에 와야 하면 `0`보다 큰 값을 반환하며, 두 요소가 현재 위치를 유지해야 하면 `0`를 반환해야 합니다.

숫자를 작은 것부터 큰 것 순으로 정렬하려면 `a`과 `b` 매개변수를 받고 `b`에서 `a`를 뺀 결과를 반환하는 콜백 함수를 전달하세요.

# --hints--

`sort` 메서드는 `a`와 `b` 매개변수를 받는 콜백 함수를 가져야 합니다.

```js
assert.match(getMedian.toString(), /(array\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>|array\.sort\(\s*function\s*\(\s*a\s*,\s*b\s*\)\s*\{)/)
```

`sort` 메서드는 `b`에서 `a`를 뺀 결과를 반환해야 합니다.

```js
assert.match(getMedian.toString(), /(array\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>|array\.sort\(\s*function\s*\(\s*a\s*,\s*b\s*\)\s*\{)\s*(return\s+)?\s*a\s*\-\s*b/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

--fcc-editable-region--
const getMedian = (array) => {
  const sorted = array.sort();
}
--fcc-editable-region--

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
  
  const mean = getMean(numbers);

  document.querySelector("#mean").textContent = mean;
}
```
