---
id: 635089e3bd3e144f2db4094f
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

숫자 목록으로 폼을 테스트하면 페이지에 평균이 표시되는 것을 볼 수 있습니다. 하지만 이것은 freeCodeCamp의 iframe에 특별한 설정이 있기 때문에 가능한 일입니다. 일반적으로 폼이 제출되면 이벤트가 페이지 새로 고침을 발생시킵니다.

이 문제를 해결하려면 `return false;` 속성에서 `calculate();` 호출 뒤에 `onsubmit`를 추가하세요.

# --hints--

`onsubmit` 속성에 `return false;` 문장이 있어야 합니다.

```js
assert.match(document.querySelector("form").getAttribute("onsubmit"), /return\s+false\s*;?/);
```

`onsubmit` 속성은 여전히 `calculate()`를 호출해야 합니다.

```js
assert.match(document.querySelector("form").getAttribute("onsubmit"), /calculate\(\s*\)/);
```

`false`를 호출한 후에는 `calculate()`를 반환해야 합니다.

```js
assert.match(document.querySelector("form").getAttribute("onsubmit"), /calculate\(\s*\)\s*;?\s*return\s+false\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
--fcc-editable-region--
    <form onsubmit="calculate();">
--fcc-editable-region--
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
  
  const mean = getMean(numbers);

  document.querySelector("#mean").textContent = mean;
}
```
