---
id: 6350805fe0fe283dd347b0dc
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

<dfn>평균</dfn>은 목록에 있는 모든 숫자의 평균값입니다. 평균을 계산하는 첫 번째 단계는 목록에 있는 모든 숫자의 합을 구하는 것입니다. 배열에는 이 상황에 완벽한 `.reduce()`라는 메서드가 있습니다. `.reduce()` 메서드는 배열을 받아 콜백 함수를 적용해 배열을 하나의 값으로 압축합니다.

`sum` 변수를 선언하고 `array.reduce()`를 할당하세요.

# --hints--

`getMean` 함수에는 `sum` 변수가 있어야 합니다.

```js
assert.match(getMean.toString(), /sum/);
```

`getMean` 함수는 `.reduce()` 매개변수의 `array` 메서드를 사용해야 합니다.

```js
assert.match(getMean.toString(), /array\.reduce\(\s*\)/);
```

`array.reduce()`의 결과를 `sum` 변수에 할당해야 합니다.

```js
assert.match(getMean.toString(), /sum\s*=\s*array\.reduce\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const getMean = (array) => {

}
--fcc-editable-region--

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
}
```
