---
id: 63507d810f1a2e38f1908fd8
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

사용자가 입력 상자에 어떤 텍스트든 입력할 수 있습니다. 숫자만 다루고 있는지 확인해야 합니다. `Number()` 생성자는 전달된 값이 숫자로 변환할 수 없으면 `NaN`("숫자가 아님")을 반환합니다.

이 값을 걸러내야 합니다. 다행히 배열에는 이를 위한 메서드가 있습니다. `.filter()` 메서드는 배열에서 요소를 걸러내 새 배열을 만듭니다.

`filtered` 변수를 선언하고 `numbers.filter()`를 할당하세요.

# --hints--

`calculate` 함수에는 `filtered` 변수가 있어야 합니다.

```js
assert.match(calculate.toString(), /filtered/);
```

`calculate` 함수는 `.filter()` 배열에서 `numbers` 메서드를 사용해야 합니다.

```js
assert.match(calculate.toString(), /numbers\.filter\(\s*\)/)
```

`numbers.filter()` 변수에 `filtered` 결과를 할당해야 합니다.

```js
assert.match(calculate.toString(), /filtered\s*=\s*numbers\.filter\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el));
}
--fcc-editable-region--
```
