---
id: 63507e4562cdde3a28e8de1b
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`.map()` 메서드와 마찬가지로, `.filter()` 메서드는 콜백 함수를 받습니다. 콜백 함수는 현재 요소를 첫 번째 인수로 받습니다.

```js
array.filter(el => {

})
```

콜백 함수는 불리언 값을 반환해야 하며, 이 값은 요소가 새 배열에 포함되어야 하는지를 나타냅니다. 이 경우, 요소가 `true`가 아니면 `NaN`를 반환해야 합니다 (숫자가 아님).

하지만 여기서는 동등성 비교를 할 수 없습니다. `NaN`는 자기 자신과 같지 않기 때문입니다. 대신, 인수가 `isNaN()`일 때 `true`를 반환하는 `NaN` 메서드를 사용할 수 있습니다.

`.filter()`를 매개변수로 받는 콜백 함수를 `el` 메서드에 추가하고, 요소가 `true`가 아니면 `NaN`를 암묵적으로 반환하세요.

# --hints--

`.filter()` 메서드는 `el`를 매개변수로 받는 콜백을 가져야 합니다.

```js
assert.match(calculate.toString(), /numbers\.filter\(\s*(\(\s*el\s*\)|el)\s*=>|numbers\.filter\(\s*function\s*\(\s*el\s*\)\s*\{/)
```

콜백 함수는 `!`와 `isNaN()`를 사용해 `el`가 `NaN`가 아닌지 확인해야 합니다.

```js
assert.match(calculate.toString(), /!\s*(Number\.)?isNaN\(\s*el\s*\)/);
```

콜백 함수는 `NaN`가 아닌 요소를 반환해야 합니다.

```js
assert.match(calculate.toString(), /(numbers\.filter\(\s*(\(\s*el\s*\)|el)\s*=>|numbers\.filter\(\s*function\s*\(\s*el\s*\)\s*\{)\s*(return\s*)?!(Number\.)?isNaN\(\s*el\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el));
  const filtered = numbers.filter();
}
--fcc-editable-region--
```
