---
id: 635060a5c03c950f46174cb5
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

통계학은 데이터를 이해하기 위해 수학을 활용하는 방법입니다. 정보를 통해 패턴과 추세를 파악하여 그 정보를 바탕으로 예측과 결정을 내릴 수 있게 도와줍니다.

이번 도전 과제에서는 숫자 집합을 입력받아 평균, 중앙값, 최빈값, 표준편차, 분산을 반환하는 통계 계산기를 만들게 됩니다.

HTML과 CSS는 이미 제공되어 있습니다. 코드를 자유롭게 살펴보세요 — 폼이 제출될 때 `calculate` 함수가 호출되는 것을 확인할 수 있을 겁니다. 준비가 되면 `calculate` 파일에서 빈 함수로 `script.js` 변수를 선언하고 할당하세요.

# --hints--

변수를 선언할 때 `const` 키워드를 사용하세요 `calculate`.

```js
assert.match(code, /const\s+calculate/);
```

`calculate` 변수는 함수여야 합니다.

```js
assert.isFunction(calculate);
```

`calculate` 함수는 비어 있어야 합니다.

```js
assert.match(calculate.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
