---
id: 65ca2d0625aa3a3201067f70
title: 34단계
challengeType: 0
dashedName: step-34
---

# --description--

`array.forEach()` 콜백 함수 안에서 현재 요소가 `counts` 객체 안에 있는지 확인하세요. 요소를 찾을 수 있으면 `counts[el]`만큼 `1` 값을 증가시키세요. 그렇지 않으면 숫자 `1`를 `counts[el]`에 할당하세요.

return 문을 `counts` 대신 `array`을 반환하도록 변경하세요.

이를 테스트하려면 숫자 `4, 4, 2, 5`을 입력하고 `Calculate`를 클릭하세요. 콘솔에 다음이 표시되어야 합니다:

```js
{ '2': 1, '4': 2, '5': 1 }
```

# --hints--

return 문을 `counts` 대신 `array`을 반환하도록 변경해야 합니다.

```js
assert.match(getMode.toString(), /return\s+counts\s*;/);
```

현재 요소가 `counts` 객체 안에 있는지 확인해야 합니다. 찾을 수 있으면 `counts[el]`만큼 `1` 값을 업데이트하세요. 그렇지 않으면 숫자 `1`를 `counts[el]`에 할당하세요. 테스트하려면 숫자 `4, 4, 2, 5`을 입력하고 `"Calculate"`를 클릭하세요.

```js
const expected = {'2': 1, '4': 2, '5': 1};
assert.deepEqual(getMode([4, 4, 2, 5]), expected)
```




# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

const getMedian = (array) => {
  const sorted = array.toSorted((a, b) => a - b);
  const median =
    sorted.length % 2 === 0
      ? getMean([sorted[sorted.length / 2], sorted[sorted.length / 2 - 1]])
      : sorted[Math.floor(sorted.length / 2)];
  return median;
}

--fcc-editable-region--
const getMode = (array) => {
  const counts = {};
  array.forEach(el => {

  })
  return array;
}
--fcc-editable-region--

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
  
  const mean = getMean(numbers);
  const median = getMedian(numbers);
  console.log(getMode(numbers));

  document.querySelector("#mean").textContent = mean;
  document.querySelector("#median").textContent = median;
}


```
