---
id: a5229172f011153519423690
title: Build an Odd Fibonacci Sum Calculator
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

이 실습에서는 주어진 양의 정수 이하인 모든 홀수 피보나치 수의 합을 계산하는 홀수 피보나치 합 계산기를 만듭니다.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**User Stories:**

1. 숫자를 인수로 받는 `sumFibs` 함수를 만들어야 합니다.
1. `sumFibs` 함수는 주어진 숫자 이하인 모든 홀수 피보나치 수의 합을 반환해야 합니다.
1. 피보나치 수열은 `0`와 `1`로 시작하며, 이후 각 숫자는 이전 두 숫자의 합입니다.
1. 합계에는 홀수인 피보나치 수만 더해야 합니다.

# --hints--

`sumFibs` 함수를 만들어야 합니다.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` should return `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` should return `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` should return `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` should return `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` should return `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
