---
id: 66ec4c8e9878d8441956516f
title: 本のカタログテーブルを作成する
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. 本の情報を一覧表示する`table`要素を作成してください。
1. テーブルには1行のテーブルヘッド要素が必要です。
1. テーブルヘッドの行には4つのテーブルヘッダー要素があり、それぞれ`Title`、`Author`、`Genre`、`Publication Year`のテキストが順に入っている必要があります。
1. テーブルには少なくとも5行のテーブルボディ要素が必要です。
1. テーブルボディの各行には、本のタイトル、著者、ジャンル、出版年を表示する4つのテーブルデータ要素が必要です。
1. テーブルには1行のテーブルフッター要素が必要です。
1. テーブルフッターの行には4列をまたぐテーブルデータ要素があり、テキストは`Total Books: N`で、`N`はテーブル内の本の数に置き換えられている必要があります。

# --hints--

`table`要素は1つだけ必要です。

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

`thead`要素は`table`要素の中に1つだけ必要です。

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

`tr`要素は`thead`要素の中に1つだけ必要です。

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

`th`要素の行の中に4つの`thead`要素が必要です。

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

4つの`th`要素には、順に`Title`、`Author`、`Genre`、`Publication Year`のテキストが入っている必要があります。

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent.trim(), 'Title');
assert.equal(ths[1]?.textContent.trim(), 'Author');
assert.equal(ths[2]?.textContent.trim(), 'Genre');
assert.equal(ths[3]?.textContent.trim(), 'Publication Year');
```

`tbody`要素の中に`table`要素が1つ必要です。

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

`tbody`要素には少なくとも5行必要です。

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

`tbody`要素の各行には、正確に4つの`td`要素が子要素として必要です。

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

テーブルボディの各`td`要素には本の情報のテキストが入っている必要があります。

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.trim().length, 1);
});
```

`tfoot`要素の中に`table`要素が1つ必要です。

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

`tr`要素の中に正確に1つの`tfoot`要素が必要です。

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

`td`要素の行の中の`tfoot`要素には、`colspan`属性が`4`に設定されている必要があります。

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

`td`要素の行の中の`tfoot`要素には、テキスト`Total Books: N`が入り、`N`はテーブル内の本の数です。

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
