---
id: 662a6bc12cde72c32fb526f0
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

<dfn>인터페이스</dfn>는 클래스의 청사진과 같습니다. 인터페이스는 클래스가 구현해야 하는 메서드와 프로프 집합을 포함합니다.

이 프로젝트는 `Equation`라는 빈 클래스를 선언하는 것으로 시작하세요. 이 클래스를 사용해 일반 방정식의 청사진인 인터페이스를 정의할 것입니다.

# --hints--

`Equation`라는 새 클래스를 정의하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
