---
id: 662bd456896f16d9bd03f1a6
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

`Equation` 클래스 내에 `solve`과 `analyze`라는 이름의 두 개의 새로운 인스턴스 메서드를 정의하세요.

# --hints--

`solve` 클래스 내에 `Equation`라는 이름의 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

`solve` 메서드는 `self`라는 하나의 매개변수를 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

`analyze` 클래스 내에 `Equation`라는 이름의 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

`analyze` 메서드는 `self`라는 하나의 매개변수를 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
