---
id: 662bde88dc84f1e249801b1a
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

`ABC`는 *추상 기본 클래스*를 의미합니다. `ABC` 클래스는 일반 클래스를 구체 클래스의 청사진 역할을 하는 추상 클래스로 바꿀 수 있게 해줍니다.

`import` 문장을 수정하여 `ABC` 모듈에서 `abc` 클래스만 가져오세요. 모듈 `x`에서 특정 객체 `y`를 가져오려면 `from y import x` 구문을 따르세요.

그다음 `Equation` 클래스를 `ABC`를 상속받아 추상 클래스로 만드세요.

# --hints--

`ABC`를 `abc` 모듈에서 가져와야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

`Equation` 클래스는 `ABC`를 상속받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
