---
id: 662bdd364bf2cde1487922a9
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

다른 프로그래밍 언어와 달리 Python은 핵심 언어에서 인터페이스를 구현하지 않지만, Python 표준 라이브러리를 사용하면 인터페이스를 간단하게 정의할 수 있습니다.

이 프로젝트에서는 `abc` 모듈의 유틸리티를 사용할 것입니다. 따라서 이 모듈을 코드에 가져오세요.

# --hints--

`abc` 모듈을 가져와야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
