---
id: 662bd552e1c1d2db1b88ba47
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

이제 `LinearEquation`라는 이름의 또 다른 클래스를 정의하고 `Equation` 클래스로부터 상속받게 하세요. 이 클래스를 선형 방정식을 나타내는 데 사용할 것입니다.

# --hints--

`LinearEquation`라는 이름의 클래스를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

`LinearEquation` 클래스는 `Equation` 클래스로부터 상속받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
