---
id: 662bd8260da84bdd5feae419
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

`LinearEquation` 클래스가 `Equation` 클래스 안에 정의된 모든 메서드를 단순히 상속받는 것이 아니라 직접 구현하도록 하세요. `Equation` 클래스는 인터페이스 역할을 해야 합니다.

현재 `Equation` 클래스는 단순히 `LinearEquation` 클래스의 부모 클래스입니다. 다음 단계에서 이를 정식 인터페이스로 만드는 방법을 배우게 됩니다.

지금은 `Equation`의 인스턴스를 생성하여 변수 `eq`에 할당하고, `LinearEquation`의 인스턴스를 생성하여 변수 `lin_eq`에 할당하세요.

# --hints--

변수 `eq`을 선언하고 `Equation`의 인스턴스를 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

변수 `lin_eq`을 선언하고 `LinearEquation`의 인스턴스를 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
