---
id: 662fbcef5f05e1b84f541a0c
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

각 방정식 객체는 방정식의 계수만큼 인수를 전달받아 인스턴스화되며, n차 항부터 0차 항까지, `0` 값을 가질 수 있는 계수를 포함합니다.

예를 들어, `LinearEquation(4, 5)`은 방정식 \\( 4x + 5 = 0 \\)을 나타내며, `4`는 첫 번째(여기서는 최고차) 항의 계수이고 `5`는 0차 항의 계수입니다.

방정식 객체를 인스턴스화할 때 올바른 수의 인수가 전달되었는지 확인해야 합니다.

`__init__` 메서드 안에서 `if` 문을 만들어 `args`의 길이가 방정식이 가져야 할 계수의 수(`degree + 1`)와 다른지 확인하세요. 다르면 `TypeError`를 발생시키고 다음 문자열을 사용해 사용자 정의 메시지를 제공하세요: `f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`.

그다음, `2`과 `3`을 전달해 `lin_eq`를 인스턴스화하여 오류를 수정하세요.

# --hints--

방정식을 인스턴스화하는 데 사용된 계수의 수가 `if`와 다른지 확인하는 `degree + 1` 문을 만들어야 합니다.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

새로운 `TypeError` 문 안에서 `if`를 발생시키고 제공된 문자열을 사용해 사용자 정의 오류 메시지를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

`2`과 `3`을 전달해 `lin_eq`를 인스턴스화해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
