---
id: 662f6d7c92381a3049e4c987
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

인터페이스는 추상 메서드만 정의할 필요가 없으며, 구체 클래스가 상속받을 메서드를 구현할 수도 있습니다.

`solve` 클래스 내에서 `analyze`와 `Equation`의 실제 구현을 처리하기 전에, `__init__` 메서드를 정의하세요. 이 메서드에는 어떤 데코레이터도 사용하지 마세요.

# --hints--

`__init__` 클래스에 `Equation` 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

`__init__` 메서드는 `self`이라는 매개변수 하나를 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
