---
id: 662fc3eba556a6bf800d48c1
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

`isinstance()` 내장 함수는 두 개의 인수를 받아 첫 번째 인수로 전달된 객체가 두 번째 인수로 전달된 클래스의 인스턴스인지 여부를 나타내는 불리언을 반환합니다.

```py
isinstance(7, int) # True
```

또 확인해야 할 사항은 모든 인수가 숫자인지입니다. 첫 번째 `if` 이후에 `for` 반복문을 만들어 `args`을 반복하며 현재 반복 중인 인수가 `int` 또는 `float`의 인스턴스가 아닌지 확인하세요. `isinstance()` 함수를 사용하고 두 번째 인수로 `int`과 `float`을 포함하는 튜플을 전달하세요.

인수가 숫자가 아니면 `TypeError`를 발생시키고 `"Coefficients must be of type 'int' or 'float'"`을 메시지로 사용하세요.

# --hints--

`for` 문장 다음에 `args`을 반복하는 `if` 반복문을 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

`if` 반복문 내에서 현재 계수가 `int` 또는 `float`의 인스턴스가 아닌지 확인하는 `for` 문장을 만들어야 합니다.

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

제공된 문자열을 사용해 `TypeError` 문장 내에서 `if`를 발생시켜야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
