---
id: 662f96576ef178927de87975
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

추상 메서드로 인식되려면, 메서드는 `@abstractmethod` 데코레이터로 장식되어야 합니다.

<dfn>데코레이터</dfn>는 Python에서 함수의 동작을 수정하는 데 사용됩니다. `spam`라는 데코레이터를 사용하는 예는 다음과 같습니다:

```py
@spam
def foo():
    pass
```

가져오기 문을 수정하여 `abstractmethod` 데코레이터를 가져오고 `solve` 클래스의 `analyze`와 `Equation` 메서드 모두를 장식하세요. 이렇게 하면 두 개의 예외가 발생합니다.

`ABC`를 상속하는 클래스가 추상 메서드를 가지면, 해당 클래스는 더 이상 인스턴스화할 수 없습니다. 따라서 오류를 없애려면 `Equation` 인스턴스를 삭제하세요.

다른 오류는 `LinearEquation` 클래스가 인터페이스에 정의된 모든 추상 메서드를 구현해야 하기 때문에 발생합니다. `LinearEquation` 클래스 안에도 반드시 정의해야 합니다. 구체 클래스에서는 `abstractmethod` 데코레이터를 사용하면 안 됩니다.

# --hints--

`abstractmethod`를 `abc` 모듈에서 가져와야 합니다.

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

`@abstractmethod` 클래스 내의 `solve` 메서드를 `Equation`으로 장식해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

`@abstractmethod` 클래스 내의 `analyze` 메서드를 `Equation`으로 장식해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

`solve` 클래스 내에 `LinearEquation`라는 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

`solve` 메서드는 `self`라는 매개변수를 하나 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

`analyze` 클래스 내에 `LinearEquation`라는 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

`solve` 메서드는 `self`라는 매개변수를 하나 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
