---
id: 662fa2e2cf27c09f21f4f5d0
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

Python에서는 데이터 타입이 실행 시점에 인식됩니다. 따라서 변수를 선언할 때 데이터 타입을 명시할 필요가 없습니다. 그럼에도 불구하고, `variable: <data type> = value` 또는 단순히 `variable: <data type>`를 사용해 변수가 특정 데이터 타입을 가질 것임을 명확히 하기 위해 주석을 달 수 있습니다. Python 인터프리터는 변수에 주석으로 표시된 타입을 강제하지 않으며, 보통 이를 위해 외부 도구가 필요하다는 점을 참고하세요.

`Equation` 클래스 내부에 클래스 속성 `degree`를 정의하세요. 값을 할당하지 마세요. 대신 `int` 타입 주석을 사용해 구체적인 클래스들이 정수형 숫자를 저장할 것임을 나타내세요.

나중에 이 클래스 속성은 방정식 객체를 인스턴스화할 때 전달되는 인수의 유효성 검사 과정의 일부로 사용하게 됩니다.

# --hints--

`degree` 클래스 내에 `int`라는 이름의 클래스 속성을 정의하고 `Equation`로 주석을 달아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
