---
id: 6639fdcc701833a54c364211
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

계수를 검증한 후에는 인스턴스 속성에 저장해야 합니다. 딕셔너리 내포를 사용해 키는 계수의 차수이고 값은 해당 계수인 딕셔너리를 만든 다음 `coefficients`라는 속성에 할당하세요.

예를 들어, `LinearEquation`와 `2`으로 생성된 `4` 객체는 다음과 같은 `coefficients` 속성을 가져야 합니다: `{1: 2, 0: 4}`. 이는 `2`가 `x`의 1차에 해당하고 `4`가 `x`의 0차에 해당하기 때문입니다.

새 딕셔너리의 키-값 쌍은 `args`에 저장된 순서와 동일하게 만드세요.

# --hints--

`coefficients` 메서드 내에 `__init__`라는 속성을 선언해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

계수를 저장할 때 딕셔너리 내포를 사용하세요.

```js
({ test: () => runPython(`
import ast
node = _Node(_code).find_class("Equation").find_function("__init__").find_variable("self.coefficients")
assert isinstance(node.tree.value, ast.DictComp)
`) })
```

`coefficients` 속성은 차수-계수 형태의 키-값 쌍을 포함하는 딕셔너리여야 합니다. 계수가 `args`에 저장된 순서를 반드시 따르세요.

```js
({ test: () => runPython(`
actual1 = list(LinearEquation(1, 6).coefficients.items())
expected1 = list({1: 1, 0: 6}.items())
actual2 = list(LinearEquation(-3.5, 0).coefficients.items())
expected2 = list({1: -3.5, 0: 0}.items())
assert actual1 == expected1
assert actual2 == expected2
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass

    
lin_eq = LinearEquation(2, 3)
```
