---
id: 6639f947d3a1818c9322c64a
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

계수를 검증하는 마지막 단계는 최고차항 계수가 0이 아닌지 확인하는 것입니다. 최고차항 계수는 객체를 생성할 때 첫 번째 인수로 전달해야 한다는 점을 기억하세요.

이를 위해 `if` 문을 추가하고, 다음 문자열을 사용해 사용자 지정 메시지를 제공하는 `ValueError`를 발생시키세요: `'Highest degree coefficient must be different from zero'`.

# --hints--

방정식을 생성할 때 전달된 첫 번째 계수가 0인지 확인하는 `if` 문을 작성해야 합니다.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]") or cond.is_equivalent("not args[0]")
`)) })
```

새로운 `ValueError` 문 안에서 `if`를 발생시키고, 제공된 문자열을 사용해 사용자 지정 오류 메시지를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_bodies()[0].has_stmt("raise ValueError('Highest degree coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
