---
id: 663a32735b317af9812eb0d7
title: 27단계
challengeType: 20
dashedName: step-27
---

# --description--

\\( ax + b = 0 \\) 형태의 일차방정식에서 기울기는 단순히 계수 \\( a \\)이고, y절편은 계수 \\( b \\)입니다.

<img alt="a plot of a linear function" src="https://cdn.freecodecamp.org/curriculum/python/linear-equation.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

`analyze` 메서드를 사용해 방정식에 대한 추가 정보를 제공하세요. `analyze` 메서드 안에서 계수를 변수 `slope`와 `intercept`에 언팩하세요.

그다음, 기울기와 y절편 값을 각각 `'slope'`와 `'intercept'` 키를 가진 딕셔너리로 반환하세요. 이후 `analyze`에 대해 `lin_eq`를 호출하고 결과를 출력하세요.


# --hints--

`coefficients` 메서드 안에서 `slope` 속성에 저장된 값을 변수 `intercept`와 `analyze`에 언팩해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_stmt("slope, intercept = self.coefficients.values()")`)) })
```

`analyze` 메서드는 기울기와 y절편 값을 각각 `'slope'`와 `'intercept'` 키를 가진 딕셔너리로 반환해야 합니다.

```js
({ test: () => runPython(`
eq = LinearEquation(2.2, 1.5)
a = eq.analyze()
assert a['slope'] == 2.2, "Expected different slope"
assert a['intercept'] == 1.5, "Expected different intercept"
`) })
```

`analyze` 객체의 `lin_eq` 메서드를 호출해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(lin_eq.analyze())")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x
--fcc-editable-region--
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)
print(lin_eq.solve())
--fcc-editable-region--
```
