---
id: 663a22ba7420c4d2f7fd2aec
title: 25단계
challengeType: 20
dashedName: step-25
---

# --description--

`solve` 메서드를 구현할 시간입니다. \\( ax + b = 0 \\) 형태의 일차 방정식이 주어졌을 때, 해는 \\(x = -\frac{b}{a}\\)입니다.

`coefficients` 속성에 저장된 계수를 `a`와 `b` 변수로 언팩하세요. 이때 `.values()` 메서드를 사용해야 한다는 점을 기억하세요.

그다음, `x` 변수를 선언하고 방정식의 해를 할당한 뒤 `solve` 메서드에서 반환하세요.

# --hints--

`coefficients` 속성에 저장된 값을 `a`와 `b` 변수로 언팩해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("a, b = self.coefficients.values()")`)) })
```

`x`라는 변수를 선언하고 일차 방정식의 해를 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("x = -b/a")`)) })
```

`x`을 `solve` 메서드에서 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_return("x")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
--fcc-editable-region--
    def solve(self):
        pass
--fcc-editable-region--
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
