---
id: 663b7fefd437bd984e091cbf
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

다음으로, `QuadraticEquation`라는 새 클래스를 만들고 `Equation`에서 상속받게 하세요. 이 새 클래스를 사용해 $ax^2 + bx + c = 0$ 형태인 2차 방정식을 나타낼 것입니다.

새 클래스 안에 2차 방정식의 차수를 나타내는 값 `degree`인 `2` 클래스 속성을 정의하세요. 또한 `solve`와 `analyze` 메서드를 정의하세요. 구현은 다음 단계에서 진행합니다.

# --hints--

`QuadraticEquation`라는 새 클래스를 만들고 `Equation` 클래스를 상속받게 하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").inherits_from("Equation")`)) })
```

`solve` 클래스 안에 `QuadraticEquation` 메서드를 정의하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("solve")`)) })
```

`solve` 메서드는 `self`이라는 단일 매개변수를 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("solve").has_args("self")`)) })
```

`analyze` 클래스 안에 `QuadraticEquation` 메서드를 정의하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("analyze")`)) })
```

`analyze` 메서드는 `self`라는 단일 매개변수를 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("analyze").has_args("self")`)) })
```

`degree` 클래스 안에 `QuadraticEquation` 클래스 속성을 정의하고 값 `2`을 할당하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_variable("degree").is_equivalent("degree = 2")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}
--fcc-editable-region--

--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
