---
id: 663b83a28943e6aa6275a514
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

여전히 `Equation` 클래스 안에서, 생성할 방정식 객체에 적절한 문자열 표현을 제공하는 `__str__` 메서드를 정의하세요.

지금은 `__str__` 메서드 안에서, 변수 `terms`을 선언하고 빈 리스트를 할당하세요. 이 변수를 사용해 방정식의 각 항(계수 곱하기 \\( x^n \\))을 저장할 것입니다.

그다음, 변수 `equation_string`를 선언하고 `terms` 리스트의 요소들을 공백으로 연결한 결과를 할당하세요. 마지막으로 `equation_string`를 반환하세요.

# --hints--

`__str__` 클래스 안에 `Equation` 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

`__str__` 메서드는 하나의 매개변수 `self`를 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

`terms` 메서드 안에서 변수 `__str__`을 선언하고 빈 리스트를 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

`equation_string` 메서드 안에서 `terms`의 요소들을 공백으로 연결한 결과를 변수 `__str__`에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

`equation_string` 메서드에서 `__str__`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
