---
id: 663c9f31306353460da54542
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

보시다시피 출력에 `+` 기호가 없습니다. 숫자 기호는 음수일 때만 기본적으로 표시됩니다. 이 동작을 변경하려면, f-string의 중괄호 안에서 평가할 표현식 뒤에 콜론을 쓰고 `+` 옵션을 지정하세요. 그러면 양수와 음수 모두에 대해 기호를 표시할 수 있습니다.

두 조건절의 문자열을 수정하여 `:+` 뒤 중괄호 안에 `coefficient`를 추가하세요.

# --hints--

`terms` 문장 내 `if` 리스트에 추가하는 문자열을 `f'{coefficient:+}'`로 수정하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[0].is_equivalent("terms.append(f'{coefficient:+}')")`)) })
```

`terms` 절 내 `elif` 리스트에 삽입하는 문자열을 `f'{coefficient:+}x'`로 수정하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[1].is_equivalent("terms.append(f'{coefficient:+}x')")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient}')
            elif n == 1:
                terms.append(f'{coefficient}x')                
        equation_string = ' '.join(terms)        
        return equation_string
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
