---
id: 664e4a590b52ba8d2adff19f
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

일반적으로 대문자 그리스 문자 델타로 표시되는 \\( ax^2 + bx + c = 0 \\) 형태의 이차방정식 판별식은 \\( Δ = b^2 - 4ac \\)와 같습니다.

`QuadraticEquation` 클래스 내에 `__init__` 메서드를 정의하세요. `super()`를 사용해 부모 클래스의 `__init__` 메서드를 호출하세요. 그런 다음 방정식 판별식 값을 저장하는 `delta`라는 새 속성을 정의하세요.

# --hints--

`__init__` 클래스 내에 `QuadraticEquation` 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("__init__")`)) })
```

`__init__` 메서드는 `self`와 `*args` 두 개의 매개변수를 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_args("self, *args")`)) })
```

`super().__init__(*args)` 메서드 내에서 `__init__`를 호출해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_call("super().__init__(*args)")`)) })
```

`delta` 메서드 내에 `__init__` 속성을 선언하고 방정식 판별식 값을 할당해야 합니다.

```js
({ test: () => runPython(`
eq = QuadraticEquation(2, -3, -4)
assert eq.delta == 41
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2
--fcc-editable-region--
    
--fcc-editable-region--    
    def solve(self):
        pass
    
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
