---
id: 664e4e1b6c35a99cbba49e84
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

이제 `QuadraticEquation` 클래스로 방정식 \\( 11x^2 - x + 1 = 0 \\)을 나타내는 인스턴스를 만드세요.

새 인스턴스를 `quadr_eq` 변수에 할당한 다음, 새 변수를 출력하세요. 이 시점에서는 방정식의 문자열 표현에서 2차 항이 빠져 있을 것입니다.

# --hints--

`quadr_eq`라는 변수를 선언하고 `QuadraticEquation`의 인스턴스를 생성하여 `11`, `-1`, `1`을 인수로 전달해 할당하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("quadr_eq = QuadraticEquation(11, -1, 1)")`)) })
```

`quadr_eq` 변수를 출력하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(quadr_eq)")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2

    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c
   
    def solve(self):
        pass
    
    def analyze(self):
        pass
--fcc-editable-region-- 
lin_eq = LinearEquation(2, 3)
print(lin_eq)

--fcc-editable-region--
```
