---
id: 664eef158d792a509e8d708a
title: 34단계
challengeType: 20
dashedName: step-34
---

# --description--

`sub` 모듈의 `re` 함수는 정규식 패턴을 기반으로 문자열 안의 텍스트를 교체할 수 있게 해줍니다.

```py
verse = 'Always look on the bright side of life'
spam = re.sub('bright', 'spam', verse)
spam == 'Always look on the spam side of life' # True
```

이 함수는 세 개의 인수를 받습니다: 일치시킬 정규식 패턴, 교체할 문자열, 그리고 교체를 수행할 문자열입니다.

`__str__` 메서드에서 `sub()` 호출을 반환하세요. 이때 인수로 문자열 `'1'`, 빈 문자열, 그리고 기존 `equation_string.strip('+')` 호출을 전달하세요. 이렇게 하면 각 `1`가 빈 문자열로 교체됩니다. 결과는 아직 완성되지 않았으며 다음 단계에서 정규식 패턴을 계속 다룰 것입니다.

# --hints--

`re.sub()` 메서드에서 `__str__` 호출을 반환해야 합니다. `'1'`에 인수로 문자열 `equation_string.strip('+')`, 빈 문자열, 그리고 기존 `re.sub()` 호출을 전달하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("re.sub('1', '', equation_string.strip('+'))")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re


class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
--fcc-editable-region--
        return equation_string.strip('+')        
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2

    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c
   
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(11, -1, 1)
print(quadr_eq)

```
