---
id: 664ef4623946e65e18d59764
title: 35단계
challengeType: 20
dashedName: step-35
---

# --description--

정규식 패턴에서 *lookaround*는 문자열에서 문자를 소비하지 않고 특정 패턴과 일치하는 어서션입니다. lookaround의 한 종류인 lookbehind는 긍정형 또는 부정형일 수 있습니다. 각각 `(?<=...)`와 `(?<!...)`로 표시합니다.

```py
spam = 'black back bat'
re.sub('(?<=l)a', 'o', spam) == 'block back bat' # True
re.sub('(?<!l)a', 'o', spam) == 'black bock bot' # True
```

위 예시에서, 패턴 `(?<=l)a`는 긍정형 lookbehind를 포함하며, `a`가 앞에 있을 때만 `l` 문자와 일치하도록 사용됩니다. 예시의 마지막 줄에서, 패턴 `(?<!l)a`는 부정형 lookbehind를 포함하며, `a`가 앞에 **없을 때만** `l` 문자와 일치하도록 사용됩니다. 두 경우 모두 lookbehind에 포함된 문자가 소비되지 않는 점에 주목하세요.

부정형 lookbehind를 사용하도록 정규식 패턴을 수정하여, 숫자가 앞에 없을 때만 `1` 문자가 대체되도록 하세요.

# --hints--

부정형 lookbehind를 사용하여 `re.sub()` 호출의 첫 번째 인수로 전달된 패턴을 수정하세요.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("Equation").find_function("__str__")
values = [
  "re.sub('(?<!\\d)1', '', equation_string.strip('+'))",
  "re.sub(r'(?<!\\d)1', '', equation_string.strip('+'))",
  "re.sub('(?<![0-9])1', '', equation_string.strip('+'))"
]
any(node.has_return(value) for value in values)
`)) })
```



# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re


class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
--fcc-editable-region--
        return re.sub('1', '', equation_string.strip('+'))        
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2

    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c
   
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(11, -1, 1)
print(quadr_eq)

```
