---
id: 664f0389424a6f7aa15fd3e5
title: 36단계
challengeType: 20
dashedName: step-36
---

# --description--

또 다른 종류의 주변 어서션은 전방 탐색입니다. 긍정 전방 탐색과 부정 전방 탐색은 각각 `(?=...)`와 `(?!...)`로 표시합니다. 이들은 특정 수열이 뒤따를 때 패턴과 일치시키지만, 그 수열은 소비하지 않습니다.

```py
spam = 'black back bat'
re.sub('a(?=t)', 'o', spam) == 'black back bot' # True
re.sub('a(?!t)', 'o', spam) == 'block bock bat' # True
```

위 예시에서, 패턴 `a(?=t)`는 긍정 전방 탐색을 포함하며, `a`가 뒤따를 때만 `t` 문자와 일치시킵니다. 예시의 마지막 줄에서, 패턴 `a(?!t)`는 부정 전방 탐색을 포함하며, `a`가 뒤따르지 않을 때만 `t` 문자와 일치시킵니다. 두 경우 모두 전방 탐색에 포함된 문자는 소비되지 않습니다.

정규식 패턴에 긍정 전방 탐색을 추가하여 `1` 문자가 `x` 문자 뒤에 있을 때만 대체되도록 하세요.

# --hints--

정규식 패턴을 수정하여 `1` 문자가 `x` 뒤에 있을 때만 대체되도록 긍정 전방 탐색을 사용하세요. 패턴에서 부정 후방 탐색은 제거하지 마세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("re.sub(r'(?<!\\d)1(?=x)', '', equation_string.strip('+'))")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re


class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
--fcc-editable-region--
        return re.sub(r'(?<!\d)1', '', equation_string.strip('+'))        
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2

    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c
   
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(11, -1, 1)
print(quadr_eq)

```
