---
id: 6650f037c017aa6855a608e3
title: 44단계
challengeType: 20
dashedName: step-44
---

# --description--

모든 이차 방정식의 그래프는 포물선 모양입니다. 포물선 꼭짓점의 \\( x \\) 좌표는 \\( x = - \frac{b}{2a} \\)에서 찾을 수 있습니다.

`analyze` 메서드에서 꼭짓점 \\( x \\)와 \\( y \\) 좌표의 해당 값을 각각 포함하는 두 개의 키 `'x'`과 `'y'`를 가진 사전을 반환하세요.

위의 관계를 사용해 \\( x \\) 좌표를 찾으세요. 그런 다음 \\( x \\) 좌표를 사용해 \\( y \\) 좌표를 계산하세요.

# --hints--

반환할 사전에는 꼭짓점 `'x'`와 `'y'` 좌표의 해당 값을 각각 포함하는 두 개의 키 `x`과 `y`가 있어야 합니다.

```js
({ test: () => runPython(`
eq = QuadraticEquation(16, 2, 1)
assert eq.analyze() == {'x': -0.0625, 'y': 0.9375}
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re


class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r"(?<!\d)1(?=x)", "", equation_string.strip("+"))        

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return [x]

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2

    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c

    def solve(self):
        if self.delta < 0:
            return []
        a, b, _ = self.coefficients.values()
        x1 = (-b + (self.delta) ** 0.5) / (2 * a)
        x2 = (-b - (self.delta) ** 0.5) / (2 * a)
        if self.delta == 0:
            return [x1]

        return [x1, x2]
--fcc-editable-region--
    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(1, 2, 1)
print(quadr_eq)
print(quadr_eq.solve())

```
