---
id: 66759e32b88fb5459b1e0234
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`__init_subclass__` 메서드는 해당 메서드를 정의한 클래스가 서브클래스화될 때마다 호출되며, 자식 클래스를 사용자 정의할 수 있게 해줍니다. 이 메서드는 관례상 `cls`(“class”를 의미)라는 이름의 매개변수를 받으며, 이는 새 자식 클래스를 나타냅니다.

`__init_subclass__` 클래스에 `Equation` 메서드를 정의하고 `cls` 매개변수를 지정하세요.

# --hints--

`__init_subclass__` 클래스에 `cls` 매개변수를 가진 `Equation` 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
