---
id: 6675aaf418b41157f6ccd692
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

이제 `__init__` 메서드로 돌아갈 시간입니다. 방정식 유형에 따라 인스턴스화할 때 가변 개수의 인수를 전달해야 합니다.

메서드에 두 번째 매개변수 `args`를 추가하고 `*` 연산자를 사용해 가변 개수의 인수를 받을 수 있게 만드세요.

# --hints--

`__init__` 메서드는 두 개의 매개변수, `self`와 `*args`를 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
