---
id: 6675a38a8b535e4ff3274520
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

`hasattr` 내장 함수는 첫 번째 인수로 객체를 받고 두 번째 인수로 속성 이름을 나타내는 문자열을 받습니다. 이 함수는 객체가 지정된 속성을 가지고 있는지 여부를 나타내는 불리언을 반환합니다.

이제 `__init_subclass__` 메서드를 사용해 자식 클래스가 인스턴스화되는 순간에 `degree` 속성을 가지고 있는지 확인할 것입니다.

`if` 문을 만들어 `cls`가 `degree` 속성을 가지고 있지 않은지 확인하세요. 그렇다면 `AttributeError`를 발생시키고 문자열 `f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"`을 사용해 사용자 지정 메시지를 제공하세요.

그다음 터미널에 나타난 오류를 고치기 위해 `degree` 클래스 안에 `LinearEquation` 클래스 속성을 선언하세요. 이 속성은 방정식의 차수, 즉 최고 차수 \\( x \\) 항의 지수를 나타내야 합니다. 따라서 정수 `1`을 `degree` 속성에 할당하세요.

# --hints--

`if` 메서드 안에서 `cls`가 `degree` 속성을 가지고 있지 않은지 확인하는 `__init_subclass__` 문을 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_conditions()[0].is_equivalent("not hasattr(cls, 'degree')")`)) })
```

제공된 문자열을 사용해 `AttributeError` 문 안에서 `if`를 발생시켜야 합니다.

```js
({ test: () => runPython(`
raise_stmt = 'raise AttributeError(f"Cannot create \\'{cls.__name__}\\' class: missing required attribute \\'degree\\'")'
node = _Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_bodies()[0]
assert node.has_stmt(raise_stmt) 
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        pass

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    
--fcc-editable-region--
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
