---
id: 667e623208053643ca9d3c6e
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

이제 이전 단계에서 추가한 `for` 반복문과 `if` 문을 `if` 내장 함수를 사용하는 `any()` 문으로 바꾸세요.

# --hints--

새로운 `if` 문의 조건은 `any()` 호출이어야 합니다.

```js
({ test: () => runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[1].find_conditions()[0]
calls = _Node(str(cond)).find_calls("any")
assert len(calls) == 1
`) })
```

`any()` 호출에 인수로 제너레이터 표현식을 전달해야 합니다.

```js
({ test: () => runPython(`
import ast
argument = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[1].find_conditions()[0].find_call_args()[0]
assert isinstance(argument.tree, ast.GeneratorExp)
`) })
```

`any()`에 전달된 제너레이터 표현식은 `args`를 순회해야 합니다.

```js
({ test: () => runPython(`
import ast
argument = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[1].find_conditions()[0].find_call_args()[0]
iters = argument.find_comp_iters()
assert len(iters) == 1
assert iters[0].is_equivalent("args")
`) })
```

`if` 문은 `args`의 인수 중 하나라도 `int` 또는 `float`의 인스턴스가 아닌지 확인해야 합니다.

```js
({ test: () => runPython(`
import ast
argument = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[1].find_conditions()[0].find_call_args()[0]
target = argument.find_comp_targets()[0]
expr = argument.find_comp_expr()
solutions = [
  f"not isinstance({target}, (int, float))",
  f"not isinstance({target}, (float, int))",
  f"not isinstance({target}, float) and not isinstance({target}, int)",
  f"not isinstance({target}, int) and not isinstance({target}, float)",
]
assert any(expr.is_equivalent(sol) for sol in solutions)
`) })
```

새로운 `TypeError` 문 안에서 제공된 문자열을 사용해 `if`을 발생시켜야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[1].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
