---
id: 587d8248367417b2b2512c3c
title: Chiedi ai browser di consultare il tuo sito solo tramite HTTPS con helmet.hsts()
challengeType: 2
forumTopicId: 301573
dashedName: ask-browsers-to-access-your-site-via-https-only-with-helmet-hsts
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

HTTP Strict Transport Security (HSTS) è una politica di sicurezza web che aiuta a proteggere i siti contro attacchi di downgrade del protocollo e furto di cookie. Se il tuo sito può essere consultato tramite HTTPS, puoi chiedere ai browser degli utenti di evitare di usare HTTP non sicuro. Impostando l’intestazione Strict-Transport-Security, comunichi ai browser di usare HTTPS per le richieste future per un periodo di tempo specificato. Questo funzionerà per le richieste successive a quella iniziale.

# --instructions--

Configura `helmet.hsts()` per usare HTTPS per i prossimi 90 giorni. Passa l’oggetto di configurazione `{maxAge: timeInSeconds, force: true}`. Puoi creare una variabile `ninetyDaysInSeconds = 90*24*60*60;` da usare per `timeInSeconds`.

Nota: configurare HTTPS su un sito personalizzato richiede l’acquisizione di un dominio e un certificato SSL/TLS.

# --hints--

Il middleware helmet.hsts() dovrebbe essere montato correttamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hsts');
assert.property(data.headers, 'strict-transport-security');
```

maxAge dovrebbe essere uguale a 7776000 s (90 giorni)

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.match(
  data.headers['strict-transport-security'],
  /^max-age=7776000;?/
);
```
