---
id: 587d8248367417b2b2512c3d
title: Disabilita il prefetching DNS con helmet.dnsPrefetchControl()
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

Come promemoria, questo progetto si basa sul seguente progetto starter clonato da <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Per migliorare le prestazioni, la maggior parte dei browser esegue il prefetching dei record DNS per i collegamenti in una pagina. In questo modo l’IP di destinazione è già noto quando l’utente clicca su un collegamento. Questo può portare a un uso eccessivo del servizio DNS (se possiedi un sito web molto grande, visitato da milioni di persone…), problemi di privacy (un intercettatore potrebbe dedurre che ti trovi su una certa pagina) o alterazioni delle statistiche della pagina (alcuni collegamenti potrebbero apparire visitati anche se non lo sono). Se hai esigenze di sicurezza elevate puoi disabilitare il prefetching DNS, a costo di una penalizzazione delle prestazioni.

# --instructions--

Usa il metodo `helmet.dnsPrefetchControl()` sul tuo server.

# --hints--

Il middleware helmet.dnsPrefetchControl() dovrebbe essere montato correttamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```
