---
id: 587d8249367417b2b2512c3e
title: Disabilita la cache lato client con helmet.noCache()
challengeType: 2
forumTopicId: 301576
dashedName: disable-client-side-caching-with-helmet-nocache
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Se stai rilasciando un aggiornamento per il tuo sito web e vuoi che gli utenti scarichino sempre la versione più recente, puoi provare a disabilitare la cache nel browser del client. Può essere utile anche durante lo sviluppo. La cache offre vantaggi in termini di prestazioni, che perderai, quindi usa questa opzione solo quando è davvero necessaria.

# --instructions--

Usa il metodo `helmet.noCache()` sul tuo server.

# --hints--

Il middleware helmet.noCache() dovrebbe essere montato correttamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nocache');
assert.equal(
  data.headers['cache-control'],
  'no-store, no-cache, must-revalidate, proxy-revalidate'
);
```
