---
id: 587d8247367417b2b2512c36
title: installare e richiamare helmet
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

Per affrontare queste sfide scriverai il tuo codice usando uno dei seguenti metodi:

- Clona <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">questo repository GitHub</a> e completa queste sfide in locale.
- Usa un generatore di siti a tua scelta per completare il progetto. Assicurati di includere tutti i file dal nostro repository GitHub.

Helmet ti aiuta a proteggere le tue app Express impostando varie intestazioni HTTP.

# --instructions--

Tutto il tuo codice per queste lezioni va nel file `myApp.js` tra le righe di codice che ti abbiamo fornito. Non modificare né eliminare il codice che abbiamo aggiunto per te.

La versione `3.21.3` di Helmet è già stata installata, quindi richiamala come `helmet` in `myApp.js`.

# --hints--

La versione `helmet` di `3.21.3` dovrebbe essere in `package.json`

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```
