---
id: 587d8248367417b2b2512c3b
title: Impedisci a IE di aprire HTML non attendibili con helmet.ieNoOpen()
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Alcune applicazioni web forniscono HTML non attendibile da scaricare. Alcune versioni di Internet Explorer per impostazione predefinita aprono quei file HTML nel contesto del tuo sito. Questo significa che una pagina HTML non attendibile potrebbe iniziare a fare cose dannose nel contesto delle tue pagine. Questo middleware imposta l’intestazione X-Download-Options su noopen. Questo impedirà agli utenti di IE di eseguire download nel contesto del sito attendibile.

# --instructions--

Usa il metodo `helmet.ieNoOpen()` sul tuo server.

# --hints--

Il middleware helmet.ieNoOpen() dovrebbe essere montato correttamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```
