---
id: 587d8249367417b2b2512c40
title: "Configura Helmet usando il middleware `helmet()` genitore"
challengeType: 2
forumTopicId: 301575
dashedName: configure-helmet-using-the-parent-helmet-middleware
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`app.use(helmet())` includerà automaticamente tutti i middleware introdotti sopra, tranne `noCache()` e `contentSecurityPolicy()`, ma questi possono essere abilitati se necessario. Puoi anche disabilitare o configurare singolarmente qualsiasi altro middleware, usando un oggetto di configurazione.

**Esempio:**

```js
app.use(helmet({
  frameguard: {         // configure
    action: 'deny'
  },
  contentSecurityPolicy: {    // enable and configure
    directives: {
      defaultSrc: ["'self'"],
      styleSrc: ['style.com'],
    }
  },
  dnsPrefetchControl: false     // disable
}))
```

Abbiamo introdotto ogni middleware separatamente per scopi didattici e per facilitare i test. Usare il middleware ‘genitore’ `helmet()` è facile da implementare in un progetto reale.

# --hints--

no tests - it's a descriptive challenge

```js
assert(true);
```
