---
id: 587d8247367417b2b2512c38
title: Mitigare il rischio di clickjacking con helmet.frameguard()
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

Come promemoria, questo progetto si basa sul seguente progetto starter clonato da <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

La tua pagina potrebbe essere inserita in un `<frame>` o `<iframe>` senza il tuo consenso. Questo può causare attacchi di clickjacking, tra le altre cose. Il clickjacking è una tecnica per ingannare un utente facendolo interagire con una pagina diversa da quella che pensa di vedere. Questo può avvenire eseguendo la tua pagina in un contesto malevolo, tramite l'uso di iframe. In quel contesto, un hacker può mettere un livello nascosto sopra la tua pagina. Pulsanti nascosti possono essere usati per eseguire script dannosi. Questo middleware imposta l'intestazione X-Frame-Options. Limita chi può inserire il tuo sito in un frame. Ha tre modalità: DENY, SAMEORIGIN e ALLOW-FROM.

Non abbiamo bisogno che la nostra app venga inserita in un frame.

# --instructions--

Usa `helmet.frameguard()` passando l'oggetto di configurazione `{action: 'deny'}`.

# --hints--

Il middleware helmet.frameguard() dovrebbe essere montato correttamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

L'azione di helmet.frameguard() dovrebbe essere impostata su 'DENY'

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```
