---
id: 587d8248367417b2b2512c3a
title: evitare di inferire il tipo mime della risposta con helmet.noSniff()
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

I browser possono usare il content sniffing o MIME sniffing per sovrascrivere l'intestazione `Content-Type` di una risposta per indovinare e processare i dati usando un tipo di contenuto implicito. Sebbene questo possa essere comodo in alcuni scenari, può anche portare a degli attacchi pericolosi. Questo middleware imposta l'intestazione `X-Content-Type-Options` su `nosniff`, indicando al browser di non bypassare il `Content-Type` fornito.

# --instructions--

Usa il metodo `helmet.noSniff()` sul tuo server.

# --hints--

Il middleware helmet.noSniff() dovrebbe essere montato correttamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```
