---
id: 58a25bcff9fc0f352b528e7d
title: Hash e confronto asincrono delle password
challengeType: 2
forumTopicId: 301578
dashedName: hash-and-compare-passwords-asynchronously
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Poiché l'hashing è progettato per essere computazionalmente intensivo, si consiglia di eseguirlo in modo asincrono sul server per evitare di bloccare le connessioni in arrivo durante l'hashing. Tutto ciò che devi fare per eseguire l'hashing di una password in modo asincrono è chiamare

```js
bcrypt.hash(myPlaintextPassword, saltRounds, (err, hash) => {
  /*Store hash in your db*/
});
```

# --instructions--

Aggiungi questa funzione di hashing al tuo server (abbiamo già definito le variabili usate nella funzione per te) e stampala nella console per poterla vedere! A questo punto normalmente salveresti l'hash nel tuo database.

Ora, quando devi capire se un nuovo input corrisponde ai dati dell'hash, ti basta usare la funzione di confronto.

```js
bcrypt.compare(myPlaintextPassword, hash, (err, res) => {
  /*res == true or false*/
});
```

Inseriscila nella tua funzione di hash esistente (dato che devi aspettare che l'hash sia completato prima di chiamare la funzione di confronto) dopo aver stampato l'hash completato e stampa 'res' nella console all'interno del confronto. Dovresti vedere nella console un hash, seguito dalla stampa di 'true'! Se cambi 'myPlaintextPassword' nella funzione di confronto in 'someOtherPlaintextPassword', allora dovrebbe stampare false.

```js
bcrypt.hash('passw0rd!', 13, (err, hash) => {
  console.log(hash);
  //$2a$12$Y.PHPE15wR25qrrtgGkiYe2sXo98cjuMCG1YwSI5rJW1DSJp0gEYS
  bcrypt.compare('passw0rd!', hash, (err, res) => {
    console.log(res); //true
  });
});

```

Invia la tua pagina quando pensi di aver fatto tutto correttamente.

# --hints--

L'hash asincrono dovrebbe essere generato e confrontato correttamente.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash.*myPlaintextPassword( |),( |)saltRounds( |),( |).*err( |),( |)hash[^]*END_ASYNC/gi,
  'You should call bcrypt.hash on myPlaintextPassword and saltRounds and handle err and hash as a result in the callback'
);
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash[^]*bcrypt.compare.*myPlaintextPassword( |),( |)hash( |),( |).*err( |),( |)res[^]*}[^]*}[^]*END_ASYNC/gi,
  'Nested within the hash function should be the compare function comparing myPlaintextPassword to hash'
);
```
