---
id: 58a25bcff9fc0f352b528e7e
title: Hash e confronto delle password in modo sincrono
challengeType: 2
forumTopicId: 301579
dashedName: hash-and-compare-passwords-synchronously
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Eseguire l'hash in modo sincrono è altrettanto semplice, ma può causare rallentamenti se usato lato server con un costo elevato o con hashing effettuato molto spesso. L'hashing con questo metodo è semplice come chiamare

```js
var hash = bcrypt.hashSync(myPlaintextPassword, saltRounds);
```

Aggiungi questo metodo di hashing al tuo codice e poi registra il risultato nella console. Ancora una volta, le variabili usate sono già definite nel server quindi non dovrai modificarle. Potresti notare che, anche se stai hashando la stessa password della funzione asincrona, il risultato nella console è diverso: questo è dovuto al fatto che il salt viene generato in modo casuale ogni volta, come si vede dai primi 22 caratteri nella terza stringa dell'hash. Ora, per confrontare un input di password con il nuovo hash sincrono, useresti il metodo compareSync:

```js
var result = bcrypt.compareSync(myPlaintextPassword, hash);
```

con il risultato che sarà un booleano true o false.

# --instructions--

Aggiungi la funzione e registra il risultato nella console per vedere che funziona.

Invia la tua pagina quando pensi di aver fatto tutto correttamente.

# --hints--

L'hash sincrono dovrebbe essere generato e confrontato correttamente.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_SYNC[^]*hash.*=.*bcrypt.hashSync.*myPlaintextPassword( |),( |)saltRounds[^]*END_SYNC/gi,
  'You should call bcrypt.hashSync on myPlaintextPassword with saltRounds'
);
assert.match(
  data,
  /START_SYNC[^]*result.*=.*bcrypt.compareSync.*myPlaintextPassword( |),( |)hash[^]*END_SYNC/gi,
  'You should call bcrypt.compareSync on myPlaintextPassword with the hash generated in the last line'
);
```
