---
id: 587d8259367417b2b2512c86
title: 實作插入排序
challengeType: 1
forumTopicId: 301613
dashedName: implement-insertion-sort
---

# --description--

接下來我們要看的排序方法是插入排序。這種方法透過在列表開頭建造一個已排序的陣列來運作。它以第一個元素開始已排序的陣列。然後檢查下一個元素，並將它向後交換到已排序陣列中，直到它處於排序位置。它持續遍歷列表，並將新的項目向後交換到已排序部分，直到到達結尾。這個演算法在平均和最壞情況下具有二次時間複雜度。

**說明：** 撰寫一個函式 `insertionSort`，該函式接受一個整數陣列作為輸入，並傳回一個由小到大排序的整數陣列。

# --hints--

`insertionSort` 應該是一個函式（程式）。

```js
assert.isFunction(insertionSort);
```

`insertionSort` 應該傳回一個已排序的陣列（由小到大）。

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    insertionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`insertionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` 應該傳回一個除了順序外未改變的陣列。

```js
assert.sameMembers(
  insertionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort([5, 4, 33, 2, 8])` 應該傳回 `[2, 4, 5, 8, 33]`。

```js
assert.deepEqual(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` 不應該使用內建的 `.sort()` 方法。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function insertionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
