---
id: 587d825c367417b2b2512c8f
title: 實作 `Merge Sort`
challengeType: 1
forumTopicId: 301614
dashedName: implement-merge-sort
---

# --description--

另一個常見的中介排序演算法是合併排序。像快速排序一樣，合併排序也使用分治法的遞迴方法來排序陣列。它利用了只要每個陣列本身已排序，排序兩個陣列相對容易的事實。但我們一開始只有一個陣列作為輸入，那麼如何從中得到兩個已排序的陣列呢？我們可以遞迴地將原始輸入分成兩半，直到達到只有一個項目的陣列這個基本情況。單一項目陣列自然是已排序的，接著我們就可以開始合併。這個合併過程會展開分割原始陣列的遞迴呼叫，最終產生包含所有元素的最終已排序陣列。合併排序的步驟如下：

**1)** 遞迴地將輸入陣列對半拆分，直到產生只包含一個元素的子陣列。

**2)** 將每個已排序的子陣列合併在一起，以產生最終排序的陣列。

合併排序是一種高效的排序方法，時間複雜度為 *O(nlog(n))*。這個演算法很受歡迎，因為它效能良好且相對容易實作。

順帶一提，這將是我們在此涵蓋的最後一個排序演算法。然而，在後面關於樹狀資料結構的章節中，我們將描述堆積排序，這是另一種高效的排序方法，其實作需要二進位堆積。

**說明：** 撰寫一個函式 `mergeSort`，該函式接受一個整數陣列作為輸入，並傳回一個從小到大排序的整數陣列。實作這個功能的一個好方法是撰寫一個函式，例如 `merge`，負責合併兩個已排序的陣列，另一個函式，例如 `mergeSort`，負責產生單一項目陣列以供 `merge` 使用的遞迴。祝你好運！

# --hints--

`mergeSort` 應該是一個函式（程式）。

```js
assert.isFunction(mergeSort);
```

`mergeSort` 應該傳回一個排序過的陣列（由小到大）。

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    mergeSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`mergeSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` 應該傳回一個除了順序外未改變的陣列。

```js
assert.sameMembers(
  mergeSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`mergeSort` 不應該使用內建的 `.sort()` 方法。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    mergeSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function mergeSort(array) {
  if (array.length === 1) {
    return array;
  } else {
    const splitIndex = Math.floor(array.length / 2);
    return merge(
      mergeSort(array.slice(0, splitIndex)),
      mergeSort(array.slice(splitIndex))
    );
  }

  // Merge two sorted arrays
  function merge(array1, array2) {
    let merged = [];
    while (array1.length && array2.length) {
      if (array1[0] < array2[0]) {
        merged.push(array1.shift());
      } else if (array1[0] > array2[0]) {
        merged.push(array2.shift());
      } else {
        merged.push(array1.shift(), array2.shift());
      }
    }

    // After looping ends, one array is empty, and other array contains only
    // values greater than all values in `merged`
    return [...merged, ...array1, ...array2];
  }
}

mergeSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]);
```
