---
id: 587d825a367417b2b2512c89
title: 實作 `Quick Sort`
challengeType: 1
forumTopicId: 301615
dashedName: implement-quick-sort
---

# --description--

這裡我們將進入一個中階排序演算法：快速排序。快速排序是一種高效的、遞迴的分治法來排序陣列。在這個方法中，會在原始陣列中選擇一個樞軸值。接著將陣列分割成兩個子陣列，分別包含小於和大於該樞軸值的值。然後我們將遞迴呼叫快速排序演算法於這兩個子陣列的結果合併。這個過程會持續直到達到空陣列或單一項目陣列的基底情況，並傳回該陣列。遞迴呼叫的回溯會傳回排序後的陣列。

快速排序是一種非常高效的排序方法，平均提供 *O(nlog(n))* 的效率。它也相對容易實作。這些屬性使它成為一種受歡迎且實用的排序方法。

**說明：** 撰寫一個函式 `quickSort`，該函式接受一個整數陣列作為輸入，並傳回一個由小到大排序的整數陣列。雖然樞軸值的選擇很重要，但在這裡任何樞軸都可以。為了簡化，可以使用第一個或最後一個元素。

# --hints--

`quickSort` 應該是一個函式（程式）。

```js
assert(typeof quickSort == 'function');
```

`quickSort` 應該傳回一個已排序的陣列（由小到大）。

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    quickSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`quickSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` 應該傳回一個除了順序外未改變的陣列。

```js
assert.sameMembers(
  quickSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quickSort` 不應該使用內建的 `.sort()` 方法。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quickSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function quickSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function quickSort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quickSort(lesser), ...equal, ...quickSort(greater)];
  }
}
```
