---
id: a3f503de51cfab748ff001aa
title: 成對的
challengeType: 1
forumTopicId: 301617
dashedName: pairwise
---

# --description--

給定一個陣列 `arr`，找出其元素配對，其和與第二個引數 `arg` 相等，並傳回它們索引的和。

你可以使用多個具有相同數字元素但不同索引的配對。每個配對應該使用最低可能的可用索引。一旦元素被使用，就不能再用來與另一個元素配對。例如，`pairwise([1, 1, 2], 3)` 使用索引 0 的 1 而非索引 1 的 1 來創建配對 `[2, 1]`，因為 0+2 ＜ 1+2。

例如 `pairwise([7, 9, 11, 13, 15], 20)` 傳回 `6`。加總為 20 的配對是 `[7, 13]` 和 `[9, 11]`。接著我們可以寫出帶有它們索引和值的陣列。

<div style='margin-left: 2em;'>

| Index | 0 | 1 | 2  | 3  | 4  |
| ----- | - | - | -- | -- | -- |
| Value | 7 | 9 | 11 | 13 | 15 |

</div>

下面我們將取它們對應的索引並將它們相加。

<div style='margin-left: 2em;'>

7 + 13 = 20 → 索引 0 + 3 = 39 + 11 = 20 → 索引 1 + 2 = 33 + 3 = 6 → 傳回 `6`

</div>

# --hints--

`pairwise([1, 4, 2, 3, 0, 5], 7)` 應該傳回 11。

```js
assert.deepEqual(pairwise([1, 4, 2, 3, 0, 5], 7), 11);
```

`pairwise([1, 3, 2, 4], 4)` 應該傳回 1。

```js
assert.deepEqual(pairwise([1, 3, 2, 4], 4), 1);
```

`pairwise([1, 1, 1], 2)` 應該傳回 1。

```js
assert.deepEqual(pairwise([1, 1, 1], 2), 1);
```

`pairwise([0, 0, 0, 0, 1, 1], 1)` 應該傳回 10。

```js
assert.deepEqual(pairwise([0, 0, 0, 0, 1, 1], 1), 10);
```

`pairwise([], 100)` 應該傳回 0。

```js
assert.deepEqual(pairwise([], 100), 0);
```

# --seed--

## --seed-contents--

```js
function pairwise(arr, arg) {
  return arg;
}

pairwise([1,4,2,3,0,5], 7);
```

# --solutions--

```js
function pairwise(arr, arg) {
  var sum = 0;
  arr.forEach(function(e, i, a) {
    if (e != null) {
      var diff = arg-e;
      a[i] = null;
      var dix = a.indexOf(diff);
      if (dix !== -1) {
        sum += dix;
        sum += i;
        a[dix] = null;
      }
    }
  });
  return sum;
}

pairwise([1,4,2,3,0,5], 7);
```
