---
id: 6909ed80e1c195f49564edac
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop imparerai a lavorare con le istruzioni condizionali e gli operatori di confronto.

Inizia dichiarando una variabile chiamata `truthyOrFalsy` e assegnale il valore booleano `true`.

# --hints--

Dovresti avere una variabile chiamata `truthyOrFalsy`.

```js
assert.exists(truthyOrFalsy);
```

Dovresti assegnare `true` alla tua variabile `truthyOrFalsy`.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
