---
id: 690a06833d3e83b5226ade78
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Un'istruzione condizionale serve per eseguire codice in base a una condizione. È composta dalla parola chiave `if`, seguita da una condizione racchiusa tra parentesi tonde, e da un corpo racchiuso tra parentesi graffe. Ecco un esempio di un'istruzione `if`:

```js
if (condition) {
  console.log("condition is truthy");
}
```

Quando `condition` è un valore truthy (stringa non vuota, numero diverso da zero, `true`, ecc.) il codice all'interno del corpo dell'istruzione `if` viene eseguito.

Sotto la tua variabile, crea un'istruzione `if` che usa `hasDeveloperJob` come condizione. All'interno del corpo della tua istruzione `if`, stampa `"Timmy is employed as a developer."` nella console.

# --hints--

Dovresti avere un'istruzione `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

La tua istruzione `if` dovrebbe usare `hasDeveloperJob` come condizione.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

Dovresti stampare `"Timmy is employed as a developer."` nella console all'interno della tua istruzione `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
