---
id: 690a110a76d0496d9e7c28eb
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Dichiara un'altra variabile chiamata `timmyAge` e impostala su `18`.

# --hints--

Dovresti avere una variabile chiamata `timmyAge`.

```js
assert.exists(timmyAge);
```

Dovresti assegnare `18` alla tua variabile `timmyAge`.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
