---
id: 690a57095c4afb27bbe751b3
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

I valori falsy sono valori che valutano a `false` in un contesto booleano. Esempi di valori falsy sono:

- `""` (stringa vuota)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

Li incontrerai tutti durante questo corso.

Per ora, cambia l'assegnazione di `truthyOrFalsy` impostando il suo valore su una stringa vuota. Dopo di che, vedrai `false` stampato nella console.

# --hints--

Dovresti assegnare una stringa vuota a `truthyOrFalsy`.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
