---
id: 690a3ad44874a928a2fe3077
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Gli operatori di confronto ti permettono di confrontare valori e scrivere istruzioni condizionali più complesse.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

Per esempio, il codice sopra stampa `"Happy Birthday!"` nella console solo quando i valori delle variabili `today` e `birthday` sono uguali.

Crea una terza istruzione `if`. Per la sua condizione, usa l'operatore `>=` per verificare se `timmyAge` è maggiore o uguale a `16`.

All'interno del corpo della tua nuova istruzione `if`, stampa `"Timmy is old enough to drive."` nella console.

# --hints--

Dovresti avere una terza istruzione `if` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

La tua terza istruzione `if` dovrebbe usare `timmyAge >= 16` come condizione.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

Dovresti stampare `"Timmy is old enough to drive."` nella console all'interno della tua terza istruzione `if`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
