---
id: 690a58cc4f828a49eab1b53a
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Dichiara una variabile chiamata `hasDeveloperJob` e assegnale il valore booleano `true`.

# --hints--

Dovresti avere una variabile chiamata `hasDeveloperJob`.

```js
assert.exists(hasDeveloperJob);
```

Dovresti assegnare `true` alla tua variabile `hasDeveloperJob`.

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
