---
id: 690a562fd611e8155599ce98
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Cambia l'assegnazione di `truthyOrFalsy` impostando il suo valore su `"freeCodeCamp"`. Vedrai comunque `true` stampato nella console perché una stringa non vuota è un valore truthy.

# --hints--

La tua variabile `truthyOrFalsy` dovrebbe avere il valore `"freeCodeCamp"`.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
