---
id: 690a51806fb743bc896c3af7
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Come hai imparato in una lezione precedente, i booleani sono un tipo di dato con solo `true` e `false` come valori.

Un valore truthy è un valore che viene valutato come `true` in un contesto booleano. Esempi di valori truthy sono:

- stringhe non vuote
- numeri diversi da zero
- il booleano `true`

E altri valori che incontrerai più avanti.

Puoi usare la funzione `Boolean()` per verificare se un valore è truthy. Per esempio, `Boolean("Hello World!")` restituirà `true` perché `"Hello World!"` è truthy.

Ora, chiama `Boolean(truthyOrFalsy)` e registra il risultato nella console.

# --hints--

Dovresti registrare `Boolean(truthyOrFalsy)` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
