---
id: 690a0f7a5c22c15222a6aa5c
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Ora crea un'altra istruzione `if` che usa `isTimmyAGamer` come condizione. All'interno del corpo della tua istruzione `if`, registra `"Timmy loves to play World of Warcraft."` nella console.

Nota che non vedrai nulla registrato nella console perché la condizione della tua nuova istruzione `if` è falsa.

# --hints--

Dovresti avere due istruzioni `if` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

La tua seconda istruzione `if` dovrebbe usare `isTimmyAGamer` come condizione.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

Dovresti registrare `"Timmy loves to play World of Warcraft."` nella console all'interno della tua seconda istruzione `if`.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
