---
id: 690a3ce03013be44d116136e
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

Un'istruzione condizionale può avere una clausola `else` che esegue codice quando la condizione `if` è falsy. Ecco un esempio di un'istruzione `if...else`:

```js
if (condition) {
  console.log("condition is truthy");
} else {
  console.log("condition is falsy");
}
```

Aggiungi una clausola `else` all'istruzione `if` esistente. All'interno del corpo della tua clausola `else`, registra `"Timmy is not old enough to drive."` nella console.

# --hints--

La tua terza istruzione `if` dovrebbe avere una clausola `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{/s);
```

Dovresti registrare `"Timmy is not old enough to drive."` nella console all'interno del corpo della tua clausola `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is not old enough to drive.\2\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
--fcc-editable-region--
}
--fcc-editable-region--
```
