---
id: 690a57fef66aac39dffc61a2
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora che sai cosa sono i valori truthy e falsy, rimuovi sia la dichiarazione della variabile sia l'istruzione `console.log` dal tuo codice.

# --hints--

Non dovresti avere una variabile `truthyOrFalsy`.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

Non dovresti avere `console.log(Boolean(truthyOrFalsy))` nel tuo codice.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
