---
id: 690a0df27ee9ee36d935d530
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Dichiara un'altra variabile chiamata `isTimmyAGamer` e assegnale il booleano `false`.

# --hints--

Dovresti dichiarare una variabile chiamata `isTimmyAGamer`.

```js
assert.exists(isTimmyAGamer);
```

Dovresti assegnare il valore booleano `false` alla tua variabile `isTimmyAGamer`.

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
