---
id: 66be24cb4144f955b6bcc550
title: Створіть сторінку з ігровими картками
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити вебсторінку, яка відображає щонайменше три ігрові картки.
1. Ви повинні мати `main` елемент з ідентифікатором `playing-cards`.
1. Усередині вашого `#playing-cards` елемента має бути щонайменше три `div` елементи, кожен з класом `card`.
1. Усередині кожного `.card` елемента має бути три `div` елементи: перший з класом `left`, другий з класом `middle`, і третій з класом `right`.
1. Ваш `#playing-cards` елемент має використовувати гнучкий контейнер, щоб горизонтально центровано розташувати дочірні елементи, дозволити їм переноситися і розмістити `20px` пробіл між ними.
1. Кожен з ваших `.card` елементів має використовувати гнучкий контейнер, щоб вирівнювати дочірні елементи за допомогою `space-between`.
1. Кожен з ваших `.left` елементів має використовувати властивості гнучкого контейнера, щоб вирівняти себе на початку батьківського контейнера.
1. Кожен з ваших `.middle` елементів має використовувати властивості гнучкого контейнера, щоб вирівняти себе по центру батьківського контейнера.
1. Кожен з ваших `.right` елементів має використовувати властивості гнучкого контейнера, щоб вирівняти себе в кінці батьківського контейнера.
1. Кожен з ваших `.middle` елементів має використовувати гнучкий контейнер, щоб відображати дочірні елементи у колонку.

Ось кілька символів, які ви можете скопіювати та вставити для використання у своїх картках, якщо хочете: `♠`, `♣`, `♥` та `♦`.

# --hints--

Ви повинні мати `main` елемент з `id` рівним `playing-cards`.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

Усередині вашого `div` елемента має бути щонайменше три `card` елементи з класом `#playing-cards`.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

Кожен з ваших `.card` елементів має мати ширину та висоту.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

Кожен з ваших `.card` елементів має мати рівно три дочірні `div` елементи.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

Усередині кожного `div` елемента має бути `left` з класом `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

Усередині кожного `div` елемента має бути `middle` з класом `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

Усередині кожного `div` елемента має бути `right` з класом `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

Ви повинні мати `#playing-cards` селектор, який встановлює `display` своїх елементів у `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

Ви повинні мати `#playing-cards` селектор, який встановлює `justify-content` своїх елементів у `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

Ви повинні мати `#playing-cards` селектор, який встановлює `flex-wrap` своїх елементів у `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

Ви повинні мати `#playing-cards` селектор, який встановлює `gap` своїх елементів у `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

Ви повинні мати `.card` селектор, який встановлює `display` своїх елементів у `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

Ви повинні мати `.card` селектор, який встановлює `justify-content` своїх елементів у `space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

Ви повинні мати `.left` селектор, який встановлює `align-self` своїх елементів у `flex-start`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

Ви повинні мати `.middle` селектор, який встановлює `align-self` своїх елементів у `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

Ви повинні мати `.right` селектор, який встановлює `align-self` своїх елементів у `flex-end`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

Ви повинні мати `.middle` селектор, який встановлює `display` своїх елементів у `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

Ви повинні мати `flex-direction` встановленим у `column` у вашому класі `.middle`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
