---
id: 662a6bc12cde72c32fb526f0
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

<dfn>Інтерфейс</dfn> — це як креслення для класу. Інтерфейс містить множину методів і пропів, які клас має реалізувати.

Почніть цей проєкт із оголошення порожнього класу з іменем `Equation`. Ви будете використовувати цей клас, щоб визначити інтерфейс, креслення для загального рівняння.

# --hints--

Ви маєте визначити новий клас з іменем `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
