---
id: 662bd552e1c1d2db1b88ba47
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Тепер визначте інший клас з назвою `LinearEquation` і зробіть його спадкоємцем `Equation`. Ви будете використовувати цей клас, щоб представити лінійні рівняння.

# --hints--

Ви повинні визначити клас з назвою `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

Ваш клас `LinearEquation` має спадкувати від класу `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
