---
id: 662bd8260da84bdd5feae419
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Ви хочете, щоб клас `LinearEquation` реалізував, а не просто успадковував усі методи, визначені всередині класу `Equation`, який має виступати як інтерфейс.

Наразі клас `Equation` є просто батьківським елементом для `LinearEquation`. У наступних кроках ви навчитеся, як перетворити його на формальний інтерфейс.

Поки що створіть екземпляр `Equation` і призначте його змінній `eq`, а також екземпляр `LinearEquation` і призначте його змінній `lin_eq`.

# --hints--

Ви повинні оголосити змінну `eq` і призначити їй екземпляр `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

Ви повинні оголосити змінну `lin_eq` і призначити їй екземпляр `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
