---
id: 662bd456896f16d9bd03f1a6
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

У класі `Equation` визначте два нові методи екземпляра з іменами `solve` та `analyze`.

# --hints--

Ви повинні визначити метод з іменем `solve` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

Ваш метод `solve` має приймати один параметр, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

Ви повинні визначити метод з іменем `analyze` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

Ваш метод `analyze` має приймати один параметр, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
