---
id: 662bde88dc84f1e249801b1a
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` означає *абстрактні базові класи*. Клас `ABC` дає змогу перетворити звичайний клас на абстрактний клас, який слугує шаблоном для конкретних класів.

Змініть свою інструкцію `import`, щоб імпортувати лише клас `ABC` з модуля `abc`. Ви можете імпортувати конкретний об’єкт `x` з модуля `y`, використовуючи конструкцію імпорту `from y import x`.

Потім зробіть свій клас `Equation` абстрактним, успадкувавши його від `ABC`.

# --hints--

Ви повинні імпортувати `ABC` з модуля `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

Ваш клас `Equation` має успадковуватися від `ABC`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
