---
id: 662fa2e2cf27c09f21f4f5d0
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

У Python типи даних визначаються під час виконання коду. Тому не потрібно вказувати тип даних змінної при її оголошенні. Проте ви можете анотовано вказати тип змінної, щоб уточнити, що вона зберігатиме певний тип даних, використовуючи `variable: <data type> = value` або просто `variable: <data type>`. Зверніть увагу, що інтерпретатор Python не примушує дотримуватися типів, зазначених в анотаціях, і зазвичай для цього потрібні зовнішні інструменти.

Всередині класу `Equation` оголосіть атрибут класу `degree`. Не призначайте йому значення. Замість цього використайте анотацію типу `int`, щоб показати, що він зберігатиме ціле число у конкретних класах.

Пізніше ви використаєте цей атрибут класу як частину процесу валідації аргументів, переданих для створення об’єктів рівняння.

# --hints--

Ви повинні оголосити атрибут класу з ім’ям `degree` і анотовано вказати для нього тип `int` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
