---
id: 662fbcef5f05e1b84f541a0c
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Кожен об’єкт рівняння буде створений, передаючи стільки аргументів, скільки коефіцієнтів у рівнянні, починаючи з коефіцієнта при степені \\( x \\) n-го ступеня і до коефіцієнта при степені нуль, включно з можливим коефіцієнтом зі значенням `0`.

Наприклад, `LinearEquation(4, 5)` представляв би рівняння \\( 4x + 5 = 0 \\), де `4` — це коефіцієнт при першому (тут найвищому) степені, а `5` — коефіцієнт при степені нуль.

Вам потрібно перевірити, що для створення об’єкта рівняння передано правильну кількість аргументів.

Всередині методу `__init__` створіть інструкцію `if`, щоб перевірити, чи довжина `args` відрізняється від кількості коефіцієнтів, які має мати рівняння (`degree + 1`). Якщо так, викличте помилку `TypeError` і використайте наступний рядок, щоб надати кастомне повідомлення: `f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`.

Потім виправте помилку, передавши `2` і `3` для створення `lin_eq`.

# --hints--

Ви маєте створити інструкцію `if`, яка перевіряє, чи кількість коефіцієнтів, що використовуються для створення рівняння, відрізняється від `degree + 1`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

Ви маєте викликати помилку `TypeError` всередині нової інструкції `if` і використати наданий рядок, щоб повернути кастомне повідомлення про помилку.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

Ви маєте передати `2` і `3` для створення `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
