---
id: 662f6d7c92381a3049e4c987
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Інтерфейс не обов’язково має визначати лише абстрактні методи, він також може реалізовувати методи, які успадковуватимуться конкретними класами.

Перед тим як займатися фактичною реалізацією `solve` і `analyze` у класі `Equation`, визначте метод `__init__`. Не використовуйте для нього жодного декоратора.

# --hints--

Ви маєте визначити метод `__init__` у своєму класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

Ваш метод `__init__` має приймати один параметр, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
