---
id: 662f96576ef178927de87975
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Щоб метод вважався абстрактним, його слід декорувати декоратором `@abstractmethod`.

<dfn>Декоратор</dfn> у Python використовується для зміни поведінки функції. Ось приклад, як використати декоратор з назвою `spam`:

```py
@spam
def foo():
    pass
```

Змініть свій імпорт, щоб імпортувати декоратор `abstractmethod`, і задекоруйте ним обидва методи `solve` та `analyze` класу `Equation`. Це спричинить дві помилки.

Якщо клас, що наслідує `ABC`, має абстрактний метод, цей клас більше не можна створити. Тому видаліть екземпляр `Equation`, щоб позбутися помилки.

Інша помилка виникає через те, що клас `LinearEquation` повинен реалізувати всі абстрактні методи, визначені в інтерфейсі. Переконайтеся, що вони визначені також у класі `LinearEquation`. Не використовуйте декоратор `abstractmethod` у конкретному класі.

# --hints--

Ви повинні імпортувати `abstractmethod` з модуля `abc`.

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

Ви повинні задекорувати методом `@abstractmethod` метод `solve` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

Ви повинні задекорувати методом `@abstractmethod` метод `analyze` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

Ви повинні визначити метод з назвою `solve` у класі `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

Ваш метод `solve` має приймати один параметр, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

Ви повинні визначити метод з назвою `analyze` у класі `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

Ваш метод `solve` має приймати один параметр, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
