---
id: 6639fdcc701833a54c364211
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Після перевірки коефіцієнтів потрібно зберегти їх в атрибуті екземпляра. Використайте словникове включення, щоб створити словник, у якому ключем буде степінь коефіцієнта, а відповідним значенням — сам коефіцієнт, і призначте його атрибуту з ім’ям `coefficients`.

Наприклад, об’єкт `LinearEquation`, створений з `2` і `4`, повинен мати такий атрибут `coefficients`: `{1: 2, 0: 4}`, тому що `2` відповідає першому степеню `x`, а `4` — нульовому степеню `x`.

Створіть пари ключ-значення у вашому новому словнику у тому ж порядку, що й у `args`.

# --hints--

Ви повинні оголосити атрибут з ім’ям `coefficients` у вашому методі `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

Використайте словникове включення, щоб зберегти ваші коефіцієнти.

```js
({ test: () => runPython(`
import ast
node = _Node(_code).find_class("Equation").find_function("__init__").find_variable("self.coefficients")
assert isinstance(node.tree.value, ast.DictComp)
`) })
```

Ваш атрибут `coefficients` має бути словником, що містить пари ключ-значення у формі степінь-коефіцієнт. Пам’ятайте дотримуватися того ж порядку, у якому коефіцієнти зберігаються у `args`.

```js
({ test: () => runPython(`
actual1 = list(LinearEquation(1, 6).coefficients.items())
expected1 = list({1: 1, 0: 6}.items())
actual2 = list(LinearEquation(-3.5, 0).coefficients.items())
expected2 = list({1: -3.5, 0: 0}.items())
assert actual1 == expected1
assert actual2 == expected2
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass

    
lin_eq = LinearEquation(2, 3)
```
