---
id: 663a32735b317af9812eb0d7
title: Крок 27
challengeType: 20
dashedName: step-27
---

# --description--

У лінійних рівняннях у вигляді \\( ax + b = 0 \\) кутовий коефіцієнт це просто коефіцієнт \\( a \\), а точка перетину y це коефіцієнт \\( b \\).

<img alt="графік лінійної функції" src="https://cdn.freecodecamp.org/curriculum/python/linear-equation.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;" />

Ви збираєтеся використовувати метод `analyze`, щоб надати додаткову інформацію про рівняння. Усередині методу `analyze` розпакуйте коефіцієнти в змінні `slope` та `intercept`.

Потім поверніть словник із ключами `'slope'` та `'intercept'` і значеннями кутового коефіцієнту та точку перетину y, відповідно. Після цього викличте `analyze` на `lin_eq` та надрукуйте результат.


# --hints--

Розпакуйте значення, збережені в атрибуті `coefficients`, у змінні `slope` та `intercept` в межах метода `analyze`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_stmt("slope, intercept = self.coefficients.values()")`)) })
```

Метод `analyze` має повернути словник із ключами `'slope'` та `'intercept'` і значеннями кутового коефіцієнту та точку перетину y, відповідно.

```js
({ test: () => runPython(`
eq = LinearEquation(2.2, 1.5)
a = eq.analyze()
assert a['slope'] == 2.2, "Expected different slope"
assert a['intercept'] == 1.5, "Expected different intercept"
`) })
```

Викличте метод `analyze` об'єкта `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(lin_eq.analyze())")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int

    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1

    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x
--fcc-editable-region--
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)
print(lin_eq.solve())
--fcc-editable-region--
```
