---
id: 663b83a28943e6aa6275a514
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Ще в межах класу `Equation` оголосіть метод `__str__`, щоб надати правильне рядкове представлення об’єктам рівнянь, які ви збираєтеся створити.

Поки що в методі `__str__` оголосіть змінну `terms` і призначте їй порожній список. Ви будете використовувати цю змінну, щоб зберігати кожен доданок (коефіцієнт, помножений на \\( x^n \\)) вашого рівняння.

Потім оголосіть змінну `equation_string`, призначте їй результат об’єднання елементів списку `terms` через пробіл. Нарешті, поверніть `equation_string`.

# --hints--

Ви повинні оголосити метод `__str__` у класі `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

Ваш метод `__str__` має приймати один параметр, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

Ви повинні оголосити змінну `terms` і призначити їй порожній список у методі `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

Ви повинні оголосити змінну `equation_string` і призначити їй результат об’єднання елементів `terms` через пробіл у методі `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

Ви повинні повернути `equation_string` з вашого методу `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
