---
id: 6675a38a8b535e4ff3274520
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Вбудована функція `hasattr` приймає об’єкт як перший аргумент і рядок, що представляє ім’я атрибута, як другий аргумент. Вона повертає булеве значення, яке вказує, чи має об’єкт вказаний атрибут.

Тепер ви використаєте метод `__init_subclass__`, щоб перевірити, чи має дочірній клас атрибут `degree` у момент створення екземпляра.

Створіть інструкцію `if`, щоб перевірити, чи `cls` не має атрибута `degree`. Якщо так, викличте `AttributeError` і використайте рядок `f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"` для надання кастомного повідомлення.

Після цього виправте помилку, яка з’явилася в терміналі, оголосивши атрибут класу `degree` всередині класу `LinearEquation`. Цей атрибут має представляти степінь рівняння, тобто показник степеня найвищого члена \\( x \\). Тому призначте цілому числу `1` атрибут `degree`.

# --hints--

Ви маєте створити інструкцію `if`, яка перевіряє, чи `cls` не має атрибута `degree` всередині методу `__init_subclass__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_conditions()[0].is_equivalent("not hasattr(cls, 'degree')")`)) })
```

Ви маєте викликати `AttributeError`, використовуючи наданий рядок всередині вашої інструкції `if`.

```js
({ test: () => runPython(`
raise_stmt = 'raise AttributeError(f"Cannot create \\'{cls.__name__}\\' class: missing required attribute \\'degree\\'")'
node = _Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_bodies()[0]
assert node.has_stmt(raise_stmt) 
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        pass

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    
--fcc-editable-region--
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
