---
id: 694648acde178bb8202d9513
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Ім’я студента має відповідати тому ж формату, що й інші деталі.

Видаліть попередні вихідні дані змінної `name`. Потім виведіть `name` і `type(name)` разом в одному рядку, розділені комою, як у попередньому кроці.

# --hints--

У вашому коді не повинно бути `print(name)`.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

У вашому коді не повинно бути `print(type(name))`.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

Ви маєте вивести `name` і `type(name)` в одному виклику `print()` з роздільником-комою.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
