---
id: 694648acde178bb8202d9515
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Тепер додайте бал студента.

Оголосіть змінну з іменем `score` і призначте їй значення `80.5`.

Хоча і `age`, і `score` є числами, вони можуть бути різних типів. Python надає функцію `isinstance()` для перевірки цього.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

Використайте `isinstance()`, щоб перевірити, чи є `score` `int`, і виведіть результат у термінал, як показано в прикладі вище.

# --hints--

У вас має бути змінна `score`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

Змінна `score` має зберігати значення `80.5`. Не беріть це значення в лапки.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

Ви маєте вивести результат `isinstance(score, int)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
