---
id: 694648acde178bb8202d9512
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

У табелі також має бути показано, чи студент зараз зарахований. Це можна представити за допомогою булевого значення.

Булеві значення представляють умову «так» або «ні» і часто використовуються для прийняття рішень у коді. Існує лише два булеві значення: `True` та `False`.

Оголосіть змінну з ім’ям `is_student` і призначте їй значення `True`.

# --hints--

У вас має бути булева змінна `is_student`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

Змінна `is_student` має зберігати значення `True`. Не беріть значення в лапки.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
