---
id: 694648abde178bb8202d9511
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Тепер ви повинні побачити ім’я студента, виведене в терміналі.

Python надає функцію з назвою `type()`, яку ви можете використати, щоб перевірити тип значення.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

У наведеному вище прикладі вихідні дані `<class 'str'>` означають, що змінна, передана у функцію `type()`, є рядком.

Використайте функцію `type()` з `name` як аргументом і виведіть вихідні дані, як у прикладі. Перевірте в терміналі, що `name` має тип `str` (рядок).

# --hints--

Ви повинні вивести результат виклику `type(name)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
