---
id: 69461e4f98594012939bc1c6
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся з типами даних Python, створюючи простий принтер табеля успішності.

Як ви дізналися в попередніх уроках, змінні призначаються таким чином:

```py
greeting = 'Hello World'
```

У наведеному прикладі змінна `greeting` має значення рядка, який є послідовністю символів, оточених лапками.

Створіть змінну з назвою `name` і призначте їй рядок `Alice`. Пам’ятайте, що значення потрібно оточувати або одинарними, або подвійними лапками, як показано в прикладі.

**Важливо:** Python використовує відступи (пробіли на початку рядка) для організації коду. Для цього практичного заняття переконайтеся, що на початку кожного рядка коду немає зайвих пробілів. Додавання зайвих пробілів спричинить `IndentationError` і не дозволить вашому коду запуститися.

# --hints--

У вас має бути змінна `name`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

Змінна `name` має зберігати значення `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
