---
id: 694648acde178bb8202d9514
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Тепер потрібно додати вік студента до табеля. Для цього ви використаєте ціле число, один із числових типів даних у Python.

Оголосіть змінну з іменем `age` і призначте їй ціле число `20`.

Потім виведіть значення та тип даних `age` разом, розділені комою. Перевірте вихідні дані в терміналі, де показано значення `age` і його тип як `int` (ціле число).

# --hints--

Ви повинні мати змінну `age`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

Змінна `age` має зберігати значення `20`. Не оточуйте значення лапками.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

Ви повинні вивести `age` і `type(age)`, використовуючи кому як роздільник.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
