---
id: 694648acde178bb8202d9517
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Вихідні дані — це `True`, що підтверджує, що `score` є `float`.

Завершіть табель, вивівши значення `score` разом із його типом даних за допомогою однієї інструкції `print()`.

# --hints--

Ви повинні вивести `score` і `type(score)` через кому.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
