---
id: 694648acde178bb8202d9516
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Вихідні дані — це `False`, що показує, що `score` не є `int`.

Ще один загальний тип числа в Python — це `float`, який представляє число з десятковими знаками. Замініть `int` на `float` у наявному виклику `isinstance()`, щоб підтвердити це.

# --hints--

Ви повинні замінити `int` на `float` у наявному виклику `isinstance(score, int)`.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
