---
id: 6947a3e516bb8a4a681274a7
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Функція `print()` може виводити більше одного значення одночасно. Розділяйте значення комою (`,`), щоб вивести їх в одному рядку.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

Виведіть одночасно `is_student` і `type(is_student)` в одному рядку, використовуючи кому `,`, як показано в прикладі. Потім перевірте вихідні дані в терміналі, де буде показано значення `is_student` та його тип як `bool` (булевий).

# --hints--

Ви повинні вивести `is_student` і `type(is_student)`, використовуючи кому як роздільник.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
