---
id: 6947a0c2b6de035dd8d3eea1
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Ви можете вивести значення змінної за допомогою функції `print()`.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

Ви дізнаєтеся більше про функції в наступних уроках. Наразі знайте, що функція — це багаторазовий блок коду, який можна викликати, щоб запустити його код, і їй можна передавати аргументи.

У наведеному вище прикладі `print(greeting)` — це виклик функції, а `greeting` — аргумент функції.

Зверніться до прикладу і виведіть змінну `name`. Перевірте вихідні дані в терміналі.

# --hints--

Ви повинні викликати `print()` з `name` як аргументом.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
