---
id: a0b5010f579e69b815e7c5d6
title: البحث والاستبدال
challengeType: 1
forumTopicId: 16045
dashedName: search-and-replace
---

# --description--

نفذ عملية بحث واستبدال في الجملة باستخدام المعطيات المقدمة وأرجع الجملة الجديدة.

المعلمة الأولى هي الجملة التي ستُجرى عليها عملية البحث والاستبدال.

المعلمة الثانية هي الكلمة التي ستُستبدل (قبل).

المعلمة الثالثة هي الكلمة التي ستُستبدل بها المعلمة الثانية (بعد).

**ملاحظة:** احفظ حالة الحرف الأول في الكلمة الأصلية عند استبدالها. على سبيل المثال، إذا أردت استبدال الكلمة `Book` بالكلمة `dog`، يجب أن تُستبدل بـ `Dog`

# --hints--

`myReplace("Let us go to the store", "store", "mall")` يجب أن تُرجع السلسلة النصية `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` يجب أن تُرجع السلسلة النصية `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` يجب أن تُرجع السلسلة النصية `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` يجب أن تُرجع السلسلة النصية `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` يجب أن تُرجع السلسلة النصية `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` يجب أن تُرجع السلسلة النصية `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js
function myReplace(str, before, after) {
  return str;
}

myReplace("A quick brown fox jumped over the lazy dog", "jumped", "leaped");
```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
