---
id: a105e963526e7de52b219be9
title: الاتحاد المرتب
challengeType: 1
forumTopicId: 16077
dashedName: sorted-union
---

# --description--

اكتب دالة تأخذ مصفوفتين أو أكثر وتُرجع مصفوفة جديدة تحتوي على القيم الفريدة بترتيب المصفوفات الأصلية المقدمة.

بمعنى آخر، يجب تضمين كل القيم الموجودة في جميع المصفوفات بترتيبها الأصلي، ولكن بدون تكرار في المصفوفة النهائية.

يجب ترتيب الأرقام الفريدة حسب ترتيبها الأصلي، لكن المصفوفة النهائية لا يجب أن تُرتب ترتيبًا رقميًا.

تحقق من اختبارات التحقق للاطلاع على الأمثلة.

# --hints--

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` يجب أن تُرجع `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` يجب أن تُرجع `[1, 2, 3, 5]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` يجب أن تُرجع `[1, 2, 3, 5, 4, 6, 7, 8]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` يجب أن تُرجع `[1, 3, 2, 5, 4, 6]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` يجب أن تُرجع `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js
function uniteUnique(arr) {
  return arr;
}

uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]);
```

# --solutions--

```js
function uniteUnique(arr) {
  return [].slice.call(arguments).reduce(function(a, b) {
    return [].concat(
      a, 
      b.filter(function(e, currentIndex) {
        return b.indexOf(e) === currentIndex && a.indexOf(e) === -1;
      }));
  }, []);
}
```
