---
id: a2f1d72d9b908d0bd72bb9f6
title: إنشاء شخص
challengeType: 1
forumTopicId: 16020
dashedName: make-a-person
---

# --description--

املأ دالة البناء للكائن بالطرق التالية أدناه:

```js
getFirstName()
getLastName()
getFullName()
setFirstName(first)
setLastName(last)
setFullName(first, last)
```

شغّل الاختبارات لترى المخرجات المتوقعة لكل طريقة. يجب أن تكون هذه الطرق الوسيلة الوحيدة للتفاعل مع الكائن. كل اختبار سيُعلن عن نسخة جديدة من `Person` كـ `new Person('Bob', 'Ross')`.

# --hints--

لا يجب تغيير توقيع الدالة.

```js
assert.match(code, /const\s+Person\s*=\s*function\s*\(\s*first\s*,\s*last\s*\)\s*{/);
```

لا يجب إعادة تعيين المعلمة `first`.

```js
assert.notMatch(code, /\bfirst\s*=\s*/);
```

لا يجب إعادة تعيين المعلمة `last`.

```js
assert.notMatch(code, /\blast\s*=\s*/);
```

لا يجب إضافة خصائص. يجب أن تعيد `Object.keys(Person).length` دائمًا القيمة 6.

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
_person.setLastName('Curry');
_person.setFullName('John', 'Smith');
assert.lengthOf(Object.keys(_person), 6);
```

يجب أن تتمكن من إنشاء نسخة من كائن `Person` الخاص بك.

```js
const _person = new Person('Bob', 'Ross');
assert.instanceOf(_person, Person);
```

يجب ألا يحتوي كائن `Person` الخاص بك على خاصية `firstName`.

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'firstName');
```

يجب ألا يحتوي كائن `Person` الخاص بك على خاصية `lastName`.

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'lastName');
```

يجب أن تُعيد الطريقة `.getFirstName()` السلسلة النصية `Bob`.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFirstName(), 'Bob');
```

يجب أن تُعيد الطريقة `.getLastName()` السلسلة النصية `Ross`.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getLastName(), 'Ross');
```

يجب أن تُعيد الطريقة `.getFullName()` السلسلة النصية `Bob Ross`.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFullName(), 'Bob Ross');
```

يجب أن تُعيد الطريقة `.getFullName()` السلسلة النصية `Haskell Ross` بعد استدعاء `.setFirstName('Haskell')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
assert.strictEqual(_person.getFullName(), 'Haskell Ross');
```

يجب أن تُعيد الطريقة `.getFullName()` السلسلة النصية `Bob Curry` بعد استدعاء `.setLastName('Curry')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setLastName('Curry');
assert.strictEqual(_person.getFullName(), 'Bob Curry');
```

يجب أن تُعيد الطريقة `.getFullName()` السلسلة النصية `Haskell Curry` بعد استدعاء `.setFullName('Haskell', 'Curry')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFullName(), 'Haskell Curry');
```

يجب أن تُعيد الطريقة `.getFirstName()` السلسلة النصية `Haskell` بعد استدعاء `.setFullName('Haskell', 'Curry')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFirstName(), 'Haskell');
```

يجب أن تُعيد الطريقة `.getLastName()` السلسلة النصية `Curry` بعد استدعاء `.setFullName('Haskell', 'Curry')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getLastName(), 'Curry');
```

يجب أن تُعيد الطريقة `.getFullName()` السلسلة النصية `Emily Martinez de la Rosa` بعد استدعاء `.setFullName('Emily Martinez', 'de la Rosa')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFullName(), 'Emily Martinez de la Rosa');
```

يجب أن تُعيد الخاصية `.getFirstName()` السلسلة النصية `Emily Martinez` بعد استدعاء `.setFullName('Emily Martinez', 'de la Rosa')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFirstName(), 'Emily Martinez');
```

يجب أن تُعيد الخاصية `.getLastName()` السلسلة النصية `de la Rosa` بعد استدعاء `.setFullName('Emily Martinez', 'de la Rosa')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getLastName(), 'de la Rosa');
```

# --seed--

## --seed-contents--

```js
const Person = function(first, last) {
  this.getFullName = function() {
    return "";
  };
  return "";
};
```

# --solutions--

```js
const Person = function(first, last) {
  let firstName = first;
  let lastName = last;

  this.getFirstName = function(){
    return firstName;
  };

  this.getLastName = function(){
    return lastName;
  };

  this.getFullName = function(){
    return firstName + " " + lastName;
  };

  this.setFirstName = function(str){
    firstName = str;
  };


  this.setLastName = function(str){
    lastName = str;
  };

  this.setFullName = function(first, last){
    firstName = first;
    lastName = last;
  };
};
```
