---
id: a3566b1109230028080c9345
title: جمع كل الأعداد في نطاق
challengeType: 1
forumTopicId: 16083
dashedName: sum-all-numbers-in-a-range
---

# --description--

سنمرر لك مصفوفة تحتوي على رقمين. أرجع مجموع هذين الرقمين بالإضافة إلى مجموع كل الأعداد بينهما. الرقم الأقل لن يكون دائمًا في البداية.

على سبيل المثال، يجب أن تُرجع الدالة `sumAll([4,1])` القيمة `10` لأن مجموع كل الأعداد بين 1 و4 (شاملة) هو `10`.

# --hints--

يجب أن تُرجع الدالة `sumAll([1, 4])` رقمًا.

```js
assert(typeof sumAll([1, 4]) === 'number');
```

يجب أن تُرجع الدالة `sumAll([1, 4])` القيمة 10.

```js
assert.deepEqual(sumAll([1, 4]), 10);
```

يجب أن تُرجع الدالة `sumAll([4, 1])` القيمة 10.

```js
assert.deepEqual(sumAll([4, 1]), 10);
```

يجب أن تُرجع الدالة `sumAll([5, 10])` القيمة 45.

```js
assert.deepEqual(sumAll([5, 10]), 45);
```

يجب أن تُرجع الدالة `sumAll([10, 5])` القيمة 45.

```js
assert.deepEqual(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
function sumAll(arr) {
  return 1;
}

sumAll([1, 4]);
```

# --solutions--

```js
function sumAll(arr) {
  var sum = 0;
  arr.sort(function(a,b) {return a-b;});
  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }
  return sum;
}
```
