---
id: a39963a4c10bc8b4d4f06d7e
title: ابحث ودمر
challengeType: 1
forumTopicId: 16046
dashedName: seek-and-destroy
---

# --description--

سيُزوَّد لديك مصفوفة أولية كالمعلمة الأولى للدالة `destroyer`، تليها معلمة أو أكثر. احذف كل العناصر من المصفوفة الأولية التي لها نفس القيمة كالمعلمات هذه.

يجب أن تقبل الدالة عددًا غير محدد من المعلمات، وتُعرف أيضًا بالدالة متعددة المعلمات. يمكنك الوصول إلى المعلمات الإضافية بإضافة معلمة استراحة إلى تعريف الدالة أو باستخدام كائن `arguments`.

# --hints--

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` يجب أن تُرجع `[1, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` يجب أن تُرجع `[1, 5, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` يجب أن تُرجع `[1]`.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` يجب أن تُرجع `[]`.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` يجب أن تُرجع `["hamburger"]`.

```js
assert.deepEqual(destroyer(['tree', 'hamburger', 53], 'tree', 53), [
  'hamburger'
]);
```

`destroyer(["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan")` يجب أن تُرجع `[12,92,65]`.

```js
assert.deepEqual(
  destroyer(
    [
      'possum',
      'trollo',
      12,
      'safari',
      'hotdog',
      92,
      65,
      'grandma',
      'bugati',
      'trojan',
      'yacht'
    ],
    'yacht',
    'possum',
    'trollo',
    'safari',
    'hotdog',
    'grandma',
    'bugati',
    'trojan'
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js
function destroyer(arr) {
  return arr;
}

destroyer([1, 2, 3, 1, 2, 3], 2, 3);
```

# --solutions--

```js
function destroyer(arr) {
  var hash = Object.create(null);
  [].slice.call(arguments, 1).forEach(function(e) {
    hash[e] = true;
  });
  return arr.filter(function(e) { return !(e in hash);});
}

destroyer([1, 2, 3, 1, 2, 3], 2, 3);
```
