---
id: a3bfc1673c0526e06d3ac698
title: جمع كل الأعداد الأولية
challengeType: 1
forumTopicId: 16085
dashedName: sum-all-primes
---

# --description--

ال<dfn>عدد الأولي</dfn> هو عدد صحيح أكبر من 1 وله مقسومان فقط: 1 ونفسه. على سبيل المثال، 2 هو عدد أولي لأنه يقبل القسمة فقط على 1 و2. بالمقابل، 4 ليس عددًا أوليًا لأنه يقبل القسمة على 1 و2 و4.

أعد كتابة `sumPrimes` بحيث تُرجع مجموع كل الأعداد الأولية التي تكون أقل من أو تساوي num.

# --hints--

يجب أن تُرجع `sumPrimes(10)` رقمًا.

```js
assert.deepEqual(typeof sumPrimes(10), 'number');
```

يجب أن تُرجع `sumPrimes(10)` القيمة 17.

```js
assert.deepEqual(sumPrimes(10), 17);
```

يجب أن تُرجع `sumPrimes(977)` القيمة 73156.

```js
assert.deepEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js
function sumPrimes(num) {
  return num;
}

sumPrimes(10);
```

# --solutions--

```js
class PrimeSeive {
  constructor(num) {
    const seive = Array(Math.floor((num - 1) / 2)).fill(true);
    const upper = Math.floor((num - 1) / 2);
    const sqrtUpper = Math.floor((Math.sqrt(num) - 1) / 2);

    for (let i = 0; i <= sqrtUpper; i++) {
      if (seive[i]) {
        // Mark value in seive array
        const prime = 2 * i + 3;
        // Mark all multiples of this number as false (not prime)
        const primeSquaredIndex = 2 * i ** 2 + 6 * i + 3;
        for (let j = primeSquaredIndex; j < upper; j += prime) {
          seive[j] = false;
        }
      }
    }

    this._seive = seive;
  }

  isPrime(num) {
    return num === 2
      ? true
      : num % 2 === 0
        ? false
        : this.isOddPrime(num);
  }

  isOddPrime(num) {
    return this._seive[(num - 3) / 2];
  }
};

function sumPrimes(num) {
  const primeSeive = new PrimeSeive(num);

  let sum = 2;
  for (let i = 3; i <= num; i += 2) {
    if (primeSeive.isOddPrime(i)) sum += i;
  }
  return sum;
}

sumPrimes(10);
```
