---
id: a5229172f011153519423690
title: جمع كل أعداد فيبوناتشي الفردية
challengeType: 1
forumTopicId: 16084
dashedName: sum-all-odd-fibonacci-numbers
---

# --description--

بالنظر إلى عدد صحيح موجب `num`، أرجع مجموع كل أعداد فيبوناتشي الفردية التي تكون أقل من أو تساوي `num`.

أول رقمين في متتالية فيبوناتشي هما 0 و 1. كل رقم إضافي في المتتالية هو مجموع الرقمين السابقين. أول سبعة أرقام في متتالية فيبوناتشي هي 0، 1، 1، 2، 3، 5 و 8.

على سبيل المثال، يجب أن تُرجع الدالة `sumFibs(10)` القيمة `10` لأن جميع أعداد فيبوناتشي الفردية التي تقل عن أو تساوي `10` هي 1، 1، 3، و5.

# --hints--

يجب أن تُرجع الدالة `sumFibs(1)` رقمًا.

```js
assert(typeof sumFibs(1) === 'number');
```

يجب أن تُرجع الدالة `sumFibs(1000)` القيمة 1785.

```js
assert(sumFibs(1000) === 1785);
```

يجب أن تُرجع الدالة `sumFibs(4000000)` القيمة 4613732.

```js
assert(sumFibs(4000000) === 4613732);
```

يجب أن تُرجع الدالة `sumFibs(4)` القيمة 5.

```js
assert(sumFibs(4) === 5);
```

يجب أن تُرجع الدالة `sumFibs(75024)` القيمة 60696.

```js
assert(sumFibs(75024) === 60696);
```

يجب أن تُرجع الدالة `sumFibs(75025)` القيمة 135721.

```js
assert(sumFibs(75025) === 135721);
```

# --seed--

## --seed-contents--

```js
function sumFibs(num) {
  return num;
}

sumFibs(4);
```

# --solutions--

```js
function sumFibs(num) {
  var a = 1;
  var b = 1;
  var s = 0;
  while (a <= num) {
    if (a % 2 !== 0) {
      s += a;
    }
    a = [b, b=b+a][0];
  }
  return s;
}
```
