---
id: a6b0bb188d873cb2c8729495
title: تحويل كيانات HTML
challengeType: 1
forumTopicId: 16007
dashedName: convert-html-entities
---

# --description--

حوّل الأحرف `&`، `<`، `>`، `"` (علامة الاقتباس المزدوجة)، و`'` (الفاصلة العليا) في سلسلة نصية إلى كيانات HTML المقابلة لها.

# --hints--

`convertHTML("Dolce & Gabbana")` يجب أن تُرجع السلسلة النصية `Dolce &amp; Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` يجب أن تُرجع السلسلة النصية `Hamburgers &lt; Pizza &lt; Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` يجب أن تُرجع السلسلة النصية `Sixty &gt; twelve`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` يجب أن تُرجع السلسلة النصية `Stuff in &quot;quotation marks&quot;`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` يجب أن تُرجع السلسلة النصية `Schindler&apos;s List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` يجب أن تُرجع السلسلة النصية `&lt;&gt;`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` يجب أن تُرجع السلسلة النصية `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js
function convertHTML(str) {
  return str;
}

convertHTML("Dolce & Gabbana");
```

# --solutions--

```js
var MAP = { '&': '&amp;',
            '<': '&lt;',
            '>': '&gt;',
            '"': '&quot;',
            "'": '&apos;'};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(c) {
    return MAP[c];
  });
}
```
