---
id: a97fd23d9b809dac9921074f
title: المعلمات اختيارية
challengeType: 1
forumTopicId: 14271
dashedName: arguments-optional
---

# --description--

أنشئ دالة تجمع معلمتين معًا. إذا تم توفير معلمة واحدة فقط، فأرجع دالة تتوقع معلمة واحدة وتُرجع المجموع.

على سبيل المثال، يجب أن تُرجع `addTogether(2, 3)` القيمة `5`، ويجب أن تُرجع `addTogether(2)` دالة.

استدعاء هذه الدالة المرجعة بمعلمة واحدة سيُرجع المجموع:

```js
var sumTwoAnd = addTogether(2);
```

`sumTwoAnd(3)` تُرجع `5`.

إذا لم تكن أي من المعلمتين رقمًا صالحًا، فأرجع undefined.

# --hints--

`addTogether(2, 3)` يجب أن تُرجع 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` يجب أن تُرجع 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` يجب أن تُرجع دالة.

```js
assert.deepEqual(typeof(addTogether(5)), 'function');
```

`addTogether(5)(7)` يجب أن تُرجع 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js
function addTogether() {
  return false;
}

addTogether(2,3);
```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
