---
id: aa7697ea2477d1316795783b
title: لغة الخنازير
challengeType: 1
forumTopicId: 16039
dashedName: pig-latin
---

# --description--

لغة الخنازير هي طريقة لتغيير الكلمات الإنجليزية. القواعد كما يلي:

\- إذا بدأت الكلمة بحرف ساكن، خذ أول حرف ساكن أو مجموعة حروف ساكنة، انقلها إلى نهاية الكلمة، وأضف إليها `ay`.

\- إذا بدأت الكلمة بحرف متحرك، فقط أضف `way` في النهاية.

# --instructions--

ترجم السلسلة النصية المقدمة إلى لغة الخنازير. السلاسل النصية المُدخلة مضمونة أن تكون كلمات إنجليزية بأحرف صغيرة.

# --hints--

`translatePigLatin("california")` يجب أن تُرجع السلسلة النصية `aliforniacay`.

```js
assert.deepEqual(translatePigLatin('california'), 'aliforniacay');
```

`translatePigLatin("paragraphs")` يجب أن تُرجع السلسلة النصية `aragraphspay`.

```js
assert.deepEqual(translatePigLatin('paragraphs'), 'aragraphspay');
```

`translatePigLatin("glove")` يجب أن تُرجع السلسلة النصية `oveglay`.

```js
assert.deepEqual(translatePigLatin('glove'), 'oveglay');
```

`translatePigLatin("algorithm")` يجب أن تُرجع السلسلة النصية `algorithmway`.

```js
assert.deepEqual(translatePigLatin('algorithm'), 'algorithmway');
```

`translatePigLatin("eight")` يجب أن تُرجع السلسلة النصية `eightway`.

```js
assert.deepEqual(translatePigLatin('eight'), 'eightway');
```

يجب أن تتعامل مع الكلمات التي يأتي فيها أول حرف متحرك في وسط الكلمة. `translatePigLatin("schwartz")` يجب أن تُرجع السلسلة النصية `artzschway`.

```js
assert.deepEqual(translatePigLatin('schwartz'), 'artzschway');
```

يجب أن تتعامل مع الكلمات التي لا تحتوي على حروف متحركة. `translatePigLatin("rhythm")` يجب أن تُرجع السلسلة النصية `rhythmay`.

```js
assert.deepEqual(translatePigLatin('rhythm'), 'rhythmay');
```

# --seed--

## --seed-contents--

```js
function translatePigLatin(str) {
  return str;
}

translatePigLatin("consonant");
```

# --solutions--

```js
function translatePigLatin(str) {
  if (isVowel(str.charAt(0))) return str + "way";
  var front = [];
  str = str.split('');
  while (str.length && !isVowel(str[0])) {
    front.push(str.shift());
  }
  return [].concat(str, front).join('') + 'ay';
}

function isVowel(c) {
  return ['a', 'e', 'i', 'o', 'u'].indexOf(c.toLowerCase()) !== -1;
}
```
