---
id: af4afb223120f7348cdfc9fd
title: تحويل الحطام
challengeType: 1
forumTopicId: 16021
dashedName: map-the-debris
---

# --description--

وفقًا لقانون كبلر الثالث، فإن الفترة المدارية $T$ لجسمين نقطيين يدوران حول بعضهما في مدار دائري أو بيضاوي هي:

$$
T = 2 \pi \sqrt{\frac{a^{3}}{\mu}}
$$

- $a$ هو المحور شبه الرئيسي للمدار
- $μ = GM$ هو المعامل الجاذبي القياسي
- $G$ هو ثابت الجاذبية،
- $M$ هو كتلة الجسم الأكثر ضخامة.

أرجع مصفوفة جديدة تحول متوسط ارتفاع العناصر إلى فترات مدارية لها (بالثواني).

ستحتوي المصفوفة على كائنات بالشكل `{name: 'name', avgAlt: avgAlt}`.

يجب تقريب القيم إلى أقرب عدد صحيح. الجسم الذي يدور حوله هو الأرض.

نصف قطر الأرض هو 6367.4447 كيلومتر، وقيمة GM للأرض هي 398600.4418 كم<sup>3</sup>ث<sup>-2</sup>.

# --hints--

`orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}])` يجب أن تُرجع `[{name: "sputnik", orbitalPeriod: 86400}]`.

```js
assert.deepEqual(orbitalPeriod([{ name: 'sputnik', avgAlt: 35873.5553 }]), [
  { name: 'sputnik', orbitalPeriod: 86400 }
]);
```

`orbitalPeriod([{name: "iss", avgAlt: 413.6}, {name: "hubble", avgAlt: 556.7}, {name: "moon", avgAlt: 378632.553}])` يجب أن تُرجع `[{name : "iss", orbitalPeriod: 5557}, {name: "hubble", orbitalPeriod: 5734}, {name: "moon", orbitalPeriod: 2377399}]`.

```js
assert.deepEqual(
  orbitalPeriod([
    { name: 'iss', avgAlt: 413.6 },
    { name: 'hubble', avgAlt: 556.7 },
    { name: 'moon', avgAlt: 378632.553 }
  ]),
  [
    { name: 'iss', orbitalPeriod: 5557 },
    { name: 'hubble', orbitalPeriod: 5734 },
    { name: 'moon', orbitalPeriod: 2377399 }
  ]
);
```

# --seed--

## --seed-contents--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  return arr;
}

orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}]);
```

# --solutions--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  const TAU = 2 * Math.PI;
  return arr.map(function(obj) {
    return {
      name: obj.name,
      orbitalPeriod: Math.round(TAU * Math.sqrt(Math.pow(obj.avgAlt+earthRadius, 3)/GM))
    };
  });
}
```
