---
id: af7588ade1100bde429baf20
title: الحروف المفقودة
challengeType: 1
forumTopicId: 16023
dashedName: missing-letters
---

# --description--

ابحث عن الحرف المفقود في نطاق الحروف الممرر وأعده.

إذا كانت كل الحروف موجودة في النطاق، أعد القيمة `undefined`.

# --hints--

يجب أن تُعيد الدالة `fearNotLetter("abce")` السلسلة النصية `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

يجب أن تُعيد الدالة `fearNotLetter("abcdefghjklmno")` السلسلة النصية `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

يجب أن تُعيد الدالة `fearNotLetter("stvwx")` السلسلة النصية `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

يجب أن تُعيد الدالة `fearNotLetter("bcdf")` السلسلة النصية `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

يجب أن تُعيد الدالة `fearNotLetter("abcdefghijklmnopqrstuvwxyz")` القيمة `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js
function fearNotLetter(str) {
  return str;
}

fearNotLetter("abce");
```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    var letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
