---
id: 69757cc0faae0152c1418aad
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utafanya mazoezi ya kutumia nambari na operesheni za kihesabu kujenga kifaa cha kugawanya bili. Kifaa hiki kitahesabu kiasi ambacho kila mtu analazimika kulipa baada ya kuongeza gharama za chakula na tip.

Kuanza, unahitaji njia ya kufuatilia jumla ya kiasi kinapoongezwa gharama.

Tengeneza kigezo kinachoitwa `running_total` na weka thamani `0`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `running_total`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

Unapaswa kuweka thamani ya nambari kamili `0` katika kigezo chako cha `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
