---
id: 6976cc0b0c135686a4620b94
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Ifuatayo, unahitaji kuhesabu idadi ya watu wanaogawiana bili. Hifadhi thamani hii katika variable, kama ulivyofanya katika hatua iliyopita.

Tengeneza variable iitwayo `num_of_friends` na weka thamani ya `4`. Hii itatumika baadaye katika warsha kuhesabu mgawanyo wa mwisho.

# --hints--

Unapaswa kuwa na variable iitwayo `num_of_friends`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

Unapaswa kuweka nambari kamili `4` katika variable yako ya `num_of_friends`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
