---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Kila kozi ina gharama. Unahitaji kuhifadhi kiasi hiki katika vigezo ili kuitumia baadaye. Kwa kuwa kiasi hiki kinajumuisha senti, utatumia aina ya float, ambayo hutumika kuwakilisha nambari za desimali. Hapa kuna mfano wa kigezo chenye thamani ya float:

```py
change = 2.35
```

Tengeneza vigezo vinne: `appetizers` weka thamani `37.89`, `main_courses` weka thamani `57.34`, `desserts` weka thamani `39.39`, na `drinks` weka thamani `64.21`. 

# --hints--

Unapaswa kufafanua kigezo kinachoitwa `appetizers` na kuwekea thamani `37.89`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

Unapaswa kufafanua kigezo kinachoitwa `main_courses` na kuwekea thamani `57.34`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

Unapaswa kufafanua kigezo kinachoitwa `desserts` na kuwekea thamani `39.39`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

Unapaswa kufafanua kigezo kinachoitwa `drinks` na kuwekea thamani `64.21`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
