---
id: 6977a8bbcbdd557886bd0718
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Sasa baada ya kujumuisha tip, una jumla ya mwisho kwa kundi lote. Lazima uamue ni kiasi gani kila mtu analazimika kulipa kwa kugawanya bili yote kwa idadi ya marafiki.

Katika Python, unatumia alama ya mstari wa mbele `/` kufanya mgawanyo. Kwa mfano:

```py
half = 10 / 2
```

Tengeneza kigezo kinachoitwa `final_bill` na ukaweke thamani ya matokeo ya kugawanya `running_total` kwa `num_of_friends`.

Mwishowe, tumia kitendakazi cha `print()` kuonyesha mfuatano wa herufi `Bill per person:` ukifuatiwa na nafasi na thamani ya `final_bill`.

# --hints--

Unapaswa kufafanua kigezo kinachoitwa `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

Unapaswa kutumia kielezi cha `/` kugawanya `running_total` kwa `num_of_friends` na kuweka matokeo kwenye kigezo chako cha `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

Unapaswa kuchapisha mfuatano wa herufi `Bill per person:` ukifuatiwa na nafasi na kigezo cha `final_bill`.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Bill per person: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - (((37.89 + 57.34 + 39.39 + 64.21) * 1.25) / 4)) < 1e-6
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
