---
id: 6977a4306bb3e856690a4890
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Sasa baada ya kuhesabu tip, unahitaji kuiongeza kwenye `running_total` ili kupata jumla ya bili ya mwisho.

Tumia opereta `+=` kuongeza thamani ya `tip` kwenye `running_total` yako. Mwishowe, tumia `print()` kuonyesha mfuatano wa herufi `Total with tip:` ukifuatiwa na nafasi na thamani ya `running_total`.

# --hints--

Unapaswa kutumia opereta `+=` kuongeza thamani ya `tip` kwenye `running_total` yako.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

Unapaswa kuchapisha mfuatano wa herufi `Total with tip:` ukifuatiwa na nafasi na variable ya `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total with tip:', running_total)") or _Node(_code).has_call("print(f'Total with tip: {running_total}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
