---
id: 6a89c69c99298438d6fbe60b
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

再帰ケースを追加しましょう。`if`文の後に、`count_list`という名前の変数を作成してください。`number - 1`を引数にして`countup`を呼び出した結果を代入します。

各呼び出しはより小さい数を使うので、再帰は基本ケースに向かって進みます。

# --hints--

`countup(number - 1)`を`count_list`という変数に代入してください。

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
