---
id: 6a89c69c91bc7365fbd4b733
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

`countup` 関数に基本ケースを追加してください。`pass` を、`number` が `1` 未満かどうかをチェックする `if` 文に置き換えてください。`if` 文の中で空のリストを返してください。

# --hints--

`countup` 関数は `number` が `1` 未満の場合、空のリストを返す必要があります。

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

`countup` 関数は `number` が `1` 以上の場合、空のリストを返してはいけません。

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
