---
id: 6a89c69c8dbfda74151fa0b6
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、再帰を使って `1` から指定された数までの数字のリストを作成する `countup` 関数を作成します。

`countup` という名前の関数を定義し、`number` という名前のパラメータを1つ持たせることから始めてください。関数の本体には `pass` キーワードを追加し、空のままでも関数が有効になるようにしてください。

# --hints--

`countup` という名前の関数があり、`number` というパラメータを持っている必要があります。

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

`countup` 関数には `pass` キーワードが含まれている必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
