---
id: 6a89c69c6af682856558d78b
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

再帰呼び出しは、より小さい数字を含むリストを返します。`append()`を使って`number`を`count_list`に追加し、その後`count_list`を返してください。

# --hints--

`number`を`count_list`に追加してください。

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

`countup`関数から`count_list`を返してください。

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

`countup`関数は`3`を渡して呼び出したときに`[1, 2, 3]`を返す必要があります。

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
