---
id: 6a89c69c111694edc6d1c628
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`print()` の呼び出し内で `countup` 関数に `5` を渡してテストしてください。出力は `[1, 2, 3, 4, 5]` になるはずです。

これで countup 関数のワークショップは完了です。

# --hints--

`print()` 関数の中で `countup(5)` を呼び出してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
