---
id: 5cd9a70215d3c4e65518328f
title: カウントダウンを作成する
challengeType: 26
dashedName: build-a-countdown
---

# --description--

このラボでは、開始番号を受け取りカウントダウンのリストを生成する`countdown`関数を作成します。

`countdown(n)`を呼び出すと、`n`から始まり`1`ずつカウントダウンして`1`に達するまでの配列を返す必要があります。

例えば：

- `countdown(5)`は`[5, 4, 3, 2, 1]`を返すべきです
- `countdown(1)`は`[1]`を返すべきです

開始番号が`1`より小さい場合はカウントダウンするものがないので、空の配列を返してください：

- `countdown(0)`は`[]`を返すべきです
- `countdown(-3)`は`[]`を返すべきです

ラボを完了するには、結果を再帰を使って構築し、ループ（`for`、`while`）、`Array.from()`、または配列の反復メソッド（`forEach`、`map`、`filter`、`reduce`）を使わないでください。各呼び出しは独自の結果配列を生成する必要があります（グローバル変数を使ってカウントダウンを保存しないでください）。

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `countdown`という名前の関数を作成してください。
2. `countdown`関数は1つの引数`n`（数値）を受け取るべきです。
3. `n`が`1`より小さい場合、`countdown`は空の配列を返すべきです。
4. それ以外の場合、`countdown`は`n`から`1`までの整数を降順で含む配列を返すべきです。
5. 再帰を使い（関数が自分自身を呼び出す）、どんな種類のループ（`for`、`while`、`Array.from()`）や高階メソッド（`forEach`、`map`、`filter`、`reduce`）も使わないでください。
6. 各再帰呼び出しは前の呼び出しより小さい値（例えば`n - 1`）を使い、関数が基本ケースに到達するようにしてください。
7. 結果の配列を保存またはキャッシュするためにグローバル変数を使わないでください。
8. `countdown`を異なる入力で複数回呼び出しても、常に正しく独立した結果を返すべきです。

# --hints--

`countdown`という名前の関数を作成してください。

```js
assert.isFunction(countdown);
```

`countdown(-1)`は空の配列を返すべきです。

```js
assert.isEmpty(countdown(-1));
```

`countdown(10)`は`[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`を返すべきです。

```js
assert.deepStrictEqual(countdown(10), [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

`countdown(5)`は`[5, 4, 3, 2, 1]`を返すべきです。

```js
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

コードはどんな種類のループ（`for`、`while`、`Array.from()`）や高階関数（`forEach`、`map`、`filter`、`reduce`）にも依存しないようにしてください。

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

この問題を解くには再帰を使ってください。

```js
assert.isAtLeast(
  __helpers.removeJSComments(countdown.toString()).match(/countdown\s*\(.+\)/g)?.length,
  2
);
```

配列をキャッシュするためにグローバル変数を使わないでください。

```js
countdown(1)
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function countdown(n) {
  if (n < 1) {
    return [];
  }

  return [n, ...countdown(n - 1)];
}
```
