---
id: 615f34948891834dd77655a6
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

為你的文件添加第二個帶有文字 `Serving size 2/3 cup (55g)` 的 `p` 元素。

# --hints--

你應該有第二個 `p` 元素。

```js
assert(document.querySelectorAll('p')?.length === 2);
```

你的第二個 `p` 元素應該在你的 `body` 元素內。

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

你的第二個 `p` 元素應該放在你現有的 `p` 元素之後。

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

你的第二個 `p` 元素應該有文字 `Serving size 2/3 cup (55g)`。

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
