---
id: 615f4172e9eec061d6456221
title: 步驟 19
challengeType: 0
dashedName: step-19
---

# --description--

營養成分表有很多粗體文字以引起人們對重要信息的注意。 與其針對每個需要加粗的元素，不如使用一個類將粗體樣式應用於每個元素更有效。

將你的第二個 `p` 元素的 `class` 屬性設置爲 `bold`。

# --hints--

你的第二個 `p` 元素應將 `class` 屬性設置爲 `bold`。

```js
assert(document.querySelectorAll('p')?.[1]?.classList?.contains('bold'));
```

你的第一個 `p` 元素不應該有設置爲 `bold` 的 `class` 屬性。

```js
assert(!document.querySelector('p')?.classList?.contains('bold'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}
```
