---
id: 615f2abbe7d18d49a1e0e1c8
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

你已經獲得一個基本的 HTML 範本。

在你的 `body` 元素中創建一個 `h1` 元素，並給它 `Nutrition Facts` 的文字。

# --hints--

你應該添加一個新的 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素應該位於你的 `body` 元素內。

```js
assert.equal(document.querySelector('h1')?.parentElement?.localName, 'body');
```

你的 `h1` 元素應該有 `Nutrition Facts` 文字。

```js
assert.equal(document.querySelector('h1')?.innerText, 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
