---
id: 615f3b091162165948e1cb33
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

如果你使用瀏覽器的開發者工具檢查你的 `.label` 元素，你可能會注意到它實際上寬度是 288 像素，而不是 270。這是因為瀏覽器在預設情況下會將邊框和內距納入元素大小的計算。

為了解決這個問題，透過建立一個 `*` 選擇器並給予它 `box-sizing` 屬性為 `border-box` 來重設盒模型。

# --hints--

你應該創建一個 `*` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

你的 `*` 選擇器應該設定 `box-sizing` 屬性為 `border-box`。

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
