---
id: 615f3ed16592445e57941aec
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

分隔線可以幫助區分並群組重要內容，特別是在空間有限時。

在你的 `h1` 元素下方創建一個 `div` 元素，並為它的 `class` 屬性設定為 `divider`。

# --hints--

你應該創建一個新的 `div`。

```js
assert(document.querySelectorAll('div')?.length === 2);
```

你的新 `div` 應該將 `class` 屬性設為 `divider`。

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('divider'));
```

你的 `.divider` 元素應該在你的 `.label` 元素內。

```js
assert(document.querySelector('.label')?.querySelector('.divider'));
```

你的 `.divider` 元素應該放在你的 `h1` 元素之後。

```js
assert(document.querySelector('.divider')?.previousElementSibling?.localName === 'h1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}
```
