---
id: 615f50473cc0196c6dd3892a
title: 步驟 28
challengeType: 0
dashedName: step-28
---

# --description--

你可能會注意到你的 `.large` 元素底部仍有一個小邊框。要重設它，請為你的 `.large, .medium` 選擇器設定一個 `border` 屬性，並將其設為 `0`。

注意：`medium` 類別將在稍後用於營養標籤中較細的條狀。 

# --hints--

你的 `.large, .medium` 選擇器應該將 `border` 屬性設為 `0`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.large, .medium')?.borderWidth === '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

--fcc-editable-region--
.large, .medium {
  background-color: black;
}
--fcc-editable-region--
```
