---
id: 6140c7e645d8e905819f1dd4
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Um zu beginnen, fügen Sie ein `link`-Element für die `./styles.css`-Datei hinzu.

# --hints--

Ihr Code sollte ein `link`-Element enthalten.

```js
assert.match(code, /<link/);
```

Sie sollten ein `link`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > link'));
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>

  </body>
</html>

```

```css

```
