---
id: 6140e0d875ec16262f26432b
title: Schritt 21
challengeType: 0
dashedName: step-21
---

# --description--

Die Eigenschaft `animation-iteration-count` legt fest, wie oft Ihre Animation wiederholt werden soll. Dies kann auf eine Zahl gesetzt werden oder auf `infinite`, um die Animation unendlich oft zu wiederholen. Ihr Riesenrad sollte niemals anhalten, also setzen Sie den Selektor `.wheel` so, dass er eine Eigenschaft `animation-iteration-count` mit dem Wert `infinite` hat.

Die Eigenschaft `animation-timing-function` bestimmt, wie die Animation im Laufe der Zeit ablaufen soll. Es gibt verschiedene Werte für diese Eigenschaft, aber Sie möchten, dass die Riesenrad-Animation von Anfang bis Ende mit der gleichen Geschwindigkeit läuft. Setzen Sie die Eigenschaft `animation-timing-function` auf `linear` in Ihrem Selektor `.wheel`.

# --hints--

Ihr Selektor `.wheel` sollte eine Eigenschaft `animation-iteration-count` mit dem Wert `infinite` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationIterationCount, 'infinite');
```

Ihr Selektor `.wheel` sollte eine Eigenschaft `animation-timing-function` mit dem Wert `linear` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationTimingFunction, 'linear');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
  animation-name: wheel;
  animation-duration: 10s;
--fcc-editable-region--
  
--fcc-editable-region--
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

@keyframes wheel {
  0% {
    transform: rotate(0deg);
  }
  100% {
    transform: rotate(360deg);
  }
}
```
