---
id: 6140f4b5c1555a2960de1e5f
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Erstellen Sie eine weitere `@keyframes`-Regel mit dem Namen `cabins`. Verwenden Sie dieselben Eigenschaften wie bei Ihrer `@keyframes wheel`, indem Sie sowohl die `0%`- als auch die `100%`-Regeln kopieren, aber setzen Sie die `transform`-Eigenschaft des `100%`-Selektors auf `rotate(-360deg)`.

# --hints--

Sie sollten eine `@keyframes`-Regel haben.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 2);
```

Ihre neue `@keyframes`-Regel sollte den Namen `cabins` tragen.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes');
assert.isTrue(rules?.[0]?.name === 'cabins' || rules?.[1]?.name === 'cabins');
```

Ihre neue `@keyframes`-Regel sollte nach Ihrer `@keyframes wheel`-Regel kommen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.name, 'cabins');
```

Sie sollten den Namen Ihrer `@keyframes wheel`-Regel nicht ändern.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'wheel');
```

Ihre `@keyframes cabins`-Regel sollte einen `0%`-Selektor haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules?.[0]?.keyText, '0%');
```

Ihr `0%`-Selektor sollte eine `transform`-Eigenschaft mit dem Wert `rotate(0deg)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules?.[0]?.style?.transform, 'rotate(0deg)');
```

Ihre `@keyframes cabins`-Regel sollte einen `100%`-Selektor haben.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

Ihr `100%`-Selektor sollte nach Ihrem `0%`-Selektor kommen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules?.[1]?.keyText, '100%')
```

Ihr `100%`-Selektor sollte eine `transform`-Eigenschaft mit dem Wert `rotate(-360deg)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules?.[1]?.style?.transform, 'rotate(-360deg)')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
  animation-name: wheel;
  animation-duration: 10s;
  animation-iteration-count: infinite;
  animation-timing-function: linear;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

@keyframes wheel {
   0% {
     transform: rotate(0deg);
   }
   100% {
     transform: rotate(360deg);
   }
}

--fcc-editable-region--

--fcc-editable-region--
```
