---
id: 6140d4bc9db3c81c51a09ab7
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Setzen Sie das Muster fort, wählen Sie die folgenden `.cabin`-Elemente aus und wenden Sie die spezifischen Regeln darauf an:

- Das zweite `.cabin` sollte die `right`-Eigenschaft auf `17%` und die `top`-Eigenschaft auf `93.5%` gesetzt haben.
- Das dritte `.cabin` sollte die `right`-Eigenschaft auf `67%` und die `top`-Eigenschaft auf `93.5%` gesetzt haben.
- Das vierte `.cabin` sollte die `left`-Eigenschaft auf `-8.5%` und die `top`-Eigenschaft auf `50%` gesetzt haben.
- Das fünfte `.cabin` sollte die `left`-Eigenschaft auf `17%` und die `top`-Eigenschaft auf `7%` gesetzt haben.
- Das sechste `.cabin` sollte die `right`-Eigenschaft auf `17%` und die `top`-Eigenschaft auf `7%` gesetzt haben.

# --hints--

Sie sollten einen `.cabin:nth-of-type(2)`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)'));
```

Ihr `.cabin:nth-of-type(2)`-Selektor sollte eine `right`-Eigenschaft mit dem Wert `17%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.right, '17%');
```

Ihr `.cabin:nth-of-type(2)`-Selektor sollte eine `top`-Eigenschaft mit dem Wert `93.5%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.top, '93.5%');
```

Sie sollten einen `.cabin:nth-of-type(3)`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)'));
```

Ihr `.cabin:nth-of-type(3)`-Selektor sollte eine `right`-Eigenschaft mit dem Wert `67%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.right,'67%');
```

Ihr `.cabin:nth-of-type(3)`-Selektor sollte eine `top`-Eigenschaft mit dem Wert `93.5%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.top, '93.5%');
```

Sie sollten einen `.cabin:nth-of-type(4)`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)'));
```

Ihr `.cabin:nth-of-type(4)`-Selektor sollte eine `left`-Eigenschaft mit dem Wert `-8.5%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.left, '-8.5%');
```

Ihr `.cabin:nth-of-type(4)`-Selektor sollte eine `top`-Eigenschaft mit dem Wert `50%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.top, '50%');
```

Sie sollten einen `.cabin:nth-of-type(5)`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)'));
```

Ihr `.cabin:nth-of-type(5)`-Selektor sollte eine `left`-Eigenschaft mit dem Wert `17%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.left, '17%');
```

Ihr `.cabin:nth-of-type(5)`-Selektor sollte eine `top`-Eigenschaft mit dem Wert `7%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.top, '7%');
```

Sie sollten einen `.cabin:nth-of-type(6)`-Selektor haben.

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s)
assert.exists(def('.cabin:nth-of-type(6)') || def('.cabin:last-of-type'));
```

Ihr `.cabin:nth-of-type(6)`-Selektor sollte eine `right`-Eigenschaft mit dem Wert `17%` haben.

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(6)') || right('.cabin:last-of-type'),'17%');
```

Ihr `.cabin:nth-of-type(6)`-Selektor sollte eine `top`-Eigenschaft mit dem Wert `7%` haben.

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(6)') || top('.cabin:last-of-type'),'7%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
--fcc-editable-region--

--fcc-editable-region--
```
