---
id: 6140d1a351e88f159ed54fca
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Verwenden Sie dasselbe Muster, um einen separaten Selektor für das dritte `.line`, das vierte `.line`, das fünfte `.line` und das sechste `.line` zu erstellen.

Setzen Sie die `transform`-Eigenschaft für das dritte `.line` auf `rotate(120deg)`, für das vierte auf `rotate(180deg)`, für das fünfte auf `rotate(240deg)` und für das sechste auf `rotate(300deg)`.

# --hints--

Sie sollten einen `.line:nth-of-type(3)`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

Ihr `.line:nth-of-type(3)`-Selektor sollte eine `transform`-Eigenschaft haben, die auf `rotate(120deg)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform, 'rotate(120deg)');
```

Sie sollten einen `.line:nth-of-type(4)`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

Ihr `.line:nth-of-type(4)`-Selektor sollte eine `transform`-Eigenschaft haben, die auf `rotate(180deg)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform, 'rotate(180deg)');
```

Sie sollten einen `.line:nth-of-type(5)`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

Ihr `.line:nth-of-type(5)`-Selektor sollte eine `transform`-Eigenschaft haben, die auf `rotate(240deg)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform, 'rotate(240deg)');
```

Sie sollten einen `.line:nth-of-type(6)`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

Ihr `.line:nth-of-type(6)`-Selektor sollte eine `transform`-Eigenschaft haben, die auf `rotate(300deg)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform, 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
