---
id: 6140cbeec34e970dfe75e710
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Erstellen Sie einen Selektor für Ihr `.wheel`-Element. Beginnen Sie damit, die `border` auf `2px solid black`, die `border-radius` auf `50%` und die `margin-left` auf `50px` zu setzen.

# --hints--

Sie sollten einen `.wheel`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Ihr `.wheel`-Selektor sollte eine `border`-Eigenschaft mit dem Wert `2px solid black` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.border, '2px solid black');
```

Ihr `.wheel`-Selektor sollte eine `border-radius`-Eigenschaft mit dem Wert `50%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius, '50%');
```

Ihr `.wheel`-Selektor sollte eine `margin-left`-Eigenschaft mit dem Wert `50px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft, '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
