---
id: 6140c9d35015ae0ba0c250e8
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie ein `div` innerhalb Ihres `body`-Elements hinzu und geben Sie ihm ein `class` von `wheel`.

Fügen Sie innerhalb Ihres neuen `div` sechs `span`-Elemente mit einem `class`, das auf `line` gesetzt ist, und sechs `div`-Elemente mit einem `class`, das auf `cabin` gesetzt ist, hinzu.

# --hints--

Sie sollten ein neues `div` innerhalb Ihres `body`-Elements erstellen.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert.lengthOf(divs, 1);
```

Ihr `div` innerhalb Ihres `body`-Elements sollte ein `class` von `wheel` haben.

```js
assert.isTrue(document.querySelector('body div')?.classList?.contains('wheel'));
```

Ihr `.wheel`-Element sollte sechs `span`-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('.wheel span'), 6);
```

Ihre sechs `span`-Elemente innerhalb des `.wheel`-Elements sollten ein `class` von `line` haben.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert.isTrue(spans?.every(span => span?.classList?.contains('line')));
assert.lengthOf(document.querySelectorAll('.line'), 6);
```

Ihr `.wheel`-Element sollte sechs `div`-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('.wheel div'), 6);
```

Ihre sechs `div`-Elemente innerhalb des `.wheel`-Elements sollten ein `class` von `cabin` haben.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert.isTrue(divs?.every(div => div?.classList?.contains('cabin')));
assert.lengthOf(document.querySelectorAll('.cabin'), 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
