---
id: 6140d0069049f5139d78da40
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Die `transform-origin`-Eigenschaft wird verwendet, um den Punkt festzulegen, um den eine CSS-Transformation angewendet wird. Zum Beispiel bestimmt bei einer `rotate`-Transformation (wie Sie sie später in diesem Projekt anwenden werden) die `transform-origin`, um welchen Punkt das Element rotiert wird.

Geben Sie dem `.line`-Selektor eine `transform-origin`-Eigenschaft mit dem Wert `0% 0%`. Dadurch wird der Ursprungspunkt um `0%` von links und `0%` von oben versetzt und auf die obere linke Ecke des Elements gesetzt.

# --hints--

Ihr `.line`-Selektor sollte eine `transform-origin`-Eigenschaft mit dem Wert `0% 0%` haben.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert.oneOf(transformOrigin, ['0% 0%', '0% 0% 0px']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
