---
id: 617ace7d831f9c16a569b38a
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Geben Sie Ihrem `.wheel`-Selektor eine `max-height`- und `max-width`-Eigenschaft, die beide auf `500px` gesetzt sind.

# --hints--

Ihr `.wheel`-Selektor sollte eine `max-height`-Eigenschaft haben, die auf `500px` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxHeight, '500px');
```

Ihr `.wheel`-Selektor sollte eine `max-width`-Eigenschaft haben, die auf `500px` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxWidth, '500px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
