---
id: a97fd23d9b809dac9921074f
title: Crea una funzione somma con argomenti opzionali
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

In questo laboratorio creerai una funzione che accetta fino a due argomenti e li somma, ma se c'è solo un argomento restituisce una funzione che aspetta il secondo numero da sommare.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Crea una funzione `addTogether`.
2. Se la funzione riceve due argomenti, dovrebbe restituire la somma dei due argomenti.
3. Se viene fornito solo un argomento, `addTogether` dovrebbe restituire una funzione.
   - Quando la funzione restituita viene chiamata con un singolo argomento, dovrebbe restituire la somma.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. Se uno dei due argomenti non è un numero valido, restituisci `undefined`.

# --hints--

Dovresti avere la funzione `addTogether`.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` dovrebbe restituire 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` dovrebbe restituire 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` dovrebbe restituire `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` dovrebbe restituire `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` dovrebbe restituire `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` dovrebbe restituire una funzione.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` dovrebbe restituire 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` dovrebbe restituire `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` dovrebbe restituire `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
