---
id: 61fd75ea7f663457612dba02
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Within each of your new `th` elements, nest a `span` element with the `class` set to `sr-only year`. Give them the following text (in order): `2019`, `2020`, and `2021`.

Dai al terzo elemento `th` l'attributo `class` con il valore `current`.

Lascia vuoto l'elemento `td`. Questo elemento esiste solo per garantire che la tabella abbia un layout a quattro colonne e associare le intestazioni alle colonne corrette.

# --hints--

Ognuno degli elementi `th` dovrebbe avere un elemento `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.lengthOf(th?.children, 1);
  assert.strictEqual(th?.children?.[0]?.localName, 'span');
});
```

Ognuno dei nuovi elementi `span` dovrebbe avere l'attributo `class` con il valore `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.isTrue(th?.children?.[0]?.classList?.contains('sr-only'));
  assert.isTrue(th?.children?.[0]?.classList?.contains('year'));
});
```

Il primo elemento `span` dovrebbe avere il testo `2019`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent.trim(),
  '2019'
);
```

Il secondo elemento `span` dovrebbe avere il testo `2020`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent.trim(),
  '2020'
);
```

Il terzo elemento `span` dovrebbe avere il testo `2021`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent.trim(),
  '2021'
);
```

Il terzo elemento `th` dovrebbe avere l'attributo `class` con il valore `current`.

```js
assert.isTrue(
  document
    .querySelector('table')
    ?.querySelectorAll('th')?.[2]
    ?.classList?.contains('current')
);
```

L'elemento `td` dovrebbe essere vuoto.

```js
const table = document.querySelector('table');
const tableData = table.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '');
assert.isEmpty(tableData?.children);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
