---
id: 61fd67a656743144844941cb
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

I lettori di schermo annunciano gli elementi HTML basandosi sul flusso del documento. Alla fine vorremo che il bilancio abbia un’intestazione "Bilancio" e una sottointestazione "AcmeWidgetCorp". Tuttavia, questo ordine non ha senso se annunciato da un lettore di schermo.

Dai al tuo `span` esistente l’attributo `class` impostato su `flex`, e aggiungi due elementi `span` al suo interno. Dai al primo il testo `AcmeWidgetCorp`. Dai al secondo il testo `Balance Sheet`. Userai il CSS per invertire l’ordine del testo sulla pagina, ma l’ordine HTML avrà più senso per un lettore di schermo.

# --hints--

Il tuo elemento `span` esistente dovrebbe avere l’attributo `class` impostato su `flex`.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

Il tuo elemento `span` esistente dovrebbe avere due nuovi elementi `span` al suo interno.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

I tuoi nuovi elementi `span` non dovrebbero avere un attributo `class`.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

Il tuo primo nuovo elemento `span` dovrebbe avere il testo `AcmeWidgetCorp`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

Il tuo secondo nuovo elemento `span` dovrebbe avere il testo `Balance Sheet`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
