---
id: 61fd5a93fd62bb35968adeab
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Inizia il tuo progetto dando al tuo elemento `head` un elemento `link` per il tuo foglio di stile.

# --hints--

Il tuo codice dovrebbe avere un elemento `link`.

```js
assert.match(code, /<link/);
```

Il tuo elemento `link` dovrebbe essere all'interno del tuo elemento `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
