---
id: 61fd6cc9475a784b7776233e
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Sotto il tuo elemento `div` esistente, aggiungi un nuovo elemento `div` con un insieme `class` impostato su `table-wrap`. Questo sarà il contenitore per le tue tabelle.

Al suo interno, aggiungi tre elementi `table`. Userai il CSS per definire lo stile in modo che appaiano come un'unica tabella, ma usare tabelle separate aiuterà i lettori di schermo a comprendere il flusso del documento.

# --hints--

Dovresti creare un nuovo elemento `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Il tuo nuovo elemento `div` dovrebbe avere l'insieme `class` impostato su `table-wrap`.

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

Il tuo elemento `.table-wrap` dovrebbe venire dopo il tuo elemento `div` esistente.

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

Il tuo elemento `.table-wrap` dovrebbe avere tre elementi `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
