---
id: 61fd6b7c83dbf54a08cf0498
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

All'interno del tuo elemento `div`, aggiungi tre elementi `span`. Dai a ciascuno un attributo `class` impostato su `year` e aggiungi il seguente testo (in ordine): `2019`, `2020` e `2021`.

# --hints--

Il tuo elemento `div` dovrebbe avere tre elementi `span`.

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

Ogni elemento `span` dovrebbe avere un attributo `class` impostato su `year`.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

Il tuo primo `span` dovrebbe avere il testo `2019`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

Il tuo secondo `span` dovrebbe avere il testo `2020`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

Il tuo terzo `span` dovrebbe avere il testo `2021`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
