---
id: 66db57ad34c7089b9b41bfd6
title: створіть переглядач lightbox
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

Lightbox використовується на вебсайтах для більш захопливого показу мультимедійного вмісту через спливаюче або модальне вікно поверх основного вмісту сторінки.

У цій лабораторній роботі ви створите галерею lightbox, яка відображає зображення повного розміру при натисканні на мініатюру. Для кожного зображення надаються дві версії: мініатюра та зображення повного розміру. Зображення повного розміру таке саме, як і мініатюра, але без суфікса `-thumbnail`.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас повинен бути `div` з класом `gallery` всередині вашого `body`.
1. Всередині елемента `.gallery` має бути три мініатюри зображень, кожна з класом `gallery-item`. Використайте такі посилання для мініатюр:
   
    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. У вас повинен бути `div` з класом `lightbox` всередині вашого `body`.
1. У вашому елементі `button` має бути `id` з `close-btn`, встановленим на `.lightbox`. Ви можете використати `&times;` як текст, якщо хочете.
1. У вашому елементі `img` має бути `id` з `lightbox-image`, встановленим на `.lightbox`.
1. Ваш елемент `.lightbox` повинен мати фіксоване позиціонування, щоб попередній перегляд відкривався поверх поточних зображень.
1. Ваш елемент `.lightbox` має покривати весь вікно перегляду, встановивши висоту та ширину на 100% контейнера. Переконайтеся, що елемент `.lightbox` починається у верхньому лівому куті контейнера.
1. `.lightbox` повинен мати колір фону. Спочатку його властивість `display` має бути встановлена на `none`, щоб приховати його.
1. Коли ви натискаєте на один із ваших елементів `.gallery-item`, властивість `.lightbox` елемента `display` має бути встановлена на `flex`, щоб зробити видимими елемент `.lightbox` та два елементи всередині нього.
1. Коли ви натискаєте на один із ваших елементів `.gallery-item`, властивість `#lightbox-image` елемента `src` має бути встановлена на версію зображення повного розміру, видаливши `-thumbnail` з атрибута `src` зображення. Зображення повного розміру знаходяться за такими посиланнями:

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`
  
1. Коли ваш елемент `.lightbox` видимий і ви натискаєте на елемент `#close-btn` або `.lightbox`, властивість `.lightbox` елемента `display` має бути повернена до `none`.

**Примітка:** Обов’язково підключіть свій файл стилів і JavaScript у вашому HTML.

# --hints--

У вашому елементі `div` має бути `gallery` з класом `body`.

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

У елементі `.gallery` має бути три елементи `img` з класом `gallery-item`.

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

У елементі `.gallery` має бути елемент `img` з `src`, встановленим на `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

У елементі `.gallery` має бути елемент `img` з `src`, встановленим на `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

У елементі `.gallery` має бути елемент `img` з `src`, встановленим на `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

У вашому елементі `div` має бути елемент `lightbox` з класом `body`.

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

У вашому елементі `.lightbox` має бути елемент `button` з `id`, встановленим на `close-btn`.

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

У вашому елементі `.lightbox` має бути елемент `img` з `id`, встановленим на `lightbox-image`.

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

Ваш елемент `.lightbox` має мати фіксоване позиціонування.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

Ваш елемент `.lightbox` має покривати весь вікно перегляду.

```js
const style = new __helpers.CSSHelp(document).getStyle(".lightbox");

const width = style?.width;
const height = style?.height;

assert.oneOf(
  width,
  ["100%", "100vw"]
);

assert.oneOf(
  height,
  ["100%", "100vh"]
);
```

Ваш елемент `.lightbox` має бути вирівняний у верхньому лівому куті контейнера.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

Ваш елемент `.lightbox` має мати колір фону.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

Ваш елемент `.lightbox` має бути прихований спочатку.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

Коли ви натискаєте на один із ваших елементів `.gallery-item`, властивість `.lightbox` елемента `display` має бути встановлена на `flex`, щоб зробити видимими `.lightbox` та два елементи всередині нього.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

Коли ви натискаєте на один із ваших елементів `.gallery-item`, властивість `#lightbox-image` елемента `src` має бути встановлена на версію зображення повного розміру, видаливши `-thumbnail` з атрибута `src` зображення.

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

Коли ваш елемент `.lightbox` видимий і ви натискаєте кнопку `#close-btn`, властивість `.lightbox` елемента `display` має бути повернена до `none`.

```js
const lightbox = document.querySelector(".lightbox");
const closeBtn = document.getElementById("close-btn"); 
const galleryItem = document.querySelector(".gallery-item");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

closeBtn.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

Коли ваш елемент `.lightbox` видимий і ви натискаєте на елемент `.lightbox`, властивість `.lightbox` елемента `display` має бути повернена до `none`.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

lightbox.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```
