---
id: 681cb1a2dab50c87ddb2e514
title: "Desafio 3: Sequência de Fibonacci"
challengeType: 29
dashedName: challenge-3
---

# --description--

A sequência de Fibonacci é uma série de números onde cada número é a soma dos dois anteriores. Ao começar com `0` e `1`, os primeiros 10 números da sequência são `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21` e `34`.

Dado um array contendo os dois primeiros números de uma sequência de Fibonacci e um inteiro representando o comprimento da sequência, retorne um array contendo a sequência do comprimento dado.

- Sua função deve lidar com sequências de qualquer comprimento maior ou igual a zero.
- Se o comprimento for zero, retorne um array vazio.
- Observe que os números iniciais fazem parte da sequência.

# --hints--


`fibonacci_sequence([0, 1], 20)` deve retornar `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181])`)
}})
```

`fibonacci_sequence([21, 32], 1)` deve retornar `[21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([21, 32], 1), [21])`)
}})
```

`fibonacci_sequence([0, 1], 0)` deve retornar `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 0), [])`)
}})
```

`fibonacci_sequence([10, 20], 2)` deve retornar `[10, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([10, 20], 2), [10, 20])`)
}})
```

`fibonacci_sequence([123456789, 987654321], 5)` deve retornar `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541])`)
}})
```

# --seed--

## --seed-contents--

```py
def fibonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def fibonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2]
        sequence.append(next_value)

    return sequence
```
