---
id: 681cb1b0dab50c87ddb2e518
title: "Desafio 7: Soma Alvo"
challengeType: 29
dashedName: challenge-7
---

# --description--

Dado um array de números e um inteiro target, encontre dois números únicos no array que somem o valor target. Retorne um array com os índices desses dois números ou `"Target not found"` se nenhum par de números somar o target.

- O array retornado deve ter os índices em ordem crescente.

# --hints--

`find_target([2, 7, 11, 15], 9)` deve retornar `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([2, 7, 11, 15], 9), [0, 1])`)
}})
```

`find_target([3, 2, 4, 5], 6)` deve retornar `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([3, 2, 4, 5], 6), [1, 2])`)
}})
```

`find_target([1, 3, 5, 6, 7, 8], 15)` deve retornar `[4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 6, 7, 8], 15), [4, 5])`)
}})
```

`find_target([1, 3, 5, 7], 14)` deve retornar `'Alvo não encontrado'`.

```js

({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 7], 14), "Target not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_target(arr, target):

    return arr
```

# --solutions--

```py
def find_target(arr, target):
    for i in range(len(arr)):
        for j in range(i + 1, len(arr)):
            if arr[i] + arr[j] == target:
                return [i, j]
    return 'Target not found'
```
