---
id: 681cb1afdab50c87ddb2e517
title: "Desafio 6: Verificador de Anagramas"
challengeType: 29
dashedName: challenge-6
---

# --description--

Dadas duas strings, determine se elas são anagramas uma da outra (contêm os mesmos caracteres em qualquer ordem).

- Ignore maiúsculas, minúsculas e espaços em branco.

# --hints--

`are_anagrams("listen", "silent")` deve retornar `true`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("listen", "silent"), True)`)
}})
```

`are_anagrams("School master", "The classroom")` deve retornar `true`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("School master", "The classroom"), True)`)
}})
```

`are_anagrams("A gentleman", "Elegant man")` deve retornar `true`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("A gentleman", "Elegant man"), True)`)
}})
```

`are_anagrams("Hello", "World")` deve retornar `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("Hello", "World"), False)`)
}})
```

`are_anagrams("apple", "banana")` deve retornar `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("apple", "banana"), False)`)
}})
```

`are_anagrams("cat", "dog")` deve retornar `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("cat", "dog"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def are_anagrams(str1, str2):

    return str1
```

# --solutions--

```py
def are_anagrams(str1, str2):
    def clean(s):
        return sorted(s.replace(" ", "").lower())
    
    return clean(str1) == clean(str2)
```
