---
id: 681cb1b0dab50c87ddb2e519
title: "Desafio 8: Factorializer"
challengeType: 29
dashedName: challenge-8
---

# --description--

Dado um inteiro de zero a 20, retorne o fatorial desse número. O fatorial de um número é o produto de todos os números entre 1 e o número dado.

- O fatorial de zero é 1.

# --hints--

`factorial(0)` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(0), 1)`)
}})
```

`factorial(5)` deve retornar `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(5), 120)`)
}})
```

`factorial(20)` deve retornar `2432902008176640000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(20), 2432902008176640000)`)
}})
```

# --seed--

## --seed-contents--

```py
def factorial(n):

    return n
```

# --solutions--

```py
def factorial(n):
    result = 1
    for i in range(1, n + 1):
        result *= i
    return result
```
