---
id: 6821ebce237de8297eaee790
title: "Desafio 14: Batalha de Personagens"
challengeType: 29
dashedName: challenge-14
---

# --description--

Dadas duas strings representando seu exército e um exército oposto, cada caractere do seu exército batalha contra o caractere na mesma posição do exército oposto usando as seguintes regras:

- Os caracteres `a-z` têm uma força de `1-26`, respectivamente.
- Os caracteres `A-Z` têm uma força de `27-52`, respectivamente.
- Dígitos `0-9` têm uma força igual ao seu valor nominal.
- Todos os outros caracteres têm valor zero.
- Cada personagem pode lutar em apenas uma batalha.

Para cada batalha, o personagem mais forte vence. O exército com mais vitórias vence a guerra. Retorne os seguintes valores:

- "`Opponent retreated` se seu exército tiver mais personagens do que o exército oposto."
- "Recuamos" se o exército adversário tiver mais personagens do que o seu.
- "Ganharam" se seu exército venceu mais batalhas.
- "We lost" se o exército adversário venceu mais batalhas.
- "Foi um empate" se ambos os exércitos venceram o mesmo número de batalhas.

# --hints--

`battle("Hello", "World")` deve retornar `"We lost"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` deve retornar `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` deve retornar `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` deve retornar `"Opponent retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` deve retornar `"We retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` deve retornar `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` deve retornar `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
