---
id: 6821ec02237de8297eaee79a
title: "Desafio 24: Pangrama"
challengeType: 29
dashedName: challenge-24
---

# --description--

Dada uma palavra ou frase e uma string de letras minúsculas, determine se a palavra ou frase usa todas as letras do conjunto fornecido pelo menos uma vez e nenhuma outra letra.

- Ignore caracteres não alfabéticos na palavra ou frase.
- Ignore a diferenciação entre maiúsculas e minúsculas na palavra ou frase.

# --hints--

`is_pangram("hello", "helo")` deve retornar `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "helo"), True)`)
}})
```

`is_pangram("hello", "hel")` deve retornar `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "hel"), False)`)
}})
```

`is_pangram("hello", "helow")` deve retornar `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "helow"), False)`)
}})
```

`is_pangram("hello world", "helowrd")` deve retornar `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello world", "helowrd"), True)`)
}})
```

`is_pangram("Hello World!", "helowrd")` deve retornar `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("Hello World!", "helowrd"), True)`)
}})
```

`is_pangram("Hello World!", "heliowrd")` deve retornar `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("Hello World!", "heliowrd"), False)`)
}})
```

`is_pangram("freeCodeCamp", "frcdmp")` deve retornar `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("freeCodeCamp", "frcdmp"), False)`)
}})
```

`is_pangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz")` deve retornar `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_pangram(sentence, letters):

    return sentence
```

# --solutions--

```py
import re
def is_pangram(sentence, letters):
    used_letters = []
    for char in sentence.lower():
        if re.match(r'[a-z]', char) and char not in used_letters:
            used_letters.append(char)

    sorted_letters = ''.join(sorted(letters.lower()))
    sorted_used_letters = ''.join(sorted(used_letters))

    return sorted_letters == sorted_used_letters
```
