---
id: 6821ebe4237de8297eaee794
title: "Desafio 18: Segundo Melhor"
challengeType: 29
dashedName: challenge-18
---

# --description--

Dado um array de inteiros representando o preço de diferentes laptops e um inteiro representando seu orçamento, retorne:

1. O segundo laptop mais caro se estiver dentro do seu orçamento, ou
2. O laptop mais caro que está dentro do seu orçamento, ou
3. `0` se nenhum laptop estiver dentro do seu orçamento.

- Preços duplicados devem ser ignorados.

# --hints--

`get_laptop_cost([1500, 2000, 1800, 1400], 1900)` deve retornar `1800`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900)` deve retornar `1800`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 2200)` deve retornar `1899`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 2200), 1899)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 1000)` deve retornar `0`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 1000), 0)`)
}})
```

`get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` deve retornar `1400`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_laptop_cost(laptops, budget):

    return laptops
```

# --solutions--

```py
def get_laptop_cost(laptops, budget):
    laptops = sorted(set(laptops), reverse=True)

    if len(laptops) >= 2 and budget >= laptops[1]:
        return laptops[1]
    if not laptops or budget < laptops[-1]:
        return 0

    for price in laptops[2:]:
        if budget >= price:
            return price

    return 0
```
