---
id: 6821ebda237de8297eaee792
title: "Desafio 16: Parênteses Invertidos"
challengeType: 29
dashedName: challenge-16
---

# --description--

Dada uma string que contém parênteses devidamente aninhados, retorne a versão decodificada da string usando as seguintes regras:

- Todos os caracteres dentro de cada parêntese devem ser invertidos.
- Os parênteses devem ser removidos do resultado final.
- Se parênteses estiverem aninhados, o par mais interno deve ser invertido primeiro e então seu resultado deve ser incluído na inversão do par externo.
- Assuma que todos os parênteses estão equilibrados e corretamente aninhados.

# --hints--

`decode("(f(b(dc)e)a)")` deve retornar `"abcdef"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` deve retornar `"Can you read this?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` deve retornar `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
