---
id: 68adce01c0e1144d0a90295e
title: "Desafio 29: Construtor de Acrônimos"
challengeType: 29
dashedName: challenge-29
---

# --description--

Dada uma string contendo uma ou mais palavras, retorne um acrônimo das palavras usando as seguintes restrições:

- A sigla deve consistir na primeira letra de cada palavra em maiúscula, a menos que indicado de outra forma.
- A sigla deve ignorar a primeira letra dessas palavras a menos que sejam a primeira palavra da string dada: `a`, `for`, `an`, `and`, `by` e `of`.
- As letras do acrônimo devem ser retornadas na ordem em que são fornecidas.
- O acrônimo não deve conter espaços.

# --hints--

`build_acronym("Search Engine Optimization")` deve retornar `"SEO"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Search Engine Optimization"), "SEO")`)
}})
```

`build_acronym("Frequently Asked Questions")` deve retornar `"FAQ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Frequently Asked Questions"), "FAQ")`)
}})
```

`build_acronym("National Aeronáutica e Administração Espacial")` deve retornar `"NASA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("National Aeronautics and Space Administration"), "NASA")`)
}})
```

`build_acronym("Federal Bureau of Investigation")` deve retornar `"FBI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Federal Bureau of Investigation"), "FBI")`)
}})
```

`build_acronym("For your information")` deve retornar `"FYI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("For your information"), "FYI")`)
}})
```

`build_acronym("By the way")` deve retornar `"BTW"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("By the way"), "BTW")`)
}})
```

`build_acronym("Uma manada imparável de pinguins cambaleantes ultrapassa as montanhas geladas e canta felizmente")` deve retornar `"UMIPUAMGECF"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH")`)
}})
```

# --seed--

## --seed-contents--

```py
def build_acronym(s):

    return s
```

# --solutions--

```py
def build_acronym(s):
    small_words = {"a", "for", "an", "and", "by", "of"}
    words = s.split()
    acronym = ""

    for i, word in enumerate(words):
        if i == 0 or word.lower() not in small_words:
            acronym += word[0].upper()

    return acronym
```
