---
id: 68adce01c0e1144d0a902956
title: "Desafio 25: Repetidor de Vogais"
challengeType: 29
dashedName: challenge-25
---

# --description--

Dada uma string, retorne uma nova versão da string onde cada vogal é duplicada uma vez a mais do que a vogal anterior que você encontrou. Por exemplo, a primeira vogal na frase deve permanecer inalterada. A segunda vogal deve aparecer duas vezes seguidas. A terceira vogal deve aparecer três vezes seguidas e assim por diante.

- As letras `a`, `e`, `i`, `o` e `u`, tanto maiúsculas quanto minúsculas, são consideradas vogais.
- A vogal original deve manter seu caso.
- Vogais repetidas devem estar em minúsculas.
- Todos os caracteres não vogais devem manter sua caixa original.

# --hints--

`repeat_vowels("hello world")` deve retornar `"helloo wooorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` deve retornar `"freeeCooodeeeeCaaaaamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` deve retornar `"AEeIiiOoooUuuuu"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` deve retornar `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
