---
id: 68adce01c0e1144d0a90295c
title: "Desafio 28: Parser de Numerais Romanos"
challengeType: 29
dashedName: challenge-28
---

# --description--

Dada uma string que representa um numeral romano, retorne seu valor inteiro.

Os numerais romanos consistem nos seguintes símbolos e valores:

| Symbol | Value |
|--------|-------|
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

- Os numerais são lidos da esquerda para a direita. Se um numeral menor aparecer antes de um maior, o valor é subtraído. Caso contrário, os valores são somados.

# --hints--

`parse_roman_numeral("III")` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("III"), 3)`)
}})
```

`parse_roman_numeral("IV")` deve retornar `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("IV"), 4)`)
}})
```

`parse_roman_numeral("XXVI")` deve retornar `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XXVI"), 26)`)
}})
```

`parse_roman_numeral("XCIX")` deve retornar `99`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XCIX"), 99)`)
}})
```

`parse_roman_numeral("CDLX")` deve retornar `460`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("CDLX"), 460)`)
}})
```

`parse_roman_numeral("DIV")` deve retornar `504`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("DIV"), 504)`)
}})
```

`parse_roman_numeral("MMXXV")` deve retornar `2025`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("MMXXV"), 2025)`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_roman_numeral(numeral):

    return numeral
```

# --solutions--

```py
def parse_roman_numeral(numeral):
    roman_map = {
        "I": 1,
        "V": 5,
        "X": 10,
        "L": 50,
        "C": 100,
        "D": 500,
        "M": 1000
    }

    total = 0
    for i in range(len(numeral)):
        current = roman_map[numeral[i]]
        next_val = roman_map[numeral[i + 1]] if i + 1 < len(numeral) else 0

        if current < next_val:
            total -= current
        else:
            total += current

    return total
```
