---
id: 68af0687ef34c76c28ffa54b
title: "Desafio 32: Inverter Sentença"
challengeType: 29
dashedName: challenge-32
---

# --description--

Dada uma string de palavras, retorne uma nova string com as palavras em ordem inversa. Por exemplo, a primeira palavra deve estar no final da string retornada e a última palavra deve estar no início da string retornada.

- Na string fornecida, as palavras podem ser separadas por um ou mais espaços.
- A string retornada deve ter apenas um espaço entre as palavras.

# --hints--

`reverse_sentence("world hello")` deve retornar `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` deve retornar `"git commit push"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm  install   apt    sudo")` deve retornar `"sudo apt install npm"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import    default   function  export")` deve retornar `"export function default import"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```
