---
id: 68af0687ef34c76c28ffa547
title: "Desafio 30: Caracteres Únicos"
challengeType: 29
dashedName: challenge-30
---

# --description--

Dada uma string, determine se todos os caracteres na string são únicos.

- Letras maiúsculas e minúsculas devem ser consideradas caracteres diferentes.

# --hints--

`all_unique("abc")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("abc"), True)`)
}})
```

`all_unique("aA")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("aA"), True)`)
}})
```

`all_unique("QwErTy123!@")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("QwErTy123!@"), True)`)
}})
```

`all_unique("~!@#$%^&*()_+")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("~!@#$%^&*()_+"), True)`)
}})
```

`all_unique("hello")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("hello"), False)`)
}})
```

`all_unique("freeCodeCamp")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("freeCodeCamp"), False)`)
}})
```

`all_unique("!@#*$%^&*()aA")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(all_unique("!@#*$%^&*()aA"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def all_unique(s):

    return s
```

# --solutions--

```py
def all_unique(s):
    seen = set()
    for ch in s:
        if ch in seen:
            return False
        seen.add(ch)
    return True
```
