---
id: 68b06e589bf2273243814777
title: "Desafio 39: Encha o Tanque"
challengeType: 29
dashedName: challenge-39
---

# --description--

Dado o tamanho de um tanque de combustível, o nível atual de combustível e o preço por galão, retorne o custo para encher o tanque completamente.

- `tankSize` é a capacidade total do tanque em galões.
- `fuelLevel` é a quantidade atual de combustível no tanque em galões.
- `pricePerGallon` é o custo de um galão de combustível.
- O valor retornado deve ser arredondado para duas casas decimais no formato: `"$d.dd"`.

# --hints--

`cost_to_fill(20, 0, 4.00)` deve retornar `"$80.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` deve retornar `"$17.50"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` deve retornar `"$29.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` deve retornar `"$0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` deve retornar `"$21.89"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
