---
id: 68b06e589bf227324381476f
title: "Desafio 35: Frequência de Palavras"
challengeType: 29
dashedName: challenge-35
---

# --description--

Dado um parágrafo, retorne um array com as três palavras que ocorrem com mais frequência.

- Palavras no parágrafo serão separadas por espaços.
- Ignore maiúsculas e minúsculas no parágrafo dado. Por exemplo, trate `Hello` e `hello` como a mesma palavra.
- Ignore a pontuação no parágrafo fornecido. A pontuação consiste em vírgulas (`,`), pontos (`.`) e pontos de exclamação (`!`).
- O array retornado deve conter todas as palavras em letras minúsculas.
- O array retornado deve estar em ordem decrescente com a palavra que ocorre com mais frequência primeiro.

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` deve retornar `["coding", "python", "in"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` deve retornar `["i", "like", "coding"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` deve retornar `["debug", "test", "deploy"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
