---
id: 68b06e589bf2273243814775
title: "Desafio 38: Gerador de Slug"
challengeType: 29
dashedName: challenge-38
---

# --description--

Dada uma string, retorne uma versão amigável para URL da string usando as seguintes restrições:

- Todas as letras devem estar em minúsculas.
- Todos os caracteres que não são letras, números ou espaços devem ser removidos.
- Todos os espaços devem ser substituídos pelo código de espaço codificado para URL `%20`.
- Espaços consecutivos devem ser substituídos por um único `%20`.
- A string retornada não deve ter `%20` no início ou no final.

# --hints--

`generate_slug("helloWorld")` deve retornar `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("helloWorld"), "helloworld")`)
}})
```

`generate_slug("hello world!")` deve retornar `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello world!"), "hello%20world")`)
}})
```

`generate_slug(" hello-world ")` deve retornar `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug(" hello-world "), "helloworld")`)
}})
```

`generate_slug("hello  world")` deve retornar `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello  world"), "hello%20world")`)
}})
```

`generate_slug("  ?H^3-1*1]0! W[0%R#1]D  ")` deve retornar `"h3110%20w0r1d"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_slug(str):

    return str
```

# --solutions--

```py
import re
def generate_slug(s):
    cleaned = re.sub(r'[^a-zA-Z0-9 ]+', '', s)
    cleaned = re.sub(r'\s+', ' ', cleaned).strip()
    return cleaned.lower().replace(' ', '%20')
```
