---
id: 68b06e589bf2273243814771
title: "Desafio 36: Ajustador de Termostato"
challengeType: 29
dashedName: challenge-36
---

# --description--

Dada a temperatura atual de uma sala e uma temperatura alvo, retorne uma string indicando como ajustar a temperatura da sala com base nessas restrições:

- Retorne `"heat"` se a temperatura atual estiver abaixo do alvo.
- Retorne `"cool"` se a temperatura atual estiver acima do alvo.
- Retorne `"hold"` se a temperatura atual for igual à temperatura alvo.

# --hints--

`adjust_thermostat(68, 72)` deve retornar `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(68, 72), "heat")`)
}})
```

`adjust_thermostat(75, 72)` deve retornar `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(75, 72), "cool")`)
}})
```

`adjust_thermostat(72, 72)` deve retornar `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 72), "hold")`)
}})
```

`adjust_thermostat(-20.5, -10.1)` deve retornar `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(-20.5, -10.1), "heat")`)
}})
```

`adjust_thermostat(100, 99.9)` deve retornar `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(100, 99.9), "cool")`)
}})
```

`adjust_thermostat(0.0, 0.0)` deve retornar `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(0.0, 0.0), "hold")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(temp, target):

    return temp
```

# --solutions--

```py
def adjust_thermostat(temp, target):
    if temp < target:
        return "heat"
    elif temp > target:
        return "cool"
    else:
        return "hold"
```
