---
id: 68b1f72371a5ac895ac70a06
title: "Desafio 42: Armazenamento de Vídeo"
challengeType: 29
dashedName: challenge-42
---

# --description--

Dado um tamanho de vídeo, uma unidade para o tamanho do vídeo, uma capacidade de disco rígido e uma unidade para o disco rígido, retorne o número de vídeos que o disco rígido pode armazenar usando as seguintes restrições:

- A unidade para o tamanho do vídeo pode ser bytes (`"B"`), kilobytes (`"KB"`), megabytes (`"MB"`), ou gigabytes (`"GB"`).
- Se não for fornecida uma das unidades de vídeo acima, retorne `"Invalid video unit"`.
- A unidade da capacidade do disco rígido pode ser gigabytes (`"GB"`) ou terabytes (`"TB"`).
- Se não for fornecida uma das unidades de disco rígido acima, retorne `"Invalid drive unit"`.
- Retorna o número de vídeos inteiros que o drive pode armazenar.
- Use as seguintes conversões:

| Unidade | Equivalente |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |
| 1 TB |   1000 GB  |

Por exemplo, dado `500`, `"MB"`, `100` e `"GB"` como argumentos, determine quantos vídeos de 500 MB podem caber em um disco rígido de 100 GB.

# --hints--

`number_of_videos(500, "MB", 100, "GB")` deve retornar `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` deve retornar `"Invalid video unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` deve retornar `"Unidade de disco inválida"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` deve retornar `4000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` deve retornar `1466`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```
