---
id: 68b1f72371a5ac895ac70a08
title: "Desafio 43: Dígitos vs Letras"
challengeType: 29
dashedName: challenge-43
---

# --description--

Dada uma string, retorne `"digits"` se a string tiver mais dígitos do que letras, `"letters"` se tiver mais letras do que dígitos e `"tie"` se tiver a mesma quantidade de dígitos e letras.

- Dígitos consistem em `0-9`.
- Letras consistem em `a-z` em maiúsculas ou minúsculas.
- Ignore quaisquer outros caracteres.

# --hints--

`digits_or_letters("abc123")` deve retornar `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123"), "tie")`)
}})
```

`digits_or_letters("a1b2c3d")` deve retornar `"letters"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("a1b2c3d"), "letters")`)
}})
```

`digits_or_letters("1a2b3c4")` deve retornar `"digits"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("1a2b3c4"), "digits")`)
}})
```

`digits_or_letters("abc123!@#DEF")` deve retornar `"letters"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123!@#DEF"), "letters")`)
}})
```

`digits_or_letters("H3110 W0R1D")` deve retornar `"digits"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("H3110 W0R1D"), "digits")`)
}})
```

`digits_or_letters("P455W0RD")` deve retornar `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("P455W0RD"), "tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def digits_or_letters(s):

    return s
```

# --solutions--

```py
def digits_or_letters(s):
    digit_count = 0
    letter_count = 0

    for char in s:
        if char.isdigit():
            digit_count += 1
        elif char.isalpha():
            letter_count += 1

    if digit_count > letter_count:
        return "digits"
    if letter_count > digit_count:
        return "letters"
    return "tie"
```
