---
id: 68b1f72371a5ac895ac70a04
title: "Desafio 41: Armazenamento de Arquivos"
challengeType: 29
dashedName: challenge-41
---

# --description--

Dado um tamanho de arquivo, uma unidade para o tamanho do arquivo e a capacidade do disco rígido em gigabytes (GB), retorne o número de arquivos que o disco rígido pode armazenar usando as seguintes restrições:

- A unidade para o tamanho do arquivo pode ser bytes (`"B"`), kilobytes (`"KB"`) ou megabytes (`"MB"`).
- Retorna o número de arquivos inteiros que o drive pode armazenar.
- Use as seguintes conversões:

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

Por exemplo, dado `500`, `"KB"` e `1` como argumentos, determine quantos arquivos de 500 KB podem caber em um disco rígido de 1 GB.

# --hints--

`number_of_files(500, "KB", 1)` deve retornar `2000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(500, "KB", 1), 2000)`)
}})
```

`number_of_files(50000, "B", 1)` deve retornar `20000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(50000, "B", 1), 20000)`)
}})
```

`number_of_files(5, "MB", 1)` deve retornar `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(5, "MB", 1), 200)`)
}})
```

`number_of_files(4096, "B", 1.5)` deve retornar `366210`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4096, "B", 1.5), 366210)`)
}})
```

`number_of_files(220.5, "KB", 100)` deve retornar `453514`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(220.5, "KB", 100), 453514)`)
}})
```

`number_of_files(4.5, "MB", 750)` deve retornar `166666`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4.5, "MB", 750), 166666)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_files(file_size, file_unit, drive_size_gb):

    return file_size
```

# --solutions--

```py
def number_of_files(file_size, file_unit, drive_size_gb):
    drive_size_bytes = drive_size_gb * 1000 * 1000 * 1000

    if file_unit == "B":
        file_size_bytes = file_size
    elif file_unit == "KB":
        file_size_bytes = file_size * 1000
    else:
        file_size_bytes = file_size * 1000 * 1000

    return int(drive_size_bytes // file_size_bytes)
```
