---
id: 68b1f72371a5ac895ac70a02
title: "Desafio 40: Armazenamento de Fotos"
challengeType: 29
dashedName: challenge-40
---

# --description--

Dado um tamanho de foto em megabytes (MB) e a capacidade do disco rígido em gigabytes (GB), retorne o número de fotos que o disco rígido pode armazenar usando as seguintes restrições:

- 1 gigabyte equivale a 1000 megabytes.
- Retorna o número de fotos inteiras que o drive pode armazenar.

# --hints--

`number_of_photos(1, 1)` deve retornar `1000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(1, 1), 1000)`)
}})
```

`number_of_photos(2, 1)` deve retornar `500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(2, 1), 500)`)
}})
```

`number_of_photos(4, 256)` deve retornar `64000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(4, 256), 64000)`)
}})
```

`number_of_photos(3.5, 750)` deve retornar `214285`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 750), 214285)`)
}})
```

`number_of_photos(3.5, 5.5)` deve retornar `1571`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 5.5), 1571)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_photos(photo_size_mb, drive_size_gb):

    return photo_size_mb
```

# --solutions--

```py
def number_of_photos(photo_size_mb, drive_size_gb):
    drive_size_mb = drive_size_gb * 1000
    return drive_size_mb // photo_size_mb
```
