---
id: 68b7cadffed0e75a517da673
title: "Desafio 52: Binário para Decimal"
challengeType: 29
dashedName: challenge-52
---

# --description--

Dada uma string que representa um número binário, retorne seu equivalente decimal como um número.

Um número binário usa apenas os dígitos `0` e `1` para representar qualquer número. Para converter binário em decimal, multiplique cada dígito por uma potência de `2` e some-os. Comece multiplicando o dígito mais à direita por `2^0`, o próximo dígito à esquerda por `2^1` e assim por diante. Depois que todos os dígitos forem multiplicados por uma potência de `2`, some o resultado.

Por exemplo, o número binário `101` equivale a `5` em decimal porque:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` deve retornar `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` deve retornar `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` deve retornar `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```
