---
id: 68b7687dded630607aceccb1
title: "Desafio 48: Detector de Spam"
challengeType: 29
dashedName: challenge-48
---

# --description--

Dado um número de telefone no formato `"+A (BBB) CCC-DDDD"`, onde cada letra representa um dígito conforme a seguir:

- `A` representa o código do país e pode ser qualquer número de dígitos.
- `BBB` representa o código de área e sempre terá três dígitos.
- `CCC` e `DDDD` representam o número local e terão sempre três e quatro dígitos, respectivamente.

Determine se é um número de spam com base nos seguintes critérios:

- O código do país tem mais de 2 dígitos ou não começa com zero (`0`).
- O código de área é maior que 900 ou menor que 200.
- A soma dos três primeiros dígitos do número local aparece dentro dos últimos quatro dígitos do número local.
- O número tem o mesmo dígito quatro ou mais vezes seguidas (ignorando os caracteres de formatação).

# --hints--

`is_spam("+0 (200) 234-0182")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (200) 234-0182"), False)`)
}})
```

`is_spam("+091 (555) 309-1922")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+091 (555) 309-1922"), True)`)
}})
```

`is_spam("+1 (555) 435-4792")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+1 (555) 435-4792"), True)`)
}})
```

`is_spam("+0 (955) 234-4364")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (955) 234-4364"), True)`)
}})
```

`is_spam("+0 (155) 131-6943")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (155) 131-6943"), True)`)
}})
```

`is_spam("+0 (555) 135-0192")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 135-0192"), True)`)
}})
```

`is_spam("+0 (555) 564-1987")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 564-1987"), True)`)
}})
```

`is_spam("+00 (555) 234-0182")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+00 (555) 234-0182"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_spam(number):

    return number
```

# --solutions--

```py
import re
def is_spam(number):
    digits = re.sub(r"\D", "", number)
    match = re.match(r"^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$", number)
    country_code, area_code, ccc, dddd = match.groups()

    if len(country_code) > 2 or not country_code.startswith("0"):
        return True

    area_num = int(area_code)
    if area_num > 900 or area_num < 200:
        return True

    sum_ccc = sum(int(d) for d in ccc)
    if str(sum_ccc) in dddd:
        return True

    if re.search(r"(\d)\1\1\1", digits):
        return True

    return False
```
