---
id: 68b7687dded630607aceccb3
title: "Desafio 49: Parser de Cabeçalho CSV"
challengeType: 29
dashedName: challenge-49
---

# --description--

Dada a primeira linha de um arquivo de valores separados por vírgula (CSV), retorne um array contendo os cabeçalhos.

- A primeira linha de um arquivo CSV contém cabeçalhos separados por vírgulas.
- Remova quaisquer espaços em branco no início ou no final de cada título.

# --hints--

`get_headings("name,age,city")` deve retornar `["name", "age", "city"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("name,age,city"), ["name", "age", "city"])`)
}})
```

`get_headings("first name,last name,phone")` deve retornar `["first name", "last name", "phone"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("first name,last name,phone"), ["first name", "last name", "phone"])`)
}})
```

`get_headings("username , email , signup date ")` deve retornar `["username", "email", "signup date"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("username , email , signup date "), ["username", "email", "signup date"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_headings(csv):

    return csv
```

# --solutions--

```py
def get_headings(csv):
    return [h.strip() for h in csv.split(",")]
```
