---
id: 68b7cadffed0e75a517da677
title: "Desafio 54: Força da Senha P@ssw0rd Str3ngth!"
challengeType: 29
dashedName: challenge-54
---

# --description--

Dada uma string de senha, retorne `"weak"`, `"medium"` ou `"strong"` com base na força da senha.

Uma senha é avaliada de acordo com as seguintes regras:

- Tem pelo menos 8 caracteres.
- Ele contém letras maiúsculas e minúsculas.
- Ele contém pelo menos um número.
- Ele contém pelo menos um caractere especial deste conjunto: `!`, `@`, `#`, `$`, `%`, `^`, `&` ou `*`.

Retorne `"weak"` se a senha atender a menos de duas das regras.  
Retorne `"medium"` se a senha atender 2 ou 3 das regras.  
Retorne `"strong"` se a senha atender todas as 4 regras.

# --hints--

`check_strength("123456")` deve retornar `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("123456"), "weak")`)
}})
```

`check_strength("pass!!!")` deve retornar `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("pass!!!"), "weak")`)
}})
```

`check_strength("Qwerty")` deve retornar `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("Qwerty"), "weak")`)
}})
```

`check_strength("PASSWORD")` deve retornar `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD"), "weak")`)
}})
```

`check_strength("PASSWORD!")` deve retornar `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("PassWord%^!")` deve retornar `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PassWord%^!"), "medium")`)
}})
```

`check_strength("qwerty12345")` deve retornar `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("qwerty12345"), "medium")`)
}})
```

`check_strength("PASSWORD!")` deve retornar `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("PASSWORD!")` deve retornar `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("S3cur3P@ssw0rd")` deve retornar `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("S3cur3P@ssw0rd"), "strong")`)
}})
```

`check_strength("C0d3&Fun!")` deve retornar `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("C0d3&Fun!"), "strong")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_strength(password):

    return password
```

# --solutions--

```py
import re
def check_strength(password):
    rules_met = 0

    if len(password) >= 8:
        rules_met += 1
    if re.search(r'[a-z]', password) and re.search(r'[A-Z]', password):
        rules_met += 1
    if re.search(r'\d', password):
        rules_met += 1
    if re.search(r'[!@#$%^&*]', password):
        rules_met += 1

    if rules_met < 2:
        return "weak"
    elif rules_met < 4:
        return "medium"
    else:
        return "strong"
```
