---
id: 68b7cadffed0e75a517da675
title: "Desafio 53: Decimal para Binário"
challengeType: 29
dashedName: challenge-53
---

# --description--

Dado um inteiro não negativo, retorne sua representação binária como uma string.

Um número binário usa apenas os dígitos `0` e `1` para representar qualquer número. Para converter um número decimal em binário, divida repetidamente o número por `2` e registre o resto. Repita até que o número seja zero. Leia os restos do último registrado para o primeiro. Por exemplo, para converter `12` em binário:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` em binário é `1100`.

# --hints--

`to_binary(5)` deve retornar `"101"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(5), "101")`)
}})
```

`to_binary(12)` deve retornar `"1100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(12), "1100")`)
}})
```

`to_binary(50)` deve retornar `"110010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(50), "110010")`)
}})
```

`to_binary(99)` deve retornar `"1100011"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(99), "1100011")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_binary(decimal):

    return decimal
```

# --solutions--

```py
def to_binary(decimal):
    if decimal == 0:
        return "0"
    binary = ""
    while decimal > 0:
        binary = str(decimal % 2) + binary
        decimal //= 2
    return binary
```
