---
id: 68c1a929005bf54d342aa8d6
title: "Desafio 59: Semana do Espaço Dia 5: Zona Goldilocks"
challengeType: 29
dashedName: challenge-59
---

# --description--

No quinto dia da Semana do Espaço, você calculará a "zona Goldilocks" de uma estrela - a região ao redor de uma estrela onde as condições são "perfeitas" para a existência de água líquida.

Dada a massa de uma estrela, retorne um array com as distâncias inicial e final da sua Zona Habitável em Unidades Astronômicas.

Para calcular a Goldilocks Zone:

1. Encontre a luminosidade da estrela elevando sua massa à potência de 3,5.
2. O início da zona é 0,95 vezes a raiz quadrada da sua luminosidade.
3. O final da zona é 1,37 vezes a raiz quadrada da sua luminosidade.

- Retorne as distâncias arredondadas para duas casas decimais.

Por exemplo, dado `1` como uma massa, retorne `[0.95, 1.37]`.

# --hints--

`goldilocks_zone(1)` deve retornar `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` deve retornar `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` deve retornar `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` deve retornar `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` deve retornar `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
