---
id: 68c1a929005bf54d342aa8d3
title: "Desafio 56: Semana do Espaço Dia 2: Busca por Exoplanetas"
challengeType: 29
dashedName: challenge-56
---

# --description--

No segundo dia da Semana do Espaço, você recebe uma string onde cada caractere representa a leitura de luminosidade de uma estrela. Determine se as leituras detectaram um exoplaneta usando o método de trânsito. O método de trânsito é quando um planeta passa na frente de uma estrela, reduzindo sua luminosidade observada.

- As leituras de luminosidade compreendem apenas os caracteres `0-9` e `A-Z` onde cada leitura corresponde aos seguintes valores numéricos:
- Os caracteres `0-9` correspondem aos níveis de luminosidade `0-9`.
- Os caracteres `A-Z` correspondem aos níveis de luminosidade `10-35`.

Uma estrela é considerada como tendo um exoplaneta se qualquer leitura única for menor ou igual a 80% da média de todas as leituras. Por exemplo, se a luminosidade média de uma estrela for 10, ela seria considerada como tendo um exoplaneta se qualquer leitura única for 8 ou menos.

# --hints--

`has_exoplanet("665544554")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```
