---
id: 68c1a929005bf54d342aa8d4
title: "Desafio 57: Semana do Espaço Dia 3: Phone Home"
challengeType: 29
dashedName: challenge-57
---

# --description--

No terceiro dia da Semana do Espaço, você recebe um array de números que representam distâncias (em quilômetros) entre você, satélites e seu planeta natal em uma rota de comunicação. Determine quanto tempo levará para uma mensagem enviada através da rota alcançar seu planeta de destino usando as seguintes restrições:

- O primeiro valor no array é a distância da sua localização até o primeiro satélite.
- Cada valor subsequente, exceto o último, é a distância até o próximo satélite.
- O último valor no array é a distância do satélite anterior ao seu planeta natal.
- A mensagem viaja a 300.000 km/s.
- Cada satélite pelo qual a mensagem **passa** adiciona um atraso de transmissão de 0,5 segundo.
- Retorna um número arredondado para 4 casas decimais, com zeros à direita removidos.

# --hints--

`send_message([300000, 300000])` deve retornar `2.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` deve retornar `3.0627`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` deve retornar `364.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` deve retornar `1674.3333`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` deve retornar `2.4086`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` deve retornar `21.1597`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
