---
id: 68c497f3aaefc9fd9f1b0e26
title: "Desafio 62: Hex para Decimal"
challengeType: 29
dashedName: challenge-62
---

# --description--

Dada uma string que representa um número hexadecimal (base 16), retorne seu valor decimal (base 10) como um inteiro.

Hexadecimal é um sistema numérico que usa 16 dígitos:

- `0-9` representam valores de `0` a `9`.
- `A-F` representam valores de `10` a `15`.

Aqui está uma tabela de conversão parcial:

| Hexadecimal | Decimal |
|:-----------:|:-------:|
| 0           | 0       |
| 1           | 1       |
| ...         | ...     |
| 9           | 9       |
| A           | 10      |
| ...         | ...     |
| F           | 15      |
| 10          | 16      |
| ...         | ...     |
| 9F          | 159     |
| A0          | 160     |
| ...         | ...     |
| FF          | 255     |
| 100         | 256     |

- A string conterá apenas os caracteres `0–9` e `A–F`.

# --hints--

`hex_to_decimal("A")` deve retornar `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A"), 10)`)
}})
```

`hex_to_decimal("15")` deve retornar `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("15"), 21)`)
}})
```

`hex_to_decimal("2E")` deve retornar `46`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("2E"), 46)`)
}})
```

`hex_to_decimal("FF")` deve retornar `255`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("FF"), 255)`)
}})
```

`hex_to_decimal("A3F")` deve retornar `2623`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A3F"), 2623)`)
}})
```

# --seed--

## --seed-contents--

```py
def hex_to_decimal(hex):

    return hex
```

# --solutions--

```py
def hex_to_decimal(hex):

    return int(hex, 16)
```
