---
id: 68c497f3aaefc9fd9f1b0e24
title: "Desafio 60: Semana do Espaço Dia 6: Fase da Lua"
challengeType: 29
dashedName: challenge-60
---

# --description--

Para o sexto dia da Semana do Espaço, você receberá uma data no formato `"YYYY-MM-DD"` e precisará determinar a fase da lua para esse dia usando as seguintes regras:

Use um ciclo lunar simplificado de 28 dias, dividido em quatro fases iguais:

- "New": dias 1 - 7
- "Waning": dias 8 - 14
- "Full": dias 15 - 21
- "Waning": dias 22 - 28

Após o dia 28, o ciclo se repete com o dia 1, uma lua nova.

- Use `"2000-01-06"` como uma lua nova de referência (dia 1 do ciclo) para determinar a fase do dia fornecido.
- Você não receberá nenhuma data antes da data de referência.
- Retorne a fase correta como uma string.

# --hints--

`moon_phase("2000-01-12")` deve retornar `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` deve retornar `"Waxing"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` deve retornar `"Full"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` deve retornar `"Waning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` deve retornar `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
