---
id: 68c497f3aaefc9fd9f1b0e25
title: "Desafio 61: Semana do Espaço Dia 7: Combustível de Lançamento"
challengeType: 29
dashedName: challenge-61
---

# --description--

No último dia da Semana do Espaço, você receberá a massa em quilogramas (kg) de uma carga útil que deseja enviar para a órbita. Determine a quantidade de combustível necessária para enviar sua carga útil para a órbita usando as seguintes regras:

- Foguetes requerem 1 kg de combustível para cada 5 kg de massa que devem levantar.
- O combustível em si tem massa. Então, quando você adiciona combustível, a massa a ser levantada aumenta, o que requer mais combustível, o que aumenta a massa e assim por diante.
- Para calcular o total de combustível necessário: comece com a massa da carga útil, calcule o combustível necessário para isso, adicione esse combustível à massa total e calcule novamente. Repita esse processo até que o combustível adicional necessário seja menor que 1 kg, então pare.
- Ignore a massa do próprio foguete. Calcule apenas o combustível necessário para levantar a carga útil e seu próprio combustível.

Por exemplo, dado uma massa de carga útil de 50 kg, você precisaria de 10 kg de combustível para levantá-la (payload / 5), o que aumenta a massa total para 60 kg, que precisa de 12 kg para levantar (2 kg adicionais), o que aumenta a massa total para 62 kg, que precisa de 12,4 kg para levantar - 0,4 kg adicionais - que é menos que 1 kg adicional, então paramos aqui. A massa total para levantar é 62,4 kg, 50 dos quais é a carga útil inicial e 12,4 de combustível.

- Retorna a quantidade de combustível necessária arredondada para uma casa decimal.

# --hints--

`launch_fuel(50)` deve retornar `12.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(50), 12.4)`)
}})
```

`launch_fuel(500)` deve retornar `124.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(500), 124.8)`)
}})
```

`launch_fuel(243)` deve retornar `60.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(243), 60.7)`)
}})
```

`launch_fuel(11000)` deve retornar `2749.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(11000), 2749.8)`)
}})
```

`launch_fuel(6214)` deve retornar `1553.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(6214), 1553.4)`)
}})
```

# --seed--

## --seed-contents--

```py
def launch_fuel(payload):

    return payload
```

# --solutions--

```py
def launch_fuel(payload):
    total_mass = payload
    additional_fuel = total_mass / 5
    total_fuel = additional_fuel

    while additional_fuel >= 1:
        total_mass += additional_fuel
        additional_fuel = (total_mass / 5) - total_fuel
        total_fuel += additional_fuel

    return round(total_fuel, 1)
```
