---
id: 68cae5b538ff798bbd4da00a
title: "Desafio 72: Ajustador de Termostato 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

Dada a temperatura atual de uma sala em Fahrenheit e uma temperatura alvo em Celsius, retorne uma string indicando como ajustar a temperatura da sala com base nessas restrições:

- Retorne `"Heat: X degrees Fahrenheit"` se a temperatura atual estiver abaixo do alvo. Com `X` sendo o número de graus em Fahrenheit para aquecer o ambiente até o alvo, arredondado para 1 casa decimal.
- Retorne `"Cool: X degrees Fahrenheit"` se a temperatura atual estiver acima do alvo. Com `X` sendo o número de graus em Fahrenheit para resfriar o ambiente até atingir o alvo, arredondado para 1 casa decimal.
- Retorne `"Hold"` se a temperatura atual for igual à temperatura alvo.

Para converter Celsius em Fahrenheit, multiplique a temperatura em Celsius por 1.8 e some 32 ao resultado (`F = (C * 1.8) + 32`).

# --hints--

`adjust_thermostat(32, 0)` deve retornar `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` deve retornar `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` deve retornar `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` deve retornar `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` deve retornar `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
