---
id: 68cae5b538ff798bbd4da006
title: "Desafio 68: Mascara de Cartão de Crédito"
challengeType: 29
dashedName: challenge-68
---

# --description--

Dada uma string de números de cartão de crédito, retorne uma versão mascarada dela usando as seguintes restrições:

- A string conterá quatro conjuntos de quatro dígitos (`0-9`), com todos os conjuntos separados por um único espaço ou um único hífen (`-`).
- Substitua todos os números, exceto os últimos quatro, por um asterisco (`*`).
- Deixe os caracteres restantes inalterados.

Por exemplo, dado `"4012-8888-8888-1881"` retorne `"****-****-****-1881"`.

# --hints--

`mask("4012-8888-8888-1881")` deve retornar `"****-****-****-1881"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("4012-8888-8888-1881"), "****-****-****-1881")`)
}})
```

`mask("5105 1051 0510 5100")` deve retornar `"**** **** **** 5100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("5105 1051 0510 5100"), "**** **** **** 5100")`)
}})
```

`mask("6011 1111 1111 1117")` deve retornar `"**** **** **** 1117"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("6011 1111 1111 1117"), "**** **** **** 1117")`)
}})
```

`mask("2223-0000-4845-0010")` deve retornar `"****-****-****-0010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("2223-0000-4845-0010"), "****-****-****-0010")`)
}})
```

# --seed--

## --seed-contents--

```py
def mask(card):

    return card
```

# --solutions--

```py
def mask(card):
    if '-' in card:
        split = card.split('-')
        return f"****-****-****-{split[3]}"
    else:
        split = card.split(' ')
        return f"**** **** **** {split[3]}"
```
