---
id: 68cae5b538ff798bbd4da009
title: "Desafio 71: Calculadora de Gorjeta"
challengeType: 29
dashedName: challenge-71
---

# --description--

Dado o preço da sua refeição e uma porcentagem personalizada de gorjeta, retorne um array com três valores de gorjeta; 15%, 20% e o valor personalizado.

- Os preços serão fornecidos no formato: `"$N.NN"`.
- Porcentagens personalizadas de dicas serão fornecidas neste formato: `"25%"`.
- Retorne os valores no mesmo formato `"$N.NN"`, arredondados para duas casas decimais.

Por exemplo, dado um preço de refeição de `"$10.00"` e um valor de gorjeta personalizada de `"25%"`, retorne `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculate_tips("$10.00", "25%")` deve retornar `["$1.50", "$2.00", "$2.50"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"])`)
}})
```

`calculate_tips("$89.67", "26%")` deve retornar `["$13.45", "$17.93", "$23.31"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"])`)
}})
```

`calculate_tips("$19.85", "9%")` deve retornar `["$2.98", "$3.97", "$1.79"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_tips(meal_price, custom_tip):

    return meal_price
```

# --solutions--

```py
def calculate_tips(meal_price, custom_tip):
    meal = float(meal_price[1:])
    custom_percent = float(custom_tip[:-1])

    tips_percent = [15, 20, custom_percent]
    tips = [f"${meal * p / 100:.2f}" for p in tips_percent]

    return tips
```
