---
id: 68cae5b538ff798bbd4da004
title: "Desafio 66: Removedor de Tags HTML"
challengeType: 29
dashedName: challenge-66
---

# --description--

Dada uma string de código HTML, remova as tags e retorne o conteúdo de texto simples.

- A string de entrada conterá apenas HTML válido.
- As tags HTML podem ser aninhadas.
- Remova as tags e quaisquer atributos.

Por exemplo, `'<a href="#">Click here</a>'` deve retornar `"Click here"`.

# --hints--

`strip_tags('<a href="#">Click here</a>')` deve retornar `"Click here"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` deve retornar `"Hello World!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` deve retornar uma string vazia (`""`).

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` deve retornar `sectionsection`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
