---
id: 68cae5b538ff798bbd4da007
title: "Desafio 69: Meias Desaparecidas"
challengeType: 29
dashedName: challenge-69
---

# --description--

Dado um inteiro representando o número de pares de meias com os quais você começou e outro inteiro representando quantos ciclos de lavagem você realizou, retorne o número de pares completos de meias que você tem atualmente usando as seguintes restrições:

- A cada 2 ciclos de lavagem, você perde uma única meia.
- A cada 3 ciclos de lavagem, você encontra uma única meia desaparecida.
- A cada 5 ciclos de lavagem, uma única meia se desgasta e deve ser descartada.
- A cada 10 ciclos de lavagem, você compra um par de meias.
- Você nunca pode ter menos que zero meias no total.
- Regras podem se sobrepor. Por exemplo, no ciclo de lavagem 10, você perderá uma meia, jogará fora uma meia e comprará um novo par de meias.
- Retorna o número de pares completos de meias.

# --hints--

`sock_pairs(2, 5)` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(2, 5), 1)`)
}})
```

`sock_pairs(1, 2)` deve retornar `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 2), 0)`)
}})
```

`sock_pairs(5, 11)` deve retornar `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(5, 11), 4)`)
}})
```

`sock_pairs(6, 25)` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(6, 25), 3)`)
}})
```

`sock_pairs(1, 8)` deve retornar `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 8), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sock_pairs(pairs, cycles):

    return pairs
```

# --solutions--

```py
def sock_pairs(pairs, cycles):
    socks = pairs * 2

    for i in range(1, cycles + 1):
        if i % 2 == 0:
            socks -= 1
        if i % 3 == 0:
            socks += 1
        if i % 5 == 0:
            socks -= 1
        if i % 10 == 0:
            socks += 2

        if socks < 0:
            socks = 0

    return socks // 2
```
