---
id: 68cae5b538ff798bbd4da002
title: "Desafio 64: 24 para 12"
challengeType: 29
dashedName: challenge-64
---

# --description--

Dada uma string que representa um horário no formato 24 horas `"HHMM"`, retorne o horário no seu equivalente formato 12 horas `"H:MM AM"` ou `"H:MM PM"`.

- A entrada fornecida será sempre uma string de quatro dígitos no formato de hora 24 horas, de `"0000"` a `"2359"`.

# --hints--

`to_12("1124")` deve retornar `"11:24 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1124"), "11:24 AM")`)
}})
```

`to_12("0900")` deve retornar `"9:00 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0900"), "9:00 AM")`)
}})
```

`to_12("1455")` deve retornar `"2:55 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1455"), "2:55 PM")`)
}})
```

`to_12("2346")` deve retornar `"11:46 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("2346"), "11:46 PM")`)
}})
```

`to_12("0030")` deve retornar `"12:30 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0030"), "12:30 AM")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_12(time):

    return time
```

# --solutions--

```py
def convert_hours(hours):
    if hours == 0:
        return 12
    elif hours > 12:
        return hours - 12
    else:
        return hours

def to_12(time):
    hours = int(time[:2])
    minutes = time[2:]
    period = "AM" if hours < 12 else "PM"

    return f"{convert_hours(hours)}:{minutes} {period}"
```
