---
id: 68cae5b538ff798bbd4da001
title: "Desafio 63: Batalha de Palavras"
challengeType: 29
dashedName: challenge-63
---

# --description--

Dadas duas sentenças representando seu time e um time adversário, onde cada palavra do seu time batalha contra a palavra correspondente do time adversário, determine qual time vence usando as seguintes regras:

- As sentenças fornecidas sempre conterão o mesmo número de palavras.
- As palavras são separadas por um único espaço e conterão apenas letras.
- O valor de cada palavra é a soma de suas letras.
- As letras `a` a `z` correspondem aos valores de `1` a `26`. Por exemplo, `a` é `1` e `z` é `26`.
- Uma letra maiúscula dobra o valor da letra. Por exemplo, `A` é `2` e `Z` é `52`.
- Palavras batalham em ordem: a primeira palavra do seu time batalha a primeira palavra do time adversário e assim por diante.
- Uma palavra vence se seu valor for maior que o valor da palavra oposta.
- A equipe com mais palavras vencedoras é a vencedora.

Retorne `"We win"` se seu time for o vencedor, `"We lose"` se seu time perder e `"Draw"` se ambos os times tiverem o mesmo número de vitórias.

# --hints--

`battle("hello world", "hello word")` deve retornar `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` deve retornar `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` deve retornar `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` deve retornar `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` deve retornar `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` deve retornar `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` deve retornar `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
