---
id: 68d2ba1468508398389487ce
title: "Desafio 73: Fale com Sabedoria, Você Deve"
challengeType: 29
dashedName: challenge-73
---

# --description--

Dada uma frase, retorne uma versão dela que soe como um conselho de um sábio mestre usando as seguintes regras:

- Palavras são separadas por um único espaço.
- Encontre a primeira ocorrência de uma das seguintes palavras na frase: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Mova todas as palavras antes e incluindo essa palavra para o final da frase e:
  - Preserve a ordem das palavras quando você as mover.
  - Deixe todos em letras minúsculas.
  - E adicione uma vírgula e espaço antes deles.
- Coloque em maiúscula a primeira letra da nova primeira palavra da frase.
- Todas as sentenças fornecidas terminarão com um único sinal de pontuação. Mantenha a pontuação original da sentença e coloque-a no final da nova sentença.
- Retorne a nova sentença, certifique-se de que haja um único espaço entre cada palavra e nenhum espaço no início e no final da sentença.

Por exemplo, dado `"You must speak wisely."` retorne `"Speak wisely, you must."`

# --hints--

`wise_speak("You must speak wisely.")` deve retornar `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` deve retornar `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` deve retornar `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` deve retornar `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` deve retornar `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
