---
id: 68d2ba1468508398389487d0
title: "Desafio 75: Tesouro Escondido"
challengeType: 29
dashedName: challenge-75
---

# --description--

Dada uma matriz 2D que representa um mapa do fundo do oceano que inclui um tesouro escondido e uma matriz com as coordenadas (`[row, column]`) para o próximo mergulho da sua busca pelo tesouro, retorne `"Empty"`, `"Found"` ou `"Recovered"` usando as seguintes regras:

- A matriz 2D fornecida conterá exatamente um tesouro não recuperado, que ocupará múltiplas células.
- Cada célula na matriz 2D conterá um dos seguintes valores:
  - `"-"`: Sem tesouro.
  - "O": Uma parte do tesouro que não foi encontrada.
  - "X": Uma parte do tesouro que já foi encontrada.

- Se o local do mergulho não tiver tesouro, retorne `"Empty"`.
- Se o local do mergulho encontrar tesouro, mas pelo menos outra parte do tesouro permanecer não encontrada, retorne `"Found"`.
- Se o local do mergulho encontrar a última parte do tesouro não encontrada, retorne `"Recovered"`.

Por exemplo, dado:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

E `[2, 1]` para as coordenadas do local do mergulho, retorna `"Recovered"` porque o mergulho encontrou a última parte do tesouro que não havia sido encontrada.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` deve retornar `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered")`)
}})
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` deve retornar `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty")`)
}})
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` deve retornar `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` deve retornar `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` deve retornar `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` deve retornar `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty")`)
}})
```

# --seed--

## --seed-contents--

```py
def dive(map, coordinates):

    return map
```

# --solutions--

```py
def dive(map, coordinates):
    row, col = coordinates
    target = map[row][col]

    if target == "-":
        return "Empty"

    if target == "O":
        for r in range(len(map)):
            for c in range(len(map[r])):
                if (r != row or c != col) and map[r][c] == "O":
                    return "Found"
        return "Recovered"

    return "Found"
```
