---
id: 68d30845cc08266018fc46be
title: "Desafio 78: Sequência Inteira"
challengeType: 29
dashedName: challenge-78
---

# --description--

Dado um número inteiro positivo, retorne uma string com todos os números inteiros de `1` até, e incluindo, o número fornecido, em ordem numérica.

Por exemplo, dado `5`, retorne `"12345"`.

# --hints--

`sequence(5)` deve retornar `"12345"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(5), "12345")`)
}})
```

`sequence(10)` deve retornar `"12345678910"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(10), "12345678910")`)
}})
```

`sequence(1)` deve retornar `"1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(1), "1")`)
}})
```

`sequence(27)` deve retornar `"123456789101112131415161718192021222324252627"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(27), "123456789101112131415161718192021222324252627")`)
}})
```

# --seed--

## --seed-contents--

```py
def sequence(n):

    return n
```

# --solutions--

```py
def sequence(n):
    result = ""
    for i in range(1, n + 1):
        result += str(i)
    return result
```
