---
id: 68d30fc57588d97fd3027b30
title: "Desafio 79: Navigator"
challengeType: 29
dashedName: challenge-79
---

# --description--

Em 28 de outubro de 1994, o Netscape Navigator foi lançado, ajudando milhões a explorar a web inicial.

Dado um array de comandos do navegador que você executou no Netscape Navigator, retorne a página atual em que você está após executar todos os comandos usando as seguintes regras:

- Você sempre começa na página `"Home"`, que não será incluída no array de comandos.
- Os comandos válidos são:
  - "Visit Page": Onde "Page" é o nome da página que você está visitando. Por exemplo, "Visit About" leva você para a página "About". Quando você visita uma nova página, certifique-se de descartar qualquer histórico de avanço que você tenha.
  - "Back": Leva você para a página anterior no seu histórico ou permanece na página atual se não houver uma.
  - "Forward": Leva você para frente no histórico até a página de onde você veio ou permanece na página atual se não houver uma.

Por exemplo, dado `["Visit About Us", "Back", "Forward"]`, retorne `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` deve retornar `"About Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Back", "Forward"]), "About Us")`)
}})
```

`navigate(["Forward"])` deve retornar `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Forward"]), "Home")`)
}})
```

`navigate(["Back"])` deve retornar `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Back"]), "Home")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery"])` deve retornar `"Gallery"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery"]), "Gallery")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` deve retornar `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home")`)
}})
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` deve retornar `"Contact"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact")`)
}})
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` deve retornar `"Visit Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate(commands):

    return commands
```

# --solutions--

```py
def navigate(commands):
    history = ["Home"]
    current_index = 0

    for command in commands:
        if command.startswith("Visit "):
            history = history[:current_index + 1]
            history.append(command[6:])
            current_index += 1
        elif command == "Back" and current_index > 0:
            current_index -= 1
        elif command == "Forward" and current_index < len(history) - 1:
            current_index += 1

    return history[current_index]
```
