---
id: 68d30845cc08266018fc46bd
title: "Desafio 77: Formatador de Duração"
challengeType: 29
dashedName: challenge-77
---

# --description--

Dado um número inteiro de segundos, retorne uma string que represente a mesma duração no formato `"H:MM:SS"`, onde `"H"` é o número de horas, `"MM"` é o número de minutos e `"SS"` é o número de segundos. Retorne o tempo usando as seguintes regras:

- Segundos: Deve sempre ter dois dígitos.
- Minutos: Deve omitir zeros à esquerda quando não forem necessários. Use `"0"` se a duração for menor que um minuto.
- Horas: Deve ser incluído apenas se for maior que zero.

# --hints--

`format(500)` deve retornar `"8:20"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(500), "8:20")`)
}})
```

`format(4000)` deve retornar `"1:06:40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(4000), "1:06:40")`)
}})
```

`format(1)` deve retornar `"0:01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(1), "0:01")`)
}})
```

`format(5555)` deve retornar `"1:32:35"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(5555), "1:32:35")`)
}})
```

`format(99999)` deve retornar `"27:46:39"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(99999), "27:46:39")`)
}})
```

# --seed--

## --seed-contents--

```py
def format(seconds):

    return seconds
```

# --solutions--

```py
def format(seconds):
    h = seconds // 3600
    m = (seconds % 3600) // 60
    s = seconds % 60

    seconds_str = f"{s:02d}"
    minutes_str = str(m)

    if h > 0:
        hours_str = str(h)
        return f"{hours_str}:{minutes_str.zfill(2)}:{seconds_str}"
    else:
        return f"{minutes_str}:{seconds_str}"
```
