---
id: 68d30845cc08266018fc46bc
title: "Desafio 76: DNA Complementar"
challengeType: 29
dashedName: challenge-76
---

# --description--

Dada uma string que representa uma sequência de DNA, retorne sua fita complementar usando as seguintes regras:

- DNA consiste nas letras `"A"`, `"C"`, `"G"` e `"T"`.
- As letras `"A"` e `"T"` se complementam.
- As letras `"C"` e `"G"` se complementam.

Por exemplo, dado `"ACGT"`, retorne `"TGCA"`.

# --hints--

`complementary_dna("ACGT")` deve retornar `"TGCA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ACGT"), "TGCA")`)
}})
```

`complementary_dna("ATGCGTACGTTAGC")` deve retornar `"TACGCATGCAATCG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ATGCGTACGTTAGC"), "TACGCATGCAATCG")`)
}})
```

`complementary_dna("GGCTTACGATCGAAG")` deve retornar `"CCGAATGCTAGCTTC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC")`)
}})
```

`complementary_dna("GATCTAGCTAGGCTAGCTAG")` deve retornar `"CTAGATCGATCCGATCGATC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC")`)
}})
```

# --seed--

## --seed-contents--

```py
def complementary_dna(strand):

    return strand
```

# --solutions--

```py
def complementary_dna(strand):
    complements = { "A": "T", "T": "A", "C": "G", "G": "C" }
    return "".join(complements[n] for n in strand)
```
