---
id: 68e39ed6106dac2f0a98fd63
title: "Desafio 81: Nth Prime"
challengeType: 29
dashedName: challenge-81
---

# --description--

Um número primo é um inteiro positivo maior que 1 que é divisível apenas por 1 e por ele mesmo. Os cinco primeiros números primos são `2`, `3`, `5`, `7` e `11`.

Dado um inteiro positivo `n`, retorne o `n`th número primo. Por exemplo, dado `5` retorne o 5th número primo: `11`.

# --hints--

`nth_prime(5)` deve retornar `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(5), 11)`)
}})
```

`nth_prime(10)` deve retornar `29`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(10), 29)`)
}})
```

`nth_prime(16)` deve retornar `53`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(16), 53)`)
}})
```

`nth_prime(99)` deve retornar `523`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(99), 523)`)
}})
```

`nth_prime(1000)` deve retornar `7919`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(1000), 7919)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_prime(n):

    return n
```

# --solutions--

```py
def nth_prime(n):
    primes = []
    num = 2

    while len(primes) < n:
        is_prime = True
        for i in range(2, int(num**0.5) + 1):
            if num % i == 0:
                is_prime = False
                break
        if is_prime:
            primes.append(num)
        num += 1

    return primes[n - 1]
```
