---
id: 68e39ed6106dac2f0a98fd64
title: "Desafio 82: SpOoKy~CaSe"
challengeType: 29
dashedName: challenge-82
---

# --description--

Dada uma string que representa um nome de variável, converta-a para "spooky case" usando as seguintes restrições:

- Substitua todos os underscores (`_`) e hífens (`-`) por um til (~).
- Coloque a primeira letra da string em maiúscula e todas as outras letras depois disso em minúscula. Ignore o caractere til ao contar.

Por exemplo, dado `hello_world`, retorne `HeLlO~wOrLd`.

# --hints--

`spookify("hello_world")` deve retornar `"HeLlO~wOrLd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("hello_world"), "HeLlO~wOrLd")`)
}})
```

`spookify("Spooky_Case")` deve retornar `"SpOoKy~CaSe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("Spooky_Case"), "SpOoKy~CaSe")`)
}})
```

`spookify("TRICK-or-TREAT")` deve retornar `"TrIcK~oR~tReAt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt")`)
}})
```

`spookify("c_a-n_d-y_-b-o_w_l")` deve retornar `"C~a~N~d~Y~~b~O~w~L"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L")`)
}})
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` deve retornar `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS")`)
}})
```

# --seed--

## --seed-contents--

```py
def spookify(boo):

    return boo
```

# --solutions--

```py
def spookify(boo):
    replaced = boo.replace("_", "~").replace("-", "~")
    
    result = []
    capitalize = True

    for char in replaced:
        if char == "~":
            result.append(char)
        else:
            result.append(char.upper() if capitalize else char.lower())
            capitalize = not capitalize

    return "".join(result)
```
