---
id: 68e39ed6106dac2f0a98fd66
title: "Desafio 84: Infectado"
challengeType: 29
dashedName: challenge-84
---

# --description--

Em 2 de novembro de 1988, o primeiro worm importante da internet foi lançado, infectando cerca de 10% dos computadores conectados à internet após apenas um dia.

Neste desafio, você recebe um número de dias que se passaram desde que um worm de internet foi liberado, e precisa determinar quantos computadores estão infectados usando as seguintes regras:

- No dia 0, o primeiro computador é infectado.
- A cada dia subsequente, o número de computadores infectados dobra.
- A cada 3º dia, um patch é aplicado após o vírus se espalhar e reduz o número de computadores infectados em 20%. Arredonde o número de computadores com patch para o próximo número inteiro.

Por exemplo, em:

- Dia 0: 1 computador no total está infectado.
- Dia 1: 2 computadores no total estão infectados.
- Dia 2: 4 computadores no total estão infectados.
- Dia 3: 8 computadores no total estão infectados. Então, aplique o patch: 8 infectados * 20% = 1,6 com patch. Arredonde 1,6 para cima para 2. 8 computadores infectados - 2 com patch = 6 computadores no total infectados após o dia 3.

Retorna o número total de computadores infectados após a quantidade de dias informada.

# --hints--

`infected(1)` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(1), 2)`)
}})
```

`infected(3)` deve retornar `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(3), 6)`)
}})
```

`infected(8)` deve retornar `152`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(8), 152)`)
}})
```

`infected(17)` deve retornar `39808`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(17), 39808)`)
}})
```

`infected(25)` deve retornar `5217638`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(25), 5217638)`)
}})
```

# --seed--

## --seed-contents--

```py
def infected(days):

    return days
```

# --solutions--

```py
import math
def infected(days):
    infected = 1
    for day in range(1, days + 1):
        infected *= 2
        if day % 3 == 0:
            patched = math.ceil(infected * 0.2)
            infected -= patched

    return infected
```
