---
id: 68e39ed6106dac2f0a98fd62
title: "Desafio 80: Organizador de Email"
challengeType: 29
dashedName: challenge-80
---

# --description--

Em 29 de outubro de 1971, o primeiro email foi enviado, introduzindo o formato `username@domain` que ainda usamos. Agora, existem bilhões de endereços de email.

Neste desafio, você recebe uma lista de endereços de email e precisa ordená-los alfabeticamente pelo nome do domínio primeiro (a parte depois do `@`) e pelo nome de usuário segundo (a parte antes do `@`).

- A ordenação deve ser case-insensitive.
- Se mais de um email tiver o mesmo domínio, ordene-os pelo nome de usuário.
- Retorne um array dos endereços ordenados.
- Os endereços retornados devem manter sua capitalização original.

Por exemplo, dado `["jill@mail.com", "john@example.com", "jane@example.com"]`, retorne `["jane@example.com", "john@example.com", "jill@mail.com"]`.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` deve retornar `["jane@example.com", "john@example.com", "jill@mail.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"])`)
}})
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` deve retornar `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"])`)
}})
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` deve retornar `["user@x.com", "user@y.com", "user@z.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"])`)
}})
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` deve retornar `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"])`)
}})
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` deve retornar `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"])`)
}})
```

# --seed--

## --seed-contents--

```py
def sort(emails):

    return emails
```

# --solutions--

```py
def sort(emails):
    return sorted(
        emails,
        key=lambda email: (email.split('@')[1].lower(), email.split('@')[0].lower())
    )
```
