---
id: 68ee9e3066cfd4eb2328e8a8
title: "Desafio 89: Contando Cartas"
challengeType: 29
dashedName: challenge-89
---

# --description--

Um baralho padrão de cartas tem 13 cartas únicas em cada naipe. Dado um inteiro que representa o número de cartas a serem escolhidas do baralho, retorne o número de combinações únicas de cartas que você pode escolher.

- A ordem não importa. Escolher a carta A e depois a carta B é o mesmo que escolher a carta B e depois a carta A.

Por exemplo, dado `52`, retorne `1`. Existe apenas uma combinação de 52 cartas para escolher de um baralho de 52 cartas. E dado `2`, retorne `1326`, Existem 1326 combinações de cartas que você pode obter ao escolher 2 cartas do baralho.

# --hints--

`combinations(52)` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` deve retornar `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` deve retornar `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` deve retornar `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` deve retornar `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` deve retornar `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
