---
id: 68ee9e3066cfd4eb2328e8a4
title: "Desafio 85: Contador de Palavras"
challengeType: 29
dashedName: challenge-85
---

# --description--

Dada uma string de sentença, retorne o número de palavras que estão na sentença.

- Palavras são qualquer sequência de caracteres sem espaço e são separadas por um único espaço.

# --hints--

`count_words("Hello world")` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("Hello world"), 2)`)
}})
```

`count_words("The quick brown fox jumps over the lazy dog.")` deve retornar `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("The quick brown fox jumps over the lazy dog."), 9)`)
}})
```

`count_words("I like coding challenges!")` deve retornar `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("I like coding challenges!"), 4)`)
}})
```

`count_words("Complete the challenge in JavaScript and Python.")` deve retornar `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("Complete the challenge in JavaScript and Python."), 7)`)
}})
```

`count_words("The missing semi-colon crashed the entire internet.")` deve retornar `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("The missing semi-colon crashed the entire internet."), 7)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_words(sentence):

    return sentence
```

# --solutions--

```py
def count_words(sentence):

    return len(sentence.split(' '))
```
