---
id: 68ee9e3066cfd4eb2328e8a7
title: "Desafio 88: Localizador de Dia da Semana"
challengeType: 29
dashedName: challenge-88
---

# --description--

Dada uma string de data no formato `YYYY-MM-DD`, retorne o dia da semana.

Os dias válidos para retorno são:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Certifique-se de ignorar os fusos horários.

# --hints--

`get_weekday("2025-11-06")` deve retornar `Thursday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2025-11-06"), "Thursday")`)
}})
```

`get_weekday("1999-12-31")` deve retornar `Friday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1999-12-31"), "Friday")`)
}})
```

`get_weekday("1111-11-11")` deve retornar `Saturday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1111-11-11"), "Saturday")`)
}})
```

`get_weekday("2112-12-21")` deve retornar `Wednesday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2112-12-21"), "Wednesday")`)
}})
```

`get_weekday("2345-10-01")` deve retornar `Monday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2345-10-01"), "Monday")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_weekday(date_string):

    return date_string
```

# --solutions--

```py
import datetime
def get_weekday(date_string):
    days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"]
    year, month, day = map(int, date_string.split("-"))
    date = datetime.date(year, month, day)

    return days[date.weekday() % 7 + 1 if date.weekday() != 6 else 0]
```
