---
id: 68ee9e3066cfd4eb2328e8a5
title: "Desafio 86: Pesquisa de Imagem"
challengeType: 29
dashedName: challenge-86
---

# --description--

Em 4 de novembro de 2001, o Google lançou sua busca de imagens, permitindo que as pessoas encontrassem imagens usando termos de busca. Neste desafio, você irá imitar a busca de imagens.

Dado um array de nomes de imagens e um termo de busca, retorne um array de nomes de imagens que contenham o termo de busca.

- Ignore a diferença entre maiúsculas e minúsculas ao corresponder os termos de busca.
- Retorne as imagens na mesma ordem em que aparecem no array de entrada.

# --hints--

`image_search(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` deve retornar `["dog.png"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"])`)
}})
```

`image_search(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` deve retornar `["Sunset.jpg", "sunflower.jpeg"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"])`)
}})
```

`image_search(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` deve retornar `["Moon.png", "stars.png"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"])`)
}})
```

`image_search(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` deve retornar `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"])`)
}})
```

# --seed--

## --seed-contents--

```py
def image_search(images, term):

    return images
```

# --solutions--

```py
def image_search(images, term):
    lower_term = term.lower()
    return [img for img in images if lower_term in img.lower()]
```
