---
id: 68f6587287ad1f4ad39b0c7e
title: "Desafio 92: Extension Extractor"
challengeType: 29
dashedName: challenge-92
---

# --description--

Dada uma string que representa um nome de arquivo, retorne a extensão do arquivo.

- A extensão é a parte do nome do arquivo que vem após o último ponto (`.`).
- Se o nome do arquivo não contiver um ponto ou terminar com um ponto, retorne `"none"`.
- A extensão deve ser retornada como está, preservando maiúsculas e minúsculas.

# --hints--

`get_extension("document.txt")` deve retornar `"txt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("document.txt"), "txt")`)
}})
```

`get_extension("README")` deve retornar `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("README"), "none")`)
}})
```

`get_extension("image.PNG")` deve retornar `"PNG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("image.PNG"), "PNG")`)
}})
```

`get_extension(".gitignore")` deve retornar `"gitignore"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension(".gitignore"), "gitignore")`)
}})
```

`get_extension("archive.tar.gz")` deve retornar `"gz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("archive.tar.gz"), "gz")`)
}})
```

`get_extension("final.draft.")` deve retornar `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("final.draft."), "none")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_extension(filename):

    return filename
```

# --solutions--

```py
def get_extension(filename):
    last_dot = filename.rfind('.')
    
    if last_dot == -1 or last_dot == len(filename) - 1:
        return "none"
    
    return filename[last_dot + 1:]
```
