---
id: 68f6587287ad1f4ad39b0c81
title: "Desafio 95: Array Shift"
challengeType: 29
dashedName: challenge-95
---

# --description--

Dado um array e um inteiro representando quantas posições deslocar o array, retorne o array deslocado.

- Um número inteiro positivo desloca o array para a esquerda.
- Um número inteiro negativo desloca o array para a direita.
- O deslocamento contorna o array.

Por exemplo, dado `[1, 2, 3]` e `1`, desloque o array 1 para a esquerda, retornando `[2, 3, 1]`.

# --hints--

`shift_array([1, 2, 3], 1)` deve retornar `[2, 3, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], 1), [2, 3, 1])`)
}})
```

`shift_array([1, 2, 3], -1)` deve retornar `[3, 1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], -1), [3, 1, 2])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], 5)` deve retornar `["charlie", "alpha", "bravo"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], -11)` deve retornar `["bravo", "charlie", "alpha"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"])`)
}})
```

`shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` deve retornar `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_array(arr, n):

    return arr
```

# --solutions--

```py
def shift_array(arr, n):
    length = len(arr)
    n = n % length
    if n < 0:
        n += length

    return arr[n:] + arr[:n]
```
