---
id: 68f6587287ad1f4ad39b0c82
title: "Desafio 96: É fim de semana?"
challengeType: 29
dashedName: challenge-96
---

# --description--

Dada uma data no formato `"YYYY-MM-DD"`, retorne o número de dias restantes até o final de semana.

- O fim de semana começa no sábado.
- Se a data fornecida for sábado ou domingo, retorne `"It's the weekend!"`.
- Caso contrário, retorne `"X days until the weekend."`, onde `X` é o número de dias até sábado.
- Se `X` for `1`, use `"day"` (singular) em vez de `"days"` (plural).
- Certifique-se de que o cálculo ignore seu fuso horário local.

# --hints--

`days_until_weekend("2025-11-14")` deve retornar `"1 day until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-11-14"), "1 day until the weekend.")`)
}})
```

`days_until_weekend("2025-01-01")` deve retornar `"3 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-01-01"), "3 days until the weekend.")`)
}})
```

`days_until_weekend("2025-12-06")` deve retornar `"It's the weekend!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-12-06"), "It's the weekend!")`)
}})
```

`days_until_weekend("2026-01-27")` deve retornar `"4 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-01-27"), "4 days until the weekend.")`)
}})
```

`days_until_weekend("2026-09-07")` deve retornar `"5 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-09-07"), "5 days until the weekend.")`)
}})
```

`days_until_weekend("2026-11-29")` deve retornar `"It's the weekend!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-11-29"), "It's the weekend!")`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_weekend(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime

def days_until_weekend(date_string):
    date = datetime.strptime(date_string, "%Y-%m-%d").date()
    day_of_week = date.weekday()

    if day_of_week == 5 or day_of_week == 6:
        return "It's the weekend!"

    days_until_saturday = 5 - day_of_week
    return f"{days_until_saturday} day{'s' if days_until_saturday != 1 else ''} until the weekend."
```
