---
id: 68f6587287ad1f4ad39b0c83
title: "Desafio 97: MDC"
challengeType: 29
dashedName: challenge-97
---

# --description--

Dado dois inteiros positivos, retorne seu maior divisor comum (GCD).

- O MDC de dois inteiros é o maior número que divide igualmente ambos os números sem deixar resto.

Por exemplo, os divisores de `4` são `1`, `2` e `4`. Os divisores de `6` são `1`, `2`, `3` e `6`. Então, dado `4` e `6`, retorne `2`, o maior número que aparece em ambos os conjuntos de divisores.

# --hints--

`gcd(4, 6)` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(4, 6), 2)`)
}})
```

`gcd(20, 15)` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(20, 15), 5)`)
}})
```

`gcd(13, 17)` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(13, 17), 1)`)
}})
```

`gcd(654, 456)` deve retornar `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(654, 456), 6)`)
}})
```

`gcd(3456, 4320)` deve retornar `864`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(3456, 4320), 864)`)
}})
```

# --seed--

## --seed-contents--

```py
def gcd(x, y):

    return x
```

# --solutions--

```py
def gcd(x, y):
    while y != 0:
        x, y = y, x % y
    return x
```
