---
id: 68f6587287ad1f4ad39b0c7c
title: "Desafio 90: Limite de Caracteres"
challengeType: 29
dashedName: challenge-90
---

# --description--

Neste desafio, você recebe uma string e precisa determinar se ela cabe em uma postagem de mídia social. Retorne as seguintes strings com base nas regras fornecidas:

- "postagem curta" se couber dentro do limite de 40 caracteres.
- "long post" se for maior que 40 caracteres e couber dentro de um limite de 80 caracteres.
- "post inválido" se for muito longo para caber dentro de qualquer limite.

# --hints--

`can_post("Hello world")` deve retornar `"short post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` deve retornar `"long post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` deve retornar `"invalid post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
