---
id: 68f6587287ad1f4ad39b0c84
title: "Desafio 98: Contagem de Retângulos"
challengeType: 29
dashedName: challenge-98
---

# --description--

Dado dois inteiros positivos representando a largura e a altura de um retângulo, determine quantos retângulos podem caber no dado.

- Conte apenas retângulos com largura e altura inteiras.

Por exemplo, dado `1` e `3`, retorne `6`. Três retângulos 1x1, dois retângulos 1x2 e um retângulo 1x3.

# --hints--

`count_rectangles(1, 3)` deve retornar `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 3), 6)`)
}})
```

`count_rectangles(3, 2)` deve retornar `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(3, 2), 18)`)
}})
```

`count_rectangles(1, 2)` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 2), 3)`)
}})
```

`count_rectangles(5, 4)` deve retornar `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(5, 4), 150)`)
}})
```

`count_rectangles(11, 19)` deve retornar `12540`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(11, 19), 12540)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_rectangles(width, height):

    return width
```

# --solutions--

```py
def count_rectangles(width, height):
    total = 0
    for w in range(1, width + 1):
        for h in range(1, height + 1):
            total += (width - w + 1) * (height - h + 1)
    return total
```
