---
id: 68f6587287ad1f4ad39b0c80
title: "Desafio 94: Gerador de Assinatura de Email"
challengeType: 29
dashedName: challenge-94
---

# --description--

Dadas strings para o nome, cargo e empresa de uma pessoa, retorne uma assinatura de email como uma única string usando as seguintes regras:

- O nome deve aparecer primeiro, precedido por um prefixo que depende da primeira letra do nome. Para nomes que começam com (case-insensitive):
  - `A-I`: Use `>>` como prefixo.
  - `J-R`: Use `--` como prefixo.
  - `S-Z`: Use `::` como prefixo.
- Uma vírgula e um espaço (`, `) devem seguir o nome.
- O título e a empresa devem seguir a vírgula e o espaço, separados por `" at "` (com espaços ao redor).

Por exemplo, dado `"Quinn Waverly"`, `"Founder and CEO"` e `"TechCo"` retorne `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generate_signature("Quinn Waverly", "Founder and CEO", "TechCo")` deve retornar `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo")`)
}})
```

`generate_signature("Alice Reed", "Engineer", "TechCo")` deve retornar `">>Alice Reed, Engineer at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo")`)
}})
```

`generate_signature("Tina Vaughn", "Developer", "example.com")` deve retornar `"::Tina Vaughn, Developer at example.com"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com")`)
}})
```

`generate_signature("B. B.", "Product Tester", "AcmeCorp")` deve retornar `">>B. B., Product Tester at AcmeCorp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp")`)
}})
```

`generate_signature("windstorm", "Cloud Architect", "Atmospheronics")` deve retornar `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_signature(name, title, company):

    return name
```

# --solutions--

```py
def generate_signature(name, title, company):
    first_letter = name[0].upper()
    if first_letter in "ABCDEFGHI":
        prefix = ">>"
    elif first_letter in "JKLMNOPQR":
        prefix = "--"
    else:
        prefix = "::"

    return f"{prefix}{name}, {title} at {company}"
```
