---
id: 68ffb91507a5b645769328cc
title: "Desafio 109: Qual é a Minha Idade Mesmo?"
challengeType: 29
dashedName: challenge-109
---

# --description--

Dada a data de aniversário de alguém no formato `YYYY-MM-DD`, retorne a idade da pessoa em 27 de novembro de 2025.

- Assuma que todos os aniversários são datas válidas antes de 27 de novembro de 2025.
- Retorna a idade como um inteiro.
- Certifique-se de considerar se a pessoa já fez aniversário em 2025.

# --hints--

`calculate_age("2000-11-20")` deve retornar `25`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-11-20"), 25)`)
}})
```

`calculate_age("2000-12-01")` deve retornar `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-12-01"), 24)`)
}})
```

`calculate_age("2014-10-25")` deve retornar `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2014-10-25"), 11)`)
}})
```

`calculate_age("1994-01-06")` deve retornar `31`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-01-06"), 31)`)
}})
```

`calculate_age("1994-12-14")` deve retornar `30`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-12-14"), 30)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_age(birthday):

    return birthday
```

# --solutions--

```py
from datetime import datetime
def calculate_age(birthday):
    today = datetime(2025, 11, 27)
    birth_date = datetime.strptime(birthday, "%Y-%m-%d")

    age = today.year - birth_date.year

    if (today.month, today.day) < (birth_date.month, birth_date.day):
        age -= 1

    return age
```
