---
id: 68ffb91507a5b645769328c4
title: "Desafio 101: Conversor de Títulos Markdown"
challengeType: 29
dashedName: challenge-101
---

# --description--

Dada uma string que representa um título Markdown, retorne o título HTML equivalente.

Um cabeçalho Markdown válido deve:

- Comece com zero ou mais espaços, seguido por
- 1 a 6 caracteres hash (`#`) em sequência, então
- Pelo menos um espaço. E finalmente,
- O texto do cabeçalho.

O número de símbolos de hash determina o nível do cabeçalho. Por exemplo, um símbolo de hash corresponde a uma tag `h1` e seis símbolos de hash correspondem a uma tag `h6`.

Se a string fornecida não tiver o formato exato acima, retorne `"Invalid format"`.

Por exemplo, dado `"# My level 1 heading"`, retorne `"<h1>My level 1 heading</h1>"`.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`convert("# My level 1 heading")` deve retornar `"<h1>My level 1 heading</h1>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>")`)
}})
```

`convert("My heading")` deve retornar `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("My heading"), "Invalid format")`)
}})
```

`convert("##### My level 5 heading")` deve retornar `"<h5>My level 5 heading</h5>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>")`)
}})
```

`convert("#My heading")` deve retornar `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("#My heading"), "Invalid format")`)
}})
```

`convert("  ###  My level 3 heading")` deve retornar `"<h3>My level 3 heading</h3>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>")`)
}})
```

`convert("####### My level 7 heading")` deve retornar `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("####### My level 7 heading"), "Invalid format")`)
}})
```

`convert("## My #2 heading")` deve retornar `"<h2>My #2 heading</h2>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("## My #2 heading"), "<h2>My #2 heading</h2>")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert(heading):

    return heading
```

# --solutions--

```py
import re
def convert(heading):
    trimmed = heading.strip()

    if not re.match(r'^#{1,6}\s+\S+', trimmed):
        return "Invalid format"

    hash_count = trimmed.find(' ')
    text = trimmed[hash_count:].lstrip()

    return f"<h{hash_count}>{text}</h{hash_count}>"
```
