---
id: 68ffb91507a5b645769328c9
title: "Desafio 106: Validador de Mensagens"
challengeType: 29
dashedName: challenge-106
---

# --description--

Dada uma string de mensagem e uma string de validação, determine se a mensagem é válida.

- Uma mensagem é válida se cada palavra na mensagem começar com a letra correspondente na string de validação, na ordem.
- Letras não diferenciam maiúsculas de minúsculas.
- As palavras na mensagem são separadas por espaços simples.

# --hints--

`is_valid_message("hello world", "hw")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("hello world", "hw"), True)`)
}})
```

`is_valid_message("ALL CAPITAL LETTERS", "acl")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("ALL CAPITAL LETTERS", "acl"), True)`)
}})
```

`is_valid_message("Coding challenge are boring.", "cca")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("Coding challenge are boring.", "cca"), False)`)
}})
```

`is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"), True)`)
}})
```

`is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_message(message, validation):

    return message
```

# --solutions--

```py
def is_valid_message(message, validation):
    words = message.split()
    
    if len(words) != len(validation):
        return False

    for word, val_letter in zip(words, validation):
        if word[0].lower() != val_letter.lower():
            return False

    return True
```
