---
id: 68ffb91507a5b645769328c7
title: "Desafio 104: Escalonador de Receitas"
challengeType: 29
dashedName: challenge-104
---

# --description--

Dado um array de ingredientes da receita e um número para escalar a receita, retorne um array com as quantidades escaladas de acordo.

- Cada item no array fornecido será uma string no formato: `"quantity unit ingredient"`. Por exemplo `"2 C Flour"`.
- Escale a quantidade pelo número fornecido. Não inclua zeros à direita e não converta nenhuma unidade.
- Retorne os itens escalonados na mesma ordem em que são fornecidos.

# --hints--

`scale_recipe(["2 C Flour", "1.5 T Sugar"], 2)` deve retornar `["4 C Flour", "3 T Sugar"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_recipe(["2 C Flour", "1.5 T Sugar"], 2), ["4 C Flour", "3 T Sugar"])`)
}})
```

`scale_recipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5)` deve retornar `["6 T Flour", "1.5 C Milk", "3 T Oil"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_recipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5), ["6 T Flour", "1.5 C Milk", "3 T Oil"])`)
}})
```

`scale_recipe(["3 C Milk", "2 C Oats"], 0.5)` deve retornar `["1.5 C Milk", "1 C Oats"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_recipe(["3 C Milk", "2 C Oats"], 0.5), ["1.5 C Milk", "1 C Oats"])`)
}})
```

`scale_recipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5)` deve retornar `["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_recipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5), ["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"])`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_recipe(ingredients, scale):

    return ingredients
```

# --solutions--

```py
def scale_recipe(ingredients, scale):
    scaled = []

    for item in ingredients:
        parts = item.split(" ")
        quantity = float(parts[0]) * scale
        unit = parts[1]
        ingredient_name = " ".join(parts[2:])

        if quantity.is_integer():
            quantity_str = str(int(quantity))
        else:
            quantity_str = str(quantity)
        
        scaled.append(f"{quantity_str} {unit} {ingredient_name}")

    return scaled
```
