---
id: 68ffb91507a5b645769328c6
title: "Desafio 103: LCM"
challengeType: 29
dashedName: challenge-103
---

# --description--

Dado dois inteiros, retorne o mínimo múltiplo comum (MMC) dos dois números.

O MMC de dois números é o menor inteiro positivo que é múltiplo de ambos os números. Por exemplo, dado `4` e `6`, retorne `12` porque:

- Múltiplos de `4` são `4`, `8`, `12` e assim por diante.
- Múltiplos de `6` são `6`, `12`, `18` e assim por diante.
- `12` é o menor número que é múltiplo de ambos.

# --hints--

`lcm(4, 6)` deve retornar `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(4, 6), 12)`)
}})
```

`lcm(9, 6)` deve retornar `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(9, 6), 18)`)
}})
```

`lcm(10, 100)` deve retornar `100`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(10, 100), 100)`)
}})
```

`lcm(13, 17)` deve retornar `221`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(13, 17), 221)`)
}})
```

`lcm(45, 70)` deve retornar `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(45, 70), 630)`)
}})
```

# --seed--

## --seed-contents--

```py
def lcm(a, b):

    return a
```

# --solutions--

```py
def lcm(a, b):
    def gcd(x, y):
        return x if y == 0 else gcd(y, x % y)

    return abs(a * b) // gcd(a, b)
```
