---
id: 68ffb91507a5b645769328c5
title: "Desafio 102: Palavra Mais Longa"
challengeType: 29
dashedName: challenge-102
---

# --description--

Dada uma string de sentença, retorne a palavra mais longa na sentença.

- Palavras são separadas por um único espaço.
- Apenas letras (`a-z`, case-insensitive) contam para o comprimento da palavra.
- Se houver várias palavras com o mesmo comprimento, retorne a primeira que aparecer.
- Retorne a palavra como aparece na string fornecida, com a pontuação removida.

# --hints--

`longest_word("The quick red fox")` deve retornar `"quick"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("The quick red fox"), "quick")`)
}})
```

`longest_word("Hello coding challenge.")` deve retornar `"challenge"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Hello coding challenge."), "challenge")`)
}})
```

`longest_word("Do Try This At Home.")` deve retornar `"This"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Do Try This At Home."), "This")`)
}})
```

`longest_word("This sentence... has commas, ellipses, and an exclamation point!")` deve retornar `"exclamation"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation")`)
}})
```

`longest_word("A tie? No way!")` deve retornar `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("A tie? No way!"), "tie")`)
}})
```

`longest_word("Wouldn't you like to know.")` deve retornar `"Wouldnt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Wouldn't you like to know."), "Wouldnt")`)
}})
```

# --seed--

## --seed-contents--

```py
def longest_word(sentence):

    return sentence
```

# --solutions--

```py
import re
def longest_word(sentence):
    words = sentence.split(" ")
    longest = ""

    for word in words:
        cleaned = re.sub(r'[^a-zA-Z]', '', word)
        if len(cleaned) > len(longest):
            longest = cleaned

    return longest
```
