---
id: 69162d64f96574d9bb629efe
title: "Desafio 113: Milhas para Quilômetros"
challengeType: 29
dashedName: challenge-113
---

# --description--

Dada uma distância em milhas como um número, retorne a distância equivalente em quilômetros.

- A entrada será sempre um número não negativo.
- 1 milha equivale a 1.60934 quilômetros.
- Arredonde o resultado para duas casas decimais.

# --hints--

`convert_to_km(1)` deve retornar `1.61`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(1), 1.61)`)
}})
```

`convert_to_km(21)` deve retornar `33.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(21), 33.8)`)
}})
```

`convert_to_km(3.5)` deve retornar `5.63`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(3.5), 5.63)`)
}})
```

`convert_to_km(0)` deve retornar `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0), 0)`)
}})
```

`convert_to_km(0.621371)` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0.621371), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_km(miles):

    return miles
```

# --solutions--

```py
def convert_to_km(miles):
    km = miles * 1.60934
    return round(km, 2)
```
