---
id: 69162d64f96574d9bb629f04
title: "Desafio 119: Compressão de String"
challengeType: 29
dashedName: challenge-119
---

# --description--

Dada uma string `sentence`, retorne uma versão comprimida da sentença onde palavras duplicadas consecutivas são substituídas pela palavra seguida do número de vezes que ela se repete entre parênteses.

- Apenas duplicatas consecutivas são comprimidas.
- Palavras são separadas por espaços simples.

Por exemplo, dado `"yes yes yes please"`, retorne `"yes(3) please"`.

# --hints--

`compress_string("yes yes yes please")` deve retornar `"yes(3) please"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("yes yes yes please"), "yes(3) please")`)
}})
```

`compress_string("I have have have apples")` deve retornar `"I have(3) apples"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("I have have have apples"), "I have(3) apples")`)
}})
```

`compress_string("one one three and to the the the the")` deve retornar `"one(2) three and to the(4)"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("one one three and to the the the the"), "one(2) three and to the(4)")`)
}})
```

`compress_string("route route route route route route tee tee tee tee tee tee")` deve retornar `"route(6) tee(6)"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)")`)
}})
```

# --seed--

## --seed-contents--

```py
def compress_string(sentence):

    return sentence
```

# --solutions--

```py
def compress_string(sentence):
    words = sentence.split(" ")
    result = []
    count = 1

    for i in range(len(words)):
        if i < len(words) - 1 and words[i] == words[i + 1]:
            count += 1
        else:
            if count > 1:
                result.append(f"{words[i]}({count})")
            else:
                result.append(words[i])
            count = 1

    return " ".join(result)
```
