---
id: 69162d64f96574d9bb629efd
title: "Desafio 112: Detector de IA"
challengeType: 29
dashedName: challenge-112
---

# --description--

O desafio de hoje é inspirado pelo lançamento do ChatGPT em 30 de novembro de 2022.

Dada uma string de uma ou mais sentenças, determine se ela foi provavelmente gerada por IA usando as seguintes regras:

- Ele contém dois ou mais traços (`-`).
- Ele contém dois ou mais conjuntos de parênteses (`()`). O texto pode estar dentro dos parênteses.
- Contém três ou mais palavras com 7 ou mais letras.

- Palavras são separadas por um único espaço e consistem apenas de letras (`A-Z`). Não inclua pontuação ou outros caracteres que não sejam letras como parte de uma palavra.

Se a frase dada atender a qualquer uma das regras acima, retorne `"AI"`, caso contrário, retorne `"Human"`.

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` deve retornar `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` deve retornar `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` deve retornar `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` deve retornar `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` deve retornar `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
