---
id: 69162d64f96574d9bb629eff
title: "Desafio 114: Camel para Snake"
challengeType: 29
dashedName: challenge-114
---

# --description--

Dada uma string em camel case, retorne a versão em snake case da string usando as seguintes regras:

- A string de entrada conterá apenas letras (`A-Z` e `a-z`) e sempre começará com uma letra minúscula.
- Cada letra maiúscula na string camel case inicia uma nova palavra.
- Converta todas as letras para minúsculas.
- Separe as palavras com um underscore (`_`).

# --hints--

`to_snake("helloWorld")` deve retornar `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` deve retornar `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` deve retornar `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
