---
id: 69162d64f96574d9bb629f00
title: "Desafio 115: Conversor de Itens de Lista Ordenada em Markdown"
challengeType: 29
dashedName: challenge-115
---

# --description--

Dada uma string que representa um item de lista ordenada em Markdown, retorne a string HTML equivalente.

Um item válido de lista ordenada em Markdown deve:

- Comece com zero ou mais espaços, seguido por
- Um número (1 ou maior) e um ponto (`.`), seguido de
- Pelo menos um espaço e então
- O texto do item da lista.

Se a string não tiver o formato exato acima, retorne `"Invalid format"`. Caso contrário, envolva o texto do item da lista em tags `li` e retorne a string.

Por exemplo, dado `"1. My item"`, retorne `"<li>My item</li>"`

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`convert_list_item("1. My item")` deve retornar `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` deve retornar `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` deve retornar `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` deve retornar `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` deve retornar `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` deve retornar `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
