---
id: 69162d64f96574d9bb629f01
title: "Desafio 116: Contagem de Permutação"
challengeType: 29
dashedName: challenge-116
---

# --description--

Dada uma string, retorne o número de permutações distintas que podem ser formadas a partir de seus caracteres.

- Uma permutação é qualquer reordenação dos caracteres na string.
- Não conte permutações duplicadas.
- Se a string contiver caracteres repetidos, arranjos repetidos devem ser contados apenas uma vez.
- A string conterá apenas letras (`A-Z`, `a-z`).

Por exemplo, dado `"abb"`, retorne `3` porque há três maneiras únicas de organizar as letras: `"abb"`, `"bab"` e `"bba"`.

# --hints--

`count_permutations("abb")` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` deve retornar `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` deve retornar `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` deve retornar `39916800`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
