---
id: 691b559495c5cb5a37b9b484
title: "Desafio 124: Atualização de Inventário"
challengeType: 29
dashedName: challenge-124
---

# --description--

Dada uma matriz 2D que representa o inventário da sua loja e outra matriz 2D que representa uma remessa que você recebeu, retorne seu inventário atualizado.

- Cada elemento nos arrays terá o formato: `[quantity, "item"]`, onde `quantity` é um inteiro e `"item"` é uma string.
- Atualize os itens no inventário adicionando a quantidade de quaisquer itens correspondentes da remessa.
- Se um item recebido não existir no inventário atual, adicione-o como uma nova entrada ao final do inventário.
- Retorne o inventário na ordem em que foi fornecido com os novos itens no final na ordem em que aparecem no envio.

Por exemplo, dado um inventário de `[[2, "apples"], [5, "bananas"]]` e um envio de `[[1, "apples"], [3, "bananas"]]`, retorne `[[3, "apples"], [8, "bananas"]]`.

# --hints--

`update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` deve retornar `[[3, "apples"], [8, "bananas"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]])`)
}})
```

`update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` deve retornar `[[3, "apples"], [8, "bananas"], [4, "oranges"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]])`)
}})
```

`update_inventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` deve retornar `[[10, "apples"], [30, "bananas"], [20, "oranges"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]])`)
}})
```

`update_inventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` deve retornar `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def update_inventory(inventory, shipment):

    return inventory
```

# --solutions--

```py
def update_inventory(inventory, shipment):
    inventory_map = {item: i for i, (qty, item) in enumerate(inventory)}

    for qty, item in shipment:
        if item in inventory_map:
            index = inventory_map[item]
            inventory[index][0] += qty
        else:
            inventory.append([qty, item])
            inventory_map[item] = len(inventory) - 1

    return inventory
```
