---
id: 691b559495c5cb5a37b9b488
title: "Desafio 128: Contagem de Consoantes"
challengeType: 29
dashedName: challenge-128
---

# --description--

Dada uma string e um número alvo, determine se a string contém exatamente o número alvo de consoantes.

- Consoantes são todos os caracteres alfabéticos exceto `"a"`, `"e"`, `"i"`, `"o"` e `"u"` em qualquer caixa.
- Ignore dígitos, pontuação, espaços e outros caracteres que não sejam letras ao contar.

# --hints--

`has_consonant_count("helloworld", 7)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("helloworld", 7), True)`)
}})
```

`has_consonant_count("eieio", 5)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("eieio", 5), False)`)
}})
```

`has_consonant_count("freeCodeCamp Rocks!", 11)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("freeCodeCamp Rocks!", 11), True)`)
}})
```

`has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24), False)`)
}})
```

`has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_consonant_count(text, target):

    return text
```

# --solutions--

```py
def has_consonant_count(text, target):
    vowels = {"a", "e", "i", "o", "u"}
    count = 0

    for ch in text.lower():
        if ch.isalpha() and ch not in vowels:
            count += 1

    return count == target
```
