---
id: 691b559495c5cb5a37b9b480
title: "Desafio 120: Libras para Quilogramas"
challengeType: 29
dashedName: challenge-120
---

# --description--

Dado um peso em libras como um número, retorne a string `"(lbs) pounds equals (kgs) kilograms."`.

- Substitua `"(lbs)"` pelo número de entrada.
- Substitua `"(kgs)"` pela entrada convertida para quilogramas, arredondada para duas casas decimais e sempre incluindo duas casas decimais no valor.
- 1 libra equivale a 0,453592 quilogramas.
- Se a entrada for `1`, use `"pound"` em vez de `"pounds"`.
- Se o valor convertido for `1`, use `"kilogram"` em vez de `"kilograms"`.

# --hints--

`convert_to_kgs(1)` deve retornar `"1 pound equals 0.45 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` deve retornar `"0 pounds equals 0.00 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` deve retornar `"100 pounds equals 45.36 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` deve retornar `"2.5 pounds equals 1.13 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` deve retornar `"2.20462 pounds equals 1.00 kilogram."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
