---
id: 691b559495c5cb5a37b9b483
title: "Desafio 123: Construtor de Numerais Romanos"
challengeType: 29
dashedName: challenge-123
---

# --description--

Dado um inteiro, retorne seu valor equivalente em numerais romanos.

Números romanos usam esses símbolos:

SímboloValor
| ------ | ----- |
I1
V5
X10
L50
C100
D500
M1000

Os números romanos são escritos do maior para o menor, da esquerda para a direita, usando as seguintes regras:

- A adição é usada quando um símbolo é seguido por outro de valor igual ou menor. Por exemplo, `18` é escrito como `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- A subtração é usada quando um símbolo menor aparece antes de um maior, para representar 4 ou 9 em qualquer valor posicional. Por exemplo, 19 é escrito como `XIX` (10 + (10 - 1)).
- Nenhum símbolo pode ser repetido mais de três vezes consecutivas. A subtração é usada quando você precisaria escrever um símbolo mais de três vezes consecutivas. Então o maior número que você pode escrever é 3999.

Aqui está mais um exemplo: dado `1464`, retorne `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`to_roman(18)` deve retornar `"XVIII"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` deve retornar `"XIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` deve retornar `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` deve retornar `"MMXXV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` deve retornar `"MMMCMXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```
