---
id: 691b559495c5cb5a37b9b487
title: "Desafio 127: Verificação de Velocidade"
challengeType: 29
dashedName: challenge-127
---

# --description--

Dada a velocidade que você está viajando em milhas por hora (MPH) e um limite de velocidade em quilômetros por hora (KPH), determine se você está acima da velocidade permitida e se receberá um aviso ou uma multa.

- 1 milha equivale a 1.60934 quilômetros.
- Se você estiver viajando a uma velocidade menor ou igual ao limite, retorne `"Not Speeding"`.
- Se você estiver viajando 5 KPH ou menos acima do limite de velocidade, retorne `"Warning"`.
- Se você estiver viajando mais de 5 KPH acima do limite de velocidade, retorne `"Ticket"`.

# --hints--

`speed_check(30, 70)` deve retornar `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` deve retornar `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` deve retornar `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` deve retornar `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` deve retornar `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` deve retornar `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
