---
id: 691b559495c5cb5a37b9b489
title: "Desafio 129: Parser de Bloco de Citação Markdown"
challengeType: 29
dashedName: challenge-129
---

# --description--

Dada uma string que inclui um blockquote em Markdown, retorne a string HTML equivalente.

Um blockquote em Markdown é qualquer linha que:

- Começa com zero ou mais espaços
- Sucedido por um sinal de maior (`>`)
- Então, um ou mais espaços
- E finalmente, o texto do blockquote.

Retorne o texto do bloco de citação cercado pelas tags HTML de abertura e fechamento `blockquote`.

Por exemplo, dado `"> This is a quote"`, retorne `<blockquote>This is a quote</blockquote>`.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parse_blockquote("> This is a quote")` deve retornar `"<blockquote>This is a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` deve retornar `"<blockquote>This is also a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` deve retornar `"<blockquote>So  Is  This</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
