---
id: 6925e2068081f40f549ced1b
title: "Desafio 137: Gerador de Snowflake"
challengeType: 29
dashedName: challenge-137
---

# --description--

Dada uma string multilinha que usa caracteres de nova linha (`\n`) para representar uma quebra de linha, retorne uma nova string onde cada linha é espelhada horizontalmente e anexada ao final da linha original.

- Espelhe uma linha invertendo todos os seus caracteres, incluindo espaços.

Por exemplo, dado `"* \n *\n* "`, que registra no console como:

```sh
* 
 *
* 
```

Retorna `"*  *\n ** \n*  *"`, que é exibido no console como:

```sh
*  *
 ** 
*  *
```

Preste muita atenção aos espaços em branco nas strings fornecidas e retornadas. Certifique-se de não remover nenhum deles.

# --hints--

`generate_snowflake("* \n *\n* ")` deve retornar `"*  *\n ** \n*  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` deve retornar `"X=~~=X"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` deve retornar `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` deve retornar `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` deve retornar `"*  --  *\n * -- * \n*  --  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
