---
id: 6925e2068081f40f549ced1c
title: "Desafio 138: Soma dos Divisores"
challengeType: 29
dashedName: challenge-138
---

# --description--

Dado um número inteiro positivo, retorne a soma de todos os seus divisores.

- Um divisor é qualquer número inteiro que divide o número de forma exata (o resto é `0`).
- Conte cada divisor apenas uma vez.

Por exemplo, dado `6`, retorne `12` porque os divisores de `6` são `1`, `2`, `3` e `6` e a soma deles é `12`.

# --hints--

`sum_divisors(6)` deve retornar `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` deve retornar `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` deve retornar `56`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` deve retornar `224`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` deve retornar `806`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` deve retornar `23520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
