---
id: 6925e2068081f40f549ced1d
title: "Desafio 139: Pedra, Papel e Tesoura"
challengeType: 29
dashedName: challenge-139
---

# --description--

Dadas duas strings, a primeira representando o Jogador 1 e a segunda representando o Jogador 2, determine o vencedor de uma partida de Pedra, Papel e Tesoura.

- As strings de entrada serão sempre `"Rock"`, `"Paper"` ou `"Scissors"`.
- `"Rock"` vence `"Scissors"`.
- `"Paper"` vence `"Rock"`.
- `"Tesoura"` vence `"Papel"`.

Retornar:

- "Jogador 1 vence" se o Jogador 1 vencer.
- "Jogador 2 vence" se o Jogador 2 vencer.
- "Tie" se ambos os jogadores escolherem a mesma opção.

# --hints--

`rock_paper_scissors("Rock", "Rock")` deve retornar `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Rock"), "Tie")`)
}})
```

`rock_paper_scissors("Rock", "Paper")` deve retornar `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Paper"), "Player 2 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Paper")` deve retornar `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Paper"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Rock", "Scissors")` deve retornar `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Scissors"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Scissors")` deve retornar `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Scissors"), "Tie")`)
}})
```

`rock_paper_scissors("Scissors", "Rock")` deve retornar `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Rock"), "Player 2 wins")`)
}})
```

# --seed--

## --seed-contents--

```py
def rock_paper_scissors(player1, player2):

    return player1
```

# --solutions--

```py
def rock_paper_scissors(player1, player2):
    if player1 == player2:
        return "Tie"
    
    if (player1 == "Rock" and player2 == "Scissors") or \
       (player1 == "Paper" and player2 == "Rock") or \
       (player1 == "Scissors" and player2 == "Paper"):
        return "Player 1 wins"
    else:
        return "Player 2 wins"
```
