---
id: 69272dcf1c24b44fd79137c4
title: "Desafio 141: Combustível de Decolagem"
challengeType: 29
dashedName: challenge-141
---

# --description--

Dado o número de galões de combustível atualmente no seu avião e o número necessário de litros de combustível para alcançar seu destino, determine quantos galões adicionais de combustível você deve adicionar.

- 1 galão equivale a 3.78541 litros.
- Se o avião já tiver combustível suficiente, retorne `0`.
- Você só pode adicionar galões inteiros.
- Não inclua decimais no número retornado.

# --hints--

`fuel_to_add(0, 1)` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(0, 1), 1)`)
}})
```

`fuel_to_add(5, 40)` deve retornar `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(5, 40), 6)`)
}})
```

`fuel_to_add(10, 30)` deve retornar `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(10, 30), 0)`)
}})
```

`fuel_to_add(896, 20500)` deve retornar `4520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(896, 20500), 4520)`)
}})
```

`fuel_to_add(1000, 50000)` deve retornar `12209`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(1000, 50000), 12209)`)
}})
```

# --seed--

## --seed-contents--

```py
def fuel_to_add(current_gallons, required_liters):

    return current_gallons
```

# --solutions--

```py
import math
def fuel_to_add(current_gallons, required_liters):
    liters_per_gallon = 3.78541
    current_liters = current_gallons * liters_per_gallon
    if current_liters >= required_liters:
        return 0
    liters_needed = required_liters - current_liters
    return math.ceil(liters_needed / liters_per_gallon)
```
