---
id: 69272dcf1c24b44fd79137c6
title: "Desafio 143: Parser de Itálico em Markdown"
challengeType: 29
dashedName: challenge-143
---

# --description--

Dada uma string que pode incluir algum texto em itálico no Markdown, retorne a string HTML equivalente.

- Texto em itálico no Markdown é qualquer texto que começa e termina com um asterisco simples (`*`) ou um sublinhado simples (`_`).
- Não pode haver espaços entre o texto e o asterisco ou sublinhado, mas pode haver espaços no próprio texto.
- Converta todas as ocorrências em itálico para tags HTML `i` e retorne a string.

Por exemplo, dado `"*This is italic*"`, retorne `"<i>This is italic</i>"`.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parse_italics("*This is italic*")` deve retornar `"<i>This is italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is italic*"), "<i>This is italic</i>")`)
}})
```

`parse_italics("_This is also italic_")` deve retornar `"<i>This is also italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_This is also italic_"), "<i>This is also italic</i>")`)
}})
```

`parse_italics("*This is not italic *")` deve retornar `"*This is not italic *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is not italic *"), "*This is not italic *")`)
}})
```

`parse_italics("_ This is also not italic_")` deve retornar `"_ This is also not italic_"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_ This is also not italic_"), "_ This is also not italic_")`)
}})
```

`parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog.")` deve retornar `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_italics(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_italics(markdown):
    pattern = r'(\*|_)([^\s][^]*?[^\s])\1'
    pattern = r'(\*|_)([^\s].*?[^\s])\1'
    return re.sub(pattern, r'<i>\2</i>', markdown)
```
