---
id: 69272dcf1c24b44fd79137c3
title: "Desafio 140: SCREAMING_SNAKE_CASE"
challengeType: 29
dashedName: challenge-140
---

# --description--

Dada uma string que representa um nome de variável, retorne o nome da variável convertido para SCREAMING_SNAKE_CASE.

Os nomes de variáveis fornecidos serão escritos em um dos seguintes formatos:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

Nos formatos acima, as palavras são separadas por um underscore (`_`), um hífen (`-`) ou uma nova palavra começa com uma letra maiúscula.

Para converter para SCREAMING_SNAKE_CASE:

- Deixe todas as letras em maiúsculas
- Separe as palavras com um underscore (`_`)

# --hints--

`to_screaming_snake_case("userEmail")` deve retornar `"USER_EMAIL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("userEmail"), "USER_EMAIL")`)
}})
```

`to_screaming_snake_case("UserPassword")` deve retornar `"USER_PASSWORD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("UserPassword"), "USER_PASSWORD")`)
}})
```

`to_screaming_snake_case("user_id")` deve retornar `"USER_ID"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user_id"), "USER_ID")`)
}})
```

`to_screaming_snake_case("user-address")` deve retornar `"USER_ADDRESS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user-address"), "USER_ADDRESS")`)
}})
```

`to_screaming_snake_case("username")` deve retornar `"USERNAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("username"), "USERNAME")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_screaming_snake_case(variable_name):

    return variable_name
```

# --solutions--

```py
import re
def to_screaming_snake_case(variable_name):
    temp = re.sub(r'[-_]+', ' ', variable_name)
    temp = re.sub(r'([a-z0-9])([A-Z])', r'\1 \2', temp)
    words = temp.strip().split()
    return '_'.join(words).upper()
```
