---
id: 69272dcf1c24b44fd79137c5
title: "Desafio 142: Some a String"
challengeType: 29
dashedName: challenge-142
---

# --description--

Dada uma string contendo dígitos e outros caracteres, retorne a soma de todos os números na string.

- Trate dígitos consecutivos como um único número. Por exemplo, `"13"` conta como 13, não 1 + 3.
- Ignore quaisquer caracteres que não sejam dígitos.

# --hints--

`string_sum("3apples2bananas")` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("3apples2bananas"), 5)`)
}})
```

`string_sum("10cats5dogs2birds")` deve retornar `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("10cats5dogs2birds"), 17)`)
}})
```

`string_sum("125344")` deve retornar `125344`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("125344"), 125344)`)
}})
```

`string_sum("a1b20c300")` deve retornar `321`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a1b20c300"), 321)`)
}})
```

`string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` deve retornar `1653`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653)`)
}})
```

# --seed--

## --seed-contents--

```py
def string_sum(s):

    return s
```

# --solutions--

```py
import re
def string_sum(s):
    numbers = re.findall(r'\d+', s)
    return sum(int(num) for num in numbers)
```
