---
id: 69306364df283fcaff2e1ada
title: "Desafio 149: vOwElcAsE"
challengeType: 29
dashedName: challenge-149
---

# --description--

Dada uma string, retorne uma nova string onde todas as vogais são convertidas para maiúsculas e todos os outros caracteres alfabéticos são convertidos para minúsculas.

- As vogais são `"a"`, `"e"`, `"i"`, `"o"` e `"u"` em qualquer caso.
- Caracteres não alfabéticos devem permanecer inalterados.

# --hints--

`vowel_case("vowelcase")` deve retornar `"vOwElcAsE"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("vowelcase"), "vOwElcAsE")`)
}})
```

`vowel_case("coding is fun")` deve retornar `"cOdIng Is fUn"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("coding is fun"), "cOdIng Is fUn")`)
}})
```

`vowel_case("HELLO, world!")` deve retornar `"hEllO, wOrld!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HELLO, world!"), "hEllO, wOrld!")`)
}})
```

`vowel_case("git cherry-pick")` deve retornar `"gIt chErry-pIck"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("git cherry-pick"), "gIt chErry-pIck")`)
}})
```

`vowel_case("HEAD~1")` deve retornar `"hEAd~1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HEAD~1"), "hEAd~1")`)
}})
```

# --seed--

## --seed-contents--

```py
def vowel_case(s):

    return s
```

# --solutions--

```py
def vowel_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char in vowels:
            result += char.upper()
        elif char.isalpha():
            result += char.lower()
        else:
            result += char

    return result
```
