---
id: 69306364df283fcaff2e1ad6
title: "Desafio 145: N-ésimo Número de Fibonacci"
challengeType: 29
dashedName: challenge-145
---

# --description--

Dado um inteiro `n`, retorne o `n`-ésimo número na sequência de fibonacci.

A sequência de Fibonacci é uma série de números onde cada número é a soma dos dois anteriores. Os 10 primeiros números da sequência são `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21` e `34`.

# --hints--

`nth_fibonacci(4)` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(4), 2)`)
}})
```

`nth_fibonacci(10)` deve retornar `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(10), 34)`)
}})
```

`nth_fibonacci(15)` deve retornar `377`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(15), 377)`)
}})
```

`nth_fibonacci(40)` deve retornar `63245986`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(40), 63245986)`)
}})
```

`nth_fibonacci(75)` deve retornar `1304969544928657`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(75), 1304969544928657)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_fibonacci(n):

    return n
```

# --solutions--

```py
def nth_fibonacci(n):
    if n == 1:
        return 0
    if n == 2:
        return 1

    a, b = 0, 1
    for _ in range(3, n + 1):
        a, b = b, a + b
    return b
```
