---
id: 69306364df283fcaff2e1ad7
title: "Desafio 146: Assento Canhoto na Mesa"
challengeType: 29
dashedName: challenge-146
---

# --description--

Dada uma matriz 4x2 (array de arrays) representando a disposição dos assentos para uma refeição, determine em quantos assentos uma pessoa canhota pode se sentar.

- Uma pessoa canhota não pode sentar onde uma pessoa destra estaria no assento imediatamente à esquerda dela.

Na matriz dada:

- Um `"R"` é um assento ocupado por uma pessoa destra.
- Um `"L"` é um assento ocupado por uma pessoa canhota.
- Um `"U"` é um assento desocupado.
- Apenas assentos desocupados estão disponíveis para sentar.
- Os assentos na fileira superior estão voltados para "baixo" e os assentos na fileira inferior estão voltados para "cima" (como uma mesa), então esquerda e direita são relativos à orientação do assento.
- Os assentos de canto têm apenas um assento ao lado deles.

Por exemplo, na matriz dada:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

O assento no canto superior esquerdo não pode ser ocupado porque há uma pessoa destra à esquerda. Os outros dois assentos disponíveis podem ser ocupados porque não há uma pessoa destra à esquerda.

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` deve retornar `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` deve retornar `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
