---
id: 69373793f5a867f769cde135
title: "Desafio 150: Parser de Lista Não Ordenada em Markdown"
challengeType: 29
dashedName: challenge-150
---

# --description--

Dada a string de uma lista não ordenada válida em Markdown, retorne a string HTML equivalente.

Uma lista não ordenada consiste em um ou mais itens de lista. Um item de lista válido aparece em sua própria linha e:

- Começa com um traço (`"-"`), seguido de
- Pelo menos um espaço e então
- O texto do item da lista.

A lista é fornecida como uma única string com linhas novas separadas pelo caractere de nova linha (`"\n"`). Não inclua os caracteres de nova linha no texto do item.

Envolva cada item da lista em tags HTML `li` e a lista inteira de itens em tags `ul`.

Por exemplo, dado `"- Item A\n- Item B"`, retorne `"<ul><li>Item A</li><li>Item B</li></ul>"`.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parse_unordered_list("- Item A\n- Item B")` deve retornar `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- Item A\\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>")`)
}})
```

`parse_unordered_list("-  JavaScript\n-  Python")` deve retornar `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("-  JavaScript\\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>")`)
}})
```

`parse_unordered_list("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` deve retornar `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- 2 C Flour\\n- 1/2 C Sugar\\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>")`)
}})
```

`parse_unordered_list("- A-1\n- A-2\n- B-1")` deve retornar `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- A-1\\n- A-2\\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_unordered_list(markdown):

    return markdown
```

# --solutions--

```py
def parse_unordered_list(markdown):
    lines = markdown.split("\n")
    list_items = []

    for line in lines:
        list_items.append(f"<li>{line[1:].strip()}</li>")

    return f"<ul>{''.join(list_items)}</ul>"
```
