---
id: 69373793f5a867f769cde138
title: "Desafio 153: Jogo da Velha"
challengeType: 29
dashedName: challenge-153
---

# --description--

Dada uma matriz 3×3 (um array de arrays) representando um jogo de Tic-Tac-Toe concluído, determine o vencedor.

- Cada elemento na matriz dada é ou um `"X"` ou um `"O"`.

Um jogador vence se tiver três de seus personagens em uma linha - horizontalmente, verticalmente ou diagonalmente.

Retornar:

- "`X wins` se o jogador X tiver três em linha."
- "`O wins` se o jogador O tiver três em linha."
- "Empate" se nenhum jogador tiver três em linha.

# --hints--

`tic_tac_toe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` deve retornar `"X wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins")`)
}})
```

`tic_tac_toe([["X", "O", "X"], ["X", "O", "X"], ["O", "O", "X"]])` deve retornar `"O wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "O", "X"], ["X", "O", "X"], ["O", "O", "X"]]), "O wins")`)
}})
```

`tic_tac_toe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` deve retornar `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw")`)
}})
```

`tic_tac_toe([["X", "X", "O"], ["X", "O", "X"], ["O", "X", "X"]])` deve retornar `"O wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "X", "O"], ["X", "O", "X"], ["O", "X", "X"]]), "O wins")`)
}})
```

`tic_tac_toe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` deve retornar `"X wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins")`)
}})
```

`tic_tac_toe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` deve retornar `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def tic_tac_toe(board):

    return board
```

# --solutions--

```py
def tic_tac_toe(board):
    lines = []

    for row in board:
        lines.append(row)

    for c in range(3):
        lines.append([board[0][c], board[1][c], board[2][c]])

    lines.append([board[0][0], board[1][1], board[2][2]])
    lines.append([board[0][2], board[1][1], board[2][0]])

    for line in lines:
        if all(cell == "X" for cell in line):
            return "X wins"
        if all(cell == "O" for cell in line):
            return "O wins"

    return "Draw"
```
