---
id: 69373793f5a867f769cde139
title: "Desafio 154: Par para o Buraco"
challengeType: 29
dashedName: challenge-154
---

# --description--

Dado dois inteiros, o par para um buraco de golfe e o número de tacadas que um golfista deu nesse buraco, retorne a pontuação do golfista usando termos de golfe.

Retornar:

- "Hole in one!" se levou uma tacada.
- "Eagle" se levou dois golpes a menos que o par.
- "Birdie" se levou uma tacada a menos que o par.
- "Par" se levou o mesmo número de tacadas que o par.
- “Bogey” se levou uma tacada a mais que o par.
- "Double bogey" se levou duas tacadas a mais que o par.

# --hints--

`golf_score(3, 3)` deve retornar `"Par"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 3), "Par")`)
}})
```

`golf_score(4, 3)` deve retornar `"Birdie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 3), "Birdie")`)
}})
```

`golf_score(3, 1)` deve retornar `"Hole in one!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 1), "Hole in one!")`)
}})
```

`golf_score(5, 7)` deve retornar `"Double bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 7), "Double bogey")`)
}})
```

`golf_score(4, 5)` deve retornar `"Bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 5), "Bogey")`)
}})
```

`golf_score(5, 3)` deve retornar `"Eagle"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 3), "Eagle")`)
}})
```

# --seed--

## --seed-contents--

```py
def golf_score(par, strokes):

    return par
```

# --solutions--

```py
def golf_score(par, strokes):
    if strokes == 1:
        return "Hole in one!"

    diff = strokes - par

    if diff == -2:
        return "Eagle"
    if diff == -1:
        return "Birdie"
    if diff == 0:
        return "Par"
    if diff == 1:
        return "Bogey"
    if diff == 2:
        return "Double bogey"
```
