---
id: 69373793f5a867f769cde13a
title: "Desafio 155: Plantar a Colheita"
challengeType: 29
dashedName: challenge-155
---

# --description--

Dado um inteiro representando o tamanho do seu campo agrícola, e `"acres"` ou `"hectares"` representando a unidade para o tamanho do seu campo agrícola, e um tipo de cultura, determine quantas plantas desse tipo você pode acomodar no seu campo.

- 1 acre equivale a 4046,86 metros quadrados.
- 1 hectare equivale a 10.000 metros quadrados.

Aqui está uma lista de culturas que serão fornecidas como entrada e quanto espaço uma única planta ocupa:

CortarEspaço por planta
|-|-|
|`"corn"`1 metro quadrado
|`"wheat"`0,1 metros quadrados
|`"soybeans"`0,5 metros quadrados
|`"tomatoes"`0,25 metros quadrados
|`"lettuce"`0,2 metros quadrados

Retorna o número de plantas que cabem no campo, arredondado para baixo até a planta inteira mais próxima.

# --hints--

`get_number_of_plants(1, "acres", "corn")` deve retornar `4046`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

`get_number_of_plants(2, "hectares", "lettuce")` deve retornar `100000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

`get_number_of_plants(20, "acres", "soybeans")` deve retornar `161874`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

`get_number_of_plants(3.75, "hectares", "tomatoes")` deve retornar `150000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

`get_number_of_plants(16.75, "acres", "tomatoes")` deve retornar `271139`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
