---
id: 6939b873185d8e00d453563d
title: "Desafio 158: Troca de Array"
challengeType: 29
dashedName: challenge-158
---

# --description--

Dado um array com dois valores, retorne um array com os valores trocados.

Por exemplo, dado `["A", "B"]` retorne `["B", "A"]`.

# --hints--

`array_swap(["A", "B"])` deve retornar `["B", "A"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["A", "B"]), ["B", "A"])`)
}})
```

`array_swap([25, 20])` deve retornar `[20, 25]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([25, 20]), [20, 25])`)
}})
```

`array_swap([True, False])` deve retornar `[False, True]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([True, False]), [False, True])`)
}})
```

`array_swap(["1", 1])` deve retornar `[1, "1"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["1", 1]), [1, "1"])`)
}})
```

# --seed--

## --seed-contents--

```py
def array_swap(arr):

    return arr
```

# --solutions--

```py
def array_swap(arr):

    return [arr[1], arr[0]]
```
