---
id: 6939b873185d8e00d453563c
title: "Desafio 157: Parser de Link Markdown"
challengeType: 29
dashedName: challenge-157
---

# --description--

Dada a string de um link em Markdown, retorne a string HTML equivalente.

Uma imagem Markdown tem o seguinte formato: `"[link_text](link_url)"`. Retorne a string da tag HTML `a` com o `href` definido para o `link_url` e o `link_text` como conteúdo da tag.

Por exemplo, dado `"[freeCodeCamp](https://freecodecamp.org/)"` retorne `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens — verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` deve retornar `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` deve retornar `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribua para o nosso repositório em https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` deve retornar `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribua para o nosso repositório em https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
