---
id: 6939b873185d8e00d453563e
title: "Desafio 159: Hipotenusa Inteira"
challengeType: 29
dashedName: challenge-159
---

# --description--

Dado dois inteiros positivos representando os comprimentos dos dois catetos (os dois lados curtos) de um triângulo retângulo, determine se a hipotenusa é um inteiro.

O comprimento da hipotenusa é calculado somando os quadrados dos comprimentos dos dois catetos e depois extraindo a raiz quadrada desse total (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

`is_integer_hypotenuse(3, 4)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(3, 4), True)`)
}})
```

`is_integer_hypotenuse(2, 3)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(2, 3), False)`)
}})
```

`is_integer_hypotenuse(5, 12)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(5, 12), True)`)
}})
```

`is_integer_hypotenuse(10, 10)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(10, 10), False)`)
}})
```

`is_integer_hypotenuse(780, 1040)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(780, 1040), True)`)
}})
```

`is_integer_hypotenuse(250, 333)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(250, 333), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_integer_hypotenuse(a, b):

    return a
```

# --solutions--

```py
import math
def is_integer_hypotenuse(a, b):
    total = a*a + b*b
    c = math.isqrt(total)
    return c*c == total
```
