---
id: 694596b0585c11170ac7c7fb
title: "Desafio 163: Caso das Consoantes"
challengeType: 29
dashedName: challenge-163
---

# --description--

Dada uma string que representa o nome de uma variável, converta-a para consonant case usando as seguintes regras:

- Todas as consoantes devem ser convertidas para maiúsculas.
- Todas as vogais (`a`, `e`, `i`, `o`, `u` em qualquer caso) devem ser convertidas para minúsculas.
- Todos os hífens (`-`) devem ser convertidos em underscores (`_`).

# --hints--

`to_consonant_case("helloworld")` deve retornar `"HeLLoWoRLD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("helloworld"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("HELLOWORLD")` deve retornar `"HeLLoWoRLD"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("HELLOWORLD"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("_hElLO-WOrlD-")` deve retornar `"_HeLLo_WoRLD_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_hElLO-WOrlD-"), "_HeLLo_WoRLD_")`)
}})
```

`to_consonant_case("_~-generic_~-variable_~-name_~-here-~_")` deve retornar `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_consonant_case(s):

    return s
```

# --solutions--

```py
def to_consonant_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char == "-":
            result += "_"
        elif char.isalpha():
            if char in vowels:
                result += char.lower()
            else:
                result += char.upper()
        else:
            result += char

    return result
```
