---
id: 694596b0585c11170ac7c7fa
title: "Desafio 162: Consumo de Energia"
challengeType: 29
dashedName: challenge-162
---

# --description--

Dado o número de Calorias queimadas durante um treino e o número de watt-horas usados pelos seus dispositivos eletrônicos durante esse treino, determine qual deles usou mais energia.

Para compará-los, converta ambos os valores para joules usando as seguintes conversões:

- 1 Caloria equivale a 4184 joules.
- 1 watt-hora equivale a 3600 joules.

Retornar:

- "Workout" se o treino usou mais energia.
- "Dispositivos" se o dispositivo usou mais energia.
- "Igual" se ambos usaram a mesma quantidade de energia.

# --hints--

`compare_energy(250, 50)` deve retornar `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` deve retornar `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` deve retornar `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` deve retornar `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` deve retornar `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` deve retornar `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
