---
id: 696655d24b614176d4c9b78b
title: "Desafio 168: Imagem Escalada"
challengeType: 29
dashedName: challenge-168
---

# --description--

Dada uma string que representa a largura e altura de uma imagem e um número para escalar a imagem, retorne a largura e altura escaladas.

- A string de entrada está no formato `"WxH"`. Por exemplo, `"800x600"`.
- A escala é um número para multiplicar a largura e a altura.

Retorne as dimensões escaladas no mesmo formato `"WxH"`.

# --hints--

`scale_image("800x600", 2)` deve retornar `"1600x1200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("800x600", 2), "1600x1200")`)
}})
```

`scale_image("100x100", 10)` deve retornar `"1000x1000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("100x100", 10), "1000x1000")`)
}})
```

`scale_image("1024x768", 0.5)` deve retornar `"512x384"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("1024x768", 0.5), "512x384")`)
}})
```

`scale_image("300x200", 1.5)` deve retornar `"450x300"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("300x200", 1.5), "450x300")`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_image(size, scale):

    return size
```

# --solutions--

```py
def scale_image(size, scale):
    width, height = map(int, size.split("x"))

    new_width = round(width * scale)
    new_height = round(height * scale)

    return f"{new_width}x{new_height}"
```
