---
id: 696655d24b614176d4c9b78a
title: "Desafio 167: Bingo! Letra"
challengeType: 29
dashedName: challenge-167
---

# --description--

Dado um número, retorne a letra do bingo associada a ele (em maiúscula). Os números do bingo são agrupados da seguinte forma:

| Letter | Number Range |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

`get_bingo_letter(75)` deve retornar `"O"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(75), "O")`)
}})
```

`get_bingo_letter(54)` deve retornar `"G"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(54), "G")`)
}})
```

`get_bingo_letter(25)` deve retornar `"I"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(25), "I")`)
}})
```

`get_bingo_letter(38)` deve retornar `"N"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(38), "N")`)
}})
```

`get_bingo_letter(11)` deve retornar `"11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(11), "B")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_bingo_letter(n):

    return n
```

# --solutions--

```py
def get_bingo_letter(n):
    if 1 <= n <= 15:
        return "B"
    if 16 <= n <= 30:
        return "I"
    if 31 <= n <= 45:
        return "N"
    if 46 <= n <= 60:
        return "G"
    if 61 <= n <= 75:
        return "O"
    return None
```
