---
id: 696655d24b614176d4c9b78d
title: "Desafio 170: Dia Ímpar ou Par"
challengeType: 29
dashedName: challenge-170
---

# --description--

Dado um timestamp (número de milissegundos desde a época Unix), retorne:

- `"odd"` se o dia do mês para esse timestamp for ímpar.
- "even" se o dia do mês para esse timestamp for par.

Por exemplo, dado `1769472000000`, um timestamp para 27 de janeiro de 2026, retorne `"odd"` porque o número do dia (`27`) é um número ímpar.

# --hints--

`odd_or_even_day(1769472000000)` deve retornar `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` deve retornar `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` deve retornar `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` deve retornar `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` deve retornar `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
