---
id: 696655d24b614176d4c9b78c
title: "Desafio 169: FizzBuzz Mini"
challengeType: 29
dashedName: challenge-169
---

# --description--

Dado um inteiro, retorne uma string baseada nas seguintes regras:

- Se o número for divisível por 3, retorne `"Fizz"`.
- Se o número for divisível por 5, retorne `"Buzz"`.
- Se o número for divisível por 3 e 5, retorne `"FizzBuzz"`.
- Caso contrário, retorne o número fornecido como uma string.

# --hints--

`fizz_buzz_mini(3)` deve retornar `"Fizz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(3), "Fizz")`)
}})
```

`fizz_buzz_mini(4)` deve retornar `"4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(4), "4")`)
}})
```

`fizz_buzz_mini(35)` deve retornar `"Buzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(35), "Buzz")`)
}})
```

`fizz_buzz_mini(75)` deve retornar `"FizzBuzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(75), "FizzBuzz")`)
}})
```

`fizz_buzz_mini(98)` deve retornar `"98"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(98), "98")`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_mini(n):

    return n
```

# --solutions--

```py
def fizz_buzz_mini(n):
    if n % 3 == 0 and n % 5 == 0:
        return "FizzBuzz"
    if n % 3 == 0:
        return "Fizz"
    if n % 5 == 0:
        return "Buzz"
    return str(n)
```
