---
id: 69738771fb5a7b8b24cca2a5
title: "Desafio 179: Troco de Bolso"
challengeType: 29
dashedName: challenge-179
---

# --description--

Dado um array de inteiros representando as moedas no seu bolso, com cada inteiro sendo o valor de uma moeda em centavos, retorne o valor total no formato `"$D.CC"`.

- 100 centavos equivalem a 1 dólar.
- No valor de retorno, inclua um zero à esquerda para valores menores que um dólar e sempre exatamente dois dígitos para os centavos.

# --hints--

`count_change([25, 10, 5, 1])` deve retornar `"$0.41"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1]), "$0.41")`)
}})
```

`count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` deve retornar `"$1.43"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43")`)
}})
```

`count_change([100, 25, 100, 1000, 5, 500, 2000, 25])` deve retornar `"$37.55"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55")`)
}})
```

`count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` deve retornar `"$0.70"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70")`)
}})
```

`count_change([1])` deve retornar `"$0.01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([1]), "$0.01")`)
}})
```

`count_change([25, 25, 25, 25])` deve retornar `"$1.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 25, 25, 25]), "$1.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_change(change):

    return change
```

# --solutions--

```py
def count_change(change):
    total_cents = sum(change)
    dollars = total_cents // 100
    cents = total_cents % 100
    return f"${dollars}.{cents:02d}"
```
