---
id: 69738771fb5a7b8b24cca2a2
title: "Desafio 176: Dia da Marmota"
challengeType: 29
dashedName: challenge-176
---

# --description--

Hoje é o Dia da Marmota, no qual uma marmota prevê o tempo com base em se ela vê ou não sua sombra.

Dado um valor que representa a aparência da marmota, retorne a previsão correta:

- Se o valor fornecido for o booleano `true` (a marmota viu sua sombra), retorne `"Looks like we'll have six more weeks of winter."`.
- Se o valor for o booleano `false` (a marmota não viu sua sombra), retorne `"It's going to be an early spring."`.
- Se o valor for qualquer outra coisa (a marmota não apareceu), retorne `"No prediction this year."`.

# --hints--

`groundhog_day_prediction(True)` deve retornar `"Looks like we'll have six more weeks of winter."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

`groundhog_day_prediction(False)` deve retornar `"It's going to be an early spring."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

`groundhog_day_prediction(None)` deve retornar `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` deve retornar `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

`groundhog_day_prediction("True")` deve retornar `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
