---
id: 69738771fb5a7b8b24cca2a4
title: "Desafio 178: Truncar o Texto"
challengeType: 29
dashedName: challenge-178
---

# --description--

Dada uma string, retorne-a como está se tiver 20 caracteres ou menos. Se for maior que 20 caracteres, trunque-a para os primeiros 17 caracteres e adicione `"..."` ao final (para que tenha 20 caracteres no total) e retorne o resultado.

# --hints--

`truncate_text("Hello, world!")` deve retornar `"Hello, world!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Hello, world!"), "Hello, world!")`)
}})
```

`truncate_text("This string should get truncated.")` deve retornar `"This string shoul..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("This string should get truncated."), "This string shoul...")`)
}})
```

`truncate_text("Exactly twenty chars")` deve retornar `"Exactly twenty chars"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Exactly twenty chars"), "Exactly twenty chars")`)
}})
```

`truncate_text(".....................")` deve retornar `"...................."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("....................."), "....................")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(text):

    return text
```

# --solutions--

```py
def truncate_text(text):
    if len(text) <= 20:
        return text
    return text[:17] + "..."
```
