---
id: 69738771fb5a7b8b24cca2a3
title: "Desafio 177: Espelho de String"
challengeType: 29
dashedName: challenge-177
---

# --description--

Dada uma string, retorne uma nova string que consiste na string dada com uma cópia invertida dela mesma anexada ao final.

# --hints--

`mirror("freeCodeCamp")` deve retornar `"freeCodeCamppmaCedoCeerf"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf")`)
}})
```

`mirror("RaceCar")` deve retornar `"RaceCarraCecaR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("RaceCar"), "RaceCarraCecaR")`)
}})
```

`mirror("helloworld")` deve retornar `"helloworlddlrowolleh"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("helloworld"), "helloworlddlrowolleh")`)
}})
```

`mirror("The quick brown fox...")` deve retornar `"The quick brown fox......xof nworb kciuq ehT"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT")`)
}})
```

# --seed--

## --seed-contents--

```py
def mirror(s):

    return s
```

# --solutions--

```py
def mirror(s):
    reversed_s = s[::-1]
    return s + reversed_s
```
