---
id: 697a49e6ff50d756c9b69366
title: "Desafio 189: Jogos de Inverno 2026 Dia 10: Esqui Alpino"
challengeType: 29
dashedName: challenge-189
---

# --description--

Dada a queda vertical, distância horizontal e tipo de uma pista de esqui, determine a classificação de dificuldade da pista.

Para determinar a classificação:

- Calcule a inclinação da colina dividindo a queda pela distância.
- Then, calculate the adjusted steepness based on the hill type:
  - "Downhill" multiplique a inclinação por 1.2
  - "Slalom": multiplique a inclinação por 0.9
  - "Giant Slalom": multiplique a inclinação por 1.0

Retornar:

- "Green" se a inclinação ajustada for menor ou igual a 0.1
- "Blue" se a inclinação ajustada for maior que 0,1 e menor ou igual a 0,25
- "Preto" se a inclinação ajustada for maior que 0,25


# --hints--

`get_hill_rating(95, 900, "Slalom")` deve retornar `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(95, 900, "Slalom"), "Green")`)
}})
```

`get_hill_rating(620, 2800, "Downhill")` deve retornar `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(620, 2800, "Downhill"), "Black")`)
}})
```

`get_hill_rating(420, 1680, "Giant Slalom")` deve retornar `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(420, 1680, "Giant Slalom"), "Blue")`)
}})
```

`get_hill_rating(250, 3000, "Downhill")` deve retornar `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(250, 3000, "Downhill"), "Green")`)
}})
```

`get_hill_rating(110, 900, "Slalom")` deve retornar `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(110, 900, "Slalom"), "Blue")`)
}})
```

`get_hill_rating(380, 1500, "Giant Slalom")` deve retornar `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(380, 1500, "Giant Slalom"), "Black")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_hill_rating(drop, distance, hill_type):

    return drop
```

# --solutions--

```py
def get_hill_rating(drop, distance, hill_type):
    steepness = drop / distance

    if hill_type == "Downhill":
        steepness *= 1.2
    elif hill_type == "Slalom":
        steepness *= 0.9
    elif hill_type == "Giant Slalom":
        steepness *= 1.0

    if steepness <= 0.1:
        return "Green"
    elif steepness <= 0.25:
        return "Blue"
    else:
        return "Black"
```
