---
id: 697a49e6ff50d756c9b69364
title: "Desafio 187: Jogos de Inverno 2026 Dia 8: Luge"
challengeType: 29
dashedName: challenge-187
---

# --description--

Dado um array de cinco números, cada um representando o tempo (em segundos) que um luger levou para completar um segmento de uma pista, determine qual segmento teve a velocidade mais rápida e qual foi essa velocidade.

A trilha está dividida nos seguintes segmentos:

- Segmento 1: 320 metros
- Segmento 2: 280 metros
- Segmento 3: 350 metros
- Segmento 4: 300 metros
- Segmento 5: 250 metros

O primeiro valor no array fornecido corresponde ao tempo para o segmento 1, o segundo valor para o segmento 2 e assim por diante.

Para calcular a velocidade (em metros por segundo) de um segmento, divida a distância pelo tempo.

Retorne `"A velocidade máxima do luger foi de X m/s no segmento Y."`. Onde `X` é a velocidade máxima, arredondada para duas casas decimais, e `Y` é o número do segmento onde a velocidade máxima ocorreu.

# --hints--

`get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128])` deve retornar `"A maior velocidade do luger foi 35.07 m/s no segmento 5."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.")`)
}})
```

`get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284])` deve retornar `"A maior velocidade do luger foi 37.75 m/s no segmento 2."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.")`)
}})
```

`get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912])` deve retornar `"A maior velocidade do luger foi 38.49 m/s no segmento 3."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.")`)
}})
```

`get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840])` deve retornar `"The luger's fastest speed was 37.69 m/s on segment 1."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.")`)
}})
```

`get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508])` deve retornar `"A maior velocidade do luger foi 39.24 m/s no segmento 4."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_fastest_speed(times):

    return times
```

# --solutions--

```py
def get_fastest_speed(times):
    distances = [320, 280, 350, 300, 250]
    fastest_speed = 0
    fastest_segment = 0

    for i, t in enumerate(times):
        speed = distances[i] / t
        if speed > fastest_speed:
            fastest_speed = speed
            fastest_segment = i + 1  # 1-indexed

    return f"The luger's fastest speed was {fastest_speed:.2f} m/s on segment {fastest_segment}."
```
