---
id: 697a49e6ff50d756c9b6935f
title: "Desafio 182: Jogos de Inverno 2026 Dia 3: Biatlo"
challengeType: 29
dashedName: challenge-182
---

# --description--

Dado um array de inteiros, onde cada valor representa o número de alvos atingidos em uma única rodada de um biatlo, retorne a distância total de penalidade que o atleta deve esquiar.

- Cada rodada consiste em 5 alvos.
- Cada alvo perdido resulta em um loop de penalidade de 150 metros.

# --hints--

`calculate_penalty_distance([4, 4])` deve retornar `300`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 4]), 300)`)
}})
```

`calculate_penalty_distance([5, 5])` deve retornar `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 5]), 0)`)
}})
```

`calculate_penalty_distance([4, 5, 3, 5])` deve retornar `450`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 5, 3, 5]), 450)`)
}})
```

`calculate_penalty_distance([5, 4, 5, 5])` deve retornar `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 4, 5, 5]), 150)`)
}})
```

`calculate_penalty_distance([4, 3, 0, 3])` deve retornar `1500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 3, 0, 3]), 1500)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_penalty_distance(rounds):

    return rounds
```

# --solutions--

```py
def calculate_penalty_distance(rounds):
    total_penalty = 0

    for hits in rounds:
        misses = 5 - hits
        total_penalty += misses * 150

    return total_penalty
```
