---
id: 697a49e6ff50d756c9b69361
title: "Desafio 184: Jogos de Inverno 2026 Dia 5: Esqui Cross-Country"
challengeType: 29
dashedName: challenge-184
---

# --description--

Dado um array de tempos de chegada para uma corrida de esqui cross-country, converta-os em tempos atrás do vencedor.

- Os horários fornecidos são strings no formato `"H:MM:SS"`.
- Os tempos fornecidos estarão em ordem do mais rápido ao mais lento.
- O tempo dos vencedores (tempo mais rápido) deve corresponder a `"0"`.
- Cada outro tempo deve mostrar o tempo atrás do vencedor, no formato `"+M:SS"`.

Por exemplo, dado `["1:25:32", "1:26:10", "1:27:05"]`, retorne `["0", "+0:38", "+1:33"]`.

# --hints--

`get_relative_results(["1:25:32", "1:26:10", "1:27:05"])` deve retornar `["0", "+0:38", "+1:33"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["1:25:32", "1:26:10", "1:27:05"]), ["0", "+0:38", "+1:33"])`)
}})
```

`get_relative_results(["1:00:01", "1:00:05", "1:00:10"])` deve retornar `["0", "+0:04", "+0:09"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["1:00:01", "1:00:05", "1:00:10"]), ["0", "+0:04", "+0:09"])`)
}})
```

`get_relative_results(["1:10:06", "1:10:23", "1:10:48", "1:12:11"])` deve retornar `["0", "+0:17", "+0:42", "+2:05"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["1:10:06", "1:10:23", "1:10:48", "1:12:11"]), ["0", "+0:17", "+0:42", "+2:05"])`)
}})
```

`get_relative_results(["0:49:13", "0:49:15", "0:50:14", "0:51:30", "0:51:58", "0:52:16", "0:53:12", "0:53:31", "0:56:19", "1:02:20"])` deve retornar `["0", "+0:02", "+1:01", "+2:17", "+2:45", "+3:03", "+3:59", "+4:18", "+7:06", "+13:07"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["0:49:13", "0:49:15", "0:50:14", "0:51:30", "0:51:58", "0:52:16", "0:53:12", "0:53:31", "0:56:19", "1:02:20"]), ["0", "+0:02", "+1:01", "+2:17", "+2:45", "+3:03", "+3:59", "+4:18", "+7:06", "+13:07"])`)
}})
```

`get_relative_results(["2:01:15", "2:10:45", "2:10:53", "2:11:04", "2:11:55", "2:13:27", "2:14:30", "2:15:10"])` deve retornar `["0", "+9:30", "+9:38", "+9:49", "+10:40", "+12:12", "+13:15", "+13:55"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["2:01:15", "2:10:45", "2:10:53", "2:11:04", "2:11:55", "2:13:27", "2:14:30", "2:15:10"]), ["0", "+9:30", "+9:38", "+9:49", "+10:40", "+12:12", "+13:15", "+13:55"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_relative_results(results):

    return results
```

# --solutions--

```py
def get_relative_results(results):
    def time_to_seconds(time_str):
        hours, minutes, seconds = map(int, time_str.split(':'))
        return hours * 3600 + minutes * 60 + seconds
    
    def seconds_to_time_format(secs):
        mins = secs // 60
        secs_remainder = secs % 60
        return f"+{mins}:{secs_remainder:02d}"
    
    winner_seconds = time_to_seconds(results[0])
    
    relative_results = []
    for i, time in enumerate(results):
        if i == 0:
            relative_results.append('0')
        else:
            current_seconds = time_to_seconds(time)
            difference = current_seconds - winner_seconds
            relative_results.append(seconds_to_time_format(difference))
    
    return relative_results
```
