---
id: 697a49e6ff50d756c9b6935e
title: "Desafio 181: Jogos de Inverno 2026 Dia 2: Snowboard"
challengeType: 29
dashedName: challenge-181
---

# --description--

Dada a posição inicial de um snowboarder e uma rotação em graus, determine a posição de aterrissagem dele.

- A posição de um snowboarder é `"Regular"` ou `"Goofy"`.
- As rotações de truque são múltiplos de 90 graus. Positivo indica rotação no sentido horário e negativo indica rotação no sentido anti-horário.
- A posição de aterrissagem inverte a cada 180 graus de rotação.

Por exemplo, dado `"Regular"` e `90`, retorne `"Regular"`. Dado `"Regular"` e `180` graus, retorne `"Goofy"`.

# --hints--

`get_landing_stance("Regular", 90)` deve retornar `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` deve retornar `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` deve retornar `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` deve retornar `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` deve retornar `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` deve retornar `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
