---
id: 697a49e9860d24853adef67c
title: "Desafio 191: Jogos de Inverno 2026 Dia 12: Bobsled"
challengeType: 29
dashedName: challenge-191
---

# --description--

Dado um array representando os pesos dos atletas em uma equipe de bobsled e um número representando o peso do bobsled, determine se a equipe está elegível para competir.

- O tamanho do array determina o tamanho da equipe: equipes de 1, 2 ou 4 pessoas.
- Todos os valores de peso fornecidos estão em quilogramas (kg).

O trenó (trenó por si só) deve ter um peso mínimo de:

- 162 kg para uma equipe de 1 pessoa
- 170 kg para uma equipe de 2 pessoas
- 210 kg para uma equipe de 4 pessoas

O peso total do trenó (atletas mais trenó) não deve exceder:

- 247 kg para uma equipe de 1 pessoa
- 390 kg para uma equipe de 2 pessoas
- 630 kg para uma equipe de 4 pessoas

Retorne `"Eligible"` se a equipe atender a todos os requisitos, ou `"Not Eligible"` se a equipe não cumprir um ou mais dos requisitos.

# --hints--

`check_eligibility([78], 165)` deve retornar `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([78], 165), "Eligible")`)
}})
```

`check_eligibility([80], 160)` deve retornar `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 160), "Not Eligible")`)
}})
```

`check_eligibility([80], 170)` deve retornar `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 170), "Not Eligible")`)
}})
```

`check_eligibility([85, 90], 170)` deve retornar `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 90], 170), "Eligible")`)
}})
```

`check_eligibility([85, 95], 168)` deve retornar `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 95], 168), "Not Eligible")`)
}})
```

`check_eligibility([112, 97], 185)` deve retornar `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([112, 97], 185), "Not Eligible")`)
}})
```

`check_eligibility([110, 102, 90, 106], 222)` deve retornar `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([110, 102, 90, 106], 222), "Eligible")`)
}})
```

`check_eligibility([106, 99, 90, 88], 205)` deve retornar `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 90, 88], 205), "Not Eligible")`)
}})
```

`check_eligibility([106, 99, 103, 96], 227)` deve retornar `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 103, 96], 227), "Not Eligible")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_eligibility(athlete_weights, sled_weight):

    return athlete_weights
```

# --solutions--

```py
def check_eligibility(athlete_weights, sled_weight):
    team_size = len(athlete_weights)

    rules = {
        1: {"min_sled": 162, "max_total": 247},
        2: {"min_sled": 170, "max_total": 390},
        4: {"min_sled": 210, "max_total": 630}
    }

    min_sled = rules[team_size]["min_sled"]
    max_total = rules[team_size]["max_total"]

    total_weight = sum(athlete_weights) + sled_weight

    if sled_weight < min_sled:
        return "Not Eligible"

    if total_weight > max_total:
        return "Not Eligible"

    return "Eligible"
```
