---
id: 697a49e6ff50d756c9b69363
title: "Desafio 186: Jogos de Inverno 2026 Dia 7: Patinação de Velocidade"
challengeType: 29
dashedName: challenge-186
---

# --description--

Dadas duas arrays representando os tempos de volta (em segundos) para dois patinadores de velocidade, retorne o número da volta onde a diferença nos tempos de volta é a maior.

O primeiro elemento de cada array corresponde à volta 1, o segundo à volta 2 e assim por diante.

# --hints--

`largest_difference([26.11, 25.80, 25.92, 26.23, 26.07], [25.93, 25.74, 26.53, 26.11, 26.30])` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([26.11, 25.80, 25.92, 26.23, 26.07], [25.93, 25.74, 26.53, 26.11, 26.30]), 3)`)
}})
```

`largest_difference([27.04, 25.94, 26.22, 26.07, 26.18], [25.59, 25.80, 26.11, 26.01, 26.23])` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([27.04, 25.94, 26.22, 26.07, 26.18], [25.59, 25.80, 26.11, 26.01, 26.23]), 1)`)
}})
```

`largest_difference([25.82, 25.90, 26.05, 26.00, 26.48], [25.85, 25.92, 26.07, 25.98, 25.95])` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([25.82, 25.90, 26.05, 26.00, 26.48], [25.85, 25.92, 26.07, 25.98, 25.95]), 5)`)
}})
```

`largest_difference([25.88, 26.10, 25.95, 26.05, 26.00], [25.90, 26.55, 25.92, 26.03, 26.01])` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([25.88, 26.10, 25.95, 26.05, 26.00], [25.90, 26.55, 25.92, 26.03, 26.01]), 2)`)
}})
```

`largest_difference([25.92, 26.01, 26.05, 25.88, 26.12], [25.95, 26.00, 26.03, 26.45, 26.10])` deve retornar `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([25.92, 26.01, 26.05, 25.88, 26.12], [25.95, 26.00, 26.03, 26.45, 26.10]), 4)`)
}})
```

# --seed--

## --seed-contents--

```py
def largest_difference(skater1, skater2):

    return skater1
```

# --solutions--

```py
def largest_difference(skater1, skater2):
    max_diff = -1
    lap_number = 1

    for i in range(len(skater1)):
        diff = abs(skater1[i] - skater2[i])
        if diff > max_diff:
            max_diff = diff
            lap_number = i + 1

    return lap_number
```
