---
id: 697a49e6ff50d756c9b69365
title: "Desafio 188: Jogos de Inverno 2026 Dia 9: Skeleton"
challengeType: 29
dashedName: challenge-188
---

# --description--

Dada uma string que representa as curvas em uma pista de esqueleto, determine a dificuldade da pista.

- The given string will only consist of the letters:
  - "L" para uma curva à esquerda
  - `"R"` para uma curva à direita
  - "S" para um segmento reto

- Cada mudança de direção adiciona 15 pontos (um `"L"` seguido de um `"R"` ou vice-versa).
- Todas as outras curvas adicionam 5 pontos cada (todos os outros caracteres `"L"` ou `"R"`).
- Segmentos retos adicionam 0 pontos.

A dificuldade da pista é baseada na pontuação total. Retorne:

- "Fácil" se o total for 0 - 100
- "Médio" se o total for 101-200
- "Difícil" se o total for superior a 200

# --hints--

`get_difficulty("SLSLLSRRLSRLRL")` deve retornar `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLSLLSRRLSRLRL"), "Easy")`)
}})
```

`get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` deve retornar `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard")`)
}})
```

`get_difficulty("SRRRRLSLLRLRSSRLSRL")` deve retornar `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRRRRLSLLRLRSSRLSRL"), "Medium")`)
}})
```

`get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` deve retornar `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard")`)
}})
```

`get_difficulty("SLLSSLRLSLSLRSLSSLRL")` deve retornar `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium")`)
}})
```

`get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` deve retornar `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_difficulty(track):

    return track
```

# --solutions--

```py
def get_difficulty(track):
  score = 0

  for i in range(len(track)):
    current = track[i]
    previous = track[i - 1] if i > 0 else None

    if current == 'S':
      score += 0
    elif current == 'L' or current == 'R':
      if previous and previous != 'S' and previous != current:
        score += 15
      else:
        score += 5

  if score <= 100:
    return 'Easy'
  elif score <= 200:
    return 'Medium'
  else:
    return 'Hard'
```
