---
id: 698a1a73ade5ac0e19180fa0
title: "Desafio 197: Compatibilidade de Tipo Sanguíneo"
challengeType: 29
dashedName: challenge-197
---

# --description--

Dado um tipo sanguíneo do doador e um tipo sanguíneo do receptor, determine se o doador pode doar sangue para o receptor.

Cada tipo sanguíneo consiste em:

- Uma letra: `"A"`, `"B"`, `"AB"` ou `"O"`
- E um fator Rh: `"+"` ou `"-"`

Os tipos sanguíneos serão uma das letras válidas seguidas por um fator Rh. Por exemplo, `"AB+"` e `"O-"` são tipos sanguíneos válidos.

Regras de Letras:

- `"O"` pode doar para outro tipo de letra.
- `"A"` pode doar para `"A"` e `"AB"`.
- `"B"` pode doar para `"B"` e `"AB"`.
- `"AB"` pode doar apenas para `"AB"`.

Regras Rh:

- Negativo (`"-"`) pode doar para ambos `"-"` e `"+"`.
- Positivo (`"+"`) pode doar apenas para `"+"`.

Tanto a regra da letra quanto a regra Rh devem ser aprovadas para que um doador possa doar para o receptor.

# --hints--

`can_donate("B+", "B+")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B+", "B+"), True)`)
}})
```

`can_donate("O-", "AB-")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "AB-"), True)`)
}})
```

`can_donate("O+", "A-")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "A-"), False)`)
}})
```

`can_donate("A+", "AB+")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A+", "AB+"), True)`)
}})
```

`can_donate("A-", "B-")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A-", "B-"), False)`)
}})
```

`can_donate("B-", "AB+")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "AB+"), True)`)
}})
```

`can_donate("B-", "A+")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "A+"), False)`)
}})
```

`can_donate("O-", "O+")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "O+"), True)`)
}})
```

`can_donate("O+", "O-")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "O-"), False)`)
}})
```

`can_donate("AB+", "AB-")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("AB+", "AB-"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_donate(donor, recipient):

    return donor
```

# --solutions--

```py
def can_donate(donor, recipient):
    donor_type = donor[:-1]
    donor_rh = donor[-1]

    recipient_type = recipient[:-1]
    recipient_rh = recipient[-1]

    abo_compatibility = {
        "O": ["A", "B", "AB", "O"],
        "A": ["A", "AB"],
        "B": ["B", "AB"],
        "AB": ["AB"]
    }

    abo_match = recipient_type in abo_compatibility[donor_type]
    rh_match = (
        donor_rh == "-" or
        (donor_rh == "+" and recipient_rh == "+")
    )

    return abo_match and rh_match
```
