---
id: 698a1a73ade5ac0e19180fa9
title: "Desafio 206: Valores das Cartas de Baralho"
challengeType: 29
dashedName: challenge-206
---

# --description--

Dado um array de cartas de baralho, retorne um novo array com o valor numérico de cada carta.

Valores das Cartas:

- Um Ás (`"A"`) tem um valor de 1.
- Cartas numeradas (`"2"` - `"10"`) têm seu valor facial: 2 - 10, respectivamente.
- Cartas de figura: Valete (`"J"`), Dama (`"Q"`) e Rei (`"K"`) valem 10 cada.

Trajes:

- Cada carta tem um naipe: Espadas (`"S"`), Paus (`"C"`), Ouros (`"D"`), ou Copas (`"H"`).

Formato de Cartão:

- Cada carta é representada como uma string: `"valueSuit"`. Por exemplo: `"AS"` é o Ás de Espadas, `"10H"` é o Dez de Copas e `"QC"` é a Dama de Paus.

# --hints--

`card_values(["3H", "4D", "5S"])` deve retornar `[3, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` deve retornar `[1, 10, 10, 6, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` deve retornar `[8, 10, 2, 10, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` deve retornar `[1, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` deve retornar `[10, 10, 10, 10, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```
