---
id: 698a1a863194f1f4e63f645e
title: "Desafio 207: Menor Gap"
challengeType: 29
dashedName: challenge-207
---

# --description--

Dada uma string, retorne a substring entre os dois caracteres idênticos que têm o menor número de caracteres entre eles (menor gap).

- Sempre haverá pelo menos um par de caracteres correspondentes.
- A substring retornada deve excluir os caracteres correspondentes.
- Se dois ou mais gaps tiverem o mesmo comprimento, retorne os caracteres do primeiro.

Por exemplo, dado `"ABCDAC"`, retorne `"DA"`.

- Apenas `"A"` e `"C"` se repetem na `string`.
- O número de caracteres entre os dois `"A"` é 3 e entre os caracteres `"C"` é 2.
- Então retorne a string entre os dois caracteres `"C"`.

# --hints--

`smallest_gap("ABCDAC")` deve retornar `"DA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` deve retornar `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` deve retornar `"#q6e&rkf(p"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` deve retornar `""`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` deve retornar `"fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
