---
id: 698a1a863194f1f4e63f6460
title: "Desafio 209: Tamanho do elemento"
challengeType: 29
dashedName: challenge-209
---

# --description--

Dado um tamanho de janela, a largura de um elemento em unidades de largura do viewport `"vw"` e a altura de um elemento em unidades de altura do viewport `"vh"`, determine o tamanho do elemento em pixels.

- O tamanho da janela fornecido e o tamanho do elemento retornado são strings no formato `"width x height"`, `"1200 x 800"` por exemplo.

- As unidades `"vw"` são a porcentagem da largura da janela. `"50vw"`, por exemplo, é 50% da largura da janela.
- As unidades `"vh"` são a porcentagem da altura da janela. `"50vh"`, por exemplo, é 50% da altura da janela.

# --hints--

`get_element_size("1200 x 800", "50vw", "50vh")` deve retornar `"600 x 400"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "50vw", "50vh"), "600 x 400")`)
}})
```

`get_element_size("320 x 480", "25vw", "50vh")` deve retornar `"80 x 240"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("320 x 480", "25vw", "50vh"), "80 x 240")`)
}})
```

`get_element_size("1000 x 500", "7vw", "3vh")` deve retornar `"70 x 15"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1000 x 500", "7vw", "3vh"), "70 x 15")`)
}})
```

`get_element_size("1920 x 1080", "95vw", "100vh")` deve retornar `"1824 x 1080"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1920 x 1080", "95vw", "100vh"), "1824 x 1080")`)
}})
```

`get_element_size("1200 x 800", "0vw", "0vh")` deve retornar `"0 x 0"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "0vw", "0vh"), "0 x 0")`)
}})
```

`get_element_size("1440 x 900", "100vw", "114vh")` deve retornar `"1440 x 1026"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1440 x 900", "100vw", "114vh"), "1440 x 1026")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_element_size(window_size, element_vw, element_vh):

    return window_size
```

# --solutions--

```py
def get_element_size(window_size, element_vw, element_vh):
    window_width, window_height = map(int, window_size.split(" x "))

    vw = int(element_vw.replace("vw", ""))
    vh = int(element_vh.replace("vh", ""))

    width_px = round((vw / 100) * window_width)
    height_px = round((vh / 100) * window_height)

    return f"{width_px} x {height_px}"
```
