---
id: 698a1a73ade5ac0e19180fa1
title: "Desafio 198: Contagem de Dias Úteis"
challengeType: 29
dashedName: challenge-198
---

# --description--

Dada uma data de início e uma data de término, retorne o número de dias úteis entre as duas.

- As datas fornecidas estão no formato `"YYYY-MM-DD"`.
- Dias úteis são dias de trabalho (segunda a sexta-feira).
- Finais de semana não são dias úteis (sábado e domingo).
- Inclua tanto as datas de início e término ao contar.

# --hints--

`count_business_days("2026-02-24", "2026-02-26")` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2026-02-26"), 3)`)
}})
```

`count_business_days("2026-02-24", "2026-02-28")` deve retornar `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2026-02-28"), 4)`)
}})
```

`count_business_days("2026-02-21", "2026-03-01")` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-21", "2026-03-01"), 5)`)
}})
```

`count_business_days("2026-03-08", "2026-03-17")` deve retornar `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-03-08", "2026-03-17"), 7)`)
}})
```

`count_business_days("2026-02-24", "2027-02-24")` deve retornar `262`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2027-02-24"), 262)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_business_days(start, end):

    return start
```

# --solutions--

```py
from datetime import datetime, timedelta

def count_business_days(start, end):
    start_date = datetime.strptime(start, "%Y-%m-%d")
    end_date = datetime.strptime(end, "%Y-%m-%d")
    
    count = 0
    current = start_date
    
    while current <= end_date:
        if current.weekday() < 5:
            count += 1
        current += timedelta(days=1)
    
    return count
```
