---
id: 6994cff2290543b3aec9f510
title: "Desafio 211: Soma de Array"
challengeType: 29
dashedName: challenge-211
---

# --description--

Dado um array de números, retorne a soma de todos os números.

# --hints--

`sum_array([1, 2, 3, 4, 5])` deve retornar `15`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([1, 2, 3, 4, 5]), 15)`)
}})
```

`sum_array([42])` deve retornar `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([42]), 42)`)
}})
```

`sum_array([5, -2, 7, -3])` deve retornar `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([5, -2, 7, -3]), 7)`)
}})
```

`sum_array([203, 145, -129, 6293, 523, -919, 845, 2434])` deve retornar `9395`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395)`)
}})
```

`sum_array([0, 0])` deve retornar `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([0, 0]), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_array(numbers):

    return numbers
```

# --solutions--

```py
def sum_array(numbers):

    return sum(numbers)
```
