---
id: 699c8e045ee7cb94ed2322dd
title: "Desafio 222: Sombras do Equinócio"
challengeType: 29
dashedName: challenge-222
---

# --description--

Hoje é o equinócio, quando o sol está diretamente acima do equador e perfeitamente no zênite ao meio-dia. Dado um horário, determine a sombra projetada por um poste vertical de 4 pés.

- O horário será uma string no formato 24 horas `"HH:MM"` (por exemplo, `"15:00"` é 15h).
- Você receberá apenas um horário em incrementos de 30 minutos.

Regras:

- O sol nasce às 6h diretamente `"east"`, e se põe às 18h diretamente `"west"`.
- Uma sombra sempre aponta na direção oposta ao sol.
- O comprimento da sombra (em pés) é o número de horas distante do meio-dia, elevado ao cubo.
- Não há sombra antes do nascer do sol (antes das 6h), depois do pôr do sol (6h ou mais tarde) ou ao meio-dia.

Retornar:

- Se uma sombra existir, retorne `"(length)ft (direction)"`. Por exemplo, `"8ft west"`.
- Caso contrário, retorne `"No shadow"`.

Por exemplo, dado `"10:00"`, retorne `"8ft west"` porque 10h é 2 horas antes do meio-dia, então 2<sup>3</sup> = 8 pés, e a sombra aponta para o oeste porque o sol está no leste às 10h.

# --hints--

`get_shadow("10:00")` deve retornar `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` deve retornar `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` deve retornar `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` deve retornar `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` deve retornar `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` deve retornar `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` deve retornar `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` deve retornar `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` deve retornar `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
