---
id: 699c8e045ee7cb94ed2322d6
title: "Desafio 215: Calculadora de Taxa de Estacionamento"
challengeType: 29
dashedName: challenge-215
---

# --description--

Dadas duas `strings` representando o horário em que você estacionou seu carro e o horário em que o retirou, calcule a taxa de estacionamento.

- As strings fornecidas estarão no formato `"HH:MM"` usando um relógio de 24 horas. Então `"14:00"` é 14h por exemplo.
- A primeira `string` será a hora que você estacionou seu carro e a segunda será a hora que você o pegou.
- Se o horário de retirada for anterior ao horário de entrada, isso significa que o estacionamento ultrapassou a meia-noite e entrou no dia seguinte.

Regras de taxa:

- Cada hora estacionada custa $3.
- Horas parciais são arredondadas para a próxima hora cheia.
- Se o estacionamento se estender durante a noite (após a meia-noite), uma taxa adicional de $10 por pernoite será aplicada.
- Há uma taxa mínima de $5 (usada apenas se o total for menor que $5).

Retorne o custo total no formato `"$cost"`, `"$5"` por exemplo.

# --hints--

`calculate_parking_fee("09:00", "11:00")` deve retornar `"$6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("09:00", "11:00"), "$6")`)
}})
```

`calculate_parking_fee("10:00", "10:30")` deve retornar `"$5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("10:00", "10:30"), "$5")`)
}})
```

`calculate_parking_fee("08:10", "10:45")` deve retornar `"$9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("08:10", "10:45"), "$9")`)
}})
```

`calculate_parking_fee("14:40", "23:10")` deve retornar `"$27"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("14:40", "23:10"), "$27")`)
}})
```

`calculate_parking_fee("18:15", "01:30")` deve retornar `"$34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("18:15", "01:30"), "$34")`)
}})
```

`calculate_parking_fee("11:11", "11:10")` deve retornar `"$82"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("11:11", "11:10"), "$82")`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_parking_fee(park_time, pickup_time):

    return park_time
```

# --solutions--

```py
import math

def calculate_parking_fee(park_time, pickup_time):
    def to_minutes(time_str):
        hours, minutes = map(int, time_str.split(":"))
        return hours * 60 + minutes

    entry_minutes = to_minutes(park_time)
    exit_minutes = to_minutes(pickup_time)

    overnight = False

    if exit_minutes < entry_minutes:
        overnight = True
        total_minutes = (24 * 60 - entry_minutes) + exit_minutes
    else:
        total_minutes = exit_minutes - entry_minutes

    hours = math.ceil(total_minutes / 60)

    cost = hours * 3

    if overnight:
        cost += 10

    if cost < 5:
        cost = 5

    return f"${cost}"
```
