---
id: 699c8e045ee7cb94ed2322d9
title: "Desafio 218: Divisível Uniformemente"
challengeType: 29
dashedName: challenge-218
---

# --description--

Dado dois inteiros, determine se você pode dividir o primeiro pelo segundo de forma exata.

# --hints--

`is_evenly_divisible(4, 2)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4, 2), True)`)
}})
```

`is_evenly_divisible(7, 3)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(7, 3), False)`)
}})
```

`is_evenly_divisible(5, 10)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(5, 10), False)`)
}})
```

`is_evenly_divisible(48, 6)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(48, 6), True)`)
}})
```

`is_evenly_divisible(3186, 9)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(3186, 9), True)`)
}})
```

`is_evenly_divisible(4192, 11)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4192, 11), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_evenly_divisible(a, b):

    return a
```

# --solutions--

```py
def is_evenly_divisible(a, b):
    return a % b == 0
```
