---
id: 699c8e045ee7cb94ed2322d7
title: "Desafio 216: Dia do Pi"
challengeType: 29
dashedName: challenge-216
---

# --description--

Feliz dia do pi (π)!

Dado um inteiro (`n`), onde `n` está entre 1 e 1000 (inclusive), retorne o `n`-ésimo decimal de π.

- Certifique-se de retornar um número e não uma string.

π com seus primeiros cinco decimais é 3.14159. Então, dado `5` por exemplo, retorne `9`, o quinto decimal.

Você pode ter que buscar os primeiros 1000 decimals de π em algum lugar.

# --hints--

`get_pi_decimal(5)` deve retornar `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pi_decimal(5), 9)`)
}})
```

`get_pi_decimal(10)` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pi_decimal(10), 5)`)
}})
```

`get_pi_decimal(22)` deve retornar `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pi_decimal(22), 6)`)
}})
```

`get_pi_decimal(39)` deve retornar `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pi_decimal(39), 7)`)
}})
```

`get_pi_decimal(76)` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pi_decimal(76), 2)`)
}})
```

`get_pi_decimal(384)` deve retornar `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pi_decimal(384), 4)`)
}})
```

`get_pi_decimal(601)` deve retornar `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pi_decimal(601), 0)`)
}})
```

`get_pi_decimal(1000)` deve retornar `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pi_decimal(1000), 9)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_pi_decimal(n):

    return n
```

# --solutions--

```py
def get_pi_decimal(n):
    pi = "31415926535897932384626433832795028841971693993751058209749445923078164062862089986280348253421170679821480865132823066470938446095505822317253594081284811174502841027019385211055596446229489549303819644288109756659334461284756482337867831652712019091456485669234603486104543266482133936072602491412737245870066063155881748815209209628292540917153643678925903600113305305488204665213841469519415116094330572703657595919530921861173819326117931051185480744623799627495673518857527248912279381830119491298336733624406566430860213949463952247371907021798609437027705392171762931767523846748184676694051320005681271452635608277857713427577896091736371787214684409012249534301465495853710507922796892589235420199561121290219608640344181598136297747713099605187072113499999983729780499510597317328160963185950244594553469083026425223082533446850352619311881710100031378387528865875332083814206171776691473035982534904287554687311595628638823537875937519577818577805321712268066130019278766111959092164201989"

    return int(pi[n])
```
