---
id: 699c8e045ee7cb94ed2322d5
title: "Desafio 214: Validador de Corrente de Dominó"
challengeType: 29
dashedName: challenge-214
---

# --description--

Dado um array 2D que representa uma sequência de dominós, determine se ele forma uma cadeia válida.

- Cada elemento no array representa um dominó e será um array de dois números de 1 a 6 (inclusive).
- Para a cadeia ser válida, o segundo número de cada dominó deve corresponder ao primeiro número do próximo dominó.
- O primeiro número do primeiro dominó e o último número do último dominó não precisam corresponder a nada.

# --hints--

`is_valid_domino_chain([[1, 3], [3, 6], [6, 5]])` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[1, 3], [3, 6], [6, 5]]), True)`)
}})
```

`is_valid_domino_chain([[6, 2], [3, 4], [4, 1]])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[6, 2], [3, 4], [4, 1]]), False)`)
}})
```

`is_valid_domino_chain([[2, 5], [5, 6], [5, 1]])` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 5], [5, 6], [5, 1]]), False)`)
}})
```

`is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]), True)`)
}})
```

`is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_domino_chain(dominoes):

    return dominoes
```

# --solutions--

```py
def is_valid_domino_chain(dominoes):
    for i in range(len(dominoes) - 1):
        if dominoes[i][1] != dominoes[i + 1][0]:
            return False

    return True
```
