---
id: 69a890af247de743333bd4ce
title: "Desafio 225: Sem Repetições Consecutivas"
challengeType: 29
dashedName: challenge-225
---

# --description--

Dada uma `string`, determine se ela não possui caracteres repetidos.

- Uma string não tem repetições se não tiver o mesmo caractere duas ou mais vezes seguidas.

# --hints--

`has_no_repeats("hi world")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hi world"), True)`)
}})
```

`has_no_repeats("hello world")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hello world"), False)`)
}})
```

`has_no_repeats("abcdefghijklmnopqrstuvwxyz")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("abcdefghijklmnopqrstuvwxyz"), True)`)
}})
```

`has_no_repeats("freeCodeCamp")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("freeCodeCamp"), False)`)
}})
```

`has_no_repeats("The quick brown fox jumped over the lazy dog.")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("The quick brown fox jumped over the lazy dog."), True)`)
}})
```

`has_no_repeats("Mississippi")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("Mississippi"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_no_repeats(s):

    return s
```

# --solutions--

```py
def has_no_repeats(s):
    for i in range(1, len(s)):
        if s[i] == s[i-1]:
            return False
    return True
```
