---
id: 69a890af247de743333bd4d0
title: "Desafio 227: Tempo de Resfriamento"
challengeType: 29
dashedName: challenge-227
---

# --description--

Dadas duas marcas de tempo, a primeira representando quando um usuário terminou um exame e a segunda representando o tempo atual, determine se o usuário pode fazer um exame novamente.

- Ambos os timestamps terão o formato: `"YYYY-MM-DDTHH:MM:SS"`, por exemplo `"2026-03-25T14:00:00"`. Note que o horário está no formato de 24 horas.
- Um usuário deve esperar pelo menos 48 horas antes de refazer um exame.

# --hints--

`can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00"), True)`)
}})
```

`can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00"), False)`)
}})
```

`can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00"), True)`)
}})
```

`can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_retake(finish_time, current_time):

    return finish_time
```

# --solutions--

```py
from datetime import datetime, timedelta

def can_retake(finish_time, current_time):
    fmt = "%Y-%m-%dT%H:%M:%S"
    finished_dt = datetime.strptime(finish_time, fmt)
    current_dt = datetime.strptime(current_time, fmt)
    diff = current_dt - finished_dt
    cooldown = timedelta(hours=48)

    return diff >= cooldown
```
