---
id: 69a890af247de743333bd4cd
title: "Desafio 224: Detector de Torra de Café"
challengeType: 29
dashedName: challenge-224
---

# --description--

Dada uma `string` representando os grãos usados para fazer uma xícara de café, determine a torra da xícara.

- A string fornecida conterá os seguintes caracteres, cada um representando um tipo de feijão:

  - Um apóstrofo (`'`) é um grão de torra clara que vale 1 ponto cada.
  - Um traço (`-`) é um grão de torra médio que vale 2 pontos cada.
  - Um ponto (`.`) é um grão de torra escura que vale 3 pontos cada.

- O nível de torra é determinado pela média de todos os grãos.

Retornar:

- `"Light"` se a média for menor que 1.75.
- `"Medium"` se a média estiver entre 1.75 e 2.5.
- `"Dark"` se a média for maior que 2.5.

# --hints--

`detect_roast("''-''''''-'-''--''''")` deve retornar `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("''-''''''-'-''--''''"), "Light")`)
}})
```

`detect_roast(".'-''-''..'''.-.-''-")` deve retornar `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".'-''-''..'''.-.-''-"), "Medium")`)
}})
```

`detect_roast("--.''--'-''.--..-.--")` deve retornar `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("--.''--'-''.--..-.--"), "Medium")`)
}})
```

`detect_roast("-...'-......-..-...-")` deve retornar `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-...'-......-..-...-"), "Dark")`)
}})
```

`detect_roast(".--.-..-......----.'")` deve retornar `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".--.-..-......----.'"), "Medium")`)
}})
```

`detect_roast("..-..-..-..-....-.-.")` deve retornar `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("..-..-..-..-....-.-."), "Dark")`)
}})
```

`detect_roast("-'-''''''..-'.''-'.'")` deve retornar `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-'-''''''..-'.''-'.'"), "Light")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_roast(beans):

    return beans
```

# --solutions--

```py
def detect_roast(beans):
    total = 0

    for bean in beans:
        if bean == "'":
            total += 1
        elif bean == "-":
            total += 2
        elif bean == ".":
            total += 3

    avg = total / len(beans)

    if avg < 1.75:
        return "Light"
    elif avg <= 2.5:
        return "Medium"
    else:
        return "Dark"
```
