---
id: 69b58ce40693f140c84c8559
title: "Desafio 240: caracteres palíndromos"
challengeType: 29
dashedName: challenge-240
---

# --description--

Dada uma string, determine se ela é um palíndromo e retorne o caractere do meio (se o comprimento for ímpar) ou os dois caracteres do meio (se for par).

- Um palíndromo é uma string que é igual de frente para trás e de trás para frente.
- Se não for um palíndromo, retorne `"none"`.

# --hints--

`palindrome_locator("racecar")` deve retornar `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("racecar"), "e")`)
}})
```

`palindrome_locator("level")` deve retornar `"v"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("level"), "v")`)
}})
```

`palindrome_locator("freecodecamp")` deve retornar `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("freecodecamp"), "none")`)
}})
```

`palindrome_locator("noon")` deve retornar `"oo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("noon"), "oo")`)
}})
```

`palindrome_locator("11100111")` deve retornar `"00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("11100111"), "00")`)
}})
```

# --seed--

## --seed-contents--

```py
def palindrome_locator(s):

    return s
```

# --solutions--

```py
def palindrome_locator(s):
    if s != s[::-1]:
        return "none"
    mid = len(s) // 2
    return s[mid] if len(s) % 2 == 1 else s[mid - 1] + s[mid]
```
