---
id: 69b559d2903b9e4afe9075f8
title: "Challenge 238: Digit Rotation Escape"
challengeType: 29
dashedName: challenge-238
---

# --description--

Dado um número inteiro positivo, determine se ele, ou qualquer uma de suas rotações, é divisível de forma exata pela quantidade de seus dígitos.

Uma rotação significa mover o primeiro dígito para o final. Por exemplo, após 1 rotação, 123 se torna 231.

- Verifique a rotação `0` (o número dado) primeiro.
- Os números fornecidos não conterão zeros.
- Retorne o primeiro número de rotação encontrado, ou `"none"` se não houver nenhum.

# --hints--

`get_rotation(123)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(123), 0)`)
}})
```

`get_rotation(13579)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(13579), 3)`)
}})
```

`get_rotation(24681)` should return `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(24681), "none")`)
}})
```

`get_rotation(84138789345)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(84138789345), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rotation(n):

    return n
```

# --solutions--

```py
def get_rotation(n):
    s = str(n)
    digit_count = len(s)
    current = s

    for i in range(digit_count):
        if int(current) % digit_count == 0:
            return i
        current = current[1:] + current[0]

    return "none"
```
