---
id: 69b5b2be76ec8135a7fbe973
title: "Desafio 244: ataque da torre e do bispo"
challengeType: 29
dashedName: challenge-244
---

# --description--

Dadas duas strings, uma para a localização de uma torre em um tabuleiro de xadrez e outra para a localização de um bispo, determine se uma peça pode atacar a outra.

Um tabuleiro de xadrez padrão tem 8x8 casas, com colunas rotuladas de `A` até `H` (da esquerda para a direita) e linhas rotuladas de `1` até `8` (de baixo para cima). Ele se parece com isto:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- Torres podem se mover quantas casas quiserem na direção horizontal ou vertical.
- Bispos podem se mover quantas casas quiserem em qualquer direção diagonal.
- Uma peça pode atacar outra se puder se mover para a localização dessa peça.

Retorne:

- `"rook"` se a torre puder atacar o bispo.
- `"bishop"` se o bispo puder atacar a torre.
- `"neither"` se nenhuma das peças puder atacar a outra.

# --hints--

`rook_bishop_attack("A1", "A5")` deve retornar `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("A1", "A5"), "rook")`)
}})
```

`rook_bishop_attack("C3", "F6")` deve retornar `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("C3", "F6"), "bishop")`)
}})
```

`rook_bishop_attack("D4", "D7")` deve retornar `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("D4", "D7"), "rook")`)
}})
```

`rook_bishop_attack("B7", "H1")` deve retornar `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B7", "H1"), "bishop")`)
}})
```

`rook_bishop_attack("B3", "C5")` deve retornar `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B3", "C5"), "neither")`)
}})
```

`rook_bishop_attack("G3", "E8")` deve retornar `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("G3", "E8"), "neither")`)
}})
```

# --seed--

## --seed-contents--

```py
def rook_bishop_attack(rook, bishop):

    return rook
```

# --solutions--

```py
def rook_bishop_attack(rook, bishop):
    col_diff = abs(ord(rook[0]) - ord(bishop[0]))
    row_diff = abs(int(rook[1]) - int(bishop[1]))

    if rook[0] == bishop[0] or rook[1] == bishop[1]:
        return "rook"
    if col_diff == row_diff:
        return "bishop"
    return "neither"
```
