---
id: 69bc6cb30c1d112a2e110a07
title: "Desafio 250: chave oculta"
challengeType: 29
dashedName: challenge-250
---

# --description--

Bem-vindo ao desafio diário número 250!

Dada uma string codificada, decodifique-a usando uma chave de criptografia e retorne-a.

Para encontrar a chave:

- Observe todos os desafios diários até hoje cujo número do desafio seja múltiplo de 25 (incluindo este).
- Pegue a primeira letra de cada um desses títulos de desafio e combine-as em uma string. Se o título começar com um caractere que não seja letra, encontre a primeira letra dele.

Para decodificar a mensagem, percorra cada letra da mensagem codificada e:

  1. Observe a letra correspondente na chave (repita a chave se a mensagem for maior que a chave).
  2. Converta a letra da chave para seu número correspondente: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. Desloque a letra codificada para trás no alfabeto por esse número.
  4. Se o deslocamento passar antes de `"A"`, volte para `"Z"`.

Por exemplo, se a mensagem codificada começar com `"Y"` e a primeira letra da chave for `"V"` (22), desloque `"Y"` 22 posições para trás para obter `"C"`. Repita esse processo para cada letra para decodificar a mensagem completa.

- Apenas letras são deslocadas, espaços são retornados como estão.
- Todas as letras fornecidas e retornadas estão em maiúsculas.

# --hints--

`decode("YAVJYNXE")` deve retornar `"CONGRATS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")` deve retornar `"CODING LEGEND"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")` deve retornar `"YOU ARE AWESOME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")` deve retornar `"KEEP GOING"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` deve retornar `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
