---
id: 69c5f3d787b1725d5f00c8ba
title: "Desafio 256: direção do tempo mais próxima"
challengeType: 29
dashedName: challenge-256
---

# --description--

Dadas duas horas, determine se é mais rápido ir da primeira para a segunda avançando ou retrocedendo.

- As horas são dadas no formato 24 horas (`"HH:MM"`)
- O relógio dá a volta (23:59 vai para 00:00 ao avançar, e 00:00 vai para 23:59 ao retroceder)

Retorne:

- `"forward"` se avançar for mais rápido
- `"backward"` se retroceder for mais rápido
- `"equal"` se as duas direções levarem o mesmo tempo

# --hints--

`get_direction("10:00", "12:00")` deve retornar `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("10:00", "12:00"), "forward")`)
}})
```

`get_direction("11:00", "05:00")` deve retornar `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("11:00", "05:00"), "backward")`)
}})
```

`get_direction("00:00", "12:00")` deve retornar `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("00:00", "12:00"), "equal")`)
}})
```

`get_direction("15:45", "01:10")` deve retornar `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("15:45", "01:10"), "forward")`)
}})
```

`get_direction("03:30", "19:50")` deve retornar `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("03:30", "19:50"), "backward")`)
}})
```

`get_direction("06:30", "18:30")` deve retornar `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("06:30", "18:30"), "equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_direction(time1, time2):

    return time1
```

# --solutions--

```py
def get_direction(time1, time2):
    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    start = to_minutes(time1)
    end = to_minutes(time2)
    forward = (end - start + 1440) % 1440
    backward = (start - end + 1440) % 1440

    if forward < backward:
        return 'forward'
    if backward < forward:
        return 'backward'
    return 'equal'
```
