---
id: 69c5f3d787b1725d5f00c8bd
title: "Desafio 259: explosão fizzbuzz"
challengeType: 29
dashedName: challenge-259
---

# --description--

Dado um inteiro, retorne o número de passos necessários para transformar a palavra `"fizzbuzz"` em uma string com pelo menos o número dado de `"z"` usando as seguintes regras:

- Comece com a string `"fizzbuzz"`.
- A cada passo, aplique as regras padrão do FizzBuzz usando a posição da letra na string (a primeira `"f"` está na posição 1).
  - Se a posição da letra for divisível por 3, substitua a letra por `"fizz"`
  - Se for divisível por 5, substitua a letra por `"buzz"`
  - Se for divisível por 3 e 5, substitua a letra por `"fizzbuzz"`

Então, após 1 passo, `"fizzbuzz"` se transforma em `"fifizzzbuzzfizzzz"`, que tem 9 `"z"`'s.

# --hints--

`explode_fizzbuzz(9)` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` deve retornar `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` deve retornar `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` deve retornar `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` deve retornar `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
