---
id: 69c6ff713a52713463aa7926
title: "Desafio 260: pontuação de palavra"
challengeType: 29
dashedName: challenge-260
---

# --description--

Dada uma palavra, retorne sua pontuação usando uma tabela padrão de valor por letra:

| Letra | Valor |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- Letras maiúsculas e minúsculas têm o mesmo valor.

# --hints--

`get_word_score("hi")` deve retornar `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("hi"), 17)`)
}})
```

`get_word_score("hello")` deve retornar `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("hello"), 52)`)
}})
```

`get_word_score("hippopotamus")` deve retornar `169`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("hippopotamus"), 169)`)
}})
```

`get_word_score("freeCodeCamp")` deve retornar `94`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("freeCodeCamp"), 94)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_word_score(word):

    return word
```

# --solutions--

```py
def get_word_score(word):
    return sum(ord(c) - 96 for c in word.lower())
```
