---
id: 69c6ff713a52713463aa7929
title: "Desafio 263: palavras cruzadas binárias"
challengeType: 29
dashedName: challenge-263
---

# --description--

Dado um caractere, determine se sua representação binária de 8 bits pode ser encontrada na grade a seguir, horizontalmente ou verticalmente em qualquer direção:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

Por exemplo, `"A"` tem a representação binária `01000001`, que aparece na primeira linha da esquerda para a direita.

# --hints--

`is_in_crossword("I")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("I"), True)`)
}})
```

`is_in_crossword("D")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("D"), True)`)
}})
```

`is_in_crossword("0")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("0"), True)`)
}})
```

`is_in_crossword("u")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("u"), True)`)
}})
```

`is_in_crossword("Y")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("Y"), False)`)
}})
```

`is_in_crossword("p")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("p"), False)`)
}})
```

`is_in_crossword("1")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("1"), False)`)
}})
```

`is_in_crossword("Q")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("Q"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_in_crossword(char):

    return char
```

# --solutions--

```py
def is_in_crossword(char):
    grid = [
        "01000001",
        "01101111",
        "01000100",
        "01100101",
        "01010010",
        "01010100",
        "01101000",
        "10101110",
    ]
    target = format(ord(char), '08b')
    cols = [''.join(grid[r][c] for r in range(8)) for c in range(8)]
    return any(target in seq or target in seq[::-1] for seq in grid + cols)
```
