---
id: 69c6ff713a52713463aa7927
title: "Desafio 261: palavras numéricas"
challengeType: 29
dashedName: challenge-261
---

# --description--

Dado um inteiro de 0 a 99, retorne sua representação em palavras em inglês.

- `0` retorna `"zero"`.
- Os números de 1 a 19 têm nomes únicos (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Múltiplos de 10 de 20 a 90 têm seus próprios nomes (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Números de 21 a 99 que não são múltiplos de 10 são escritos como duas palavras unidas por um hífen. Por exemplo `"forty-two"` e `"fifty-three"`.

# --hints--

`get_number_words(0)` deve retornar `"zero"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(0), "zero")`)
}})
```

`get_number_words(10)` deve retornar `"ten"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(10), "ten")`)
}})
```

`get_number_words(19)` deve retornar `"nineteen"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(19), "nineteen")`)
}})
```

`get_number_words(30)` deve retornar `"thirty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(30), "thirty")`)
}})
```

`get_number_words(53)` deve retornar `"fifty-three"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(53), "fifty-three")`)
}})
```

`get_number_words(7)` deve retornar `"seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(7), "seven")`)
}})
```

`get_number_words(12)` deve retornar `"twelve"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(12), "twelve")`)
}})
```

`get_number_words(60)` deve retornar `"sixty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(60), "sixty")`)
}})
```

`get_number_words(67)` deve retornar `"sixty-seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(67), "sixty-seven")`)
}})
```

`get_number_words(98)` deve retornar `"ninety-eight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(98), "ninety-eight")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_words(n):

    return n
```

# --solutions--

```py
def get_number_words(n):
    ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
            "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
            "seventeen","eighteen","nineteen"]
    tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"]

    if n < 20:
        return ones[n]
    if n % 10 == 0:
        return tens[n // 10]
    return tens[n // 10] + "-" + ones[n % 10]
```
