---
id: 69cfca90e8a0a6d4d6871c4e
title: "Desafio 264: grupos de anagramas"
challengeType: 29
dashedName: challenge-264
---

# --description--

Dado um array de palavras, retorne um array 2d com as palavras agrupadas em anagramas.

- Palavras são anagramas se contêm as mesmas letras em qualquer ordem.
- Cada palavra pertence a exatamente um grupo.
- A ordem de retorno não importa.

Por exemplo, dado `["listen", "silent", "hello", "enlist", "world"]`, retorne `[["listen", "silent", "enlist"], ["hello"], ["world"]]`.

# --hints--

`group_anagrams(["listen", "silent", "hello", "enlist", "world"])` deve retornar `[["listen", "silent", "enlist"], ["hello"], ["world"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["listen", "silent", "hello", "enlist", "world"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["enlist", "listen", "silent"], ["hello"], ["world"]])`)
}})
```

`group_anagrams(["eat", "tea", "tan", "ate", "nat", "bat"])` deve retornar `[["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["eat", "tea", "tan", "ate", "nat", "bat"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["ate", "eat", "tea"], ["bat"], ["nat", "tan"]])`)
}})
```

`group_anagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])` deve retornar `[["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]])`)
}})
```

`group_anagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])` deve retornar `[["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def group_anagrams(words):

    return words
```

# --solutions--

```py
def group_anagrams(words):
    groups = {}
    for word in words:
        key = ''.join(sorted(word))
        if key not in groups:
            groups[key] = []
        groups[key].append(word)
    return list(groups.values())
```
