---
id: 69cfca90e8a0a6d4d6871c52
title: "Desafio 268: número narcísico"
challengeType: 29
dashedName: challenge-268
---

# --description--

Dado um inteiro positivo, determine se ele é um número narcísico.

- Um número é narcísico se a soma de cada um de seus dígitos elevados à potência do total de dígitos for igual ao próprio número.

Por exemplo, 153 tem 3 dígitos, e 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153, então ele é narcísico.

# --hints--

`is_narcissistic(153)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(153), True)`)
}})
```

`is_narcissistic(154)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(154), False)`)
}})
```

`is_narcissistic(371)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(371), True)`)
}})
```

`is_narcissistic(512)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(512), False)`)
}})
```

`is_narcissistic(9)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9), True)`)
}})
```

`is_narcissistic(11)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(11), False)`)
}})
```

`is_narcissistic(9474)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9474), True)`)
}})
```

`is_narcissistic(6549)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(6549), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_narcissistic(n):

    return n
```

# --solutions--

```py
def is_narcissistic(n):
    digits = str(n)
    power = len(digits)
    return sum(int(d) ** power for d in digits) == n
```
