---
id: 69e2383af7832c8032603b94
title: "Desafio 278: analisador de pedido de café"
challengeType: 29
dashedName: challenge-278
---

# --description--

Dada uma string com um pedido de café, identifique os itens do menu e retorne um pedido formatado.

Use os seguintes itens do menu e preços:

| Item | Preço |
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

Retorne uma string com os itens correspondentes unidos por `" + "`, seguidos de dois pontos e espaço (`": "`), e o preço total.

Por exemplo, dado `"I'd like an oat latte with vanilla syrup and an extra shot please."`, retorne `"oat latte + vanilla syrup + extra shot: $6.25"`

Os itens devem aparecer na ordem em que aparecem no menu e o preço total deve sempre ter duas casas decimais.

# --hints--

`format_coffee_order("I'd like an oat latte with vanilla syrup and an extra shot please.")` deve retornar `"oat latte + vanilla syrup + extra shot: $6.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25")`)
}})
```

`format_coffee_order("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` deve retornar `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85")`)
}})
```

`format_coffee_order("Can I get a cold brew with some cream and an extra shot.")` deve retornar `"cold brew + extra shot + cream: $5.75"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75")`)
}})
```

`format_coffee_order("Just an espresso please.")` deve retornar `"espresso: $3.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Just an espresso please."), "espresso: $3.00")`)
}})
```

`format_coffee_order("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` deve retornar `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_coffee_order(order):

    return order
```

# --solutions--

```py
def format_coffee_order(order):
    menu = [
        ("cold brew", 450),
        ("oat latte", 500),
        ("cappuccino", 475),
        ("espresso", 300),
        ("vanilla syrup", 75),
        ("caramel drizzle", 60),
        ("extra shot", 50),
        ("oat milk", 75),
        ("cream", 75),
    ]
    found = [(name, price) for name, price in menu if name in order.lower()]
    total = sum(price for _, price in found)
    items = " + ".join(name for name, _ in found)
    return f"{items}: ${total / 100:.2f}"
```
