---
id: 69f8c998d78ad3171a0713bf
title: "Desafio 292: proporção de aspecto mais ampla"
challengeType: 29
dashedName: challenge-292
---

# --description--

Dadas duas strings para diferentes dimensões de imagem, retorne a proporção de aspecto da imagem com a maior razão largura-altura.

- As strings fornecidas estarão no formato `"WxH"`, por exemplo, `"1920x1080"`.
- A proporção de aspecto é a razão entre largura e altura, reduzida aos menores números inteiros. Por exemplo, `"1920x1080"` reduz para `"16:9"`.
- Retorne uma string no formato `"W:H"`, por exemplo, `"16:9"`.

# --hints--

`get_wider_aspect_ratio("1920x1080", "800x600")` deve retornar `"16:9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1920x1080", "800x600"), "16:9")`)
}})
```

`get_wider_aspect_ratio("1080x1350", "2048x1536")` deve retornar `"4:3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1080x1350", "2048x1536"), "4:3")`)
}})
```

`get_wider_aspect_ratio("640x480", "2440x1220")` deve retornar `"2:1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("640x480", "2440x1220"), "2:1")`)
}})
```

`get_wider_aspect_ratio("360x640", "1080x1920")` deve retornar `"9:16"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("360x640", "1080x1920"), "9:16")`)
}})
```

`get_wider_aspect_ratio("3440x1440", "2048x858")` deve retornar `"43:18"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("3440x1440", "2048x858"), "43:18")`)
}})
```

`get_wider_aspect_ratio("12345x61234", "12534x51234")` deve retornar `"2089:8539"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("12345x61234", "12534x51234"), "2089:8539")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_wider_aspect_ratio(a, b):

    return a
```

# --solutions--

```py
from math import gcd

def get_wider_aspect_ratio(a, b):
    def parse(s):
        w, h = map(int, s.split("x"))
        g = gcd(w, h)
        return w / h, f"{w // g}:{h // g}"

    ratio_a, simplified_a = parse(a)
    ratio_b, simplified_b = parse(b)

    return simplified_a if ratio_a >= ratio_b else simplified_b
```
