---
id: 69f8c998d78ad3171a0713bd
title: "Desafio 290: festa de pizza"
challengeType: 29
dashedName: challenge-290
---

# --description--

Dado um array com as horas trabalhadas hoje por pessoa, retorne o número de pizzas para pedir em uma festa de pizza.

- Divida as horas trabalhadas de cada pessoa por 3 para obter a quantidade de fatias.
- Você não pode comer uma fatia parcial, então arredonde a quantidade de fatias de cada pessoa para cima até o número inteiro mais próximo.
- Cada pessoa recebe no mínimo duas fatias.
- Cada pizza tem 8 fatias. Arredonde o número total de pizzas para cima até a pizza inteira mais próxima.

# --hints--

`get_pizzas_to_order([8, 8, 8])` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8]), 2)`)
}})
```

`get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10])` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10]), 3)`)
}})
```

`get_pizzas_to_order([1, 2, 3, 4, 5])` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([1, 2, 3, 4, 5]), 2)`)
}})
```

`get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8])` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8]), 3)`)
}})
```

`get_pizzas_to_order([9, 9, 6])` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([9, 9, 6]), 1)`)
}})
```

`get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0])` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_pizzas_to_order(hours_worked):

    return hours_worked
```

# --solutions--

```py
import math

def get_pizzas_to_order(hours_worked):
    total_slices = sum(max(math.ceil(h / 3), 2) for h in hours_worked)
    return math.ceil(total_slices / 8)
```
