---
id: 69f8c998d78ad3171a0713b9
title: "Desafio 286: issues abertas"
challengeType: 29
dashedName: challenge-286
---

# --description--

Dado um array de números de issues e outro array de números de pull requests (PR), retorne um array das issues que permanecem abertas depois que todos os PRs foram mesclados.

- Um PR fecha uma issue se os dígitos deles forem rotações um do outro. Por exemplo, a issue 123 seria fechada pelo PR 231 ou 312.
- Um PR não fecha uma issue com o mesmo número exato. Por exemplo, o PR 123 não fecha a issue 123. Então uma issue com todos os dígitos iguais não pode ser fechada.
- Qualquer número pode ter zeros à esquerda removidos. Por exemplo, o PR 201 fecharia a issue 12 (012, uma rotação de 201). Da mesma forma, a issue 201 seria fechada pelo PR 12.

Retorne as issues abertas restantes na ordem em que foram dadas.

# --hints--

`get_open_issues([123, 234], [231])` deve retornar `[234]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` deve retornar `[345, 16]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` deve retornar `[456, 332, 15]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` deve retornar `[115, 296, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` deve retornar `[95, 395, 754, 296, 709, 237, 1802]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```
