---
id: 69f8c998d78ad3171a0713bc
title: "Desafio 289: soma das diferenças"
challengeType: 29
dashedName: challenge-289
---

# --description--

Dado um array de números, retorne a soma das diferenças entre cada número e o que vem depois dele.

Por exemplo, dado `[1, 3, 4]`, retorne `3` (2 + 1).

# --hints--

`sum_of_differences([1, 3, 4])` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([1, 3, 4]), 3)`)
}})
```

`sum_of_differences([5, -3, 3, 9, 10])` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([5, -3, 3, 9, 10]), 5)`)
}})
```

`sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7])` deve retornar `-16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16)`)
}})
```

`sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])` deve retornar `25`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_of_differences(arr):

    return arr
```

# --solutions--

```py
def sum_of_differences(arr):
    return sum(arr[i + 1] - arr[i] for i in range(len(arr) - 1))
```
