---
id: 69f90c1329a94b37e2a2086d
title: "Desafio 294: combinações de parênteses"
challengeType: 29
dashedName: challenge-294
---

# --description--

Dado um inteiro, `n`, retorne o número de combinações válidas de `n` pares de parênteses.

- Uma combinação válida é uma string onde todo parêntese de abertura tem um parêntese de fechamento correspondente, e nenhum parêntese de fechamento aparece antes do seu parêntese de abertura correspondente.

Por exemplo, dado `2`, existem 2 combinações válidas:

```json
(())
()()
```

# --hints--

`get_combinations(2)` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(2), 2)`)
}})
```

`get_combinations(3)` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(3), 5)`)
}})
```

`get_combinations(5)` deve retornar `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(5), 42)`)
}})
```

`get_combinations(8)` deve retornar `1430`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(8), 1430)`)
}})
```

`get_combinations(13)` deve retornar `742900`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(13), 742900)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_combinations(n):

    return n
```

# --solutions--

```py
def get_combinations(n):
    count = 0

    def generate(op, cl):
        nonlocal count
        if op == 0 and cl == 0:
            count += 1
            return
        if op > 0:
            generate(op - 1, cl)
        if cl > op:
            generate(op, cl - 1)

    generate(n, n)
    return count
```
