---
id: 69f90c1329a94b37e2a2086c
title: "Desafio 293: melhor mão"
challengeType: 29
dashedName: challenge-293
---

# --description--

Dado um array de cinco strings representando cartas de baralho, retorne o nome da melhor mão.

- Cada carta é representada como uma string de dois caracteres: o valor seguido pelo naipe, `"2h"` por exemplo.
  - Os valores, do mais baixo ao mais alto, são: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"` e `"A"`.
  - Os naipes são: `"h"`, `"d"`, `"c"` e `"s"`.
- Ases (`"A"`) podem ser usados como alto ou baixo em uma sequência.

As mãos, em ordem da pior para a melhor, são:

| Nome | Descrição |
| - | - |
| `"High Card"` | Nenhum par ou melhor |
| `"Pair"` | Dois de um valor |
| `"Two Pair"` | Dois de um valor e dois de outro |
| `"Three of a Kind"` | Três de um valor |
| `"Straight"` | Cinco valores em sequência |
| `"Flush"` | Cinco do mesmo naipe |
| `"Full House"` | Três de um valor e dois de outro |
| `"Four of a Kind"` | Quatro de um valor |
| `"Straight Flush"` | Cinco valores em sequência do mesmo naipe |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` do mesmo naipe |

Retorne o nome da melhor mão.

# --hints--

`get_best_hand(["7s", "7h", "7d", "2c", "5h"])` deve retornar `"Three of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind")`)
}})
```

`get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"])` deve retornar `"Full House"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House")`)
}})
```

`get_best_hand(["2h", "5h", "7h", "9h", "Jh"])` deve retornar `"Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["2h", "5h", "7h", "9h", "Jh"]), "Flush")`)
}})
```

`get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"])` deve retornar `"Four of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind")`)
}})
```

`get_best_hand(["Ts", "Th", "9d", "9c", "8h"])` deve retornar `"Two Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair")`)
}})
```

`get_best_hand(["9c", "8c", "7c", "6c", "5c"])` deve retornar `"Straight Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush")`)
}})
```

`get_best_hand(["As", "Kh", "Jd", "8c", "5h"])` deve retornar `"High Card"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Kh", "Jd", "8c", "5h"]), "High Card")`)
}})
```

`get_best_hand(["As", "2h", "3d", "4c", "5h"])` deve retornar `"Straight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "2h", "3d", "4c", "5h"]), "Straight")`)
}})
```

`get_best_hand(["Ts", "Th", "7c", "6d", "5h"])` deve retornar `"Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "7c", "6d", "5h"]), "Pair")`)
}})
```

`get_best_hand(["As", "Ks", "Qs", "Js", "Ts"])` deve retornar `"Royal Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_best_hand(cards):

    return cards
```

# --solutions--

```py
def get_best_hand(cards):
    rank_order = "23456789TJQKA"

    ranks = [c[0] for c in cards]
    suits = [c[1] for c in cards]

    is_flush = len(set(suits)) == 1
    rank_indices = sorted(rank_order.index(r) for r in ranks)

    def is_straight(indices):
        is_regular = all(indices[i] == indices[i - 1] + 1 for i in range(1, len(indices)))
        is_low_ace = indices == [0, 1, 2, 3, 12]
        return is_regular or is_low_ace

    straight = is_straight(rank_indices)

    counts = {}
    for r in ranks:
        counts[r] = counts.get(r, 0) + 1
    freq = sorted(counts.values(), reverse=True)

    if is_flush and rank_indices == [8, 9, 10, 11, 12]:
        return "Royal Flush"
    if is_flush and straight:
        return "Straight Flush"
    if freq[0] == 4:
        return "Four of a Kind"
    if freq[0] == 3 and freq[1] == 2:
        return "Full House"
    if is_flush:
        return "Flush"
    if straight:
        return "Straight"
    if freq[0] == 3:
        return "Three of a Kind"
    if freq[0] == 2 and freq[1] == 2:
        return "Two Pair"
    if freq[0] == 2:
        return "Pair"
    return "High Card"
```
