---
id: 6a0dcd03ee4e68698080ef6b
title: "Desafio 307: regulamentos de zoneamento"
challengeType: 29
dashedName: challenge-307
---

# --description--

Dada uma grade 2D (array de arrays) que representa o layout dos prédios de uma cidade, retorne as coordenadas de todos os prédios que estão violando as regras de zoneamento.

Cada célula na grade contém um dos rótulos da tabela abaixo. Um prédio está em violação se qualquer um de seus (até) 4 vizinhos, horizontal ou vertical, for de um tipo ao qual ele não pode estar adjacente.

| Rótulo | Tipo | Não pode estar adjacente a |
| - | - | - |
| `"i"` | industrial | `"R"`, `"I"` |
| `"A"` | agrícola | `"C"` |
| `"R"` | residencial | `"i"`, `"C"` |
| `"I"` | institucional | `"i"` |
| `"C"` | comercial | `"R"`, `"A"` |
| `""` (string vazia) | não desenvolvido | sem restrições |

Retorne as coordenadas de todas as células em violação como um array de pares `[row, col]`, em qualquer ordem. Se não houver violações, retorne um array vazio.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` deve retornar `[[0, 0], [0, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` deve retornar `[[0, 1], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` deve retornar `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` deve retornar `[[0, 1], [0, 2], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` deve retornar `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
