---
id: 6a0dcd03ee4e68698080ef68
title: "Desafio 304: arranjos de itinerário"
challengeType: 29
dashedName: challenge-304
---

# --description--

Dado um array com pelo menos duas paradas opcionais para um passeio de um dia, retorne o número de arranjos válidos de itinerário.

O itinerário sempre inclui `"breakfast"`, `"lunch"` e `"dinner"`, que não serão passados como argumentos. As paradas opcionais podem ser colocadas em qualquer lugar do itinerário, sujeitas às seguintes regras:

- `"breakfast"` é sempre a primeira, com pelo menos uma parada antes de `"lunch"`.
- `"lunch"` deve aparecer antes de `"dinner"`, com pelo menos uma parada entre elas.
- No máximo, uma parada opcional pode aparecer depois de `"dinner"`.

Retorne o número de arranjos válidos.

# --hints--

`get_itinerary_count(["library", "park"])` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` deve retornar `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` deve retornar `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` deve retornar `840`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` deve retornar `55440`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
