---
id: 6a0dcd03ee4e68698080ef6c
title: "Desafio 308: validador de cartão de crédito"
challengeType: 29
dashedName: challenge-308
---

# --description--

Dada uma string de dígitos para um número de cartão de crédito, determine se é um número válido usando o seguinte método:

- Começando pelo penúltimo dígito, dobre cada outro dígito movendo para a esquerda.
- Se dobrar um dígito resultar em um número maior que 9, subtraia 9.
- Some todos os dígitos (dobrados e não dobrados).
- Se o total for divisível por 10, o número é válido.

# --hints--

`is_valid_card("4532015112830366")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830366"), True)`)
}})
```

`is_valid_card("5425233430109903")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("5425233430109903"), True)`)
}})
```

`is_valid_card("371449635398431")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("371449635398431"), True)`)
}})
```

`is_valid_card("6011111111111117")` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("6011111111111117"), True)`)
}})
```

`is_valid_card("4532015112830367")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830367"), False)`)
}})
```

`is_valid_card("1234567890123456")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("1234567890123456"), False)`)
}})
```

`is_valid_card("4532015112830368")` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830368"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_card(number):

    return number
```

# --solutions--

```py
def is_valid_card(number):
    total = 0
    for i, digit in enumerate(reversed(number)):
        d = int(digit)
        if i % 2 == 1:
            d *= 2
            if d > 9:
                d -= 9
        total += d
    return total % 10 == 0
```
