---
id: 6a0dcc730cb92a616f86f0bf
title: "Desafio 295: validador de esquema parte 1"
challengeType: 29
dashedName: challenge-295
---

# --description--

Dado um objeto (JavaScript) ou dicionário (Python), determine se ele corresponde ao seguinte esquema:

```json
{
  username: string
}
```

- Chaves extras são permitidas

# --hints--

`is_valid_schema({"username": "bob"})` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bob"}), True)`)
}})
```

`is_valid_schema({"username": "jen", "posts": 30})` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "jen", "posts": 30}), True)`)
}})
```

`is_valid_schema({"username": ""})` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": ""}), True)`)
}})
```

`is_valid_schema({"username": 7})` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 7}), False)`)
}})
```

`is_valid_schema({"posts": 25})` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"posts": 25}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return isinstance(obj.get("username"), str)
```
