---
id: 6a0dcc730cb92a616f86f0c5
title: "Desafio 301: última carga"
challengeType: 29
dashedName: challenge-301
---

# --description--

Dado o número de colheres de detergente para roupa que você tem restantes e um array com quantas colheres você usou em cada um dos dias anteriores, retorne o número de dias completos de detergente que ainda restam.

Calcule sua média diária de uso a partir do histórico de uso e assuma essa quantidade de uso para cada dia daqui para frente.

# --hints--

`last_load_date(10, [2, 2, 2, 2, 2, 2, 2])` deve retornar `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(10, [2, 2, 2, 2, 2, 2, 2]), 5)`)
}})
```

`last_load_date(16, [2, 3, 0, 3, 4, 2, 1])` deve retornar `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(16, [2, 3, 0, 3, 4, 2, 1]), 7)`)
}})
```

`last_load_date(33, [5, 0, 4, 3, 3, 2])` deve retornar `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(33, [5, 0, 4, 3, 3, 2]), 11)`)
}})
```

`last_load_date(50, [2, 0, 2, 9, 12, 0, 2])` deve retornar `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(50, [2, 0, 2, 9, 12, 0, 2]), 12)`)
}})
```

`last_load_date(20, [13, 9, 12, 10, 8])` deve retornar `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(20, [13, 9, 12, 10, 8]), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def last_load_date(scoops, usage):

    return scoops
```

# --solutions--

```py
def last_load_date(scoops, usage):
    avg = sum(usage) / len(usage)
    return int(scoops / avg)
```
