---
id: 6a15cadf5f240d05a2649558
title: "Desafio 314: fatoração em números primos"
challengeType: 29
dashedName: challenge-314
---

# --description--

Dado um inteiro maior que 1, retorne sua fatoração em números primos como um array de números em ordem crescente.

Uma fatoração em números primos é o conjunto de números primos que se multiplicam para produzir o inteiro dado. Cada número tem exatamente um conjunto. Por exemplo, a fatoração em números primos de 20 é `[2, 2, 5]` porque 2 * 2 * 5 = 20.

Se o inteiro dado for ele mesmo primo, retorne-o em um array de um único elemento.

# --hints--

`prime_factorization(20)` deve retornar `[2, 2, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(20), [2, 2, 5])`)
}})
```

`prime_factorization(17)` deve retornar `[17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(17), [17])`)
}})
```

`prime_factorization(15)` deve retornar `[3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(15), [3, 5])`)
}})
```

`prime_factorization(35)` deve retornar `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(35), [5, 7])`)
}})
```

`prime_factorization(999)` deve retornar `[3, 3, 3, 37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(999), [3, 3, 3, 37])`)
}})
```

`prime_factorization(360)` deve retornar `[2, 2, 2, 3, 3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(360), [2, 2, 2, 3, 3, 5])`)
}})
```

`prime_factorization(510510)` deve retornar `[2, 3, 5, 7, 11, 13, 17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(510510), [2, 3, 5, 7, 11, 13, 17])`)
}})
```

# --seed--

## --seed-contents--

```py
def prime_factorization(n):

    return n
```

# --solutions--

```py
def prime_factorization(n):
    factors = []
    divisor = 2
    while divisor * divisor <= n:
        while n % divisor == 0:
            factors.append(divisor)
            n //= divisor
        divisor += 1
    if n > 1:
        factors.append(n)
    return factors
```
