---
id: 6a151d32d772271dd2448c2d
title: "Desafio 311: conjurador"
challengeType: 29
dashedName: challenge-311
---

# --description--

Dada uma string de códigos de feitiços que você está lançando, calcule a pontuação total.

Cada caractere na string representa um feitiço:

| Código | Feitiço | Categoria | Pontuação base |
| - | - | - | - |
| `"f"` | Fogo | Destruição | 3 |
| `"l"` | Raio | Destruição | 3 |
| `"i"` | Gelo | Controle | 2 |
| `"w"` | Vento | Controle | 2 |
| `"h"` | Cura | Restauração | 1 |
| `"s"` | Escudo | Restauração | 1 |

Um multiplicador de combo é aplicado com base em quantos feitiços consecutivos foram lançados de categorias diferentes:

- O primeiro feitiço sempre pontua o valor base.
- Cada feitiço consecutivo de uma categoria diferente da anterior aumenta o multiplicador em 1.
- Lançar um feitiço da mesma categoria do anterior reseta o multiplicador para 1.
- A pontuação de cada feitiço é sua pontuação base multiplicada pelo multiplicador atual.

Retorne a pontuação total da sequência de feitiços.

# --hints--

`cast("fihwl")` deve retornar `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` deve retornar `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` deve retornar `37`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` deve retornar `50`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` deve retornar `101`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
