---
id: 6a15cadf5f240d05a264955d
title: "Desafio 319: analisador de frontmatter"
challengeType: 29
dashedName: challenge-319
---

# --description--

Dada uma string que representa um bloco de frontmatter, analise-a e retorne um objeto (JavaScript) ou dicionário (Python) com as chaves e valores.

O frontmatter é delimitado por `---` e contém pares `key: value` dentro dele, um por linha. Por exemplo:

```md
---
title: My Post
draft: false
views: 100
---
```

Deve retornar:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- Números, booleanos e strings devem ser retornados como seus respectivos tipos.
- A string fornecida terá quebras de linha separadas pelo caractere de nova linha (`"\n"`). O exemplo acima seria dado como: `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`.

# --hints--

`parse_frontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` deve retornar `{ title: "My Post", draft: False, views: 100 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\ntitle: My Post\\ndraft: false\\nviews: 100\\n---"), {"title": "My Post", "draft": False, "views": 100})`)
}})
```

`parse_frontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` deve retornar `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: False }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nid: 6a174db57256a112f932195c\\ntitle: My Book\\nlocale: en\\nwordCount: 10000\\npublished: false\\n---"), {"id": "6a174db57256a112f932195c", "title": "My Book", "locale": "en", "wordCount": 10000, "published": False})`)
}})
```

`parse_frontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` deve retornar `{ version: "1.0.0", url: "https://example.com", private: True }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nversion: 1.0.0\\nurl: https://example.com\\nprivate: true\\n---"), {"version": "1.0.0", "url": "https://example.com", "private": True})`)
}})
```

`parse_frontmatter("---\nrating: 4.5\nprice: 9.99\n---")` deve retornar `{ rating: 4.5, price: 9.99 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nrating: 4.5\\nprice: 9.99\\n---"), {"rating": 4.5, "price": 9.99})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_frontmatter(s):

    return s
```

# --solutions--

```py
def parse_frontmatter(s):
    lines = [line for line in s.split("\n") if line != "---" and line != ""]
    result = {}
    for line in lines:
        key, *rest = line.split(":")
        value = ":".join(rest).strip()
        if value == "true":
            result[key.strip()] = True
        elif value == "false":
            result[key.strip()] = False
        else:
            try:
                n = float(value)
                result[key.strip()] = int(n) if n == int(n) else n
            except ValueError:
                result[key.strip()] = value
    return result
```
