---
id: 6a15cadf5f240d05a2649557
title: "Desafio 313: custo do aluguel"
challengeType: 29
dashedName: challenge-313
---

# --description--

Dado um timestamp de aluguel, um timestamp de devolução e um nível de aluguel, retorne o custo total do aluguel incluindo quaisquer taxas por atraso.

- Os timestamps fornecidos são strings ISO UTC, por exemplo: `"2026-06-18T18:30:00Z"`.
- O nível de aluguel é o número de dias antes do aluguel vencer: `1`, `3` ou `7`.
- Os aluguéis devem ser devolvidos até as 12:00 PM UTC ou antes no último dia do período de aluguel. Por exemplo, um aluguel de 1 dia retirado a qualquer hora em 15 de março deve ser devolvido até as 12:00 PM UTC em 16 de março.
- Cada dia após a data e hora de devolução incorre em uma taxa por atraso.

Os preços são os seguintes:

| Nível | Custo base | Taxa por atraso por dia |
| - | - | - |
| 1 dia | $4.99 | $3.99 |
| 3 dias | $3.99 | $2.99 |
| 7 dias | $2.99 | $0.99 |

Retorne o custo total arredondado para duas casas decimais no formato `"$D.CC"`.

# --hints--

`get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` deve retornar `"$4.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99")`)
}})
```

`get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` deve retornar `"$12.97"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97")`)
}})
```

`get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` deve retornar `"$3.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99")`)
}})
```

`get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` deve retornar `"$9.97"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97")`)
}})
```

`get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` deve retornar `"$2.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99")`)
}})
```

`get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` deve retornar `"$358.40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rental_cost(rented, returned, tier):

    return rented
```

# --solutions--

```py
from datetime import datetime, timedelta
import math

def get_rental_cost(rented, returned, tier):
    pricing = {
        1: {"base": 4.99, "late": 3.99},
        3: {"base": 3.99, "late": 2.99},
        7: {"base": 2.99, "late": 0.99},
    }
    rented_dt = datetime.fromisoformat(rented.replace("Z", "+00:00"))
    returned_dt = datetime.fromisoformat(returned.replace("Z", "+00:00"))
    due = rented_dt + timedelta(days=tier)
    due = due.replace(hour=12, minute=0, second=0, microsecond=0)
    base = pricing[tier]["base"]
    late = pricing[tier]["late"]
    if returned_dt <= due:
        return f"${base:.2f}"
    days_late = math.ceil((returned_dt - due).total_seconds() / 86400)
    return f"${base + days_late * late:.2f}"
```
