---
id: 6a19b1062d1b153d8ac76d71
title: "Desafio 323: buscador de humor da música"
challengeType: 29
dashedName: challenge-323
---

# --description--

Dada uma string de gênero e um número de BPM para uma música, determine o humor usando a seguinte tabela:

| Humor | Gênero | Faixa de BPM |
|------|-------|-----------|
| `"focus"` | `"classical"` | 60–109 |
| `"focus"` | `"electronic"` | 60–89 |
| `"happy"` | `"pop"` | 60–180 |
| `"happy"` | `"classical"` | 110–180 |
| `"happy"` | `"rock"` | 60–129 |
| `"happy"` | `"electronic"` | 90–134 |
| `"hype"` | `"rock"` | 130–180 |
| `"hype"` | `"electronic"` | 135–180 |

# --hints--

`get_mood("rock", 111)` deve retornar `"happy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_mood("rock", 111), "happy")`)
}})
```

`get_mood("electronic", 74)` deve retornar `"focus"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_mood("electronic", 74), "focus")`)
}})
```

`get_mood("classical", 180)` deve retornar `"happy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_mood("classical", 180), "happy")`)
}})
```

`get_mood("rock", 155)` deve retornar `"hype"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_mood("rock", 155), "hype")`)
}})
```

`get_mood("electronic", 90)` deve retornar `"happy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_mood("electronic", 90), "happy")`)
}})
```

`get_mood("classical", 67)` deve retornar `"focus"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_mood("classical", 67), "focus")`)
}})
```

`get_mood("pop", 100)` deve retornar `"happy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_mood("pop", 100), "happy")`)
}})
```

`get_mood("electronic", 135)` deve retornar `"hype"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_mood("electronic", 135), "hype")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_mood(genre, bpm):

    return genre
```

# --solutions--

```py
def get_mood(genre, bpm):
    table = [
        {"mood": "focus",  "genre": "classical",  "min": 60,  "max": 109},
        {"mood": "focus",  "genre": "electronic", "min": 60,  "max": 89},
        {"mood": "happy",  "genre": "pop",        "min": 60,  "max": 180},
        {"mood": "happy",  "genre": "classical",  "min": 110, "max": 180},
        {"mood": "happy",  "genre": "rock",       "min": 60,  "max": 129},
        {"mood": "happy",  "genre": "electronic", "min": 90,  "max": 134},
        {"mood": "hype",   "genre": "rock",       "min": 130, "max": 180},
        {"mood": "hype",   "genre": "electronic", "min": 135, "max": 180},
    ]

    return next(row["mood"] for row in table if row["genre"] == genre and row["min"] <= bpm <= row["max"])
```
