---
id: 6a1d9f98e819ed70a0e994d9
title: "Desafio 327: migração de banco de dados"
challengeType: 29
dashedName: challenge-327
---

# --description--

Dado dois objetos de banco de dados, retorne o segundo objeto com quaisquer propriedades faltantes do primeiro preenchidas.

- Campos que já existem no registro não devem ser sobrescritos.

# --hints--

`migrate_record({ "username": "", "posts": 0 }, { "verified": True })` deve retornar `{ "username": "", "posts": 0, "verified": True }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "verified": True }), { "username": "", "posts": 0, "verified": True })`)
}})
```

`migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "posts": 5 })` deve retornar `{ "username": "camper", "posts": 5 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "posts": 5 }), { "username": "camper", "posts": 5 })`)
}})
```

`migrate_record({ "username": "", "posts": 0, "verified": False }, { "username": "camper" })` deve retornar `{ "username": "camper", "posts": 0, "verified": False }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0, "verified": False }, { "username": "camper" }), { "username": "camper", "posts": 0, "verified": False })`)
}})
```

`migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "role": "admin" })` deve retornar `{ "username": "camper", "role": "admin", "posts": 0 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "role": "admin" }), { "username": "camper", "role": "admin", "posts": 0 })`)
}})
```

`migrate_record({ "username": "", "email": "", "posts": 0, "verified": False, "role": "user", "banned": False }, { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin" })` deve retornar `{ "username": "camper", "email": "camper@freecodecamp.org", "role": "admin", "posts": 0, "verified": False, "banned": False }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "email": "", "posts": 0, "verified": False, "role": "user", "banned": False }, { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin" }), { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin", "posts": 0, "verified": False, "banned": False })`)
}})
```

# --seed--

## --seed-contents--

```py
def migrate_record(schema, record):

    return schema
```

# --solutions--

```py
def migrate_record(schema, record):
    return {**schema, **record}
```
