---
id: 6a1d9f98e819ed70a0e994da
title: "Desafio 328: a rotina de Kaprekar"
challengeType: 29
dashedName: challenge-328
---

# --description--

Dado um número de 4 dígitos, retorne quantas vezes você precisa aplicar a rotina de Kaprekar até chegar em 6174.

A rotina de Kaprekar funciona da seguinte forma:

- Organize os dígitos em ordem decrescente para formar o maior número
- Organize os dígitos em ordem crescente para formar o menor número (complete com zeros à esquerda se necessário)
- Subtraia o menor do maior
- Repita com o novo número

# --hints--

`kaprekar(1234)` deve retornar `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(1234), 3)`)
}})
```

`kaprekar(2025)` deve retornar `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(2025), 6)`)
}})
```

`kaprekar(7173)` deve retornar `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(7173), 4)`)
}})
```

`kaprekar(3164)` deve retornar `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(3164), 7)`)
}})
```

`kaprekar(8082)` deve retornar `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(8082), 2)`)
}})
```

# --seed--

## --seed-contents--

```py
def kaprekar(n):

    return n
```

# --solutions--

```py
def kaprekar(n):
    steps = 0
    current = n
    while current != 6174:
        digits = list(str(current).zfill(4))
        desc = int("".join(sorted(digits, reverse=True)))
        asc = int("".join(sorted(digits)))
        current = desc - asc
        steps += 1
    return steps
```
