---
id: 6a1d9f98e819ed70a0e994e2
title: "Desafio 336: Compatibilidade do horóscopo"
challengeType: 29
dashedName: challenge-336
---

# --description--

Dadas duas strings de signos, retorne a porcentagem de compatibilidade entre eles.

Os signos estão organizados em uma roda com 12 posições nesta ordem: `"Aries"`, `"Taurus"`, `"Gemini"`, `"Cancer"`, `"Leo"`, `"Virgo"`, `"Libra"`, `"Scorpio"`, `"Sagittarius"`, `"Capricorn"`, `"Aquarius"`, `"Pisces"`, voltando para `"Aries"` após `"Pisces"`. Encontre a menor distância entre os dois signos e retorne a compatibilidade:

| Distância | Compatibilidade |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscope_match("Libra", "Sagittarius")` deve retornar `"80%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Libra", "Sagittarius"), "80%")`)
}})
```

`horoscope_match("Gemini", "Scorpio")` deve retornar `"20%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Gemini", "Scorpio"), "20%")`)
}})
```

`horoscope_match("Pisces", "Aries")` deve retornar `"40%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Pisces", "Aries"), "40%")`)
}})
```

`horoscope_match("Capricorn", "Cancer")` deve retornar `"50%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Capricorn", "Cancer"), "50%")`)
}})
```

`horoscope_match("Aquarius", "Aquarius")` deve retornar `"100%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Aquarius", "Aquarius"), "100%")`)
}})
```

`horoscope_match("Virgo", "Taurus")` deve retornar `"90%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Virgo", "Taurus"), "90%")`)
}})
```

`horoscope_match("Leo", "Scorpio")` deve retornar `"30%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Leo", "Scorpio"), "30%")`)
}})
```

# --seed--

## --seed-contents--

```py
def horoscope_match(sign1, sign2):

    return sign1
```

# --solutions--

```py
def horoscope_match(sign1, sign2):
    signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"]
    a = signs.index(sign1)
    b = signs.index(sign2)
    diff = abs(a - b)
    distance = min(diff, 12 - diff)
    compatibility = {0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%"}
    return compatibility[distance]
```
