---
id: 6a1d9f98e819ed70a0e994dc
title: "Desafio 330: palavras em minúsculas"
challengeType: 29
dashedName: challenge-330
---

# --description--

Dada uma string, retorne apenas as palavras que estão inteiramente em minúsculas, na ordem original e com um espaço entre cada palavra.

# --hints--

`get_lowercase_words("hello GOOD world")` deve retornar `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("hello GOOD world"), "hello world")`)
}})
```

`get_lowercase_words("these are all lowercase")` deve retornar `"these are all lowercase"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("these are all lowercase"), "these are all lowercase")`)
}})
```

`get_lowercase_words("less is NoT more")` deve retornar `"less is more"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("less is NoT more"), "less is more")`)
}})
```

`get_lowercase_words("DonT eat pizza every OTHER day")` deve retornar `"eat pizza every day"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("DonT eat pizza every OTHER day"), "eat pizza every day")`)
}})
```

`get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` deve retornar `"the quick brown fox jumped over the lazy dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_lowercase_words(s):

    return s
```

# --solutions--

```py
def get_lowercase_words(s):
    return " ".join(word for word in s.split(" ") if word == word.lower())
```
