---
id: 6a1d9f98e819ed70a0e994de
title: "Desafio 332: triagem de issue"
challengeType: 29
dashedName: challenge-332
---

# --description--

Dado um número de milissegundos desde a última postagem em uma issue e a última mensagem postada na issue, determine o que você deve fazer com a issue de acordo com estas regras:

- Se a última mensagem foi há menos de 7 dias, retorne `"leave it"`
- Se a última mensagem foi há 7 ou mais dias e seu conteúdo contém `"bump"` (sem diferenciar maiúsculas de minúsculas), retorne `"close it"`
- Caso contrário, retorne `"bump it"`

# --hints--

`triage_issue(86400000, "Lets fix it")` deve retornar `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(86400000, "Lets fix it"), "leave it")`)
}})
```

`triage_issue(1209600000, "still waiting")` deve retornar `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(1209600000, "still waiting"), "bump it")`)
}})
```

`triage_issue(864000000, "bump")` deve retornar `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(864000000, "bump"), "close it")`)
}})
```

`triage_issue(604800000, "Do we still want this?")` deve retornar `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Do we still want this?"), "bump it")`)
}})
```

`triage_issue(604800000, "Bumping this")` deve retornar `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Bumping this"), "close it")`)
}})
```

`triage_issue(345600000, "I'll make a PR")` deve retornar `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(345600000, "I'll make a PR"), "leave it")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(ms, message):

    return ms
```

# --solutions--

```py
def triage_issue(ms, message):
    seven_days = 7 * 24 * 60 * 60 * 1000
    if ms < seven_days:
        return "leave it"
    if "bump" in message.lower():
        return "close it"
    return "bump it"
```
