---
id: 6a2037a68a0bc2aef0006005
title: "Desafio 343: paradas do elevador"
challengeType: 29
dashedName: challenge-343
---

# --description--

Dado um número para o andar atual de um elevador e um array de andares solicitados, retorne um array com a ordem em que o elevador deve visitá-los para minimizar o número de andares percorridos.

- Em caso de empate, suba primeiro
- Andares com uma requisição devem ser visitados quando o elevador passar por eles pela primeira vez

# --hints--

`elevator_stops(5, [2, 8, 3, 9])` deve retornar `[3, 2, 8, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(5, [2, 8, 3, 9]), [3, 2, 8, 9])`)
}})
```

`elevator_stops(6, [2, 10, 8, 3, 1, 9])` deve retornar `[8, 9, 10, 3, 2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(6, [2, 10, 8, 3, 1, 9]), [8, 9, 10, 3, 2, 1])`)
}})
```

`elevator_stops(1, [4, 8, 3, 6, 9])` deve retornar `[3, 4, 6, 8, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(1, [4, 8, 3, 6, 9]), [3, 4, 6, 8, 9])`)
}})
```

`elevator_stops(12, [6, 10, 7, 3, 1, 4])` deve retornar `[10, 7, 6, 4, 3, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(12, [6, 10, 7, 3, 1, 4]), [10, 7, 6, 4, 3, 1])`)
}})
```

`elevator_stops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19])` deve retornar `[10, 9, 8, 6, 5, 2, 12, 19, 23]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19]), [10, 9, 8, 6, 5, 2, 12, 19, 23])`)
}})
```

# --seed--

## --seed-contents--

```py
def elevator_stops(current_floor, stops):

    return current_floor
```

# --solutions--

```py
def elevator_stops(current_floor, stops):
    stops = sorted(stops)
    above = [s for s in stops if s > current_floor]
    below = [s for s in stops if s < current_floor][::-1]
    up_distance = above[-1] - current_floor if above else float('inf')
    down_distance = current_floor - below[-1] if below else float('inf')
    return below + above if down_distance < up_distance else above + below
```
