---
id: 6a22d77ddf034bc4e35b1d56
title: "Desafio 349: sinal de celular"
challengeType: 29
dashedName: challenge-349
---

# --description--

Dada uma grade contendo três leituras de torres de celular, determine a localização do telefone.

- Cada célula na grade é ou `0` (sem torre) ou um inteiro positivo que representa o número de células até o telefone, medido em linha reta: horizontal, vertical ou diagonal.
- Retorne o `[row, col]` da célula que está na quantidade correta de células de todas as três torres.
- Sempre existe exatamente uma solução.

# --hints--

`find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` deve retornar `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2])`)
}})
```

`find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` deve retornar `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1])`)
}})
```

`find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` deve retornar `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2])`)
}})
```

`find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` deve retornar `[3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4])`)
}})
```

`find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` deve retornar `[3, 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_signal(grid):

    return grid
```

# --solutions--

```py
def find_signal(grid):
    rows = len(grid)
    cols = len(grid[0])

    towers = []
    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                towers.append((r, c, grid[r][c]))

    def is_valid(r, c, tr, tc, dist):
        dr = abs(r - tr)
        dc = abs(c - tc)
        return (dr == dist or dc == dist) and (dr == 0 or dc == 0 or dr == dc)

    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                continue
            if all(is_valid(r, c, tr, tc, dist) for tr, tc, dist in towers):
                return [r, c]
```
