---
id: 6a22d77ddf034bc4e35b1d58
title: "Desafio 351: número prônico"
challengeType: 29
dashedName: challenge-351
---

# --description--

Dado um número, determine se ele é um número prônico.

Um número prônico é o produto de dois inteiros consecutivos. Por exemplo, 6 é prônico porque 2 * 3 = 6.

# --hints--

`is_pronic(6)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(6), True)`)
}})
```

`is_pronic(15)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(15), False)`)
}})
```

`is_pronic(12)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(12), True)`)
}})
```

`is_pronic(132)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(132), True)`)
}})
```

`is_pronic(80)` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(80), False)`)
}})
```

`is_pronic(0)` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(0), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_pronic(n):

    return n
```

# --solutions--

```py
def is_pronic(n):
    k = int(n ** 0.5)
    return k * (k + 1) == n
```
