---
id: 6a22d77ddf034bc4e35b1d54
title: "Desafio 347: teoria dos jogos"
challengeType: 29
dashedName: challenge-347
---

# --description--

Dadas duas strings de mesmo comprimento que representam as estratégias de dois jogadores para um jogo, retorne as pontuações como um array `[player1, player2]`.

- As strings dadas conterão apenas uma de duas letras: `"C"` (cooperar) ou `"D"` (trair).
- Cada caractere representa uma rodada, pontuada da seguinte forma:
  - Se ambos os jogadores cooperarem, cada um marca 3.
  - Se ambos os jogadores traírem, cada um marca 1.
  - Se um jogador trair e o outro cooperar, o traidor marca 5 e o cooperador marca 0.

# --hints--

`play_game("CCCC", "CCCC")` deve retornar `[12, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` deve retornar `[4, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` deve retornar `[5, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` deve retornar `[24, 34]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` deve retornar `[66, 21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
