---
id: 6a22d77ddf034bc4e35b1d5b
title: "Desafio 354: avaliação de contraste 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

Dado dois arrays representando valores RGB e um booleano indicando se o texto é grande, retorne a avaliação de contraste WCAG usando o seguinte método:

Primeiro, converta cada valor RGB para luminância relativa:

- Divida cada canal `[R, G, B]` por 255 para obter um valor entre `0` e `1`
- Aplique a fórmula de correção gama em cada canal:
  - Se o valor do canal for menor ou igual a `0.04045`: `channel / 12.92`
  - Caso contrário: `((channel + 0.055) / 1.055) ^ 2.4`
- Calcule a luminância: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Depois, calcule a taxa de contraste somando 0.05 a cada valor de luminância, depois dividindo o valor mais claro pelo mais escuro. O mais claro será sempre o primeiro argumento.

Retorne a avaliação baseada na taxa de contraste usando a tabela a seguir:

| Avaliação | Texto normal | Texto grande |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | abaixo de 4.5 | abaixo de 3.0 |

# --hints--

`get_contrast_rating([255, 255, 255], [0, 0, 0], False)` deve retornar `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

`get_contrast_rating([215, 188, 188], [55, 55, 55], False)` deve retornar `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

`get_contrast_rating([143, 144, 210], [46, 47, 61], False)` deve retornar `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

`get_contrast_rating([167, 167, 210], [53, 10, 53], True)` deve retornar `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

`get_contrast_rating([135, 147, 155], [60, 70, 90], True)` deve retornar `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

`get_contrast_rating([125, 210, 195], [105, 130, 90], True)` deve retornar `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
