---
id: 6a26df95efa55a2524399743
title: "Desafio 362: validador de nonogram"
challengeType: 29
dashedName: challenge-362
---

# --description--

Dado um array de números de pistas e um array de células, determine se as células satisfazem a pista do nonogram.

- A pista é um array de números que representam os comprimentos de células preenchidas consecutivas, em ordem. Por exemplo, uma pista de `[3, 2]` significa que deve haver 3 células preenchidas consecutivas seguidas por 2 células preenchidas consecutivas, separadas por pelo menos uma célula vazia.
- A linha é um array de 1s (preenchidos) e 0s (vazios).

# --hints--

`is_valid_nonogram([3, 2], [1, 1, 1, 0, 1, 1])` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2], [1, 1, 1, 0, 1, 1]), True)`)
}})
```

`is_valid_nonogram([3, 2], [0, 1, 1, 1, 1, 1])` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2], [0, 1, 1, 1, 1, 1]), False)`)
}})
```

`is_valid_nonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]), False)`)
}})
```

`is_valid_nonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]), True)`)
}})
```

`is_valid_nonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])` deve retornar `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]), True)`)
}})
```

`is_valid_nonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])` deve retornar `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_nonogram(clue, cells):

    return clue
```

# --solutions--

```py
def is_valid_nonogram(clue, cells):
    runs = []
    count = 0

    for i in range(len(cells) + 1):
        if i < len(cells) and cells[i] == 1:
            count += 1
        elif count > 0:
            runs.append(count)
            count = 0

    if len(runs) != len(clue):
        return False
    return all(runs[i] == clue[i] for i in range(len(runs)))
```
