---
id: 6a26df3e2988bcdded204894
title: "Desafio 358: tradutor de emoji"
challengeType: 29
dashedName: challenge-358
---

# --description--

Dada uma string de emojis, retorne a frase usando a tabela a seguir:

| Emoji | Palavra |
|-------|---------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

Retorne as palavras separadas por espaços.

# --hints--

`get_emoji_phrase("🪨⭐")` deve retornar `"rock star"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🪨⭐"), "rock star")`)
}})
```

`get_emoji_phrase("🥵🐕")` deve retornar `"hot dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🥵🐕"), "hot dog")`)
}})
```

`get_emoji_phrase("👶🦈")` deve retornar `"baby shark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("👶🦈"), "baby shark")`)
}})
```

`get_emoji_phrase("⭐🐟")` deve retornar `"star fish"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("⭐🐟"), "star fish")`)
}})
```

`get_emoji_phrase("🧊🧊👶")` deve retornar `"ice ice baby"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🧊🧊👶"), "ice ice baby")`)
}})
```

`get_emoji_phrase("🐱🐟🍲")` deve retornar `"cat fish soup"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🐱🐟🍲"), "cat fish soup")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_emoji_phrase(s):

    return s
```

# --solutions--

```py
def get_emoji_phrase(s):
    table = {
        '👶': 'baby',
        '🐱': 'cat',
        '🐕': 'dog',
        '🐟': 'fish',
        '🥵': 'hot',
        '🧊': 'ice',
        '🪨': 'rock',
        '🦈': 'shark',
        '🍲': 'soup',
        '⭐': 'star',
    }
    return ' '.join(table[emoji] for emoji in s)
```
