---
id: 655cd899f8de09431eabb40c
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このプロジェクトでは、マージソートアルゴリズムを作成しながらデータ構造を学びます。

これは分割統治法の原理を使ってデータの集合をソートするアルゴリズムです。つまり、集合をより小さな部分に「分割」し、それぞれの部分を独立してソートして「制覇」し、最後にソート済みの部分をマージします。

`merge_sort` という名前の関数を作成してください。この関数は数値のリストをソートする役割を担います。

関数の本体では `pass` キーワードを使ってください。

# --hints--

`merge_sort` という名前の関数を宣言してください。関数の本体で `pass` キーワードを使うのを忘れないでください。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
