---
id: 6564aee9c077774ea49c3faf
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

マージソートアルゴリズムは主に3つの動作を行います：

- ソートされていないアイテムのシーケンスをサブパーツに分割する
- サブパーツ内のアイテムをソートする
- ソートされたサブパーツをマージする

これらは再帰的に行われ、サブパーツが完全なソート済みシーケンスにマージされるまで続きます。まずはシーケンスを分割するところから始めましょう。

最初に、`pass` キーワードを変数 `middle_point` に置き換えてください。整数除算演算子 (`//`) を使って `array` リストの長さを2で割り、その結果を新しい `middle_point` 変数に代入してください。コードのインデントを忘れないでください。

# --hints--

関数の本体に `pass` を含めてはいけません。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

`middle_point` 変数を宣言し、`len(array) // 2` を代入してください。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
