---
id: 6564a9fe51964c229d5b7f4c
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

ソートするデータコレクションを示すパラメータが必要です。`array` 関数に `merge_sort` という名前のパラメータを作成してください。

# --hints--

`merge_sort` 関数は1つのパラメータを受け取る必要があります：`array`。

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
