---
id: 656639771fed09b5c0e8fe71
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`array` リストを2つの別々のリストに分割したので、今度は各リストをさらに分割して、すべての要素がそれぞれ単独のリストになるまで続けます。1つの数字だけのリストは常にソート済みです。

そのために、関数内で再帰的に `merge_sort` を呼び出してください。今は引数を渡さないでください。

# --hints--

関数本体の最後で `merge_sort` を呼び出す必要があります。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
