---
id: 656627b47bd2d2a4afbc945d
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

前のステップで、中央の位置を取得しました。それを使って `array` を2つに分け、それぞれを新しい変数に代入できます。

スライス構文を使って `array` の左半分を抽出し、`left_part` という名前の変数に代入してください。


# --hints--

`left_part` 関数内に `merge_sort` という名前の変数があるはずです。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

`array[:middle_point]` を `left_part` 変数に代入してください。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
