---
id: 656680b0fc79f2c38a34d90e
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

いよいよリスト（`left_part` と `right_part`）をソートして、元の `array` にマージする時です。

両方のリストの要素を比較し、小さい方の要素をメインのリストにマージすることでこれを行います。`left_part` と `right_part` のすべてのインデックスでこの比較を行います。

3つの変数 `left_array_index`、`right_array_index`、`sorted_index` を作成し、それぞれに `0` を設定してください。これらの変数はソート処理中に各インデックスを追跡するのに役立ちます。

# --hints--

`left_array_index` 関数の中に `merge_sort` という名前の変数があるはずです

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index/);
  }
})
```

`left_array_index` の値は 0 であるべきです

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index\s*=\s*0/);
  }
})
```

`right_array_index` 関数の中に `merge_sort` という名前の変数があるはずです

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index/);
  }
})
```

`right_array_index` の値は 0 であるべきです

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index\s*=\s*0/);
  }
})

```

`sorted_index` 関数の中に `merge_sort` という名前の変数があるはずです

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index/);
  }
})
```

`sorted_index` の値は 0 であるべきです

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index\s*=\s*0/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)

--fcc-editable-region--

--fcc-editable-region--
```
