---
id: 656706afd65547d22bee0b68
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

`while` ループ内で、`pass` を `if` 文に置き換えて、`left_part` のインデックスが `right_part` のインデックスより小さいかどうかをチェックしてください。

`pass` 文の本体に `if` キーワードを使ってください。

# --hints--

`pass` キーワードを `if` 文に置き換え、この条件を使ってください：`left_part[left_array_index] < right_part[right_array_index]`。`pass` 文の本体に `if` キーワードを追加してください。

```js
({
 test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_body().is_equivalent("if left_part[left_array_index] < right_part[right_array_index]:\\n  pass")`))
 })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        pass
--fcc-editable-region--
```
