---
id: 656710d1e0ec62253426db24
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

作成した`while`ループは、`left_part`の要素と`right_part`の要素を比較し、小さい方の要素をメインの`array`リストに追加します。

この操作は比較する要素がなくなるまで続きます。しかし、`left_part`にはまだ要素が残っている場合があり、`right_part`には要素がない場合もありますし、その逆もあります。

残っている要素を`while`から`left_part`リストにコピーするために、別の`array`ループを作成してください。`left_array_index < len(left_part)`の条件には`while`を使ってください。

# --hints--

`while`ループの条件は`left_array_index < len(left_part)`にしてください。ループの本体には`pass`キーワードを使うのを忘れないでください。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):/m);
  }
})
```

`pass`ループの本体には`while`キーワードを使う必要があります。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
