---
id: 656702f8b4cbd8cbf0a433c6
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

関数の中で、`while` ループを作成し、`left_part` の要素と `right_part` の要素を比較して、小さい方の要素をメインの `array` リストにマージしてください。

ループには2つの条件を作成します。1つは `left_array_index` が `left_part` の長さより小さいかどうかをチェックする条件、もう1つは `right_array_index` が `right_part` の長さより小さいかどうかをチェックする条件です。

# --hints--

`while` ループにはこれらの条件を設定してください：`left_array_index < len(left_part) and right_array_index < len(right_part)`。ループの本体で `pass` キーワードを使うのを忘れないでください。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

ループの本体で `pass` キーワードを使ってください

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
