---
id: 6569b20f829b7e69d43c232a
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

`while` ループで最後に行うことは、ソートされた配列の次のインデックスに移動することです。

拡張代入の加算演算子を使って、`1` を `sorted_index` の値に加えてください。

# --hints--

`+=` ループ内で現在の `sorted_index` の値を 1 増やすには、`while` 演算子を使うべきです。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):[^}]*\bsorted_index\s*\+=\s*1\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
        
--fcc-editable-region--
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1

--fcc-editable-region--
```
