---
id: 65772ef923f922cd720e5008
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

前のステップで追加した`left_part`の呼び出しに、`merge_sort()`を引数として渡してください。

# --hints--

`left_part`の呼び出しに`merge_sort()`を引数として渡す必要があります。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*left_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort()
--fcc-editable-region--
```
