---
id: 657f59751d5a5c9b069ae0f3
title: ステップ 28
challengeType: 20
dashedName: step-28
---

# --description--

`__name__` 変数を使って、Pythonスクリプトがメインプログラムとして実行されているのか、別のPythonファイルで書かれたモジュールとしてインポートされているのかを判別できます。

`__name__` の値が `'__main__'` に設定されている場合、現在のスクリプトはモジュールではなくメインプログラムであることを意味します。

このプロジェクトでは、現在のスクリプトをメインプログラムとして使います。

`if` の値が `__name__` であるかどうかを確認する `'__main__'` 文を作成してください。

`pass` 文の本体には `if` キーワードを使ってください。

# --hints--

`if` 文は `__name__` が `'__main__'` の値を持つかどうかを確認する必要があります。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
