---
id: 641cd18eb67c661d8a9e11f3
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

프로젝트를 시작하려면 `.getElementById()` 메서드를 사용해 HTML 문서에서 `#message-input`, `#result`, `#check-message-btn` 요소를 가져와 각각 `messageInput`, `result`, `checkMessageButton` 변수에 할당하세요.

# --hints--

`const`를 사용해 `messageInput` 변수를 선언하세요.

```js
assert.match(code, /const\s+messageInput\s*=/);
```

`messageInput` 변수는 `document.getElementById("message-input")` 값을 가져야 합니다.

```js
assert.deepEqual(messageInput, document.getElementById('message-input'));
```

`const`를 사용해 `result` 변수를 선언하세요.

```js
assert.match(code, /const\s+result\s*=/);
```

`result` 변수는 `document.getElementById("result")` 값을 가져야 합니다.

```js
assert.match(code, /const\s+result\s*=\s*document\.getElementById\(\s*('|")result\1\s*\)/);
```

`const`를 사용해 `checkMessageButton` 변수를 선언하세요.

```js
assert.match(code, /const\s+checkMessageButton\s*=/);
```

`checkMessageButton` 변수는 `document.getElementById("check-message-btn")` 값을 가져야 합니다.

```js
assert.deepEqual(checkMessageButton, document.getElementById('check-message-btn'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
