---
id: 64221007887f38213fa57827
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

이 표현식에 대해 마지막으로 알아볼 것이 있습니다. 캡처 그룹에서 일치하는 값을 실제로 사용할 필요가 없으니, 이를 <dfn>비캡처 그룹</dfn>으로 바꿀 수 있습니다. 이렇게 하면 결과를 저장하지 않고도 문자들을 그룹화할 수 있습니다.

정규 표현식에서 비캡처 그룹을 만들려면 그룹의 여는 괄호 뒤에 `?:`를 추가하세요. 예를 들어, `(?:a|b)`는 `a` 또는 `b` 중 하나와 일치하지만 결과를 캡처하지는 않습니다.

정규 표현식을 비캡처 그룹을 사용하도록 업데이트하세요.

# --hints--

`dollarRegex`는 `?:`를 사용해야 합니다.

```js
assert.match(dollarRegex.source, /\?:/);
```

`dollarRegex`는 비캡처 그룹을 사용해야 합니다.

```js
assert.match(dollarRegex.source, /\(\?:.*\)/);
```

`(hundred|thousand|million|billion)`는 비캡처 그룹이어야 합니다.

```js
assert.match(dollarRegex.source, /\(\?:hundred\|thousand\|million\|billion\)/);
```

`(hundred|thousand|million|billion)` 그룹은 여전히 선택 사항이어야 합니다.

```js
assert.match(dollarRegex.source, /\(\?:hundred\|thousand\|million\|billion\)\?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
--fcc-editable-region--
const dollarRegex = /[0-9]+\s*(hundred|thousand|million|billion)?\s+dollars/i;
--fcc-editable-region--

const denyList = [helpRegex, dollarRegex];

const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
