---
id: 64233d08f234a310e73f9496
title: 29단계
challengeType: 0
dashedName: step-29
---

# --description--

다음 정규 표현식은 `stock alert` 같은 문자열과 일치할 것입니다. `stockRegex` 변수를 선언하고 `stock alert` 문자열과 일치하는 정규 표현식을 할당하세요. 대소문자를 구분하지 않도록 만드세요.

또한 `denyList` 배열에 추가하세요.

# --hints--

`const`를 사용해 `stockRegex` 변수를 선언하세요.

```js
assert.match(code, /const\s+stockRegex\s*=/);
```

`stockRegex` 변수에는 정규 표현식을 할당해야 합니다.

```js
assert.instanceOf(stockRegex, RegExp);
```

`stockRegex`는 `stock alert`와 일치해야 합니다.

```js
assert.match('stock alert', stockRegex);
```

`stockRegex`는 대소문자를 구분하지 않아야 합니다.

```js
assert.include(stockRegex.flags, 'i');
```

`denyList` 배열에는 `stockRegex`이 포함되어야 합니다.

```js
assert.deepInclude(denyList, stockRegex);
```

`denyList` 배열에는 `freeRegex`이 포함되어야 합니다.

```js
assert.deepInclude(denyList, freeRegex);
```

`denyList` 배열에는 `dollarRegex`이 포함되어야 합니다.

```js
assert.deepInclude(denyList, dollarRegex);
```

`denyList` 배열에는 `helpRegex`이 포함되어야 합니다.

```js
assert.deepInclude(denyList, helpRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
const dollarRegex = /[0-9]+\s*(?:hundred|thousand|million|billion)?\s+dollars/i;
const freeRegex = /(?:^|\s)fr[e3][e3] m[o0]n[e3]y(?:$|\s)/i;
--fcc-editable-region--


const denyList = [helpRegex, dollarRegex, freeRegex];
--fcc-editable-region--

const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
