---
id: 671fa82dc39e4c320352f371
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

さまざまなタイプのストーリーを保持するために、`storyObj`という名前のオブジェクトを作成してください。そのオブジェクトは3つのプロパティ、`scary`、`funny`、`adventure`を持つ必要があります。各プロパティの値は、`story`と`borderColor`というプロパティを持つオブジェクトにしてください。

個々のオブジェクトに以下の値を代入してください：

- `scary`:
  
  - `story`: `In the dark woods, a group of friends stumbled upon an old, abandoned cabin. They enter the cabin and awaken something malevolent that had been dormant for centuries.`
  - `borderColor`: `#ee4b2b`

- `funny`:
  
  - `story`: `During a camping trip, Mark decided to show off his culinary skills by cooking dinner over an open fire. However, his attempt caused him to burn the dinner as well as his eyebrows off.`
  - `borderColor`: `#f1be32`

- `adventure`:
  
  - `story`: `Lost in the heart of the Amazon rain forest, Sarah and Jake stumbled upon an ancient temple. They braved deadly traps and encountered strange wildlife, all while deciphering cryptic clues left behind by a mysterious civilization.`
  - `borderColor`: `#acd157`

# --hints--

`storyObj`というオブジェクトがあるはずです。

```js
assert.isObject(storyObj);
```

`storyObj`オブジェクトは3つのプロパティ、`scary`、`funny`、`adventure`を持つべきです。

```js
assert.containsAllKeys(storyObj, ['scary', 'funny', 'adventure']);
```

各プロパティは`story`と`borderColor`というプロパティを持つべきです。

```js
assert.containsAllKeys(storyObj.scary, ['story', 'borderColor']);
assert.containsAllKeys(storyObj.funny, ['story', 'borderColor']);
assert.containsAllKeys(storyObj.adventure, ['story', 'borderColor']);
```

`story`の`scary`プロパティは`"In the dark woods, a group of friends stumbled upon an old, abandoned cabin. They enter the cabin and awaken something malevolent that had been dormant for centuries."`という文字列の値を持つべきです。

```js
assert.equal(
  storyObj.scary.story.replace(/\s+/g, ' '),
  "In the dark woods, a group of friends stumbled upon an old, abandoned cabin. They enter the cabin and awaken something malevolent that had been dormant for centuries."
);
```

`borderColor`の`scary`プロパティは`"#ee4b2b"`という文字列の値を持つべきです。

```js
assert.equal(storyObj.scary.borderColor, "#ee4b2b");
```

`story`の`funny`プロパティは`"During a camping trip, Mark decided to show off his culinary skills by cooking dinner over an open fire. However, his attempt caused him to burn the dinner as well as his eyebrows off."`という文字列の値を持つべきです。

```js
assert.equal(storyObj.funny.story, "During a camping trip, Mark decided to show off his culinary skills by cooking dinner over an open fire. However, his attempt caused him to burn the dinner as well as his eyebrows off.");
```

`borderColor`の`funny`プロパティは`"#f1be32"`という文字列の値を持つべきです。

```js
assert.equal(storyObj.funny.borderColor, "#f1be32");
```

`story`の`adventure`プロパティは`"Lost in the heart of the Amazon rain forest, Sarah and Jake stumbled upon an ancient temple. They braved deadly traps and encountered strange wildlife, all while deciphering cryptic clues left behind by a mysterious civilization."`という文字列の値を持つべきです。

```js
assert.equal(storyObj.adventure.story, "Lost in the heart of the Amazon rain forest, Sarah and Jake stumbled upon an ancient temple. They braved deadly traps and encountered strange wildlife, all while deciphering cryptic clues left behind by a mysterious civilization.");
```

`borderColor`の`adventure`プロパティは`"#acd157"`という文字列の値を持つべきです。

```js
assert.equal(storyObj.adventure.borderColor, "#acd157");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Storytelling App</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>

    <h1>Want to hear a short story?</h1>

    <main class="story-container">
        <div class="btn-container">
            <button class="btn" id="scary-btn">Scary Story</button>
            <button class="btn" id="funny-btn">Funny Story</button>
            <button class="btn" id="adventure-btn">Adventure Story</button>
        </div>
        <p id="result"></p>
    </main>
    <script src="./script.js"></script>
</body>
</html>
```

```js
const storyContainer = document.querySelector(".story-container");

const scaryStoryBtn = document.getElementById("scary-btn");
const funnyStoryBtn = document.getElementById("funny-btn");
const adventureStoryBtn = document.getElementById("adventure-btn");

const resultParagraph = document.getElementById("result");

--fcc-editable-region--

--fcc-editable-region--
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --dark-grey: #1b1b32;
    --black: #000;
    --white: #fff;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
}

body {
    background-color: var(--dark-grey);
    color: var(--white);
}

h1,
#result {
    text-align: center;
}

h1 {
    margin: 10px 0 15px;
}

.story-container {
    margin: auto;
    padding: 10px;
    width: 80%;
    border-style: double;
    border-width: 14px;
    border-color: var(--white);
}

.btn-container {
    display: flex;
    flex-direction: column;
    justify-content: center;
    align-items: center;
}

@media (min-width: 760px) {
    .btn-container {
        flex-direction: row;
    }
}

#result {
    margin-top: 15px;
    font-size: 1.2rem;
    line-height: 30px;
}

.btn {
    cursor: pointer;
    width: 200px;
    margin: 10px 0 10px 0.5rem;
    color: var(--black);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}
```
