---
id: 696920c0c98a1ed58eb86293
title: Jenga Maonyesho ya Bidhaa
challengeType: 25
dashedName: lab-product-showcase
demoType: onClick
---

# --description--

Kwa zoezi hili, umepewa HTML na CSS zote. Utakamilisha maonyesho ya bidhaa kwa kutumia TypeScript.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza zoezi.

**Hadithi za watumizi:**

1. Unapaswa kuunda kiolesura kinachoitwa `Item` chenye:
   - Sifa ya `type` ambayo inaweza kuwa `"book"`, `"electronics"`, au `"clothing"` tu.
   - Sifa ya `id` yenye aina ya `string`.
   - Sifa ya `price` yenye aina ya `number`.

1. Unapaswa kuunda kiolesura kinachoitwa `Book` kinachopanua `Item`. Sifa yake ya `type` inapaswa kuwekwa kuwa `"book"` halisi. Pia inapaswa kuwa na sifa za `title` na `author` zenye aina ya `string`.

1. Unapaswa kuunda kiolesura kinachoitwa `Electronics` kinachopanua `Item`. Sifa yake ya `type` inapaswa kuwekwa kuwa `"electronics"` halisi. Pia inapaswa kuwa na sifa za `item` na `model` zenye aina ya `string`, na sifa ya hiari ya `warranty` yenye aina ya `number`.

1. Unapaswa kuunda kiolesura kinachoitwa `Clothing` kinachopanua `Item`. Sifa yake ya `type` inapaswa kuwekwa kuwa `"clothing"` halisi. Pia inapaswa kuwa na sifa za `item` na `brand` zenye aina ya `string`, na sifa ya hiari ya `size` ambayo inaweza kuwa `"S"`, `"M"`, au `"L"` tu.

1. Unapaswa kuunda aina ya muungano inayoitwa `Product` inayochanganya kiolesura za `Book`, `Electronics`, na `Clothing`.

1. Unapaswa kuunda darasa la jumla linaitwa `Collection` lenye kigezo cha aina `T`. Kijenzi kinapaswa kupokea safu ya vitu vya aina `T` na kuzihifadhi katika sifa inayoitwa `items`.

1. Darasa la `Collection` linapaswa kuwa na njia ya `getAll()` inayorejesha vitu vyote katika mkusanyiko.

1. Darasa la `Collection` linapaswa kuwa na njia ya `filter()` inayokubali kitendakazi cha mwitiko na kurejesha tu vipengele kutoka kwa safu ya `items` vinavyokidhi sharti lililotolewa. Kitendakazi cha mwitiko kinapaswa kupokea kipengele cha aina `T` na kurudisha thamani ya kweli au siyo kweli.

1. Unapaswa kuunda kitendakazi kinachoitwa `renderProduct` chenye kigezo cha aina `Product`. Kinapaswa kurudisha mfuatano wa herufi wa HTML unaowakilisha bidhaa yenye:
   - Kipengele chenye darasa la `item` na kitambulisho chenye thamani ya `id` ya bidhaa.
   - Kipengele chenye darasa la `price` na maandishi ya `price` ya bidhaa.
   - Kipengele chenye taarifa za ziada za bidhaa.

1. Kitendakazi cha `renderProduct` kinapaswa kutumia upunguzaji wa aina kuonyesha HTML tofauti kulingana na aina ya bidhaa:

   - Kwa vitabu, HTML inayotolewa inapaswa kuwa na muundo wa `Book: [title] by [author]`, ambapo `[title]` na `[author]` zinapaswa kubadilishwa na kichwa na mwandishi, mtawalia.
   - Kwa vifaa vya elektroniki, HTML inayotolewa inapaswa kuwa na muundo wa `Electronics: [item] - [model]`. Ikiwa bidhaa ina dhamana, ongeza ` - Warranty: [warranty] year(s)` kwenye muundo. Badilisha `[item]`, `[model]`, na `[warranty]` na kipengele, mfano, na dhamana mtawalia.
   - Kwa mavazi, HTML inayotolewa inapaswa kuwa na muundo wa `Clothing: [item] by [brand]`. Ikiwa bidhaa ina ukubwa, ongeza ` - Size [size]` kwenye muundo. Badilisha `[item]`, `[brand]`, na `[size]` na kipengele, chapa, na ukubwa mtawalia.
   - Ikiwa aina ya bidhaa si halali, tupa kosa lenye muundo: `Unknown product type: [product]`, ambapo `[product]` inapaswa kubadilishwa na bidhaa isiyo halali iliyobadilishwa kuwa mfuatano wa herufi kwa kutumia `JSON.stringify`.

1. Unapaswa kuunda mfano wa `Collection<Product>` unaojumuisha angalau kipengele kimoja kwa kila aina ya bidhaa na kuuweka katika kigezo kinachoitwa `products`.

1. Unapaswa kuunda kitendakazi cha `showProducts` kinachoonyesha bidhaa katika kipengele cha `#output`. Kinapaswa kupokea kichujio cha hiari kinacho wakilisha aina ya bidhaa; ukiachwa wazi, onyesha bidhaa zote; ukitolewa, onyesha bidhaa za aina hiyo tu.

1. Unapaswa kuonyesha bidhaa kwa kubadilisha kila bidhaa kuwa mfuatano wa herufi wa HTML kwa kutumia kitendakazi cha `renderProduct` na kuingiza matokeo yaliyojumuishwa katika kipengele cha `#output`.

1. Kubofya vifungo vya `#all`, `#books`, `#electronics`, na `#clothing` kunapaswa kuonyesha seti husika ya bidhaa katika kipengele cha `#output`.

1. Wakati ukurasa unapopakia, bidhaa zote zinapaswa kuonyeshwa kwa chaguo-msingi.

# --hints--

Unapaswa kuwa na kiolesura kinachoitwa `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Item);
```

Kiolesura chako cha `Item` kinapaswa kuwa na sifa ya `type` ambayo inaweza kuwa `"book"`, `"electronics"`, au `"clothing"` tu.

```js
const explorer = await __helpers.Explorer(code);
const { Item } = explorer.interfaces;
const typeProp = Item.typeProps.type;
assert.isTrue(typeProp.isUnionOf(["'book'", "'electronics'", "'clothing'"]));
```

Kiolesura chako cha `Item` kinapaswa kuwa na sifa ya `id` yenye aina ya `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "id", type: "string"}));
```

Kiolesura chako cha `Item` kinapaswa kuwa na sifa ya `price` yenye aina ya `number`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "price", type: "number"}));
```

Unapaswa kuwa na kiolesura kinachoitwa `Book`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Book);
```

Kiolesura chako cha `Book` kinapaswa kupanua `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.doesExtend("Item"));
```

Kiolesura chako cha `Book` kinapaswa kuwa na sifa ya `type` iliyowekwa kuwa `"book"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "type", type: "'book'"}));
```

Kiolesura chako cha `Book` kinapaswa kuwa na sifa ya `title` yenye aina ya `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "title", type: "string"}));
```

Kiolesura chako cha `Book` kinapaswa kuwa na sifa ya `author` yenye aina ya `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "author", type: "string"}));
```

Unapaswa kuwa na kiolesura kinachoitwa `Electronics`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Electronics);
```

Kiolesura chako cha `Electronics` kinapaswa kupanua `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.doesExtend("Item"));
```

Kiolesura chako cha `Electronics` kinapaswa kuwa na sifa ya `type` iliyowekwa kuwa `"electronics"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "type", type: "'electronics'"}));
```

Kiolesura chako cha `Electronics` kinapaswa kuwa na sifa ya `item` yenye aina ya `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "item", type: "string"}));
```

Kiolesura chako cha `Electronics` kinapaswa kuwa na sifa ya `model` yenye aina ya `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "model", type: "string"}));
```

Kiolesura chako cha `Electronics` kinapaswa kuwa na sifa ya hiari ya `warranty` yenye aina ya `number`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "warranty", type: "number", isOptional: true}));
```

Unapaswa kuwa na kiolesura kinachoitwa `Clothing`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Clothing);
```

Kiolesura chako cha `Clothing` kinapaswa kupanua `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.doesExtend("Item"));
```

Kiolesura chako cha `Clothing` kinapaswa kuwa na sifa ya `type` iliyowekwa kuwa `"clothing"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "type", type: "'clothing'"}));
```

Kiolesura chako cha `Clothing` kinapaswa kuwa na sifa ya `item` yenye aina ya `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "item", type: "string"}));
```

Kiolesura chako cha `Clothing` kinapaswa kuwa na sifa ya `brand` yenye aina ya `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "brand", type: "string"}));
```

Kiolesura chako cha `Clothing` kinapaswa kuwa na sifa ya hiari ya `size` ambayo inaweza kuwa `"S"`, `"M"`, au `"L"` tu.

```js
const explorer = await __helpers.Explorer(code);
const { Clothing } = explorer.interfaces;
assert.isTrue(Clothing.hasTypeProps({name: "size", isOptional: true}));
assert.isTrue(Clothing.typeProps.size.isUnionOf(["'S'", "'M'", "'L'"]));
```

Unapaswa kuunda aina ya muungano inayoitwa `Product` inayochanganya kiolesura za `Book`, `Electronics`, na `Clothing`.

```js
const explorer = await __helpers.Explorer(code);
const { Product } = explorer.types;
assert.isTrue(Product.isUnionOf(["Book", "Electronics", "Clothing"]));
```

Unapaswa kuwa na darasa la `Collection` lenye kigezo cha aina `T`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.isTrue(Collection.typeParameters[0].matches("T"));
```

Kijenzi cha darasa lako la `Collection` kinapaswa kupokea safu ya vitu vya aina `T` na kuzihifadhi katika sifa inayoitwa `items`.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.items, arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.items, arr);
```

Darasa lako la `Collection` linapaswa kuwa na njia ya `getAll`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
console.log(Collection.methods)
console.log(Collection.methods.getAll)
assert.exists(Collection.methods.getAll);
```

Njia yako ya `getAll` inapaswa kurejesha vitu vyote katika mkusanyiko.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.getAll(), arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.getAll(), arr);
```

Darasa la `Collection` linapaswa kuwa na njia ya `filter`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.exists(Collection.methods.filter);
```

Njia yako ya `filter` inapaswa kupokea kitendakazi cha mwitiko na kurejesha safu inayojumuisha vipengele vya sifa ya `items` vinavyokidhi sharti lililotolewa.

```js
let arr = [1, 2, 3, 4, 5, 6];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.filter(n => n % 2 === 0), [2, 4, 6]);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.filter(s => s === "c"), ["c"]);
```

Unapaswa kuwa na kitendakazi kinachoitwa `renderProduct`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.renderProduct);
```

`renderProduct(product)` inapaswa kurejesha mfuatano wa herufi unaojumuisha kipengele chenye darasa la `item` na kitambulisho sawa na `product.id`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  assert.isNotNull(testEl.querySelector(`[class="item"][id="${p.id}"]`));
})
```

`renderProduct(product)` inapaswa kurejesha mfuatano wa herufi unaojumuisha kipengele chenye darasa la `price` na maandishi sawa na `product.price`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const price = testEl.querySelector(`[class="price"]`);
  assert.include(price.innerText, p.price.toString());
})
```

Wakati `renderProduct` inapoitwa na bidhaa ya `Book` `p`, inapaswa kurejesha mfuatano wa herufi unaojumuisha kipengele chenye maandishi yenye muundo wa `Book: [title] by [author]`. Badilisha `[title]` na `[author]` na sifa zinazolingana za `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const expectedText = `Book: ${p.title} by ${p.author}`;
  assert.include(testEl.innerText, expectedText);
})
```

Wakati `renderProduct` inapoitwa na bidhaa ya `Electronics` `p`, inapaswa kurejesha mfuatano wa herufi unaojumuisha kipengele chenye maandishi yenye muundo wa `Electronics: [item] - [model]` au `Electronics: [item] - [model] - Warranty: [warranty] year(s)` ikiwa `p.warranty` haiko wazi. Badilisha `[item]`, `[model]`, na `[warranty]` na sifa zinazolingana za `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Electronics: ${p.item} - ${p.model}`;
  const expectedText = !p.warranty ? baseFormat : baseFormat + ` - Warranty: ${p.warranty} year(s)`;
  assert.include(testEl.innerText, expectedText);
  if (!p.warranty) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

Wakati `renderProduct` inapoitwa na bidhaa ya `Clothing` `p`, inapaswa kurejesha mfuatano wa herufi unaojumuisha kipengele chenye maandishi yenye muundo wa `Clothing: [item] by [brand]` au `Clothing: [item] by [brand] - Size [size]` ikiwa `p.size` haiko wazi. Badilisha `[item]`, `[brand]`, na `[size]` na sifa zinazolingana za `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Clothing: ${p.item} by ${p.brand}`;
  const expectedText = !p.size ? baseFormat : baseFormat + ` - Size ${p.size}`;
  assert.include(testEl.innerText, expectedText);
  if (!p.size) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

Wakati `renderProduct` inapoitwa na bidhaa isiyo halali, inapaswa kutupa kosa lenye muundo wa `Unknown product type: [p]`. Badilisha `[p]` na bidhaa isiyo halali iliyobadilishwa kuwa mfuatano wa herufi kwa kutumia `JSON.stringify`.

```js
const invalidProduct = { item: "T-Shirt", price: 19.99 };
const expected = `Unknown product type: ${JSON.stringify(invalidProduct)}`;
try {
  renderProduct(invalidProduct);
  assert.fail();
} catch({name, message}) {
  if (name == "AssertionError") {
    assert.fail("renderProduct should throw with invalid product");
  } else {
    assert.equal(message, expected);
  }
}
```

Unapaswa kuunda mfano wa `Collection<Product>` na kuuweka katika kigezo kinachoitwa `products`.

```js
assert.instanceOf(products, Collection);
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.products.value.typeArguments[0].matches("Product"));
```

Kigezo chako cha `products` kinapaswa kuwa na angalau kipengele kimoja kwa kila aina ya bidhaa.

```js
assert.isTrue(products.items.some(p => p.type === "book"));
assert.isTrue(products.items.some(p => p.type === "electronics"));
assert.isTrue(products.items.some(p => p.type === "clothing"));
```

Unapaswa kuwa na kitendakazi kinachoitwa `showProducts`.

```js
assert.isFunction(showProducts);
```

Kubofya kitufe cha `#all` kunapaswa kuonyesha bidhaa zote katika kipengele cha `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
books.dispatchEvent(new Event("click"));
all.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

Kubofya kitufe cha `#books` kunapaswa kuonyesha seti husika ya bidhaa katika kipengele cha `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "book") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
books.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Kubofya kitufe cha `#electronics` kunapaswa kuonyesha seti husika ya bidhaa katika kipengele cha `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "electronics") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
electronics.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Kubofya kitufe cha `#clothing` kunapaswa kuonyesha seti husika ya bidhaa katika kipengele cha `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "clothing") {
    count++;
  }
});
books.dispatchEvent(new Event("click"));
clothing.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Bidhaa zote zinapaswa kuonekana kwa chaguo-msingi wakati ukurasa unapopakia. Tumia tukio la `DOMContentLoaded` kuhakikisha hilo.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
document.dispatchEvent(new Event("DOMContentLoaded"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts
interface Item {
  type: "book" | "electronics" | "clothing";
  id: string;
  price: number;
}

interface Book extends Item {
  type: "book";
  title: string;
  author: string;
}

interface Electronics extends Item {
  type: "electronics";
  item: string;
  model: string;
  warranty?: number;
}

interface Clothing extends Item {
  type: "clothing";
  item: string;
  brand: string;
  size?: "S" | "M" | "L";
}

type Product = Book | Electronics | Clothing;

class Collection<T> {
  items: T[];
  constructor(items: T[]) {
    this.items = items;
  }

  getAll(): T[] {
    return this.items;
  }

  filter(callback: (item: T) => boolean): T[] {
    return this.items.filter(callback);
  }
}

function renderProduct(p: Product): string {
  const createProductCard = (id: string, content: string, price: number): string => {
    return `
     <div class="item" id="${id}">
       ${content}
       <div class="price">$${price}</div>
     </div>
   `;
  };

  if (p.type === "book") {
    const content = `<strong>Book:</strong> ${p.title} by ${p.author}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "electronics") {
    const warranty = p.warranty ? ` - Warranty: ${p.warranty} year(s)` : "";
    const content = `<strong>Electronics:</strong> ${p.item} - ${p.model}${warranty}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "clothing") {
    const size = p.size ? ` - Size ${p.size}` : "";
    const content = `<strong>Clothing:</strong> ${p.item} by ${p.brand}${size}`;
    return createProductCard(p.id, content, p.price);
  }

  const _never: never = p;
  throw new Error(`Unknown product type: ${JSON.stringify(p)}`);
}


const products = new Collection<Product>([
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },

  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "b2", type: "book", title: "1984", author: "George Orwell", price: 9.99 },

  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 },

  { id: "b3", type: "book", title: "Brave New World", author: "Aldous Huxley", price: 11.99 },
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "e5", type: "electronics", item: "Smartwatch", model: "Chronex Pulse-Q", warranty: 1, price: 199.99 },

  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 },
]);

const output = document.getElementById("output");

// Track currently active filter category (or undefined for all)
let currentFilter: Product["type"] | undefined = undefined;

function showProducts(filter?: Product["type"]): void {
  if (!output) return;

  let itemsToShow: Product[];

  // Toggle logic
  if (currentFilter === filter || filter === undefined) {
    // Reset to all if same button clicked again or filter is undefined
    itemsToShow = products.getAll();
    currentFilter = undefined;
    setActiveButton(undefined);
  } else {
    itemsToShow = products.filter((p: Product): boolean => p.type === filter);
    currentFilter = filter;
    setActiveButton(filter);
  }

  output.innerHTML = itemsToShow.map(renderProduct).join("");
}

// Active Button Highlighting
function setActiveButton(filter: Product["type"] | undefined): void {
  const buttons: { [key: string]: HTMLElement | null } = {
    all: document.getElementById("all"),
    book: document.getElementById("books"),
    electronics: document.getElementById("electronics"),
    clothing: document.getElementById("clothing")
  };

  for (const key in buttons) {
    buttons[key]?.classList.remove("active");
  }

  if (filter === undefined) {
    buttons["all"]?.classList.add("active");
  } else {
    buttons[filter]?.classList.add("active");
  }
}

// Button Event Handlers
document.getElementById("all")?.addEventListener("click", (): void => showProducts());
document.getElementById("books")?.addEventListener("click", (): void => showProducts("book"));
document.getElementById("electronics")?.addEventListener("click", (): void => showProducts("electronics"));
document.getElementById("clothing")?.addEventListener("click", (): void => showProducts("clothing"));

// Initial Render
document.addEventListener("DOMContentLoaded", (): void => showProducts());
```
