---
id: 66da326c02141df538f29ba5
title: Crie uma página de check-out
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**Objetivo:** atender às histórias de usuário abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve ter um elemento `h1` com o texto `Checkout`.
1. Você deve ter dois elementos `section` imediatamente após o elemento `h1`.
1. Você deve ter um elemento `h2` com o texto `Your Cart` dentro da primeira seção.
1. Você deve ter uma imagem de um item na primeira seção com texto alternativo apropriado. Você pode usar esta imagem se quiser: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. Você deve ter um elemento `h2` com o texto `Payment Information` dentro da segunda seção.
1. Você deve ter um elemento `form` dentro da segunda seção.
1. Você deve ter um input com um `id` e `name` de `card-name` dentro do form e um `label` associado a ele.
1. Você deve ter um input com um `id` e `name` de `card-number` dentro do form e um `label` associado a ele.
1. Pelo menos dois dos elementos input devem ser required e ter um atributo `aria-required` definido como `true`.

# --hints--

Você deve ter um elemento `h1` com o texto `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

Você deve ter apenas um elemento `h1` na página.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Você deve ter pelo menos dois elementos `section` após o elemento `h1`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

Você deve ter um elemento `h2` com o texto `Your Cart` dentro da primeira seção.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

Você deve ter pelo menos uma imagem com texto alternativo na primeira seção.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

Você deve ter um elemento `h2` com o texto `Payment Information` dentro da segunda seção.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

Você deve ter um elemento `form` dentro do segundo elemento `section`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

Você deve ter um elemento `input` com um `id` e `name` de `card-name` dentro do form.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

Você deve ter um elemento `input` com um `id` e `name` de `card-number` dentro do form.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

Todos os elementos `input` que não tenham `type` `submit` devem ter um elemento `label` associado a eles.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => assert.exists(document.querySelector(`label[for="${input.id}"]`)));
```

Pelo menos dois dos elementos `input` devem ser required e ter um atributo `aria-required` definido como `true`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required && input.getAttribute('aria-required') === 'true');
assert.isAtLeast(requiredInputs?.length, 2)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name</label>
                    <input type="text" id="card-name" name="card-name" required aria-required="true">
                </div>
                <div>
                    <label for="card-number">Card Number</label>
                    <input type="text" id="card-number" name="card-number" required aria-required="true">
                </div>
                <div>
                    <label for="expiry-date">Expiry Date</label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required aria-required="true">
                </div>
                <div>
                    <label for="cvv">CVV</label>
                    <input type="text" id="cvv" name="cvv" required aria-required="true" aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
