---
id: 66da326c02141df538f29ba5
title: Crie uma página de check-out
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um elemento `h1` com o texto `Checkout`.
1. Você deve ter dois elementos `section` imediatamente após o elemento `h1`.
1. Você deve ter um elemento `h2` com o texto `Your Cart` dentro da primeira seção.
1. Você deve ter uma imagem de um item na primeira seção com texto alternativo apropriado. Você pode usar esta imagem se desejar: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. Você deve ter um elemento `h2` com o texto `Payment Information` dentro da segunda seção.
1. Você deve ter um elemento `form` dentro da segunda seção.
1. Você deve ter um input com um `id` e `name` de `card-name` e um `type` de `text` dentro do seu formulário e um `label` associado a ele.
1. Você deve ter um input com um `id` e `name` de `card-number`, um `type` de `text` dentro do seu formulário e um `label` associado a ele.
1. Você deve ter pelo menos dois elementos `input` com o atributo `required`.
1. Você deve incluir um elemento `span` com o texto `*` e `aria-hidden` definido como `true` dentro do elemento `label` para cada entrada obrigatória, para que os campos obrigatórios sejam indicados visualmente.
1. Você deve ter um elemento `p` com um texto de ajuda que explica o formato obrigatório do número do cartão, colocado imediatamente após o campo de entrada do número do cartão. O `p` deve ter um `id` de `card-number-help` e ser referenciado pelo campo de entrada do número do cartão usando `aria-describedby`.

# --hints--

Você deve ter um elemento `h1` com o texto `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

Você deve ter apenas um elemento `h1` na página.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Você deve ter pelo menos dois elementos `section` após o elemento `h1`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

Você deve ter um elemento `h2` com o texto `Your Cart` dentro da primeira seção.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

Você deve ter pelo menos uma imagem com texto alternativo na primeira seção.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

Você deve ter um elemento `h2` com o texto `Payment Information` dentro da segunda seção.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

Você deve ter um elemento `form` dentro do segundo elemento `section`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

Você deve ter um elemento `input` com um `id` e `name` de `card-name` dentro do form.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

O campo de nome do seu cartão deve ter um `type` igual a `text`.

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

Você deve ter um elemento `input` com um `id` e `name` de `card-number` dentro do form.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

A entrada do número do seu cartão deve ter um `type` de `text`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

Todos os elementos `input` que não tenham `type` `submit` devem ter um elemento `label` associado a eles.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => assert.exists(document.querySelector(`label[for="${input.id}"]`)));
```

Você deve ter pelo menos dois elementos `input` com o atributo `required`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

Você deve incluir um elemento `span` com o texto `*` dentro do elemento `label` para cada entrada obrigatória.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`);
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

Seus elementos `span` devem ter `aria-hidden` definido como `true`.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`);
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

Você deve ter um elemento `p` com um `id` de `card-number-help` imediatamente após o campo de entrada do número do cartão.

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
assert.equal(cardNumberInput.nextElementSibling, cardNumberHelp);
```

O texto de ajuda do seu número de cartão não deve estar vazio.

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

A entrada do número do seu cartão deve ter `aria-describedby` definido como `card-number-help`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
