---
id: 6a89c69cbe1bb69845992d8b
title: 构建数字范围生成器
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

在本实验中，你将构建一个使用递归的 `range_of_numbers` 函数，用于生成指定范围内的数字列表。

例如：

- `range_of_numbers(3, 9)` 应该返回 `[3, 4, 5, 6, 7, 8, 9]`。
- `range_of_numbers(5, 5)` 应该返回 `[5]`。

完成下面的用户故事并通过所有测试以完成本实验。

**用户故事：**

1. 你应该创建一个名为 `range_of_numbers` 的函数，带有两个名为 `start_num` 和 `end_num` 的参数。
1. 该函数应返回一个连续整数的列表，起始于 `start_num`，结束于 `end_num`，包括这两个值。
1. `start_num` 的值总是小于或等于 `end_num`。
1. 你的函数应通过调用自身来使用递归。
1. 你的函数不应使用 `for` 或 `while` 循环、推导式或 `range()` 函数。
1. 当 `start_num` 等于 `end_num` 时，函数应返回包含 `start_num` 的列表。这是基本情形。
1. 对于递归情形，函数应调用自身，传入一个向基本情形靠近的参数，然后将当前数字添加到返回的列表中。
1. 函数不应使用全局变量来保存返回的列表。

# --hints--

你应该有一个名为 `range_of_numbers` 的函数，带有名为 `start_num` 和 `end_num` 的参数。

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` 应该返回 `[1, 2, 3, 4, 5]`。

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` 应该返回 `[6, 7, 8, 9]`。

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` 应该返回 `[4]`。

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` 应该返回 `[10, 11, 12, 13, 14, 15]`。

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

你的 `range_of_numbers` 函数应该调用自身。

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

你的 `range_of_numbers` 函数不应使用循环、推导式或 `range()` 函数。

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

调用 `range_of_numbers(1, 3)` 不应影响后续调用。之后调用 `range_of_numbers(6, 9)` 仍应返回 `[6, 7, 8, 9]`。

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
