---
id: 6626a060c4006f793e10cb33
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

الآن يحتاج المعلم إلى مساعدتك في تحويل درجة الطالب إلى تقدير حرفي.

أكمل دالة `getGrade` التي تأخذ رقمًا `score` كمعلمة. يجب أن تُرجع دالتك سلسلة نصية تمثل التقدير الحرفي بناءً على الدرجة.

إليك الدرجات والتقديرات الحرفية المقابلة لها:

| نطاق الدرجة | التقدير   |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**نصائح**

- تذكر أنك تعلمت عن عبارات الشرط (`if`، `else if`، و`else`).
- تذكر أنك تعلمت عن عوامل المقارنة (`>`، `<`، `>=`، `<=`، `===`).
- تذكر أنك تعلمت عن عامل AND المنطقي (`&&`).

# --hints--

يجب أن تُرجع دالة `getGrade` القيمة `"A++"` إذا كانت الدرجة `100`.

```js
assert.strictEqual(getGrade(100), "A++");
```

يجب أن تُرجع دالة `getGrade` القيمة `"A"` إذا كانت الدرجة `94`.

```js
assert.strictEqual(getGrade(94), "A");
```

يجب أن تُرجع دالة `getGrade` القيمة `"B"` إذا كانت الدرجة بين `80` و`89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

يجب أن تُرجع دالة `getGrade` القيمة `"C"` إذا كانت الدرجة `78`.

```js
assert.strictEqual(getGrade(75), "C");
```

يجب أن تُرجع دالة `getGrade` القيمة `"D"` إذا كانت الدرجة بين `60` و`69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

يجب أن تُرجع دالة `getGrade` القيمة `"F"` إذا كانت الدرجة `57`.

```js
assert.strictEqual(getGrade(57), "F");
```

يجب أن تُرجع دالة `getGrade` القيمة `"F"` إذا كانت الدرجة بين `0` و`59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
