---
id: 662693f82c91a66be46c881b
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

أنهى المعلم تصحيح اختبارات الطلاب ويحتاج إلى مساعدتك لحساب متوسط الدرجة للفصل.

أكمل دالة `getAverage` التي تستقبل مصفوفة من درجات الاختبار وتُرجع متوسط الدرجة.

يُحسب المتوسط بجمع كل الدرجات وقسمتها على العدد الإجمالي للدرجات.

```md
average = sum of all scores / total number of scores
```

تم توفير بعض استدعاءات الدالة لتتمكن من اختبار شفرتك.

**نصائح**

- يمكنك استخدام حلقة للتكرار على مصفوفة `scores` وجمع كل الدرجات.
- يمكنك استخدام خاصية `length` للحصول على العدد الإجمالي للدرجات.

# --hints--

يجب أن تُرجع دالة `getAverage` رقمًا.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

يجب أن تُرجع `getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` القيمة `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

يجب أن تُرجع `getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` القيمة `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

يجب أن تُرجع `getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` القيمة `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

يجب أن تُرجع دالة `getAverage` متوسط درجات الاختبار.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
