---
id: 6626b8dcf5057f896f948440
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

بما أن المعلم لديه كل المعلومات التي يحتاجها، يريد أن يرسل رسالة إلى الطالب بنتائج الدورة.

أكمل دالة `studentMsg` باستخدام `totalScores` و`studentScore` كمعلمات. يجب أن تُرجع الدالة سلسلة نصية تمثل رسالة للطالب.

إذا نجح الطالب في الدورة، يجب أن تتبع السلسلة النصية التنسيق التالي:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

إذا رسب الطالب في الدورة، يجب أن تتبع السلسلة النصية التنسيق التالي:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

استبدل `average-goes-here` بمتوسط الدرجات الكلي.
استبدل `grade-goes-here` بدرجة الطالب.

**نصائح**

- استخدم الدالة `getAverage` للحصول على متوسط الصف.
- استخدم الدالة `getGrade` للحصول على درجة الطالب.
- استخدم ربط السلاسل النصية (`+`) لبناء الرسالة.
- انتبه إلى علامات الترقيم والمسافات في الرسالة.

# --hints--

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` يجب أن تُرجع الرسالة التالية: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

يجب أن تُرجع الدالة `studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` الرسالة التالية: `"Class average: 50.8. Your grade: A++. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

يجب أن تُرجع دالة `studentMsg` الرسالة الصحيحة بناءً على درجة الطالب ومتوسط الصف.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
