---
id: 6626b4c58c027d86478ff5eb
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

المعلم سعيد جدًا بالبرنامج الذي أنشأته حتى الآن. لكنه يريد الآن طريقة سهلة للتحقق مما إذا كان الطالب لديه درجة ناجحة. الدرجة الناجحة هي أي شيء ليس `"F"`.

أكمل الدالة `hasPassingGrade` التي تأخذ درجة الطالب كمعلمة. يجب أن تُرجع دالتك `true` إذا كان لدى الطالب درجة ناجحة و`false` إذا لم تكن كذلك.

**نصائح**

- استخدم الدالة `getGrade` للحصول على درجة الطالب. ثم تحقق مما إذا كانت الدرجة ناجحة أم لا.

# --hints--

يجب أن تُرجع دالتك `hasPassingGrade` قيمة منطقية.

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

يجب أن تُرجع دالتك `hasPassingGrade` القيمة `true` إذا كانت الدرجة `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

يجب أن تُرجع دالتك `hasPassingGrade` القيمة `false` إذا كانت الدرجة `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

يجب أن تُرجع دالتك `hasPassingGrade` القيمة `true` لجميع الدرجات الناجحة.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
