---
id: 66f26c32ec6f90df01a44f60
title: Diseñar un Formulario de Contacto
challengeType: 25
dashedName: design-a-contact-form
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio. 

**Historias de usuario:**

1. Debes tener un `div` con una clase de `form-container` para contener tu formulario.

2. Debes tener un elemento `form` dentro del elemento `.form-container`.

3. Within the `form` element, you should have the following elements and input fields:

   - Un elemento `h2` con algo de texto.
   - Un campo de entrada de texto para el nombre con el `type` establecido en `text` y los atributos `id`, `name`, `required`.
   - Un campo de entrada de correo electrónico con el `type` establecido en `email` y los atributos `id`, `name`, `required`.
   - Un `textarea` para el mensaje con los atributos `id`, `name` y `required`.
   - Un elemento `button` con el texto `"Submit"` y un atributo `type` de `submit`.

4. Cada elemento `input` y `textarea` debe tener su correspondiente elemento `label` con un atributo `for` configurado al `id` del elemento.

5. El elemento `.form-container` debe tener un color de fondo, y tener valores para `border-radius`, `padding` y `width`.

6. Los elementos `label` deben tener un margen y un color de fuente.

7. The `input` and `textarea` elements should have values for:
   
   - `width`.
   - `padding`.
   - Margen en la parte inferior.

8. The `button` element should have a:
   
   - Color de fondo
   - Tamaño de fuente diferente al predeterminado.

9.  El botón debe tener un efecto hover que cambie el color de fondo.

**Nota:** Asegúrate de vincular tu hoja de estilos en tu HTML y aplica tu CSS.

# --hints--

Debes tener un `div` con una clase de `form-container`.

```js
const el = document.querySelector('div.form-container');
assert.isNotNull(el);
```

Debes tener un elemento `form` dentro del elemento `.form-container`.

```js
const el = document.querySelector('.form-container form');
assert.isNotNull(el);
```

Dentro del `form`, debes tener un elemento `h2` con algo de texto.

```js
const el = document.querySelector('.form-container form h2');
assert.isNotEmpty(el.innerText);
// assert.isAtLeast(document.querySelector('h1').innerText.length, 1);

```

Dentro del `form`, debes tener un campo de entrada de texto con el `type` establecido en `text`.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.isNotNull(el);
```

Tu campo de entrada `type=text` debe tener un `id`, un `name` y un atributo `required`.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.exists(el);
assert.isTrue(el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

Dentro del `form`, debes tener un campo de entrada de correo electrónico con el `type` establecido en `email`.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isNotNull(el);
```

Tu campo de entrada `type=email` debe tener un `id`, un `name` y un atributo `required`.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isTrue(el?.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

Dentro del `form`, debes tener un `textarea` con un `id`, `name` y un atributo `required`.

```js
const el = document.querySelector('.form-container form textarea');
assert.isNotNull(el);
assert.isNotNull(document.querySelector('.form-container form textarea[required]'));
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

Dentro del `form`, debes tener un elemento `button` con un atributo `type` de `submit`.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(el?.getAttribute('type'), 'submit');
```

Tu botón debe tener un texto de `"Submit"`.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(document.querySelector('.form-container form button').textContent, 'Submit');
```

Debes tener un elemento `label` para el elemento `#name` con un atributo `for` que coincida con el valor `id` en el `input` correspondiente.

```js
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value
assert.isNotEmpty(nameInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === nameInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

Debes tener un elemento `label` para el elemento `#email` con un atributo `for` que coincida con el valor `id` en el `input` correspondiente.

```js
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);
// check if id has a value
assert.isNotEmpty(emailInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === emailInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

Debes tener un elemento `label` para el elemento `#message` con un atributo `for` que coincida con el valor `id` en el `textarea`.

```js
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value

assert.isNotEmpty(messageInput);

assert.exists(document.querySelector(`label[for="${messageInput}"]`));
```

Las etiquetas no deben estar vacías.

```js
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

for (let label of labels) {
    const labelText = label.textContent.trim();
    assert.isNotEmpty(labelText);
}

```

El elemento `.form-container` debe tener un color de fondo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

El elemento `.form-container` debe tener un `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

El elemento `.form-container` debe tener relleno.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

El elemento `.form-container` debe tener un ancho.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

Tus etiquetas deben tener un margen y color de fuente.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

Tus elementos `input` deben tener un ancho, relleno y un margen en la parte inferior.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('input'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);

```

Tu elemento `textarea` debe tener un ancho, relleno, margen en la parte inferior.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('textarea'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);
```

Tu elemento `button` debe tener un color de fondo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

Tu elemento `button` debe tener un tamaño de fuente.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

Tu elemento `button` debe tener un efecto hover que cambie el color de fondo.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const btn_hover = gs('button:hover')?.backgroundColor;
assert.isNotEmpty(btn_hover);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled Form Project</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="form-container">
        <form>
            <h2>Contact Us</h2>
            <label for="name">Name:</label>
            <input type="text" id="name" name="name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" required>

            <label for="message">Message:</label>
            <textarea id="message" name="message" required></textarea>

            <button type="submit">Submit</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background: linear-gradient(45deg, #83a4d4, #b6fbff);
    font-family: Arial, sans-serif;
}

.form-container {
    background-color: white;
    border-radius: 10px;
    padding: 20px;
    width: 300px;
    text-align: center;
}

form h2 {
    margin-bottom: 20px;
    color: #333;
}

label {
    display: block;
    margin: 10px 0 5px;
    color: #555;
}

input{
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

textarea {
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

button {
    background-color: #4CAF50;
    color: white;
    padding: 10px 15px;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    font-size: 1em;
}

button:hover {
    background-color: #45a049;
}
```
