---
id: 68ce56d54d3cad1149a1c4cf
title: تصميم صفحة اختيار ميزة
challengeType: 25
dashedName: design-a-feature-selection-page
demoType: onClick
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تحتوي صفحتك على عنصر `h1` بنص `Feature Selection`.
2. يجب أن تحتوي صفحتك على عنصر `div` بالصنف `feature-card-container`.
3. يجب أن تحتوي صفحتك على عنصرين `label` على الأقل، كل منهما بالصنف `feature-card` داخل الـ `div`.
4. يجب أن يحتوي كل عنصر `label` على نص مسمى وعنصر `input` من النوع `checkbox` بداخله.
5. أنشئ محددًا يستهدف مربعات الاختيار، وطبق التنسيق التالي:
    - يجب تعيين خاصية `appearance: none;` لجميع مربعات الاختيار في CSS.
    - يجب أن تحتوي جميع مربعات الاختيار على عرض حدود، ولون، ونمط تختاره.
6. عند تحديد مربع الاختيار، يجب أن يظهر علامة صح `✓` في وسط مربع الاختيار.
7. عند تحديد مربع الاختيار، يجب أن يتغير لون خلفية مربع الاختيار إلى لون تختاره.
8. عند تحديد مربع الاختيار، يجب أن يتغير لون حدود مربع الاختيار إلى لون تختاره.

**ملاحظة:** تأكد من ربط ملف CSS في ملف HTML الخاص بك.

# --hints--

يجب أن يكون لديك عنصر `h1`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يحتوي عنصر `h1` على نص `Feature Selection`.

```js
assert.equal(document.querySelector('h1').innerText.trim(), `Feature Selection`);
```

يجب أن يكون لديك عنصر `div` بالصنف `feature-card-container`.

```js
assert.exists(document.querySelector('div.feature-card-container'));
```

يجب أن يكون لديك على الأقل عنصران `label` داخل `.feature-card-container`.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isAtLeast(labels.length, 2);
```

يجب أن يحتوي كل عنصر `label` داخل `.feature-card-container` على الصنف `feature-card`.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isTrue(label.classList.contains('feature-card')));
```

يجب أن يحتوي كل عنصر `label` على نص مسمى.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isNotEmpty(label.innerText.trim()));
```

داخل كل عنصر مسمى يجب أن يكون هناك عنصر `input` من نوع `checkbox` واحد.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.exists(label.querySelector('input[type="checkbox"]')));
```

يجب أن يكون لديك محدد لمربعات الاختيار.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

يجب تعيين مربعات الاختيار إلى `appearance: none`.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.equal(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('appearance'), 'none');
```

يجب أن يكون لمربعات الاختيار عرض حدود تختاره.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-width'));
```

يجب أن يكون لمربعات الاختيار لون حدود تختاره.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-color'));
```

يجب أن يكون لمربعات الاختيار نمط حدود تختاره.

```js
const selectors = [  
  '.feature-card [type="checkbox"]',
  '.feature-card input[type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input', // there aren't other input elements, so this is a valid approach
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-style'));
```

عند تحديد مربعات الاختيار، يجب أن تظهر علامة صح `✓` في وسط مربع الاختيار.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;

    checkbox.checked = true;
    const after = window.getComputedStyle(checkbox, ':after');
    assert.equal(after.content, '"✓"');
});
```

عند تحديد مربع الاختيار، يجب أن يتغير لون خلفية مربع الاختيار إلى لون تختاره.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBG = window.getComputedStyle(checkbox).backgroundColor;
    checkbox.checked = true;
    const checkedBG = window.getComputedStyle(checkbox).backgroundColor;

    assert.notEqual(checkedBG, uncheckedBG);
});
```

عند تحديد مربع الاختيار، يجب أن يتغير لون حدود مربع الاختيار إلى لون تختاره.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBC = window.getComputedStyle(checkbox).borderColor;
    checkbox.checked = true;
    const checkedBC = window.getComputedStyle(checkbox).borderColor;

    assert.notEqual(checkedBC, uncheckedBC);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Selection Feature Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Feature Selection</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Feature Selection</h1>
    <div class="feature-card-container">
      <label class="feature-card">
        <input type="checkbox" />
        <h4>Cloud Storage</h4>
        <p>100 Gigabyte secure storage</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Dedicated Support</h4>
        <p>24/7 customer help</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Advanced Analytics</h4>
        <p>Insights & reports</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Custom User Themes</h4>
        <p>Personalized dashboard design</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Multi-User Collab</h4>
        <p>Team access and sharing</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>API Access</h4>
        <p>Integrate with your custom built tools</p>
      </label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  flex-direction: column;
  height: 100vh;
}

.feature-card-container {
  display: flex;
  flex-wrap: wrap;
  gap: 20px;
  max-width: 600px;
  justify-content: center;
}

.feature-card-container > * {
  flex: 1 1 180px;
}

.feature-card {
  border: 2px solid #ccc;
  border-radius: 12px;
  padding: 20px;
  text-align: center;
  position: relative;
  cursor: pointer;
}

.feature-card:hover {
  border-color: #f1be32;
  box-shadow: 0 4px 12px rgba(255, 201, 53, 0.6);
}

.feature-card input[type='checkbox'] {
  position: absolute;
  top: 5px;
  right: 5px;
  width: 24px;
  height: 24px;
  cursor: pointer;
  appearance: none;
  -webkit-appearance: none;
  border: 2px solid #ccc;
  border-radius: 4px;
  background-color: white;
}

.feature-card input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #f1be32;
}

.feature-card input[type='checkbox']:checked::after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  line-height: 20px;
}

.feature-card h4 {
  margin: 10px 0 5px;
  font-size: 1.2rem;
}

@media (max-width: 600px) {
  body {
    height: auto;
  }
}
```
