---
id: 656470d517833a39bb8b5608
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Das `random`-Modul enthält einen Pseudozufallszahlengenerator. Die meisten seiner Funktionen basieren auf der `random()`-Funktion, die eine Fließkommazahl im Bereich von `0.0` (einschließlich) bis `1.0` (ausschließlich) zurückgibt.

Rufen Sie die `random()`-Funktion aus dem `random`-Modul auf und geben Sie das Ergebnis aus.

# --hints--

Sie sollten das Ergebnis des Aufrufs von `random.random()` ausgeben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
