---
id: 656494269ccce754411a2112
title: Schritt 27
challengeType: 20
dashedName: step-27
---

# --description--

Das `re`-Modul ermöglicht es Ihnen, *reguläre Ausdrücke* in Ihrem Code zu verwenden. Sie werden sehr bald mehr über reguläre Ausdrücke lernen.

Fügen Sie vorerst eine `import`-Anweisung am Anfang Ihres Codes hinzu, um das `re`-Modul zu importieren.

# --hints--

Sie sollten das `re`-Modul importieren.

```js
({ test: () => assert.match(code, /^import\s+re/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

import secrets
import string
--fcc-editable-region--

def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
