---
id: 6564c67db26c417561ab510d
title: Schritt 31
challengeType: 20
dashedName: step-31
---

# --description--

Wie Sie an der Ausgabe sehen können, stimmt Ihr Regex jetzt mit dem ersten `l` im String überein.

In Ihrem Muster können Sie nach einem Zeichen einen Quantifizierer hinzufügen, um anzugeben, wie oft dieses Zeichen wiederholt werden soll. Zum Beispiel bedeutet der `+`-Quantifizierer, dass es ein- oder mehrmals wiederholt werden soll.

Fügen Sie Ihrem Muster einen `+`-Quantifizierer hinzu.

# --hints--

Sie sollten Ihre `pattern`-Variable in `re.compile('l+')` ändern.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
