---
id: 6564d97a5ef99783216a6229
title: Schritt 44
challengeType: 20
dashedName: step-44
---

# --description--

Python bietet einen besonderen Typ von String, den sogenannten *raw* String. Raw Strings werden mit einem `r` als Präfix versehen. Der entscheidende Unterschied zu regulären Strings liegt darin, wie sie das Backslash-Zeichen behandeln: In raw Strings werden Backslashes als buchstäbliche Zeichen und nicht als Escape-Zeichen interpretiert. Beim Schreiben von regulären Ausdrücken ist die Verwendung von raw Strings eine gute Praxis, da sie oft viele `\`-Zeichen enthalten.

Machen Sie aus Ihrem `pattern`-String einen raw String, indem Sie ihn mit einem `r` als Präfix versehen.

# --hints--

Sie sollten Ihre `pattern`-Variable in einen `r'\.'` ändern.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
