---
id: 65646ab6bced52360f44670f
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Deklarieren Sie nun eine Variable namens `all_characters` und weisen Sie ihr das Ergebnis der Verkettung Ihrer vorhandenen Variablen zu.

# --hints--

Sie sollten eine Variable namens `all_characters` deklarieren.

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

Sie sollten `letters`, `digits` und `symbols` verketten und das Ergebnis Ihrer `all_characters`-Variable zuweisen.

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
