---
id: 6564683821b2ee3174e7250e
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Sie können auf die im importierten Modul definierten Hilfsfunktionen über die Punktnotation zugreifen. Die Punktnotation funktioniert, indem Sie einen Punkt gefolgt vom Namen der Hilfsfunktion an den Modulnamen anhängen. Zum Beispiel so greifen Sie auf die Konstante `ascii_lowercase` zu:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

In diesem Projekt werden Sie verschiedene Konstanten aus dem `string`-Modul verwenden.

Deklarieren Sie eine neue Variable namens `letters` und weisen Sie dieser Variable `string.ascii_letters` zu.

# --hints--

Sie sollten eine Variable namens `letters` deklarieren.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

Sie sollten `string.ascii_letters` Ihrer `letters`-Variable zuweisen.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
