---
id: 656469e62377ab34a5bcb8d1
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Deklarieren Sie zwei neue Variablen mit den Namen `digits` und `symbols` und weisen Sie ihnen jeweils `string.digits` und `string.punctuation` zu.

# --hints--

Sie sollten eine Variable namens `digits` deklarieren.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

Sie sollten `string.digits` Ihrer Variable `digits` zuweisen.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

Sie sollten eine Variable namens `symbols` deklarieren.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

Sie sollten `string.punctuation` Ihrer Variable `symbols` zuweisen.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
