---
id: 656475bbf1c2573de1d2c69c
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Jedes Mal, wenn der Code ausgeführt wird, sollten Sie ein zufälliges Zeichen aus dem `all_characters` String sehen. Genau das wollen Sie erreichen, um ein zufälliges Passwort zu erstellen.

Der Algorithmus, auf dem `random` basiert, macht die erzeugten Pseudozufallszahlen jedoch vorhersehbar. Deshalb ist das `random` Modul zwar für die meisten Anwendungen geeignet, kann aber aufgrund seiner deterministischen Natur nicht für kryptografische Zwecke verwendet werden.

Importieren Sie statt `random` das `secrets` Modul. Ändern Sie dann den `print()` Aufruf so, dass `secrets.choice(all_characters)` verwendet wird.

# --hints--

Sie sollten das `secrets` Modul anstelle des `random` Moduls importieren.

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

Sie sollten Ihren bestehenden `print()` Aufruf so ändern, dass `secrets.choice(all_characters)` ausgegeben wird.

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
