---
id: 65649122c7f77f519aaf0975
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Deklarieren Sie schließlich eine Variable `new_password` und weisen Sie ihr das Ergebnis des Aufrufs von `generate_password` zu. Übergeben Sie `8` als Argument an Ihren `generate_password`-Aufruf.

# --hints--

Sie sollten `generate_password` aufrufen und `8` als Argument übergeben.

```js
({ test: () => assert.match(code, /generate_password\s*\(\s*8\s*\)/) })
```

Sie sollten `generate_password(8)` der Variable `new_password` zuweisen.

```js
({
  test: () => {
    assert(__userGlobals.has("new_password"));
    assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*\)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--

--fcc-editable-region--
```
