---
id: 6564cbb6311a40783b3f5de6
title: Schritt 36
challengeType: 20
dashedName: step-36
---

# --description--

Zeichenklassen erlauben es Ihnen auch, einen Bereich von Zeichen anzugeben, die übereinstimmen sollen. Sie müssen die Anfangs- und Endzeichen durch einen Bindestrich trennen, `-`. Die Zeichen folgen der Unicode-Reihenfolge.

Ändern Sie Ihre `pattern`-Variable so, dass sie jeden Buchstaben `t` erkennt, der in der `quote`-Variable von einem Kleinbuchstaben vorausgeht. Verwenden Sie dafür den Bereich von `a` bis `z`.

# --hints--

Sie sollten Ihre `pattern`-Variable so ändern, dass sie jeden Buchstaben `t` in `quote` erkennt, der von einem Kleinbuchstaben mit der `[a-z]`-Klasse vorausgeht.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
