---
id: 656474314bf37d3c83a85143
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Die `choice()`-Funktion aus dem `random`-Modul nimmt eine Folge und gibt ein zufälliges Element der Folge zurück.

Ändern Sie Ihren `print()`-Aufruf so, dass die `choice()`-Funktion verwendet wird und übergeben Sie `all_characters` als Argument.

# --hints--

Sie sollten Ihren bestehenden `print()`-Aufruf ändern, um `random.choice(all_characters)` auszugeben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
