---
id: 65646ffeaed2d238c562a014
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Es ist eine gängige Konvention, `import`-Anweisungen am Anfang Ihres Codes zu platzieren. Und zusätzlich, falls mehrere `import`-Anweisungen vorhanden sind, sortieren Sie diese alphabetisch, um die Lesbarkeit zu verbessern.

Importieren Sie am Anfang Ihres Codes das `random`-Modul.

# --hints--

Sie sollten das `random`-Modul importieren.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
