---
id: 656485a8a3496d4a36b1496a
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Diese drei Variablen stellen die möglichen Zeichen dar, aus denen das Passwort generiert wird.

Fügen Sie direkt darüber einen Kommentar mit dem Text `Define the possible characters for the password` ein.

# --hints--

Sie sollten den Kommentar direkt über der `letters`-Variablen hinzufügen.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
