---
id: 6564642ce403fe29cc23fee7
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Ein Python-Modul ist eine Datei, die Code enthält, der dazu bestimmt ist, bestimmte Aufgaben auszuführen. Die Python-Standardbibliothek enthält viele Module, die Sie importieren und in Ihrem Code verwenden können. Dies erreichen Sie, indem Sie die `import`-Anweisung gefolgt vom Namen des Moduls verwenden.

Beginnen Sie dieses Projekt, indem Sie das `string`-Modul importieren.

# --hints--

Sie sollten das `string`-Modul importieren.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
