---
id: 6564dd65c3c2fa873a83d213
title: Schritt 47
challengeType: 20
dashedName: step-47
---

# --description--

In einer Zeichenklasse können Sie mehrere Bereiche kombinieren, indem Sie einen Bereich nach dem anderen innerhalb der eckigen Klammern schreiben (ohne zusätzliche Zeichen). Zum Beispiel ist `[a-d3-6]` die Kombination aus `[a-d]` und `[3-6]`.

Ändern Sie nun das letzte Regex-Muster so, dass es auf jedes nicht alphanumerische Zeichen passt. Kombinieren Sie die `a-z`, `A-Z` und `0-9` Bereiche zu einer einzigen Zeichenklasse und fügen Sie `^` als erstes Zeichen hinzu, um das Muster zu negieren.

# --hints--

Sie sollten Ihr viertes Muster vervollständigen, damit es auf jedes nicht alphanumerische Zeichen passt.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\[\^a-zA-Z0-9\]\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
