---
id: 6564b8c9349bd76dc037967b
title: Schritt 29
challengeType: 20
dashedName: step-29
---

# --description--

Die `search()`-Funktion aus dem `re`-Modul analysiert den als Argument übergebenen String und sucht die erste Stelle, an der das Regex-Muster mit dem String übereinstimmt.

Deklarieren Sie eine Variable namens `quote` und weisen Sie dieser den String `'Not all those who wander are lost.'` zu. Geben Sie dann das Ergebnis von `pattern.search(quote)` aus.

# --hints--

Sie sollten eine `quote`-Variable haben.

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

Sie sollten den bereitgestellten String Ihrer neuen `quote`-Variable zuweisen.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

Sie sollten `pattern.search(quote)` ausgeben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
