---
id: 6564883669b5af4b69f794cc
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Ihre `all_characters`-Variable ist ein String, der aus allen Klein- und Großbuchstaben, allen 10 Ziffern und mehreren Sonderzeichen besteht.

Fügen Sie direkt davor einen Kommentar mit dem Text `Combine all characters` hinzu.

# --hints--

Sie sollten den Kommentar direkt über Ihrer `all_characters`-Variable hinzufügen.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
