---
id: 65649a631440b757f9f7b428
title: Schritt 28
challengeType: 20
dashedName: step-28
---

# --description--

Ein regulärer Ausdruck oder Regex ist ein Muster, das verwendet wird, um eine bestimmte Kombination von Zeichen innerhalb eines Strings zu finden. Er ist eine gültige Alternative zu `if`/`else` bedingten Anweisungen, wenn Sie Muster innerhalb eines Strings für Validierungszwecke, Zeichenersetzung und andere Anwendungen finden oder abgleichen müssen.

Die `compile()`-Funktion aus dem `re`-Modul kompiliert den als Argument übergebenen String in ein reguläres Ausdrucksobjekt, das von anderen `re`-Methoden verwendet werden kann.

Deklarieren Sie eine neue `pattern`-Variable und weisen Sie dieser Variable den Wert von `re.compile('i')` zu.

# --hints--

Sie sollten eine `pattern`-Variable deklarieren.

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

Sie sollten `re.compile('i')` Ihrer `pattern`-Variable zuweisen.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
