---
id: 656477845006313fbfea0ad1
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Obwohl die Wirkung gleich wie bei `random.choice()` erscheinen mag, garantiert `secrets` Ihnen die sicherste Quelle für Zufälligkeit, die Ihr Betriebssystem bereitstellen kann.

Löschen Sie nun Ihre beiden `print()`-Aufrufe.

# --hints--

Sie sollten Ihre beiden `print()`-Aufrufe löschen.

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
