---
id: 6565c13fdb798865c161d8f8
title: Schritt 65
challengeType: 20
dashedName: step-65
---

# --description--

Jetzt ist es Zeit, Ihre Funktion zu testen. Kommentieren Sie die letzten zwei Zeilen in Ihrem Code aus und ändern Sie den Funktionsaufruf, sodass Sie 5 Argumente übergeben. Verwenden Sie `8` für die Länge und `1` für die anderen vier Einschränkungen.

# --hints--

Sie sollten `generate_password` mit den angegebenen Argumenten aufrufen.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*,\s*1\s*,\s*1\s*,\s*1\s*,\s*1\s*\)/m) })
```

Sie sollten Ihre `new_password`-Variable ausgeben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*new_password\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
        
        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break

    return password
    
--fcc-editable-region--
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
```
