---
id: 6565c084627071646f94c4b0
title: Schritt 60
challengeType: 20
dashedName: step-60
---

# --description--

Erstellen Sie schließlich nach der `for`-Schleife eine `if`-Anweisung, um zu prüfen, ob `count` gleich `4` ist, und brechen Sie die `while`-Schleife mit der `break`-Anweisung ab.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, die überprüft, ob `count` gleich `4` nach der `for`-Schleife ist.

```js
({ test: () => assert.match(code, /^(\s*)for.+:\s*^\1(\s{4})if\s+constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*:\s*^\1\2\2count\s*\+=\s*1\s*^\1if\s+count\s*==\s*4\s*:/m) })
```

Sie sollten `break` innerhalb Ihrer neuen `if` verwenden, um die `while`-Schleife abzubrechen.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+count\s*==\s*4\s*/);
    assert(block_body.match(/^\s+break\s*$/m));  
  }
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
--fcc-editable-region--         
        # Check constraints
        count = 0
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
