---
id: 656ef54070c72161e6feb90f
title: Schritt 22
challengeType: 20
dashedName: step-22
---

# --description--

Es sieht alles gut aus, aber es wäre schön, eine Möglichkeit zu haben, zu überprüfen, ob das generierte Passwort bestimmten Funktionen entspricht. Zum Beispiel ein Minimum an Sonderzeichen, Ziffern oder Groß- und Kleinbuchstaben. Darum kümmern Sie sich sehr bald.

Kommentieren Sie vorerst die letzten zwei Zeilen Ihres Codes aus.

# --hints--

Sie sollten die letzten zwei Zeilen Ihres Codes auskommentieren.

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
