---
id: 656ef89b4b486b7a16077864
title: Schritt 26
challengeType: 20
dashedName: step-26
---

# --description--

Ein Tuple ist eine weitere eingebaute Datenstruktur in Python. Tuples ähneln Listen sehr, aber sie werden mit runden Klammern `()` definiert, statt mit eckigen Klammern. Außerdem sind Tuples unveränderlich, im Gegensatz zu Listen.

```py
my_tuple = ('larch', 1, True)
```

Ihre `constraints` Liste wird Tuples speichern. Das erste Element jedes Tuples wird ein Parameter für eine Einschränkung sein.

Ändern Sie die `constraints` Listen-Zuweisung, indem Sie ein Tuple zu Ihrer Liste hinzufügen. Verwenden Sie `nums` als erstes Element und einen leeren String als zweites Element.

# --hints--

Sie sollten ein Tuple zur `constraints` Liste hinzufügen.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

Sie sollten `(nums, '')` zur `constraints` Liste hinzufügen.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
