---
id: 656ef5bc5c0cd464be1df675
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Als Nächstes werden Sie Ihrer Funktion weitere Parameter hinzufügen, die als Einschränkungen für das generierte Passwort dienen.

Ändern Sie die Deklaration Ihrer Funktion, indem Sie `nums`, `special_chars`, `uppercase` und `lowercase` in dieser Reihenfolge nach dem vorhandenen `length`-Parameter hinzufügen.

# --hints--

Ihre Funktion sollte `length`, `nums`, `special_chars`, `uppercase` und `lowercase` als Parameter erhalten. Die Reihenfolge ist wichtig.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password

# new_password = generate_password(8)
# print(new_password)
```
