---
id: 656eec2f40d18056cc58b229
title: Schritt 30
challengeType: 20
dashedName: step-30
---

# --description--

Der Wert `None` wird zurückgegeben, da `'i'` im geparsten String nicht gefunden wird.

Ändern Sie nun den an `re.compile()` übergebenen String in `'l'` und sehen Sie sich das Ergebnis an.

# --hints--

Sie sollten Ihre `pattern`-Variable in `re.compile('l')` ändern.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)    
--fcc-editable-region--
pattern = re.compile('i')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
