---
id: 656f0bc259c1f6b57486ed68
title: Schritt 33
challengeType: 20
dashedName: step-33
---

# --description--

Um zu überprüfen, ob das generierte Passwort die erforderlichen Funktionen erfüllt, werden Sie die `findall()`-Funktion aus dem `re`-Modul verwenden. Sie ist ähnlich wie `search`, aber sie gibt eine Liste mit allen Vorkommen des übereinstimmenden Musters zurück.

Ersetzen Sie den `search()`-Aufruf durch `findall()` und überprüfen Sie die Ausgabe.

# --hints--

Sie sollten Ihren bestehenden `print()`-Aufruf ändern, indem Sie `search()` durch `findall()` ersetzen.

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.search(pattern, quote))
--fcc-editable-region--
```
