---
id: 656f04b20397d9a574cc7eb2
title: Schritt 32
challengeType: 20
dashedName: step-32
---

# --description--

Sie können dasselbe Ergebnis erzielen, ohne die `compile()`-Funktion zu verwenden. Ändern Sie Ihre `pattern`-Variable in den Literalstring `'l+'`. Dann ändern Sie den `print()`-Aufruf, sodass `re.search(pattern, quote)` ausgegeben wird.

# --hints--

Sie sollten Ihre `pattern`-Variable in den Literalstring `'l+'` ändern.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "l+") })
```

Sie sollten `re.search(pattern, quote)` ausgeben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.search\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l+')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
