---
id: 657732654845d8e2fb1217e6
title: Schritt 62
challengeType: 20
dashedName: step-62
---

# --description--

Im Moment nimmt `all()` eine leere Liste als Argument. Füllen Sie diese leere Liste mit der Comprehension-Syntax, sodass die Liste die Ergebnisse der Auswertung des Ausdrucks `constraint <= len(re.findall(pattern, password))` für jedes `constraint`-`pattern`-Tuple in der `constraints`-Liste speichert.

Auf diese Weise verlassen Sie die `while`-Schleife erst, nachdem alle Anforderungen erfüllt sind.

# --hints--

Sie sollten `[constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints]` an die `all()`-Funktion übergeben.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("all([constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints])")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if all([]):
--fcc-editable-region--        
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
