---
id: 6577319039f4f7de9251b822
title: Schritt 61
challengeType: 20
dashedName: step-61
---

# --description--

Anstatt eine Schleife und eine Zählvariable zu verwenden, können Sie dasselbe Ergebnis mit einem anderen Ansatz erzielen, den Sie in den nächsten Schritten implementieren werden.

`all()` ist eine eingebaute Python-Funktion, die `True` zurückgibt, wenn alle Elemente in einem gegebenen iterierbaren Objekt zu `True` ausgewertet werden. Andernfalls gibt sie `False` zurück.

Ersetzen Sie Ihre bestehende `for`-Schleife und zwei `if`-Anweisungen durch eine einzelne `if`-Anweisung. Für die `if`-Bedingung verwenden Sie einen Aufruf der `all()`-Funktion und übergeben eine leere Liste als Argument an den Funktionsaufruf.

# --hints--

Sie sollten Ihre bestehende `for`-Schleife und zwei `if`-Anweisungen durch eine einzelne `if`-Anweisung ersetzen.

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

Ihre neue `if`-Bedingung sollte `all([])` sein.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

Sie sollten `break` innerhalb des Körpers Ihrer neuen `if` haben.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
