---
id: 6577333feab1e8e927014f03
title: Schritt 63
challengeType: 20
dashedName: step-63
---

# --description--

Wenn Sie `all([expression for i in iterable])` haben, bedeutet das, dass eine neue Liste erstellt wird, indem `expression` für jedes `i` in `iterable` ausgewertet wird. Nachdem die `all()`-Funktion über die neu erstellte Liste iteriert hat, wird die Liste automatisch gelöscht, da sie nicht mehr benötigt wird.

Speicher kann eingespart werden, indem Sie einen Generatorausdruck verwenden. Generatorausdrücke folgen der Syntax von Listen-Komprehension, verwenden jedoch runde Klammern anstelle von eckigen Klammern.

Ändern Sie Ihre Listen-Komprehension in einen Generatorausdruck, indem Sie die eckigen Klammern entfernen.

# --hints--

Sie sollten Ihre Listen-Komprehension in einen Generatorausdruck umwandeln, indem Sie die eckigen Klammern entfernen.

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        if all(
            [
                constraint <= len(re.findall(pattern, password))
                for constraint, pattern in constraints
            ]
        ):
--fcc-editable-region--
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
