---
id: 657db4cb77190e33a20e852a
title: Schritt 69
challengeType: 20
dashedName: step-69
---

# --description--

Wenn Sie Standardargumente mit Schlüsselwortargumenten kombinieren, können Sie explizit weniger Argumente übergeben, als die Funktion verlangt. Die Argumente, die nicht explizit beim Funktionsaufruf übergeben werden, erhalten ihre Standardwerte.

Ändern Sie Ihren `generate_password()`-Aufruf so, dass er nur `length=8` übernimmt.

# --hints--

Sie sollten nur `length=8` an Ihren `generate_password()`-Aufruf übergeben.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region--
new_password = generate_password(nums=1, length=8, special_chars=1, uppercase=1, lowercase=1)
print(new_password)
--fcc-editable-region--
```
