---
id: 614385513d91ae5c251c2052
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 이전 버전의 freeCodeCamp 커리큘럼 보도 자료를 사용해 잡지 페이지를 만드세요.

`link` 폰트 스타일시트용 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` 요소, `link` FontAwesome 스타일시트용 `https://use.fontawesome.com/releases/v5.8.2/css/all.css` 요소, 그리고 `link` 스타일시트용 `./styles.css` 요소를 추가하는 것부터 시작하세요.

폰트 스타일시트는 `Anton`, `Baskervville`, `Raleway` 세 가지 폰트를 불러옵니다.

# --hints--

코드에는 `link` 요소가 세 개 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

`link` 요소들은 `head` 요소 안에 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

세 개의 `link` 요소에는 `rel` 속성이 `stylesheet` 값으로 있어야 합니다.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

링크 요소 중 하나는 `href`가 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`로 설정되어야 합니다.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

링크 요소 중 하나는 `href`가 `https://use.fontawesome.com/releases/v5.8.2/css/all.css`로 설정되어야 합니다.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

`link` 요소 중 하나는 `href`가 `styles.css`로 설정되어야 합니다.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
