---
id: 6143869bb45bd85e3b1926aa
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

`.heading` 요소 안에 `header` 요소를 만들고 `class` 속성을 `hero`로 설정하세요.

그 요소 안에 `img` 요소를 만들고 `src` 속성을 `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`로, `alt` 속성을 `freecodecamp logo`로, `class` 속성을 `hero-img`로 설정하세요.

`loading` 요소의 `img` 속성을 `lazy`로 설정하면 브라우저가 이미지 리소스를 필요할 때까지 가져오지 않도록 할 수 있습니다(예: 사용자가 이미지를 화면에 스크롤할 때). 추가 이점으로, 지연 로드된 요소는 지연되지 않은 요소가 로드될 때까지 로드되지 않으므로 인터넷 속도가 느린 사용자가 이미지가 로드될 때까지 기다리지 않고 페이지 내용을 볼 수 있습니다.

새로운 `img` 요소에 `loading` 속성을 `lazy`로 설정하세요.

# --hints--

`header` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('header'));
```

`header` 요소는 `.heading` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

`header` 요소는 `class` 속성을 `hero`로 설정해야 합니다.

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

`img` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('img'));
```

`img` 요소는 `header` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

`img` 요소는 `src` 속성을 `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`로 설정해야 합니다.

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

`img` 요소는 `alt` 속성을 `freecodecamp logo`로 설정해야 합니다.

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'freecodecamp logo');
```

`img` 요소는 `loading` 속성을 `lazy`로 설정해야 합니다.

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

`img` 요소는 `class` 속성을 `hero-img`로 설정해야 합니다.

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
