---
id: 614387cbefeeba5f3654a291
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

`header` 요소 다음에 `div`가 `class`로 설정된 `author`을 만드세요.

그 `div` 안에 `p`가 `class`로 설정된 `author-name` 요소를 만들고 텍스트 `By freeCodeCamp`을 넣으세요. `freeCodeCamp` 부분을 `a`이 `href`로, `https://freecodecamp.org`이 `target`로 설정된 `_blank` 요소로 감싸세요.


그 아래에 클래스 `p`와 텍스트 `publish-date`을 가진 두 번째 `March 7, 2019` 요소를 추가하세요.

# --hints--

새로운 `div` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('div'));
```

`div` 요소는 `header` 요소 다음에 와야 합니다.

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

`div` 요소는 `class`가 `author`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

새로운 `p` 요소 두 개를 만들어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

새로운 `p` 요소 두 개는 `div` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

첫 번째 새로운 `p` 요소는 `class`가 `author-name`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

첫 번째 새로운 `p` 요소는 텍스트 `By freeCodeCamp`을 가져야 합니다.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By freeCodeCamp');
```

두 번째 새로운 `p` 요소는 `class`가 `publish-date`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

두 번째 새로운 `p` 요소는 텍스트 `March 7, 2019`을 가져야 합니다.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

새로운 `a` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('a'));
```

`a` 요소는 첫 번째 새로운 `p` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

`a` 요소는 `href`가 `https://freecodecamp.org`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://freecodecamp.org');
```

`a` 요소는 `target`가 `_blank`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

`a` 요소는 텍스트 `freeCodeCamp`을 감싸야 합니다.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
