---
id: 614389f601bb4f611db98563
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

`.author` 요소 아래에 `div` 클래스를 가진 새로운 `social-icons` 요소를 만드세요.

그 새로운 `a` 안에 다섯 개의 `div` 요소를 추가하고, 각각에 다음 `href` 속성을 부여하세요.

- 첫 번째 `a` 요소는 `href`이 `https://www.facebook.com/freecodecamp`로 설정되어야 합니다.
- 두 번째 `a` 요소는 `href`이 `https://twitter.com/freecodecamp`로 설정되어야 합니다.
- 세 번째 `a` 요소는 `href`이 `https://instagram.com/freecodecamp`로 설정되어야 합니다.
- 네 번째 `a` 요소는 `href`이 `https://www.linkedin.com/school/free-code-camp`로 설정되어야 합니다.
- 다섯 번째 `a` 요소는 `href`이 `https://www.youtube.com/freecodecamp`로 설정되어야 합니다.

# --hints--

새로운 `div` 요소를 만들어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

새로운 `div` 요소는 `.author` 요소 뒤에 와야 합니다.

```js
assert.equal(document.querySelector('.author')?.nextElementSibling?.localName, 'div');
```

새로운 `div` 요소는 `social-icons` 클래스를 가져야 합니다.

```js
assert.isTrue(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

`.social-icons` 요소는 다섯 개의 `a` 요소를 가져야 합니다.

```js
assert.lengthOf(document.querySelector('.social-icons')?.querySelectorAll('a'), 5);
```

첫 번째 `a` 요소는 `href`이 `https://www.facebook.com/freecodecamp`로 설정되어야 합니다.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href'), 'https://www.facebook.com/freecodecamp');
```

두 번째 `a` 요소는 `href`이 `https://twitter.com/freecodecamp`로 설정되어야 합니다.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

세 번째 `a` 요소는 `href`이 `https://instagram.com/freecodecamp`로 설정되어야 합니다.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

네 번째 `a` 요소는 `href`이 `https://www.linkedin.com/school/free-code-camp`로 설정되어야 합니다.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

다섯 번째 `a` 요소는 `href`이 `https://www.youtube.com/freecodecamp`로 설정되어야 합니다.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
