---
id: 6143862a5e54455d262c212e
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

`body` 안에 `main` 요소를 만드세요. 그런 다음 그 요소 안에 `section`가 `class`로 설정된 `heading`를 만드세요.

# --hints--

`main` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('main'));
```

`main` 요소는 `body` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelector('main')?.parentElement?.localName, 'body');
```

`section` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('section'));
```

`section` 요소는 `main` 요소 안에 있어야 합니다.

```js
assert.equal(document.querySelector('section')?.parentElement?.localName, 'main');
```

`section` 요소는 `class`가 `heading`로 설정되어야 합니다.

```js
assert.isTrue(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
