---
id: 61438b5b66d76a6264430f2a
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

새로 만든 각 `a` 요소 안에 `i` 요소를 추가하고 다음 클래스를 지정하세요:

- 첫 번째 `i` 요소에는 `fab fa-facebook-f` 클래스를 지정하세요
- 두 번째 `i` 요소에는 `fab fa-twitter` 클래스를 지정하세요
- 세 번째 `i` 요소에는 `fab fa-instagram` 클래스를 지정하세요
- 네 번째 `i` 요소에는 `fab fa-linkedin-in` 클래스를 지정하세요
- 다섯 번째 `i` 요소에는 `fab fa-youtube` 클래스를 지정하세요

# --hints--

`i` 요소가 다섯 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('i'), 5);
```

각 `a` 요소에는 하나의 `i` 요소만 있어야 합니다.

```js
const aaaaa = document.querySelectorAll('.social-icons a');
assert.lengthOf(aaaaa, 5);
aaaaa?.forEach(
  a => {
    assert.lengthOf(a?.children, 1);
    assert.equal(a?.children?.[0]?.localName, 'i');
  }
);
```

각 `i` 요소에는 `class` 속성이 포함되어 있어야 하며, 그 속성에는 `fab`이 포함되어야 합니다.

```js
const iiiii = document.querySelectorAll('i');
assert.lengthOf(iiiii, 5);
iiiii?.forEach(
  i => {
    assert.isTrue(i?.classList?.contains('fab'));
  }
);
```

첫 번째 `i` 요소에는 `class` 속성이 포함되어 있어야 하며, 그 속성에는 `fa-facebook-f`이 포함되어야 합니다.

```js
assert.isTrue(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

두 번째 `i` 요소에는 `class` 속성이 포함되어 있어야 하며, 그 속성에는 `fa-twitter`이 포함되어야 합니다.

```js
assert.isTrue(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

세 번째 `i` 요소에는 `class` 속성이 포함되어 있어야 하며, 그 속성에는 `fa-instagram`이 포함되어야 합니다.

```js
assert.isTrue(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

네 번째 `i` 요소에는 `class` 속성이 포함되어 있어야 하며, 그 속성에는 `fa-linkedin-in`이 포함되어야 합니다.

```js
assert.isTrue(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

다섯 번째 `i` 요소에는 `class` 속성이 포함되어 있어야 하며, 그 속성에는 `fa-youtube`이 포함되어야 합니다.

```js
assert.isTrue(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
          </a>
          <a href="https://twitter.com/freecodecamp/">
          </a>
          <a href="https://instagram.com/freecodecamp">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/freecodecamp">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
