---
id: 614e0e588f0e8a772a8a81a6
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

현재 이미지가 너무 많은 공간을 차지하고 있습니다. 작업 중인 내용을 더 잘 볼 수 있도록 `width` 요소에 `img` 속성을 추가하고, 값은 `400`로 설정하세요.

나중에 CSS 작업을 할 때 이 속성은 제거할 것입니다.

# --hints--

`img` 요소에 `width` 속성이 `400`로 설정되어 있어야 합니다.

```js
assert.equal(document.querySelector('img')?.getAttribute('width'), '400');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
