---
id: 6a89c69cbe1bb69845992d8b
title: Construa um Gerador de Intervalo de Números
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

Neste laboratório, você vai construir uma função `range_of_numbers` que usa recursão para gerar uma lista de números dentro de um intervalo especificado.

Por exemplo:

- `range_of_numbers(3, 9)` deve retornar `[3, 4, 5, 6, 7, 8, 9]`.
- `range_of_numbers(5, 5)` deve retornar `[5]`.

Complete as user stories abaixo e faça todos os testes passarem para concluir o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `range_of_numbers` com dois parâmetros chamados `start_num` e `end_num`.
1. A função deve retornar uma lista de inteiros consecutivos que começa com `start_num` e termina com `end_num`, incluindo ambos os valores.
1. O valor de `start_num` será sempre menor ou igual a `end_num`.
1. Sua função deve usar recursão chamando a si mesma.
1. Sua função não deve usar loops `for` ou `while`, compreensões ou a função `range()`.
1. Quando `start_num` for igual a `end_num`, a função deve retornar uma lista contendo `start_num`. Este é o caso base.
1. Para o caso recursivo, a função deve se chamar com um argumento que se aproxime do caso base, depois adicionar o número atual à lista retornada.
1. A função não deve usar variáveis globais para armazenar a lista retornada.

# --hints--

Você deve ter uma função chamada `range_of_numbers` com parâmetros chamados `start_num` e `end_num`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` deve retornar `[1, 2, 3, 4, 5]`.

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` deve retornar `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` deve retornar `[4]`.

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` deve retornar `[10, 11, 12, 13, 14, 15]`.

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

Sua função `range_of_numbers` deve se chamar.

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

Sua função `range_of_numbers` não deve usar loops, compreensões ou a função `range()`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

Chamar `range_of_numbers(1, 3)` não deve afetar uma chamada posterior. Depois dessa chamada, `range_of_numbers(6, 9)` ainda deve retornar `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
