---
id: 66d9af3897e7d75a895b72c2
title: Crie um gerador do enésimo número de Fibonacci
challengeType: 14
dashedName: lab-nth-fibonacci-number-generator
---

# --description--

Neste laboratório, você usará uma abordagem de programação dinâmica para calcular e retornar o número `n`da sequência de Fibonacci, na qual cada número é a soma dos dois números anteriores.

Fulfill the user stories below and get all the tests to pass to complete the lab.

**Histórias de usuário:**

1. Você deve definir um array chamado `sequence` no escopo global. Ele deve conter os números `0` e `1`.
1. Você deve criar uma função chamada `fibonacci`.
1. A função `fibonacci` deve receber um único parâmetro: um número inteiro positivo.
1. Se o valor de entrada não for um número inteiro positivo, a função deve retornar `"Invalid input"`.
1. Uma chamada para `fibonacci(n)` deve usar uma abordagem de programação dinâmica para calcular e retornar o número `n`da sequência de Fibonacci.
1. Cada número calculado na posição `n` na sequência de Fibonacci deve ser armazenado no índice `n - 1` do array `sequence`.
1. Se o número de entrada solicitar um número já calculado nenhum novo cálculo deve acontecer e o número deve ser recuperado do array `sequence`.

# --hints--

Você deve ter um array `sequence`.

```js
assert.isArray(sequence);
```

O array `sequence` deve conter `0` e `1` em sua definição.

```js
assert.match(__helpers.removeJSComments(code), /(let|const)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]\s*;?/)
```

Você deve criar uma função `fibonacci`.

```js
assert.isFunction(fibonacci);
```

`fibonacci(3.5)` deve retornar `"Invalid input"`.

```js
assert.strictEqual(fibonacci(3.5), "Invalid input")
```

`fibonacci(0)` deve retornar `"Invalid input"`.

```js
assert.strictEqual(fibonacci(0), "Invalid input")
```

`fibonacci(-1)` deve retornar `"Invalid input"`.

```js
assert.strictEqual(fibonacci(-1), "Invalid input")
```

`fibonacci("three")` deve retornar `"Invalid input"`.

```js
assert.strictEqual(fibonacci("three"), "Invalid input")
```

`fibonacci(1)` deve retornar `0`.

```js
assert.strictEqual(fibonacci(1), 0)

```

`fibonacci(32)` deve retornar `1346269`.

```js
assert.strictEqual(fibonacci(32), 1346269)

```

`fibonacci(5)` deve retornar `3`.

```js
assert.strictEqual(fibonacci(5), 3);

```

`sequence` deve conter números que vão até a chamada mais alta até o momento.

```js
sequence.length = 2
fibonacci(2)
assert.lengthOf(sequence, 2)
fibonacci(5)
assert.lengthOf(sequence, 5)
fibonacci(3)
assert.lengthOf(sequence, 5)
fibonacci(10)
assert.lengthOf(sequence, 10)
assert.deepEqual(sequence, [0, 1, 1, 2, 3, 5, 8, 13, 21, 34])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const sequence = [0, 1];

const fibonacci = (n) => {
    if (typeof n !== "number" || n < 1 || n % 1 !== 0) {
        console.log("Invalid input")
        return "Invalid input"
    }
    if (n <= sequence.length) {
        return sequence[n - 1];
    }
    for (let i = sequence.length; i < n; i++) {
        sequence.push(sequence[i - 1] + sequence[i - 2]);
    }
    return sequence[n - 1]
}

```
