---
id: 687d7deea6a2a22fa8816a73
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

これまでのレッスンで、いくつかの一般的なデータ構造について学びました。このワークショップでは、連結リストを実装します。

まず、リストのノードを保持する構造を表す `LinkedList` という名前のクラスを作成してください。テストに合格するために、クラス内に `pass` キーワードを追加してください。

# --hints--

`LinkedList` という名前のクラスを作成してください。`pass` キーワードを追加することを忘れないでください。

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
