---
id: 687d8e05a529a4838f1c961c
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

次に、`pass` クラス内の `Node` キーワードを削除し、`Node` クラスに `__init__` メソッドを2つのパラメータ、`self` と `element` で追加してください。

メソッド内で、`element` に `self.element` を代入してください。これはノードを作成するときに渡された値をストアします。

# --hints--

`__init__` クラス内に `Node` メソッドを定義してください。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

`__init__` メソッドは2つのパラメータ、`self` と `element` を持つ必要があります。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

`element` メソッド内で `self.element` に `__init__` を代入してください。

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
