---
id: 687d8c6f066877768f7ee441
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

`LinkedList` クラスの中で、`pass` キーワードを `Node` という名前の内部クラスに置き換えてください。これはリスト内のデータを格納する個々のノードを表します。

# --hints--

`Node` クラスの中に `LinkedList` クラスを作成してください。`pass` クラスの中に `Node` キーワードを追加することを忘れないでください。

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
