---
id: 687e366b3004a9831e4cbd41
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`is_empty` という名前のメソッドを追加して、`self.length == 0` を返してください。

これはリスト内のノード数を追跡するために使います。`self.length == 0` の場合、`True` を返し、リストにノードがないことを意味します。

`self.length` が `0` より大きい場合は、`False` を返し、リストに少なくとも1つのノードがあることを意味します。

# --hints--

`is_empty` メソッドを作成して、`self.length == 0` を返すようにしてください。

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
