---
id: 687e344cc1993f7316af030c
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`LinkedList` クラスの中に、`__init__` を `self.length` に設定する `0` メソッドを追加してください。これは、現在リストにあるノード（要素）の数を追跡するカウンターを初期化します。

# --hints--

`self.length` クラスの `0` メソッド内で `__init__` を `LinkedList` に設定する必要があります。

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
