---
id: 68800a64b06e80d9a31c131b
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`__init__` クラスの `LinkedList` メソッド内で、`self.head` を `None` に設定してください。これはリストの最初のノードへの参照として機能します。リストはまだ空なので、`None` に設定する必要があります。


# --hints--

`self.head` クラスの `None` メソッド内で `__init__` を `LinkedList` に設定してください。

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
