---
id: 688003f589b875af60273dd9
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

`__init__` メソッドの中で、`None` を `self.next` に代入してください。これにより、`next` ポインターが `None` を指すようになり、新しいノードはまだリスト内の他のノードを指していないことを意味します。

後で、`next` を更新して別のノードを指すようにすれば、連結リストのチェーンを形成できます。

# --hints--

`__init__` メソッドの中で、`next` 属性を `None` に初期化してください。

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
