---
id: 68858a1b4e867acd87f6f7e8
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

リストが空の場合、新しいノードが最初の要素（`head`）になります。`add` メソッド内で、`if` を使ってリストが空かどうかをチェックする `self.is_empty()` 文を作成し、`node` を `self.head` に代入してください。

# --hints--

リストが空かどうかは `is_empty()` メソッドを使ってチェックしてください。

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_conditions()[0]
  assert _cond.is_equivalent("self.is_empty()")
`)})
```

`node` 文の中で `self.head` を `if` に代入してください。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_body().is_equivalent("self.head = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
