---
id: 68857e0a49baf4751f737a2a
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

`is_empty` メソッドの後に、`add` クラスの中に `LinkedList` というメソッドを作成してください。新しいメソッドは `self` と `element` の2つのパラメータを持つ必要があります。

新しいメソッドの中で、`node` という変数を作成し、`Node` を使って `self.Node(element)` の新しいインスタンスを代入してください。これにより、連結リストに追加される新しいノードが作成されます。

# --hints--

`add` クラスの中に `LinkedList` メソッドを作成してください。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

`add` メソッドは `self` と `element` の2つのパラメータを持つ必要があります。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

`self.Node(element)` メソッドの中で `node` に `add` を代入してください。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
