---
id: 688649c8673f842f49228a5e
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

リストが空でない場合、メソッドは最後のノードを見つけて、新しいノードをその`next`ポインターにリンクする必要があります。

これを行うには、`else`ブロックの後に`if`ブロックを作成してください。`else`ブロックの中で、`current_node`変数を作成し、`self.head`を代入してください。

# --hints--

`self.head`ブロックの中で`current_node`を`else`に代入する必要があります。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
