---
id: 68868f502659139bed8e5c90
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

リストに新しいノードを追加した後、その長さを更新する必要があります。`add` メソッドの最後で、`self.length` を `1` だけ増やしてください。

# --hints--

`self.length` メソッドの最後で `1` を `add` だけ増やす必要があります。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").has_stmt("self.length += 1")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
