---
id: 6888f9827637322d4f309192
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

`remove` メソッドの中で、`current_node` という名前の変数を作成し、`self.head` の値を代入してください。

これで連結リストの最初のノードからの走査が始まります。`current_node` は、指定された要素を含むノードを探す際に、現在調べているノードを追跡するために使います。

# --hints--

`self.head` メソッドの中で `current_node` に `remove` を代入する必要があります。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").has_stmt("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
