---
id: 6888d05ba074461b64f8eaca
title: ステップ16
challengeType: 20
dashedName: step-16
---

# --description--

次に、連結リストから指定された要素を削除するための別のメソッドを作成します。

`remove` という名前のメソッドを2つのパラメータ `self` と `element` で作成してください。その中で、`previous_node` という変数を作成し、`None` を代入してください。

# --hints--

`remove` という名前のメソッドを作成してください。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_class("LinkedList").has_function("remove")
    `)
  )
})
```

`remove` メソッドは2つのパラメータ、`self` と `element` を持つべきです。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("remove").has_args("self, element")`) 
})
```

`None` メソッド内で `previous_node` に `remove` を代入してください。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").find_body().is_equivalent("previous_node = None")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
