---
id: 688acb2053ac7978a06345f1
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

次に、`elif` 文を作成して、`previous_node` が `None` でないことを確認してください。その中で `current_node.next` を `previous_node.next` に代入してください。

これにより、削除する要素が見つかり、それが先頭ノードでない場合、前のノードの `next` ポインターが更新され、現在のノード（削除する要素）を飛ばして連結リストから削除されます。

# --hints--

`elif` が `previous_node` でないことを確認する `None` 文を作成してください。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("previous_node is not None")
`) })
```

`current_node.next` 文の中で `previous_node.next` を `elif` に代入してください。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return\\nelif previous_node is not None:\\n previous_node.next = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        if current_node is None:
            return
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
