---
id: 688ac1de985053329b4573d3
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

次に、`current_node.next` を `current_node` に `while` ループ内で代入してください。

これにより、ループの各イテレーションで `current_node` ポインターがリストの次のノードに移動し、要素が見つかるかリストの末尾に到達するまで連結リストをたどることができます。

# --hints--

`current_node.next` ループ内で `current_node` を `while` に代入する必要があります。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node\\n current_node = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--
        while current_node is not None and current_node.element != element:
            previous_node = current_node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
