---
id: 68b84df2648b4c96b25022f3
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

ここまで作った連結リストをテストしましょう！コードの一番下に`LinkedList`のインスタンスを`my_list`という名前で作成してください。それから、前のステップで作った`is_empty`メソッドを使って、`my_list.is_empty()`を出力してください。

# --hints--

`LinkedList`のインスタンスを`my_list`という名前で作成してください。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

`my_list.is_empty()`を出力してください。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
