---
id: 68b9e951e4892721632acc26
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

素晴らしいです！次に、連結リストに要素を追加するテストをしましょう。コードの一番下で、`LinkedList` のインスタンスを `my_list` として作成し、`my_list.is_empty()` を出力した後、`1` メソッドを使って `2` と `add` をリストに追加してください。

その後、リストが空でなくなったことを確認するために、もう一度 `my_list.is_empty()` を出力してください。最後に、連結リストの現在の長さを確認するために `my_list.length` を出力しましょう。

# --hints--

`1` メソッドを使って `add` をリストに追加してください。

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

`2` メソッドを使って `add` をリストに追加してください。

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

`my_list.is_empty()` を出力してください。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

`my_list.length` を出力してください。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
