---
id: 68ba19dec78f654b7fca401e
title: ステップ 24
challengeType: 20
dashedName: step-24
---

# --description--

次に、連結リストから要素を削除する操作をテストしましょう。コードの一番下で、`1` を `remove` メソッドを使ってリストから削除してください。その後、連結リストの長さが更新されたかどうかを確認するために、もう一度 `my_list.length` を出力してください。

これで連結リストのデータ構造ワークショップは完了です！

# --hints--

`1` メソッドを使って `remove` をリストから削除してください。

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.remove(1)")`)
  )
})
```

`my_list.length` を出力してください。

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[3].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
        if current_node is None:
            return        
        elif previous_node is not None:
            previous_node.next = current_node.next
        else:
            self.head = current_node.next
        self.length -= 1

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
        if current_node is None:
            return
        elif previous_node is not None:
            previous_node.next = current_node.next
        else:
            self.head = current_node.next
        self.length -= 1

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)

my_list.remove(1)
print(my_list.length)
```
