---
id: a3bfc1673c0526e06d3ac698
title: Construa um Calculadora de Soma de Números Primos
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

Neste laboratório, você irá construir uma calculadora que recebe um número e retorna a soma de todos os números primos que são menores ou iguais a esse número.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função `sumPrimes` que aceita um número como argumento.
1. A função `sumPrimes` deve retornar a soma de todos os números primos menores ou iguais ao número fornecido.
1. Se o número de entrada for menor que `2`, a função deve retornar `0`.

# --hints--

Você deve ter uma função `sumPrimes`.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` deve retornar `17`.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` deve retornar `10`.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` deve retornar `2`.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` deve retornar `0`.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` deve retornar `73156`.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
