---
id: 6626a060c4006f793e10cb33
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이제 학생 점수를 학점으로 변환하는 데 도움이 필요합니다.

숫자 `getGrade`를 매개변수로 받는 `score` 함수를 완성하세요. 함수는 점수에 따라 학점을 나타내는 문자열을 반환해야 합니다.

다음은 점수와 해당하는 학점입니다:

| 점수 범위 | 학점   |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**팁**

- 조건문(`if`, `else if`, `else`)에 대해 배운 것을 기억하세요.
- 비교 연산자(`>`, `<`, `>=`, `<=`, `===`)에 대해 배운 것을 기억하세요.
- 논리 AND 연산자(`&&`)에 대해 배운 것을 기억하세요.

# --hints--

점수가 `getGrade`일 때 `"A++"` 함수는 `100`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(100), "A++");
```

점수가 `getGrade`일 때 `"A"` 함수는 `94`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(94), "A");
```

점수가 `getGrade`과 `"B"` 사이일 때 `80` 함수는 `89`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

점수가 `getGrade`일 때 `"C"` 함수는 `78`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(75), "C");
```

점수가 `getGrade`과 `"D"` 사이일 때 `60` 함수는 `69`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

점수가 `getGrade`일 때 `"F"` 함수는 `57`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(57), "F");
```

점수가 `getGrade`과 `"F"` 사이일 때 `0` 함수는 `59`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
