---
id: 662693f82c91a66be46c881b
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

선생님이 학생들의 시험 채점을 마치고 학급의 평균 점수를 계산하는 데 도움이 필요합니다.

시험 점수 배열을 입력받아 평균 점수를 반환하는 `getAverage` 함수를 완성하세요.

평균은 모든 점수를 더한 후 점수의 총 개수로 나누어 계산합니다.

```md
average = sum of all scores / total number of scores
```

코드를 테스트할 수 있도록 몇 가지 함수 호출이 제공되어 있습니다.

**팁**

- 루프를 사용해 `scores` 배열을 순회하며 모든 점수를 더할 수 있습니다.
- `length` 속성을 사용해 점수의 총 개수를 얻을 수 있습니다.

# --hints--

`getAverage` 함수는 숫자를 반환해야 합니다.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])`는 `71.7`를 반환해야 합니다.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])`는 `85.4`를 반환해야 합니다.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])`는 `92.4`를 반환해야 합니다.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage` 함수는 시험 점수의 평균 점수를 반환해야 합니다.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
