---
id: 6626b4c58c027d86478ff5eb
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

선생님은 지금까지 만든 프로그램에 정말 만족하고 계세요. 하지만 이제 학생이 합격 점수를 받았는지 쉽게 확인할 수 있는 방법을 원하세요. 합격 점수는 `"F"`가 아닌 모든 점수입니다.

학생 점수를 매개변수로 받는 `hasPassingGrade` 함수를 완성하세요. 학생이 합격 점수를 받았다면 `true`를, 그렇지 않다면 `false`를 반환해야 합니다.

**팁**

- `getGrade` 함수를 사용해 학생의 점수를 가져오세요. 그런 다음 점수가 합격인지 아닌지 확인하세요.

# --hints--

`hasPassingGrade` 함수는 불리언 값을 반환해야 합니다.

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

`hasPassingGrade` 함수는 점수가 `true`일 경우 `"A"`를 반환해야 합니다.

```js
assert.isTrue(hasPassingGrade(100));
```

`hasPassingGrade` 함수는 점수가 `false`일 경우 `"F"`를 반환해야 합니다.

```js
assert.isFalse(hasPassingGrade(53));
```

`hasPassingGrade` 함수는 모든 합격 점수에 대해 `true`를 반환해야 합니다.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
