---
id: 6626b8dcf5057f896f948440
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

교사가 필요한 모든 정보를 얻었으니, 학생에게 결과를 메시지로 보낼 수 있기를 원합니다.

`studentMsg` 함수를 `totalScores`과 `studentScore` 매개변수로 완성하세요. 이 함수는 학생에게 보낼 메시지를 나타내는 문자열을 반환해야 합니다.

학생이 과정을 통과했다면, 문자열은 다음 형식을 따라야 합니다:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

학생이 과정을 실패했다면, 문자열은 다음 형식을 따라야 합니다:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

`average-goes-here`을 총점 평균으로 바꾸세요.  
`grade-goes-here`를 학생의 성적으로 바꾸세요.

**팁**

- `getAverage` 함수를 사용해 반 평균을 구하세요.  
- `getGrade` 함수를 사용해 학생의 성적을 구하세요.  
- 문자열 연결(`+`)을 사용해 메시지를 만드세요.  
- 메시지의 구두점과 공백에 주의하세요.

# --hints--

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)`은 다음 메시지를 반환해야 합니다: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)`은 다음 메시지를 반환해야 합니다: `"Class average: 50.8. Your grade: A++. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

`studentMsg` 함수는 학생 점수와 반 평균에 따라 올바른 메시지를 반환해야 합니다.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
