---
id: 616971b3cd990186b66c99a1
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Ahora que tu archivo CSS externo está configurado, puedes comenzar a darle estilo a la página.

Como recordatorio, aquí está cómo dirigir un elemento de párrafo y alinearlo a la derecha:

```css
p {
  text-align: right;
}
```

Crea una nueva regla CSS que apunte al elemento `h1` y configure su propiedad `text-align` a `center`.

# --hints--

Debes utilizar un selector `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

La regla CSS `h1` debe tener una propiedad `text-align` configurada en `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
