---
id: 616d50b93ba424d6282c99cf
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

El color de fondo se aplicó, pero dado que el elemento marcador `div` no tiene contenido, su altura por defecto es nula.

En la regla CSS `.marker`, establezca la propiedad `height` a `25px` y la propiedad `width` a `200px`.

# --hints--

La regla CSS `.marker` debe tener una propiedad `width` establecida en `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.width, '200px');
```

La regla CSS `.marker` debe tener una propiedad `height` establecida en `25px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
