---
id: 616d525007b8c5d8b3308b1c
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Tu marcador se vería mejor si estuviera centrado en la página. Una manera fácil de hacer esto es utilizando la `propiedad abreviada` `margin`.

Puedes establecer el área de margen de los elementos por separado con propiedades como `margin-top` y `margin-left`, la `propiedad abreviada` `margin` facilita establecer múltiples áreas de margen al mismo tiempo.

Para centrar tu marcador en la página, establece su `propiedad` `margin` en `auto`. Esto establece `margin-top`, `margin-right`, `margin-bottom`, y `margin-left` todos en `auto`.

# --hints--

Tu regla CSS `.marker` debe tener una `propiedad` `margin` configurada en `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
