---
id: 616d55bd160a95e22453a081
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Ya que tienes un marcador centrado y con color, es hora de añadir los otros marcadores.

En el elemento `.container`, añade dos elementos más `div` y dales a cada uno una clase de `marker`.

# --hints--

Tu primer nuevo elemento `div` debe tener una etiqueta de apertura.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

Tu primer nuevo elemento `div` debe tener una etiqueta de cierre.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Tu segundo nuevo elemento `div` debe tener una etiqueta de apertura.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

Tu segundo nuevo elemento `div` debe tener una etiqueta de cierre.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Tus nuevos elementos `div` deben estar dentro del `div` con la clase `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

Tus nuevos elementos `div` deben ambos tener una clase de `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
