---
id: 616d3a67ccf800ad94ec89ae
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Ahora añadirás algunos elementos, a los que darás forma de marcadores o plumones en los siguientes pasos.

Primero, dentro del elemento `body`, añade un elemento `div` con un atributo `class` establecido en `container`. Asegúrate de que el elemento `div` esté colocado debajo del elemento `h1`.

# --hints--

Debe haber una etiqueta de apertura `div`.

```js
assert.match(code, /<div.*>/i);
```

Debe haber una etiqueta de cierre `div`.

```js
assert.match(code, /<\/div\s*>/i);
```

El elemento `div` debe estar dentro del elemento `body`.

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

El elemento `div` debe estar después del elemento `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

El elemento `div` debe tener un atributo `class` establecido en `container`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
