---
id: 6176519636a76b810ab1219a
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

Crea una regla CSS que apunta a la clase `two` y establece su propiedad `background-color` en `green`. 

Además, cree una regla CSS independiente que apunte a la clase `three` y establezca su `background-color` en `blue`.

# --hints--

Debe usar un selector de clase para apuntar a la clase `two`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.two'));
```

Su regla CSS `.two` debe tener una propiedad `background-color` establecida en `green`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Debe usar un selector de clase para apuntar a la clase `three`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.three'));
```

Su regla CSS `.three` debe tener una propiedad `background-color` establecida en `blue`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
