---
id: 61764c602bee6974e7790f35
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Para dar colores diferentes a los marcadores, necesitará añadir una clase única a cada uno. Puede añadir varias clases a un elemento listándolas en el atributo `class` y separándolas con un espacio. Por ejemplo, el siguiente código añade tanto la clase `animal` como la clase `dog` a un elemento `div`:

```html
<div class="animal dog">
```

Si añade varias clases a un elemento HTML, los estilos de las primeras clases que liste en la hoja de estilos pueden ser sobrescritos por las clases posteriores.

Para empezar, agregue la clase `one` al primer elemento marcador `div`.


# --hints--

Debería añadir la clase `one` al primer elemento `div` del marcador.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

El primer marcador `div` debe tener las clases `marker` y `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
