---
id: 617b8e0d93a8d10d9a90e720
title: Paso 18
challengeType: 0
dashedName: step-18
---

# --description--

Observa que el `background-color` (color de fondo) de tu marcador sigue siendo rojo. Esto es porque has establecido el rojo de la función `rgb` en su valor máximo de `255`, o 100% rojo, y tanto el verde como el azul con valores de `0`.

Ahora use la función `rgb` para establecer los otros colores.

En la regla CSS `.two`, use la función `rgb` para establecer el `background-color` al valor máximo para el verde, y `0` para los otros valores. Y en la regla CSS `.three`, use la función `rgb` para establecer el `background-color` al valor máximo para el azul, y `0` para los otros valores.

# --hints--

Su regla CSS `.two` no debe usar la palabra clave de color `green` para establecer la propiedad `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Su regla CSS `.two` debe tener una propiedad `background-color` establecida en `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

Su regla CSS `.three` no debe usar la palabra clave de color `blue` para establecer la propiedad `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

Su regla CSS `.three` debe tener una propiedad `background-color` establecida en `rgb(0, 0, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
