---
id: 617b92b9de349513466f6156
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Mientras que los marcadores rojo y azul se ven iguales, el verde es mucho más claro de lo que era antes. Esto se debe a que la palabra clave del color `green` (verde) es, en realidad, una sombra más oscura, y está a medias, entre el negro y el valor máximo del verde.

En la regla CSS `.two`, establece el valor verde en la función `rgb` a `127` para reducir su intensidad.

# --hints--

Su regla CSS `.two` debe tener una propiedad `background-color` establecida en `rgb(0, 127, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
