---
id: 617b97abd9f3f61d1590b815
title: Paso 21
challengeType: 0
dashedName: step-21
---

# --description--

En el modelo de color aditivo RGB, los <dfn>colores primarios</dfn> son colores que, al combinarse, crean blanco puro. Pero para que esto suceda, cada color debe estar en su máxima intensidad.

Antes de combinar colores, establece tu marcador verde de nuevo al verde puro. Para la función `rgb` en la regla CSS `.two`, establece el verde de nuevo al valor máximo de `255`.

# --hints--

Su regla CSS `.two` debe tener una propiedad `background-color` establecida en `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
