---
id: 617bc3386dc7d07d6469bf20
title: Paso 27
challengeType: 0
dashedName: step-27
---

# --description--

Nota que, para crear el naranja, incrementaste la intensidad del rojo y disminuiste la intensidad del verde en los valores `rgb`. Esto se debe a que el naranja se crea combinando el rojo con el amarillo.

Para crear el color terciario verde primavera, combina el cian con el verde. Actualiza la función `rgb` en la regla CSS `.two` configurando el verde a su valor máximo y el azul a `127`.

# --hints--

Tu regla CSS `.two` debe tener una propiedad `background-color` establecida en `rgb(0, 255, 127)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
