---
id: 619b7fd56aa2253778bcf5f7
title: Paso 43
challengeType: 0
dashedName: step-43
---

# --description--

Otra forma de representar colores es el modelo <dfn>HSL</dfn>, siglas de hue, saturation y lightness (matiz, saturación y luminosidad).

La función CSS `hsl` acepta 3 valores: un número del 0 al 360 para el tono, un porcentaje del 0 al 100 para la saturación y un porcentaje del 0 al 100 para la luminosidad.

Si imaginas una rueda de colores, el tono rojo está a 0 grados, el verde a 120 grados y el azul a 240 grados.

La saturación (saturation) es la intensidad de un color desde 0%, o gris, hasta 100% para el color puro. Debes añadir el signo de porcentaje `%` a los valores de saturación y luminosidad.

La luminosidad es cuánta luz aparenta tener un color, desde un 0%, o negro total, hasta un 100%, blanco total, siendo 50% neutral.

En la regla CSS `.blue`, utiliza la función `hsl` para cambiar la propiedad `background-color` a azul puro. Configura el tono a `240`, la saturación a `100%` y la luminosidad a `50%`.

# --hints--

La regla CSS `.blue` debe tener una propiedad `background-color` configurada en `hsl(240, 100%, 50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

.blue {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

```
