---
id: 619b7c3c83de403126b69c1e
title: Paso 42
challengeType: 0
dashedName: step-42
---

# --description--

Seguramente estés familiarizado con valores en base 10, o decimales, los cuales van de 0 a 9. Hexadecimales, o valores en base 16, van de 0 a 9 y a continuación de A a F:

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

Con colores hexadecimales, `00` es un 0% de ese color, y `FF` es 100%. Así que `#00FF00` se traduce a 0% rojo, 100% verde y 0% azul, y es lo mismo que `rgb(0, 255, 0)`.

Reduce la intensidad del verde estableciendo el valor verde del color hexadecimal en `7F`.

# --hints--

La regla CSS `.green` debe tener la propiedad `background-color` establecida en `#007F00`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
--fcc-editable-region--
  background-color: #00FF00;
--fcc-editable-region--
}

.blue {
  background-color: rgb(0, 0, 0);
}

```
