---
id: 67051431a73f1ca25d9a6f25
title: 建置一個專案點子看板
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。 

**使用者故事：**

1. 你應該定義一個名為 `projectStatus` 的物件常數，包含三個鍵：`PENDING`、`SUCCESS` 和 `FAILURE`。每個狀態應該指定一個物件，該物件具有一個 `description` 鍵，其值分別為 `Pending Execution`、`Executed Successfully` 和 `Execution Failed`。
2. 你應該定義一個名為 `ProjectIdea` 的類別，並且有一個接受 `title` 和 `description` 作為參數的 `constructor`。使用提供的參數初始化 `title` 和 `description` 屬性。該類別還應該有一個名為 `status` 的屬性，預設值設為 `projectStatus.PENDING`。
3. 你應該在 `ProjectIdea` 類別中定義一個名為 `updateProjectStatus` 的方法。此方法應接受一個 `newStatus` 參數，並將 `status` 屬性更新為給定的值。
4. 你應該定義一個 `ProjectIdeaBoard` 類別，其 `constructor` 接受一個 `title` 並初始化一個名為 `ideas` 的空陣列，用來存放 `ProjectIdea` 類別的實體。
5. 你應該在 `ProjectIdeaBoard` 類別內定義一個名為 `pin` 的方法，該方法接受一個 `ProjectIdea` 類別的實體，並將該實體推入 `ideas` 陣列。
6. 你應該在 `ProjectIdeaBoard` 類別中定義一個名為 `unpin` 的方法。此方法應接受一個 `ProjectIdea` 類別的實體，並將它從 `ideas` 陣列中移除。
7. 你應該定義一個名為 `count` 的方法，該方法傳回給定 `ProjectIdeaBoard` 陣列中專案點子的數量。
8. 你應該定義一個名為 `formatToString` 的方法，該方法以以下格式傳回專案想法的名稱、描述和狀態：

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

你應該定義一個名為 `projectStatus` 的物件常數，並包含正確的專案狀態和描述。

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

你應該有一個 `ProjectIdea` 類別。

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

你的 `ProjectIdea` 類別應該根據傳入的參數初始化 `title` 屬性和 `description` 屬性。它還應該將狀態設為 `projectStatus.PENDING`。

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

你的 `ProjectIdea` 類別應該有一個 `updateProjectStatus` 方法。

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

你應該有一個 `ProjectIdeaBoard` 類別。

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

你的 `ProjectIdeaBoard` 應該根據傳入的參數初始化 `title` 屬性，並初始化一個名為 `ideas` 的空陣列來保存 `ProjectIdea` 類別的實體。

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

你的 `ProjectIdeaBoard` 類別應該有一個 `pin` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

你的 `ProjectIdeaBoard` 類別應該有一個 `unpin` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

你的 `ProjectIdeaBoard` 類別應該有一個 `count` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

你的 `ProjectIdeaBoard` 類別應該有一個 `formatToString` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")` 應該傳回 `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }`。

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

在 `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")` 上呼叫 `updateProjectStatus(projectStatus.SUCCESS)` 應該會將狀態更新為 `{ description: 'Executed Successfully' }`。

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

呼叫 `updateProjectStatus(projectStatus.FAILURE)` 在 `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")` 應該會將狀態更新為 `{ description: 'Execution Failed' }`。

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

呼叫 `updateProjectStatus(projectStatus.SUCCESS)` 在 `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")` 應該會將狀態更新為 `{ description: 'Executed Successfully' }`。

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

你應該能夠使用 `pin` 方法，將 `ProjectIdea` 物件釘選到你的 `ProjectIdeaBoard`。

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

你應該能夠使用 `unpin` 方法，將 `ProjectIdea` 物件從你的 `ProjectIdeaBoard` 解除固定。

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

當 `new ProjectIdeaBoard("Empty Board")` 是空的時候，`emptyBoard.formatToString()` 應該傳回 `Empty Board has 0 idea(s)\n`。

```js

let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

// Create an empty board
let emptyBoard = new ProjectIdeaBoard("Empty Board");

// Attempt to unpin a project from an empty board
emptyBoard.unpin(smartHome);

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

當你將 `new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")` 固定到 `new ProjectIdeaBoard("Tech Projects Board")` 時，`techProjects.formatToString()` 應該傳回 `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n`。

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
