---
id: 694c84330cbe604821b9deef
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

ミッションコントロールでは、メンバーを新たにリストに追加するメソッドが必要です。`addCrewMember` という名前の空の関数を作成し、`crew` と `astronaut` の2つのパラメータを受け取るようにしてください。

# --hints--

`addCrewMember` という名前の関数を作成してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

`addCrewMember` 関数は `crew` と `astronaut` のパラメータを持つべきです。

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
