---
id: 694c84106e9023f289c03b84
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

あなたのロスターは1人の指揮官宇宙飛行士から始まります。次のプロパティを持つ`firstAstronaut`という名前のオブジェクトを作成してください。

| キー | 値 |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

`firstAstronaut`という変数を作成してください。

```js
assert.isDefined(firstAstronaut);
```

`firstAstronaut`はオブジェクトである必要があります。

```js
assert.isObject(firstAstronaut);
```

`firstAstronaut`オブジェクトには正しいプロパティの値を含めてください。

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
