---
id: 694c40d86fe6a78d20b6559c
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、JavaScriptを活用して宇宙ミッションのための宇宙飛行士の名簿を作成します。

まず、名簿用の空の配列を`squad`という名前で作成してください。

# --hints--

`squad`という変数を作成してください。

```js
assert.isDefined(squad);
```

`squad`変数は配列である必要があります。

```js
assert.isArray(squad);
```

`squad`配列は空である必要があります。

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
