---
id: 694c9c186c089f7bdff0f3c6
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

一般的に、関数内で入力を検証することは良い習慣です。`for` ループを使って `crew` 配列を繰り返し処理し、入力の `id` と同じ `astronaut` を持つメンバーがすでに存在するかどうかを確認してください。重複が見つかった場合は、エラーメッセージ `console.log("Duplicate ID: " + astronaut.id)` をログに出力し、`return` を呼び出して関数を早期に終了させてください。

これはよく使われるこの手法の例です：

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

関数は `for` ループを使って既存の `crew` メンバーをスキャンする必要があります（有効な `for (...)` 形式ならどれでも構いません）。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

`for` ループ内で、重複する `id` をチェックしてください。見つかった場合は、`"Duplicate ID: " + astronaut.id"` をコンソールにログ出力し、`return` を呼び出して関数を終了させてください。

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

有効な入力が与えられた場合、`addCrewMember` 関数は `console.log()` エラーメッセージや早期の `return` 終了を呼び出してはいけません。

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
