---
id: 694ce8609fb805ba21cebc02
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

ミッションコントロールでは、乗組員名簿の中で2人の乗組員の位置を入れ替えるメソッドが必要です。3つのパラメータを持つ空の関数`swapCrewMembers`を作成してください。

- `crew`: 宇宙飛行士オブジェクトの配列

- `fromIndex`: 入れ替える最初の宇宙飛行士のインデックス

- `toIndex`: 入れ替える2番目の宇宙飛行士のインデックス

# --hints--

`swapCrewMembers`という名前の関数を作成してください。

```js
assert.isFunction(swapCrewMembers);
```

`swapCrewMembers`関数は`crew`、`fromIndex`、`toIndex`というパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('swapCrewMembers', ['crew', 'fromIndex', 'toIndex']);
assert.match(__helpers.removeJSComments(swapCrewMembers.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

--fcc-editable-region--

--fcc-editable-region--
```
