---
id: 694d8ddbfc4bb6940b921bbc
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

`crew` メソッドを使って `slice()` 配列の浅いコピーを作成してください。新しい配列の名前は `updatedCrew` にしてください。

引数なしで `slice()` を呼び出すと、配列全体のコピーが返されることを思い出してください。

```js
const copyArray = originalArray.slice();
```

# --hints--

`slice()` メソッドを使って入力の `crew` 配列をコピーし、その結果を `updatedCrew` という新しい変数に保存してください。

```js
const funcStr = __helpers.removeJSComments(swapCrewMembers.toString());
assert.match(
  funcStr,
  /(var|let|const)\s+updatedCrew/,
  "You must have at least one space between the declaration keyword (let/const) and 'updatedCrew'"
);

const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
assert.match(cleaned, /(var|let|const)updatedCrew=crew\.slice\(\);?\}$/);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }
  
--fcc-editable-region--

--fcc-editable-region--
}
```
