---
id: 694d8ddd106340237efb0c80
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

`splice()` メソッドは、配列の任意の位置、つまり途中でも要素を追加または削除して変更できます。`splice()` は削除された要素を含む配列を返すため、元の配列を変更せずに2つの要素を入れ替えるのに使えます。この手法は次のように1行で実行できます。

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

例を示します。

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

3行目で使われているこの手法は次のように動作します。

- `splice(3, 1, copyArray[1])` はインデックス `3`（`55`）の要素を削除します

- `copyArray[1]`（`97`）の要素をインデックス `3` に挿入します

- `splice()` は削除された要素を含む配列 `[55]` を返します

- `[0]` はその配列から `55` を取り出します

- その値（`55`）を `copyArray[1]` に代入し、入れ替えが完了します

この手法を `updatedCrew` 配列に使い、`fromIndex` と `toIndex` の宇宙飛行士を入れ替えてください。

# --hints--

`splice()` メソッドを呼び出して入れ替えを行ってください。

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

入れ替えを行うには `splice()` 配列に対して `updatedCrew` を呼び出してください。1行の手法を使う必要があります。

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

`swapCrewMembers` 関数は元の入力 `crew` 配列を変更してはいけません。

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
