---
id: 6391d1a4f7ac71efd0621380
title: レシピページのプロジェクトを作成する
challengeType: 25
dashedName: top-build-a-recipe-project
---

# --description--

ウェブサイトは、いくつかのレシピへのリンクがあるメインのインデックスページで構成されます。完成時にはウェブサイトはあまり見栄えが良くないかもしれません。

**ユーザーストーリー：**

1. レシピページには`DOCTYPE`タグが含まれている必要があります。
1. レシピページには`html`要素があり、その子要素として`head`と`body`要素が含まれている必要があります。
1. `title`要素の中に`head`要素があり、テキストは`The Odin Recipes`である必要があります。
1. `h1`要素があり、テキストは`Creamy Chocolate Fudge`である必要があります。
1. `alt`属性を持つ関連画像が表示される必要があります。
1. 画像の下に`h2`要素があり、テキストは`Description`である必要があります。
1. `Description`の下にレシピを説明するいくつかの段落が表示される必要があります。
1. `h2`要素があり、テキストは`Ingredients`である必要があります。
1. `Ingredients`見出しの下に、レシピに必要な材料の番号なしリストがある必要があります。
1. 材料のリストの下に、`Steps`という別の見出しを追加してください。
1. レシピを完成させるために必要なステップがいくつかある番号付きリストが表示される必要があります。
1. ステップの下に`h2`要素があり、テキストは`More Recipes`である必要があります。
1. 番号なしリストの中に、いくつかのリスト項目があり、その中に他のレシピへのリンクを含むアンカー要素がいくつかある必要があります。
1. これらのアンカー要素には、値が`href`に設定された`#`属性が必要です。

# --hints--

`DOCTYPE`タグが必要です。

```js
assert(code.match(/<!DOCTYPE\s+?html\s*?>/gi));
```

`html`要素があり、その中に`head`と`body`要素が含まれている必要があります。

```js
const html = document.querySelectorAll('html')[0];
const head = document.querySelectorAll('html > head')[0];
const body = document.querySelectorAll('html > body')[0];

assert(html && head && body);
```

`title`要素が`head`要素の中にあり、テキストは`The Odin Recipes`である必要があります。

```js
assert(document.querySelectorAll('HEAD > TITLE')[0]?.innerText == 'The Odin Recipes');
```

`h1`要素が`body`要素の中にあり、テキストは`Creamy Chocolate Fudge`である必要があります。

```js
assert(document.querySelectorAll('BODY > H1')[0]?.innerText == 'Creamy Chocolate Fudge');
```

`alt`属性を持つ画像が必要です。

```js
const img = document.querySelectorAll('IMG')[0];

assert(img && img.alt !='' && img.src != '')
```

`h2`要素があり、テキストは`Description`である必要があります。

```js
const h2 = document.querySelectorAll('H2')[0];

assert(h2?.innerText == 'Description');
```

レシピを説明する`p`要素が少なくとも2つ必要です。

```js
const paragraphs = document.querySelectorAll('P');

assert(paragraphs?.length > 1);
```

`h2`要素があり、テキストは`Ingredients`である必要があります。

```js
const h2 = document.querySelectorAll('H2')[1];

assert(h2?.innerText == 'Ingredients');
```

`ul`要素があり、その中に材料を含む`li`要素がネストされています。

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "Ingredients") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);
        } else {
            assert(false);
        }
    }
});
```

`h2`要素があり、テキストは`Steps`である必要があります。

```js
const h2 = document.querySelectorAll('H2')[2];

assert(h2?.innerText == 'Steps');
```

`ol`要素があり、その中にレシピのステップを含む`li`要素があります。

```js
const orderedList = document.querySelectorAll('OL')[0];
const listItems = document.querySelectorAll('OL > LI');

assert(orderedList && listItems && listItems.length > 1);
```

`h2`要素があり、テキストは`More Recipes`である必要があります。

```js
const h2 = document.querySelectorAll('H2')[3];

assert(h2?.innerText == 'More Recipes');
```

`ul`要素があり、その中に他のレシピへのリンクを含む`li`要素を持つ`a`要素がネストされています。

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "More Recipes") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);

            listItems.forEach(listItem => {
                const anchor = listItem.querySelector("A");

                assert(anchor);
            });
        } else {
            assert(false);
        }
    }
});
```

レシピにリンクするアンカー要素には、値が`href`に設定された`#`属性が必要です。

```js
const anchorTags = document.querySelectorAll("a");

const allAnchorsHaveHrefHash = [...anchorTags].every(function(anchorTag) {
  return anchorTag.hasAttribute("href") && anchorTag.getAttribute("href") === "#";
});

assert(allAnchorsHaveHrefHash && anchorTags.length > 0); 
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>The Odin Recipes</title>
  </head>
  <body>
    <h1>Creamy Chocolate Fudge</h1>
    <img src="https://cdn.freecodecamp.org/curriculum/odin-project/build-a-recipe-page/build-a-recipe-page-01.jpg" alt="A delicious chocolate fudge dessert">
    <h2>Description</h2>
    <p>This recipe is for a rich and creamy chocolate fudge that is sure to satisfy your sweet tooth. It's perfect for a special occasion or as a tasty treat for any time of the year.</p>
    <p>This recipe is easy to follow and only requires a few simple ingredients. With just a few steps, you'll be able to create a delicious dessert that everyone will love.</p>
    <h2>Ingredients</h2>
    <ul>
      <li>1 cup sugar</li>
      <li>1/2 cup unsalted butter</li>
      <li>1/4 cup milk</li>
      <li>1/4 cup cocoa powder</li>
      <li>1/4 cup chocolate chips</li>
      <li>1/4 tsp salt</li>
      <li>1 tsp vanilla extract</li>
    </ul>
    <h2>Steps</h2>
    <ol>
      <li>In a medium saucepan, melt the butter over medium heat.</li>
      <li>Add the sugar, milk, cocoa powder, and salt to the saucepan and stir until well combined.</li>
      <li>Bring the mixture to a boil, stirring constantly, and then reduce the heat to low and simmer for 5 minutes.</li>
      <li>Remove the saucepan from the heat and stir in the chocolate chips and vanilla extract until the chocolate is melted and the mixture is smooth.</li>
      <li>Pour the fudge into a greased 8-inch square pan and let it cool completely before cutting into squares.</li>
    </ol>
    <h2>More Recipes</h2>
    <ul>
      <li><a href="#">Peanut Butter Cookies</a></li>
      <li><a href="#">Lemon Bars</a></li>
      <li><a href="#">Chocolate Chip Pancakes</a></li>
    </ul>
  </body>
</html>
```

```css

```
