---
id: ab306dbdcc907c7ddfc30830
title: 깊은 배열 평탄화 도구 만들기
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

이 실습에서는 배열 평탄화 알고리즘을 구현할 것입니다.

배열 평탄화란 중첩된 배열을 깊이에 상관없이 하나의 1차원 배열로 바꾸는 것을 의미합니다. 이 과정은 배열만 풀어서 모든 요소를 순서대로 추출합니다. 다른 타입은 변경하지 않습니다.

예를 들어:

| 원본 | 평탄화된 배열 |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `steamrollArray`라는 이름의 함수를 만들어야 합니다.
1. `steamrollArray` 함수는 하나의 인수, 즉 중첩된 배열을 받아야 합니다.
1. 함수는 중첩 수준이 다양한 배열을 평탄화해야 합니다.
1. 해결책에는 `Array.prototype.flat()` 또는 `Array.prototype.flatMap()` 메서드를 사용하지 않아야 합니다.
1. 전역 변수를 사용하지 않아야 합니다.

# --hints--

`steamrollArray` 함수를 만들어야 합니다.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])`은 `["a", "b"]`를 반환해야 합니다.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])`은 `[1, 2, 3, 4]`를 반환해야 합니다.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])`은 `[1, 3, 4]`를 반환해야 합니다.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])`은 `[1, {}, 3, 4]`를 반환해야 합니다.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

해결책에는 `Array.prototype.flat()` 또는 `Array.prototype.flatMap()` 메서드를 사용하지 않아야 합니다.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

전역 변수를 사용하지 않아야 합니다.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
