---
id: 587d8247367417b2b2512c36
title: Sakinisha na Ita Helmet
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

Kufanya kazi kwenye changamoto hizi kutahusisha kuandika msimbo wako kwa kutumia mojawapo ya njia zifuatazo:

- Nakili <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">hii repo ya GitHub</a> na ukamilishe changamoto hizi kwa ndani.
- Tumia mtengenezaji wa tovuti unayempenda kukamilisha mradi. Hakikisha unajumuisha mafaili yote kutoka kwenye repo yetu ya GitHub.

Helmet hukusaidia kulinda programu zako za Express kwa kuweka vichwa mbalimbali vya HTTP.

# --instructions--

Msimbo wako wote kwa mafundisho haya unaweka kwenye faili la `myApp.js` kati ya mistari ya msimbo tuliyoanza kwa ajili yako. Tafadhali usibadilishe wala usifute msimbo tuliouongeza.

Toleo la Helmet `3.21.3` tayari limewekwa, hivyo ita kama `helmet` katika `myApp.js`.

# --hints--

Toleo la `helmet` la `3.21.3` linapaswa kuwa katika `package.json`

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```
