---
id: 587d8249367417b2b2512c40
title: Sanidi Helmet kwa kutumia programu ya kati ya ‘parent’ helmet()
challengeType: 2
forumTopicId: 301575
dashedName: configure-helmet-using-the-parent-helmet-middleware
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`app.use(helmet())` itajumuisha moja kwa moja programu zote za kati zilizotangazwa hapo juu, isipokuwa `noCache()`, na `contentSecurityPolicy()`, lakini hizi zinaweza kuwezeshwa ikiwa ni lazima. Pia unaweza kuzima au kusanidi programu nyingine yoyote ya kati kwa mtu mmoja mmoja, kwa kutumia kitu cha usanidi.

**Mfano:**

```js
app.use(helmet({
  frameguard: {         // configure
    action: 'deny'
  },
  contentSecurityPolicy: {    // enable and configure
    directives: {
      defaultSrc: ["'self'"],
      styleSrc: ['style.com'],
    }
  },
  dnsPrefetchControl: false     // disable
}))
```

Tulitambulisha kila programu ya kati kwa tofauti kwa madhumuni ya kufundisha na kwa urahisi wa uhakiki. Kutumia programu ya kati ya ‘parent’ `helmet()` ni rahisi kutekeleza katika mradi halisi.

# --hints--

hakuna mtihani - ni changamoto ya maelezo

```js
assert(true);
```
