---
id: 587d8248367417b2b2512c3d
title: Zima upataji wa awali wa DNS kwa helmet.dnsPrefetchControl()
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Ili kuboresha utendaji, vivinjari vingi hupata awali rekodi za DNS kwa viungo vilivyoko kwenye ukurasa. Kwa njia hiyo, anwani ya ip ya mwisho inajulikana tayari wakati mtumizi anapobofya kiungo. Hii inaweza kusababisha matumizi kupita kiasi ya huduma ya DNS (kama una tovuti kubwa, inayotembelewa na mamilioni ya watu…), matatizo ya faragha (mfuatiliaji mmoja anaweza kubaini kuwa uko kwenye ukurasa fulani), au mabadiliko ya takwimu za ukurasa (viungo vingine vinaweza kuonekana kutembelewa hata kama havikutembelewa). Ikiwa unahitaji usalama wa hali ya juu unaweza kuzima upataji wa awali wa DNS, kwa gharama ya kupungua kwa utendaji.

# --instructions--

Tumia njia ya `helmet.dnsPrefetchControl()` kwenye seva yako.

# --hints--

Programu ya kati helmet.dnsPrefetchControl() inapaswa kuwekwa ipasavyo

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```
