---
id: 587d8248367417b2b2512c3b
title: Zuia IE kufungua HTML isiyoaminika kwa helmet.ieNoOpen()
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliokopwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Baadhi ya programu za mtandao hutoa HTML isiyoaminika kwa ajili ya kupakuliwa. Baadhi ya matoleo ya Internet Explorer kwa chaguo-msingi hufungua mafaili hayo ya HTML katika muktadha wa tovuti yako. Hii inamaanisha kuwa ukurasa wa HTML usioaminika unaweza kuanza kufanya mambo mabaya katika muktadha wa kurasa zako. Programu ya kati hii inaweka kichwa cha X-Download-Options kuwa noopen. Hii itazuia watumiaji wa IE kutekeleza upakuaji katika muktadha wa tovuti inayotegemewa.

# --instructions--

Tumia njia ya `helmet.ieNoOpen()` kwenye seva yako.

# --hints--

Programu ya kati helmet.ieNoOpen() inapaswa kuwekwa ipasavyo

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```
