---
id: 587d8248367417b2b2512c3c
title: Waombe vivinjari kutumia tovuti yako kupitia HTTPS pekee kwa helmet.hsts()
challengeType: 2
forumTopicId: 301573
dashedName: ask-browsers-to-access-your-site-via-https-only-with-helmet-hsts
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

HTTP Strict Transport Security (HSTS) ni sera ya usalama wa mtandao inayosaidia kulinda tovuti dhidi ya mashambulizi ya kushusha itifaki na wizi wa vidakuzi (cookie hijacking). Ikiwa tovuti yako inaweza kupatikana kupitia HTTPS unaweza kuomba vivinjari vya watumiaji kuepuka kutumia HTTP isiyo salama. Kwa kuweka kichwa Strict-Transport-Security, unaambia vivinjari kutumia HTTPS kwa maombi yajayo kwa kipindi kilichobainishwa cha muda. Hii itafanya kazi kwa maombi yanayofuata baada ya ombi la awali.

# --instructions--

Sanidi `helmet.hsts()` kutumia HTTPS kwa siku 90 zijazo. Pitia kitu cha usanidi `{maxAge: timeInSeconds, force: true}`. Unaweza kuunda kigezo `ninetyDaysInSeconds = 90*24*60*60;` kutumia kwa `timeInSeconds`.

Kumbuka: Kusanidi HTTPS kwenye tovuti maalum kunahitaji ununuzi wa kikoa, na Cheti cha SSL/TLS.

# --hints--

programu ya kati helmet.hsts() inapaswa kuwekwa ipasavyo

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hsts');
assert.property(data.headers, 'strict-transport-security');
```

maxAge inapaswa kuwa sawa na 7776000 s (siku 90)

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.match(
  data.headers['strict-transport-security'],
  /^max-age=7776000;?/
);
```
