---
id: 587d8249367417b2b2512c3e
title: Zima kuhifadhi nakala upande wa mteja kwa helmet.noCache()
challengeType: 2
forumTopicId: 301576
dashedName: disable-client-side-caching-with-helmet-nocache
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Ikiwa unatoa sasisho kwa tovuti yako, na unataka watumizi kila mara wapakue toleo jipya zaidi, unaweza (jaribu) kuzima kuhifadhi nakala kwenye kivinjari cha mteja. Hii inaweza kuwa na manufaa pia katika ukuzaji. Kuhifadhi nakala kuna faida za utendaji, ambazo utapoteza, kwa hivyo tumia chaguo hili tu wakati kuna haja halisi.

# --instructions--

Tumia njia ya `helmet.noCache()` kwenye seva yako.

# --hints--

programu ya kati ya helmet.noCache() inapaswa kuwekwa ipasavyo

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nocache');
assert.equal(
  data.headers['cache-control'],
  'no-store, no-cache, must-revalidate, proxy-revalidate'
);
```
