---
id: 587d8247367417b2b2512c37
title: Ficha taarifa zinazoweza kuwa hatari kwa kutumia helmet.hidePoweredBy()
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Wahacki wanaweza kutumia udhaifu unaojulikana katika Express/Node ikiwa wataona kuwa tovuti yako inaendeshwa na Express. `X-Powered-By: Express` hutumwa katika kila ombi linalotoka Express kwa chaguo-msingi. Tumia programu ya kati ya `helmet.hidePoweredBy()` kuondoa kichwa cha X-Powered-By.

# --hints--

Programu ya kati helmet.hidePoweredBy() inapaswa kuwekwa ipasavyo

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```
