---
id: 58a25bcff9fc0f352b528e7e
title: Ficha na Linganisha Nywila kwa Msimbo wa Msimamo
challengeType: 2
forumTopicId: 301579
dashedName: hash-and-compare-passwords-synchronously
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Kuficha kwa msimbo wa msimamo ni rahisi kufanya pia lakini kunaweza kusababisha ucheleweshaji ikiwa unatumika upande wa seva kwa gharama kubwa au kuficha mara nyingi sana. Kuficha kwa njia hii ni rahisi kama wito la

```js
var hash = bcrypt.hashSync(myPlaintextPassword, saltRounds);
```

Ongeza njia hii ya kuficha kwenye msimbo wako kisha andika matokeo kwenye konsoli. Tena, vigezo vilivyotumika tayari vimefafanuliwa kwenye seva hivyo hutahitaji kubadilisha chochote. Unaweza kugundua ingawa unaficha nywila ile ile kama katika kitendakazi cha async, matokeo kwenye konsoli ni tofauti — hii ni kwa sababu chumvi hutengenezwa kwa bahati nasibu kila mara kama inavyoonekana kwenye herufi 22 za kwanza katika mfuatano wa tatu wa msimbo wa kuficha. Sasa kulinganisha ingizo la nywila na msimbo mpya wa msimamo, utatumia njia ya compareSync:

```js
var result = bcrypt.compareSync(myPlaintextPassword, hash);
```

na matokeo yatakuwa thamani ya kweli au siyo kweli.

# --instructions--

Ongeza kitendakazi na andika matokeo kwenye konsoli ili kuona kinavyofanya kazi.

Wasilisha ukurasa wako unapoona umefanya vizuri.

# --hints--

Msimbo wa msimamo unapaswa kutengenezwa na kulinganishwa kwa usahihi.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_SYNC[^]*hash.*=.*bcrypt.hashSync.*myPlaintextPassword( |),( |)saltRounds[^]*END_SYNC/gi,
  'You should call bcrypt.hashSync on myPlaintextPassword with saltRounds'
);
assert.match(
  data,
  /START_SYNC[^]*result.*=.*bcrypt.compareSync.*myPlaintextPassword( |),( |)hash[^]*END_SYNC/gi,
  'You should call bcrypt.compareSync on myPlaintextPassword with the hash generated in the last line'
);
```
