---
id: 58a25bcff9fc0f352b528e7d
title: Ficha na Linganisha Nywila kwa Njia Isiyo ya Wakati Mmoja
challengeType: 2
forumTopicId: 301578
dashedName: hash-and-compare-passwords-asynchronously
---

# --description--

Kama ukumbusho, mradi huu unajengwa juu ya mradi wa msingi uliochukuliwa kutoka <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Kwa kuwa kuficha (hashing) kunakusudiwa kuwa na mzigo mkubwa wa kihesabu, inashauriwa kufanywa kwa njia isiyo ya wakati mmoja kwenye seva yako ili kuepuka kuzuia miunganisho inayokuja wakati unaficha. Kila unachotakiwa kufanya kuficha nywila kwa njia isiyo ya wakati mmoja ni kufanya wito

```js
bcrypt.hash(myPlaintextPassword, saltRounds, (err, hash) => {
  /*Store hash in your db*/
});
```

# --instructions--

Ongeza kitendakazi hiki cha kuficha kwenye seva yako (tume tayari kufafanua vigezo vinavyotumika kwenye kitendakazi ili utumie) na kiandike kwenye konsoli ili uone! Hapo kawaida ungehifadhi hash kwenye hifadhidata yako.

Sasa unapotaka kubaini kama ingizo jipya ni data ile ile kama hash, utatumia tu kitendakazi cha kulinganisha.

```js
bcrypt.compare(myPlaintextPassword, hash, (err, res) => {
  /*res == true or false*/
});
```

Ongeza hii ndani ya kitendakazi chako cha kuficha kilichopo (kwa kuwa unahitaji kusubiri hash ikamilike kabla ya kuita kitendakazi cha kulinganisha) baada ya kuandika hash iliyokamilika na andika 'res' kwenye konsoli ndani ya kulinganisha. Unapaswa kuona kwenye konsoli hash, kisha 'true' ikachapishwa! Ukibadilisha 'myPlaintextPassword' kwenye kitendakazi cha kulinganisha kuwa 'someOtherPlaintextPassword', basi itasema false.

```js
bcrypt.hash('passw0rd!', 13, (err, hash) => {
  console.log(hash);
  //$2a$12$Y.PHPE15wR25qrrtgGkiYe2sXo98cjuMCG1YwSI5rJW1DSJp0gEYS
  bcrypt.compare('passw0rd!', hash, (err, res) => {
    console.log(res); //true
  });
});

```

Wasilisha ukurasa wako unapoona umefanya vizuri.

# --hints--

Hash isiyo ya wakati mmoja inapaswa kuzalishwa na kulinganishwa kwa usahihi.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash.*myPlaintextPassword( |),( |)saltRounds( |),( |).*err( |),( |)hash[^]*END_ASYNC/gi,
  'You should call bcrypt.hash on myPlaintextPassword and saltRounds and handle err and hash as a result in the callback'
);
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash[^]*bcrypt.compare.*myPlaintextPassword( |),( |)hash( |),( |).*err( |),( |)res[^]*}[^]*}[^]*END_ASYNC/gi,
  'Nested within the hash function should be the compare function comparing myPlaintextPassword to hash'
);
```
