---
id: 6814d8e1516e86b171929de4
title: "도전 과제 1: 모음 균형"
challengeType: 28
dashedName: challenge-1
---

# --description--

문자열이 주어졌을 때, 문자열 앞부분 절반의 모음 개수가 뒷부분 절반의 모음 개수와 같은지 확인하세요.

- 문자열에는 어떤 문자도 포함될 수 있습니다.
- `a`, `e`, `i`, `o`, `u` 문자는 대소문자 구분 없이 모음으로 간주합니다.
- 문자열 길이가 홀수일 경우 가운데 문자는 무시하세요.

# --hints--

`isBalanced("racecar")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const half = Math.floor(s.length / 2);

  let firstHalf = s.slice(0, half);
  let secondHalf = s.length % 2 === 0 ? s.slice(half) : s.slice(half + 1);

  const countVowels = str =>
    str
      .toLowerCase()
      .split('')
      .filter(c => vowels.includes(c))
      .length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
