---
id: 681cb1afdab50c87ddb2e517
title: "도전 과제 6: 애너그램 검사기"
challengeType: 28
dashedName: challenge-6
---

# --description--

두 개의 문자열이 주어졌을 때, 서로 애너그램인지(같은 문자를 어떤 순서로든 포함하는지) 판단하세요.

- 대소문자와 공백은 무시하세요.

# --hints--

`areAnagrams("listen", "silent")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
