---
id: 681cb1b0dab50c87ddb2e51b
title: "도전 과제 10: 3번의 기회"
challengeType: 28
dashedName: challenge-10
---

# --description--

1부터 10,000 사이의 정수가 주어질 때, 1부터 그 정수까지의 수 중 제곱한 값에 숫자 3이 적어도 하나 포함된 수의 개수를 반환하세요.

# --hints--

`squaresWithThree(1)`는 `0`을 반환해야 합니다.

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)`는 `1`을 반환해야 합니다.

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)`는 `19`를 반환해야 합니다.

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)`는 `326`을 반환해야 합니다.

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)`는 `4531`를 반환해야 합니다.

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
