---
id: 681cb1afdab50c87ddb2e516
title: "도전 과제 5: Jbelmud 텍스트"
challengeType: 28
dashedName: challenge-5
---

# --description--

문자열이 주어지면, 각 단어가 다음 조건에 따라 변환된 문자열을 반환하세요:

- 단어의 첫 글자와 마지막 글자는 그대로 유지합니다.
- 첫 글자와 마지막 글자 사이의 모든 글자는 알파벳 순서로 정렬합니다.
- 입력 문자열에는 구두점이 없으며 모두 소문자로만 이루어져 있습니다.

# --hints--

`jbelmu("hello world")`는 `"hello wlord"`를 반환해야 합니다.

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")`는 `"i love jbelmud text"`를 반환해야 합니다.

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")`는 `"faccdeeemorp is my faiortve pacle to laern to cdoe"`를 반환해야 합니다.

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")`는 `"the qciuk borwn fox jmpus oevr the lazy dog"`를 반환해야 합니다.

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
