---
id: 681cb1a2dab50c87ddb2e514
title: "도전 과제 3: 피보나치 수열"
challengeType: 28
dashedName: challenge-3
---

# --description--

피보나치 수열은 각 숫자가 바로 앞의 두 숫자의 합인 숫자들의 수열입니다. `0`와 `1`로 시작할 때, 수열의 처음 10개 숫자는 `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`입니다.

피보나치 수열의 처음 두 숫자를 담은 배열과 수열의 길이를 나타내는 정수가 주어질 때, 주어진 길이의 수열을 담은 배열을 반환하세요.

- 함수는 길이가 0 이상인 모든 수열을 처리해야 합니다.
- 길이가 0이면 빈 배열을 반환하세요.
- 시작 숫자들도 수열의 일부임을 기억하세요.

# --hints--

`fibonacciSequence([0, 1], 20)`는 `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`를 반환해야 합니다.

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)`는 `[21]`를 반환해야 합니다.

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)`는 `[10, 20]`를 반환해야 합니다.

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)`는 `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`를 반환해야 합니다.

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
