---
id: 681cb1b0dab50c87ddb2e519
title: "도전 과제 8: 팩토리얼 계산기"
challengeType: 28
dashedName: challenge-8
---

# --description--

0부터 20 사이의 정수가 주어지면, 그 숫자의 팩토리얼을 반환하세요. 팩토리얼은 1부터 주어진 숫자까지 모든 숫자의 곱입니다.

- 0의 팩토리얼은 1입니다.

# --hints--

`factorial(0)`는 `1`를 반환해야 합니다.

```js
assert.equal(factorial(0), 1);
```

`factorial(5)`는 `120`를 반환해야 합니다.

```js
assert.equal(factorial(5), 120);
```

`factorial(20)`는 `2432902008176640000`를 반환해야 합니다.

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
