---
id: 681cb1b0dab50c87ddb2e51a
title: "도전 과제 9: 제곱합"
challengeType: 28
dashedName: challenge-9
---

# --description--

1부터 주어진 최대 1,000까지의 양의 정수에 대해, 1부터 그 숫자까지의 모든 정수를 제곱한 값의 합을 반환하세요.

# --hints--

`sumOfSquares(5)`는 `55`을 반환해야 합니다.

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)`는 `385`을 반환해야 합니다.

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)`는 `5525`을 반환해야 합니다.

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)`는 `41791750`을 반환해야 합니다.

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)`는 `333833500`을 반환해야 합니다.

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
