---
id: 681cb05adab50c87ddb2e513
title: "도전 과제 2: 진법 검사"
challengeType: 28
dashedName: challenge-2
---

# --description--

숫자를 나타내는 문자열과 2부터 36까지의 정수 진법이 주어졌을 때, 해당 숫자가 그 진법에서 유효한지 판단하세요.

- 문자열에는 정수와 대문자 또는 소문자가 포함될 수 있습니다.
- 검사는 대소문자를 구분하지 않아야 합니다.
- 진법은 2에서 36 사이의 어떤 숫자도 될 수 있습니다.
- 모든 문자가 주어진 진법에서 유효한 숫자이면 숫자는 유효합니다.
- 진법별 유효 숫자 예시:
  - 진법 2: 0-1
  - 진법 8: 0-7
  - 진법 10: 0-9
  - 진법 16: 0-9 및 A-F
  - 진법 36: 0-9 및 A-Z

# --hints--

`isValidNumber("10101", 2)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` 는 `false` 을 반환해야 합니다.

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` 는 `false` 을 반환해야 합니다.

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` 는 `false` 을 반환해야 합니다.

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` 는 `false` 을 반환해야 합니다.

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` 는 `false` 을 반환해야 합니다.

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidNumber("z", 36))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
