---
id: 681cb1afdab50c87ddb2e515
title: "도전 과제 4: S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--

문자열이 주어지면 문자열에서 모든 공백을 제거하고, 각 문자 사이에 두 칸의 공백을 넣고, 모든 알파벳 문자를 대문자로 변환한 후 결과를 반환하세요.

- 알파벳이 아닌 문자는 변경하지 않아야 합니다(공백 제외).

# --hints--

`spaceJam("freeCodeCamp")`는 `"F  R  E  E  C  O  D  E  C  A  M  P"`를 반환해야 합니다.

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")`는 `"F  R  E  E  C  O  D  E  C  A  M  P"`를 반환해야 합니다.

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")`는 `"H  E  L  L  O  W  O  R  L  D  ?  !"`를 반환해야 합니다.

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")`는 `"C  @  T  $  &  D  0  G  $"`를 반환해야 합니다.

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")`는 `"A  L  L  Y  O  U  R  B  A  S  E"`를 반환해야 합니다.

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
