---
id: 681cb1b0dab50c87ddb2e518
title: "도전 과제 7: 목표 합계"
challengeType: 28
dashedName: challenge-7
---

# --description--

숫자 배열과 정수 목표값이 주어졌을 때, 배열에서 두 개의 고유한 숫자를 찾아 그 합이 목표값이 되도록 하세요. 두 숫자의 인덱스를 배열로 반환하거나, 두 숫자의 합이 목표값이 되지 않으면 `"Target not found"`를 반환하세요.

- 반환하는 배열은 인덱스가 오름차순이어야 합니다.

# --hints--

`findTarget([2, 7, 11, 15], 9)`는 `[0, 1]`를 반환해야 합니다.

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)`는 `[1, 2]`를 반환해야 합니다.

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)`는 `[4, 5]`를 반환해야 합니다.

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)`는 `"Target not found"`를 반환해야 합니다.

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
