---
id: 68216eb60f957572e7c340c4
title: "도전 과제 11: 마일 페이스"
challengeType: 28
dashedName: challenge-11
---

# --description--

달린 마일 수와 그 마일을 달리는 데 걸린 `"MM:SS"` (분:초) 시간이 주어지면, 각 마일을 달리는 데 걸린 평균 시간을 `"MM:SS"` 형식의 문자열로 반환하세요.

- 필요할 때는 앞에 0을 붙이세요.

# --hints--

`milePace(3, "24:00")`는 `"08:00"`를 반환해야 합니다.

```js
assert.equal(milePace(3, "24:00"), "08:00");
```

`milePace(1, "06:45")`는 `"06:45"`를 반환해야 합니다.

```js
assert.equal(milePace(1, "06:45"), "06:45");
```

`milePace(2, "07:00")`는 `"03:30"`를 반환해야 합니다.

```js
assert.equal(milePace(2, "07:00"), "03:30");
```

`milePace(26.2, "120:35")`는 `"04:36"`를 반환해야 합니다.

```js
assert.equal(milePace(26.2, "120:35"), "04:36");
```

# --seed--

## --seed-contents--

```js
function milePace(miles, duration) {

  return miles;
}
```

# --solutions--

```js
function milePace(miles, duration) {
  const [minutes, seconds] = duration.split(':').map(Number);
  const totalSeconds = minutes * 60 + seconds;
  const avgSecondsPerMile = totalSeconds / miles;

  const avgMinutes = Math.floor(avgSecondsPerMile / 60);
  const avgSeconds = Math.round(avgSecondsPerMile % 60);

  const paddedMinutes = avgMinutes.toString().padStart(2, '0');
  const paddedSeconds = avgSeconds.toString().padStart(2, '0');

  return `${paddedMinutes}:${paddedSeconds}`;
}
```
