---
id: 6821ebe4237de8297eaee794
title: "도전 과제 18: 두 번째로 비싼 것"
challengeType: 28
dashedName: challenge-18
---

# --description--

서로 다른 노트북 가격을 나타내는 정수 배열과 예산을 나타내는 정수가 주어졌을 때, 다음을 반환하세요:

1. 예산 내에 있다면 두 번째로 비싼 노트북,
2. 그렇지 않으면 예산 내에서 가장 비싼 노트북,
3. 예산 내에 노트북이 하나도 없다면 `0`.

- 중복된 가격은 무시하세요.

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)`은 `1800`를 반환해야 합니다.

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)`은 `1800`를 반환해야 합니다.

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)`는 `1899`를 반환해야 합니다.

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)`은 `0`를 반환해야 합니다.

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)`는 `1400`를 반환해야 합니다.

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
