---
id: 6821ebfd237de8297eaee799
title: "도전 과제 23: RGB를 16진수로 변환하기"
challengeType: 28
dashedName: challenge-23
---

# --description--

CSS `rgb(r, g, b)` 색상 문자열이 주어지면, 그 16진수 값을 반환하세요.

다음은 주어진 입력에 대한 몇 가지 예시 출력입니다:

| 입력   | 출력   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- 모든 문자는 소문자로 만드세요.
- `#` 다음에 6자리 문자를 반환하세요. 축약형 값은 사용하지 마세요.

# --hints--

`rgbToHex("rgb(255, 255, 255)")`는 `"#ffffff"`을 반환해야 합니다.

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")`는 `"#010b6f"`을 반환해야 합니다.

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")`는 `"#add8e6"`을 반환해야 합니다.

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")`는 `"#4f7bc9"`을 반환해야 합니다.

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
