---
id: 6821ebf8237de8297eaee798
title: "도전 과제 22: Tribonacci 수열"
challengeType: 28
dashedName: challenge-22
---

# --description--

Tribonacci 수열은 각 숫자가 바로 앞의 세 숫자의 합인 수열입니다. `0`, `0`, `1`로 시작할 때, 수열의 처음 10개 숫자는 `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24`, `44`입니다.

Tribonacci 수열의 처음 세 숫자가 들어 있는 배열과 수열의 길이를 나타내는 정수가 주어지면, 주어진 길이의 수열을 포함하는 배열을 반환하세요.

- 함수는 0 이상인 모든 길이의 수열을 처리해야 합니다.
- 길이가 0이면 빈 배열을 반환하세요.
- 시작 숫자들도 수열의 일부임을 유의하세요.

# --hints--

`tribonacciSequence([0, 0, 1], 20)`는 `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`를 반환해야 합니다.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)`는 `[21]`를 반환해야 합니다.

```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)`는 `[10, 20]`를 반환해야 합니다.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)`는 `[10, 20, 30]`를 반환해야 합니다.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)`는 `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`를 반환해야 합니다.

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
