---
id: 6821ebce237de8297eaee790
title: "도전 과제 14: 캐릭터 전투"
challengeType: 28
dashedName: challenge-14
---

# --description--

당신의 군대와 상대 군대를 나타내는 두 개의 문자열이 주어졌을 때, 각 군대의 같은 위치에 있는 캐릭터들이 다음 규칙에 따라 전투를 벌입니다:

- 문자 `a-z`는 각각 `1-26`의 힘을 가집니다.
- 문자 `A-Z`는 각각 `27-52`의 힘을 가집니다.
- 숫자 `0-9`는 자신의 숫자 값만큼 힘을 가집니다.
- 그 외 모든 문자는 힘이 0입니다.
- 각 캐릭터는 한 번의 전투만 할 수 있습니다.

각 전투에서 더 강한 캐릭터가 승리합니다. 승리 횟수가 더 많은 군대가 전쟁에서 이깁니다. 다음 값을 반환하세요:

- 당신의 군대가 상대 군대보다 캐릭터 수가 더 많으면 `"Opponent retreated"`를 반환하세요.
- 상대 군대가 당신의 군대보다 캐릭터 수가 더 많으면 `"We retreated"`를 반환하세요.
- 당신의 군대가 더 많은 전투에서 이겼으면 `"We won"`를 반환하세요.
- 상대 군대가 더 많은 전투에서 이겼으면 `"We lost"`를 반환하세요.
- 두 군대가 같은 수의 전투에서 이겼으면 `"It was a tie"`를 반환하세요.

# --hints--

`battle("Hello", "World")`는 `"We lost"`를 반환해야 합니다.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")`는 `"We won"`를 반환해야 합니다.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")`는 `"We won"`를 반환해야 합니다.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")`는 `"Opponent retreated"`를 반환해야 합니다.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")`는 `"We retreated"`를 반환해야 합니다.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")`는 `"It was a tie"`를 반환해야 합니다.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")`는 `"It was a tie"`를 반환해야 합니다.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
