---
id: 6821ebda237de8297eaee792
title: "도전 과제 16: 소괄호 뒤집기"
challengeType: 28
dashedName: challenge-16
---

# --description--

올바르게 중첩된 소괄호를 포함하는 문자열이 주어졌을 때, 다음 규칙을 사용해 문자열의 디코딩된 버전을 반환하세요:

- 각 소괄호 쌍 안의 모든 문자는 뒤집혀야 합니다.
- 최종 결과에는 소괄호가 제거되어야 합니다.
- 소괄호가 중첩된 경우, 가장 안쪽 쌍부터 먼저 뒤집고 그 결과를 바깥 쌍의 뒤집기에 포함해야 합니다.
- 모든 소괄호가 균형 있게 맞춰져 있고 올바르게 중첩되어 있다고 가정하세요.

# --hints--

`decode("(f(b(dc)e)a)")`는 `"abcdef"`를 반환해야 합니다.

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")`는 `"Can you read this?"`를 반환해야 합니다.

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")`는 `"freeCodeCamp"`를 반환해야 합니다.

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```
