---
id: 6821ebdf237de8297eaee793
title: "도전 과제 17: 연산 순서 무시하기"
challengeType: 28
dashedName: challenge-17
---

# --description--

정수 배열과 문자열 연산자 배열이 주어지면, 왼쪽에서 오른쪽으로 순서대로 숫자에 연산을 적용하세요. 모든 숫자가 사용될 때까지 필요에 따라 연산을 반복하세요. 최종 결과를 반환하세요.

예를 들어 `[1, 2, 3, 4, 5]`와 `['+', '*']`이 주어지면, 표준 연산 순서를 무시하고 왼쪽에서 오른쪽으로 `1 + 2 * 3 + 4 * 5`를 평가한 결과를 반환하세요.

- 유효한 연산자는 `+`, `-`, `*`, `/`, 그리고 `%`입니다.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])`는 `3`를 반환해야 합니다.

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])`는 `38`를 반환해야 합니다.

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])`는 `60`를 반환해야 합니다.

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])`는 `30`를 반환해야 합니다.

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])`는 `-2`를 반환해야 합니다.

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
