---
id: 6821ebd4237de8297eaee791
title: "도전 과제 15: 카멜 케이스"
challengeType: 28
dashedName: challenge-15
---

# --description--

문자열이 주어지면 다음 규칙을 사용해 카멜 케이스 버전을 반환하세요:

- 문자열 인수의 단어들은 공백(` `), 대시(`-`), 밑줄(`_`) 중 하나 이상으로 구분됩니다. 이들 연속된 문자를 단어 구분으로 처리하세요.
- 첫 번째 단어는 모두 소문자로 만드세요.
- 그 다음 단어들은 첫 글자는 대문자, 나머지는 소문자로 시작해야 합니다.
- 모든 공백과 구분자는 제거하세요.

# --hints--

`toCamelCase("hello world")`는 `"helloWorld"`를 반환해야 합니다.

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")`는 `"helloWorld"`를 반환해야 합니다.

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")`는 `"secretAgentX"`를 반환해야 합니다.

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")`는 `"freeCodeCamp"`를 반환해야 합니다.

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")`는 `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`를 반환해야 합니다.

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
