---
id: 68216ef80f957572e7c340c5
title: "도전 과제 12: 메시지 디코더"
challengeType: 28
dashedName: challenge-12
---

# --description--

비밀 메시지 문자열과 메시지를 인코딩하기 위해 문자를 이동시킨 횟수를 나타내는 정수가 주어지면, 디코딩된 문자열을 반환하세요.

- 양수는 메시지가 알파벳에서 앞으로 이동되었음을 의미합니다.
- 음수는 메시지가 알파벳에서 뒤로 이동되었음을 의미합니다.
- 대소문자를 구분하며, 디코딩된 문자는 인코딩된 문자와 같은 대소문자를 유지해야 합니다.
- 알파벳이 아닌 문자는 디코딩하지 않아야 합니다.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` should return `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` should return `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` should return `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` should return `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
