---
id: 6821ebee237de8297eaee796
title: "도전 과제 20: 배열 중복값"
challengeType: 28
dashedName: challenge-20
---

# --description--

정수 배열이 주어지면, 처음 배열에서 두 번 이상 나타나는 정수들로 이루어진 배열을 오름차순으로 정렬하여 반환하세요. 두 번 이상 나타나는 값이 없으면 빈 배열을 반환하세요.

- 반환하는 배열에는 각 값이 한 번만 포함되어야 합니다.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])`는 `[]`을 반환해야 합니다.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])`는 `[1, 2]`을 반환해야 합니다.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])`는 `[-6, 0, 2, 4, 5, 23]`를 반환해야 합니다.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
