---
id: 6821ebc9237de8297eaee78f
title: "도전 과제 13: 비정상 소수"
challengeType: 28
dashedName: challenge-13
---

# --description--

정수가 주어졌을 때, 그 숫자가 소수인지 아니면 음수 소수인지 판단하세요.

- 소수는 1보다 큰 양의 정수로, 1과 자기 자신만으로 나누어지는 수입니다.
- 음수 소수는 양의 소수의 음수 버전입니다.
- `1`와 `0`는 소수로 간주하지 않습니다.

# --hints--

`isUnnaturalPrime(1)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
