---
id: 68adce01c0e1144d0a90295a
title: "도전 과제 27: 행렬 90도 회전"
challengeType: 28
dashedName: challenge-27
---

# --description--

행렬(배열의 배열)이 주어지면, 행렬을 시계 방향으로 90도 회전시켜 반환하세요. 예를 들어, `[[1, 2], [3, 4]]`가 다음과 같다면:

| 1 | 2 |
|---|---|
| 3 | 4 |

`[[3, 1], [4, 2]]`처럼 보여야 합니다:

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

`rotate([[1]])`는 `[[1]]`를 반환해야 합니다.

```js
assert.deepEqual(rotate([[1]]), [[1]]);
```

`rotate([[1, 2], [3, 4]])`는 `[[3, 1], [4, 2]]`를 반환해야 합니다.

```js
assert.deepEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]]);
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])`는 `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`를 반환해야 합니다.

```js
assert.deepEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]]);
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])`는 `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`를 반환해야 합니다.

```js
assert.deepEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function rotate(matrix) {

  return matrix;
}
```

# --solutions--

```js
function rotate(matrix) {
  const n = matrix.length;
  const result = Array.from({ length: n }, () => Array(n).fill(0));

  for (let i = 0; i < n; i++) {
    for (let j = 0; j < n; j++) {
      result[j][n - 1 - i] = matrix[i][j];
    }
  }

  return result;
}
```
