---
id: 68adce01c0e1144d0a90295c
title: "도전 과제 28: 로마 숫자 파서"
challengeType: 28
dashedName: challenge-28
---

# --description--

로마 숫자를 나타내는 문자열이 주어지면, 그 정수 값을 반환하세요.

로마 숫자는 다음 기호와 값으로 구성됩니다:

| 기호 | 값   |
|------|------|
| I    | 1    |
| V    | 5    |
| X    | 10   |
| L    | 50   |
| C    | 100  |
| D    | 500  |
| M    | 1000 |

- 숫자는 왼쪽에서 오른쪽으로 읽습니다. 더 작은 숫자가 더 큰 숫자 앞에 나오면, 그 값을 뺍니다. 그렇지 않으면 값을 더합니다.

# --hints--

`parseRomanNumeral("III")` 는 `3` 을 반환해야 합니다.

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` 는 `4` 을 반환해야 합니다.

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` 는 `26` 을 반환해야 합니다.

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` 는 `99` 을 반환해야 합니다.

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` 는 `460` 을 반환해야 합니다.

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` 는 `504` 을 반환해야 합니다.

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` 는 `2025` 을 반환해야 합니다.

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
