---
id: 68adce01c0e1144d0a902958
title: "도전 과제 26: IPv4 유효성 검사기"
challengeType: 28
dashedName: challenge-26
---

# --description--

문자열이 주어졌을 때, 유효한 IPv4 주소인지 확인하세요. 유효한 IPv4 주소는 점(`.`)으로 구분된 네 개의 정수로 구성됩니다. 각 숫자는 다음 조건을 만족해야 합니다:

- 0부터 255 사이여야 합니다(포함).
- 선행 0이 없어야 합니다(예: 0은 허용되지만 01은 허용되지 않습니다).
- 숫자 문자만 허용됩니다.

# --hints--

`isValidIPv4("192.168.1.1")`은 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")`은 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")`은 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")`은 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")`은 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")`은 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")`은 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
