---
id: 68adce01c0e1144d0a90295e
title: "도전 과제 29: 약어 생성기"
challengeType: 28
dashedName: challenge-29
---

# --description--

하나 이상의 단어가 포함된 문자열이 주어지면, 다음 조건에 따라 단어들의 약어를 반환하세요:

- 약어는 각 단어의 첫 글자를 대문자로 사용해야 하며, 별도의 지시가 없는 한 그렇습니다.
- 약어는 주어진 문자열의 첫 단어가 아닌 경우 `a`, `for`, `an`, `and`, `by`, `of` 단어들의 첫 글자를 무시해야 합니다.
- 약어의 글자는 주어진 순서대로 반환해야 합니다.
- 약어에는 공백이 포함되어서는 안 됩니다.

# --hints--

`buildAcronym("Search Engine Optimization")`는 `"SEO"`를 반환해야 합니다.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")`는 `"FAQ"`를 반환해야 합니다.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")`는 `"NASA"`를 반환해야 합니다.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")`는 `"FBI"`를 반환해야 합니다.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")`는 `"FYI"`를 반환해야 합니다.

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")`는 `"BTW"`를 반환해야 합니다.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")`는 `"AUHWPOTIMSH"`를 반환해야 합니다.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
