---
id: 68af0687ef34c76c28ffa54d
title: "도전 과제 33: 화면 사용 시간"
challengeType: 28
dashedName: challenge-33
---

# --description--

일주일 동안의 시간을 나타내는 7개의 정수로 이루어진 입력 배열이 주어집니다. 각 정수는 해당 요일에 휴대폰 사용 시간(시간 단위)을 나타냅니다. 다음 조건에 따라 화면 사용 시간이 너무 많은지 판단하세요:

- 하루라도 10시간 이상 사용하면 너무 많습니다.
- 연속된 3일의 평균 사용 시간이 8시간 이상이면 너무 많습니다.
- 7일 전체 평균 사용 시간이 6시간 이상이면 너무 많습니다.

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` should return `false`.

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` should return `false`.

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` should return `false`.

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` should return `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` should return `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` should return `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` should return `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```
