---
id: 68af0687ef34c76c28ffa547
title: "도전 과제 30: 고유한 문자"
challengeType: 28
dashedName: challenge-30
---

# --description--

문자열이 주어졌을 때, 문자열 내 모든 문자가 고유한지 확인하세요.

- 대문자와 소문자는 서로 다른 문자로 간주해야 합니다.

# --hints--

`allUnique("abc")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
