---
id: 68af0687ef34c76c28ffa54b
title: "도전 과제 32: 문장 뒤집기"
challengeType: 28
dashedName: challenge-32
---

# --description--

단어들로 이루어진 문자열이 주어지면, 단어들의 순서를 거꾸로 하여 새로운 문자열을 반환하세요. 예를 들어, 첫 번째 단어는 반환된 문자열의 끝에 와야 하고, 마지막 단어는 반환된 문자열의 시작에 와야 합니다.

- 주어진 문자열에서 단어들은 하나 이상의 공백으로 구분될 수 있습니다.
- 반환된 문자열에서는 단어 사이에 공백이 하나만 있어야 합니다.

# --hints--

`reverseSentence("world hello")`는 `"hello world"`을 반환해야 합니다.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")`는 `"git commit push"`을 반환해야 합니다.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")`는 `"sudo install npm"`를 반환해야 합니다.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")`는 `"export function default import"`을 반환해야 합니다.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
