---
id: 68af0687ef34c76c28ffa549
title: "도전 과제 31: 배열 차이"
challengeType: 28
dashedName: challenge-31
---

# --description--

문자열 값을 가진 두 배열이 주어지면, 두 배열 중 하나에만 나타나는 모든 값을 포함하는 새 배열을 반환하세요.

- 반환된 배열은 알파벳 순서로 정렬되어야 합니다.

# --hints--

`arrayDiff(["apple", "banana"], ["apple", "banana", "cherry"])`는 `["cherry"]`를 반환해야 합니다.

```js
assert.deepEqual(arrayDiff(["apple", "banana"], ["apple", "banana", "cherry"]), ["cherry"]);
```

`arrayDiff(["apple", "banana", "cherry"], ["apple", "banana"])`는 `["cherry"]`를 반환해야 합니다.

```js
assert.deepEqual(arrayDiff(["apple", "banana", "cherry"], ["apple", "banana"]), ["cherry"]);
```

`arrayDiff(["one", "two", "three", "four", "six"], ["one", "three", "eight"])`는 `["eight", "four", "six", "two"]`를 반환해야 합니다.

```js
assert.deepEqual(arrayDiff(["one", "two", "three", "four", "six"], ["one", "three", "eight"]), ["eight", "four", "six", "two"]);
```

`arrayDiff(["two", "four", "five", "eight"], ["one", "two", "three", "four", "seven", "eight"])`는 `["five", "one", "seven", "three"]`를 반환해야 합니다.

```js
assert.deepEqual(arrayDiff(["two", "four", "five", "eight"], ["one", "two", "three", "four", "seven", "eight"]), ["five", "one", "seven", "three"]);
```

`arrayDiff(["I", "like", "freeCodeCamp"], ["I", "like", "rocks"])`는 `["freeCodeCamp", "rocks"]`를 반환해야 합니다.

```js
assert.deepEqual(arrayDiff(["I", "like", "freeCodeCamp"], ["I", "like", "rocks"]), ["freeCodeCamp", "rocks"]);
```

# --seed--

## --seed-contents--

```js
function arrayDiff(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function arrayDiff(arr1, arr2) {
  let unique1 = arr1.filter((v, i) => arr1.indexOf(v) === i);
  let unique2 = arr2.filter((v, i) => arr2.indexOf(v) === i);

  let only1 = unique1.filter(v => !unique2.includes(v));
  let only2 = unique2.filter(v => !unique1.includes(v));

  return [...only1, ...only2].sort();
}
```
