---
id: 68af0687ef34c76c28ffa54f
title: "도전 과제 34: 누락된 숫자"
challengeType: 28
dashedName: challenge-34
---

# --description--

1부터 `n`까지의 정수로 이루어진 배열이 주어질 때, 1부터 `n`까지(여기서 `n`는 주어진 배열에서 가장 큰 숫자임) 누락된 모든 정수를 배열로 반환하세요.

- 주어진 배열은 정렬되어 있지 않을 수 있고 중복된 값이 포함될 수 있습니다.
- 반환하는 배열은 오름차순이어야 합니다.
- 누락된 정수가 없으면 빈 배열을 반환하세요.

# --hints--

`findMissingNumbers([1, 3, 5])`는 `[2, 4]`를 반환해야 합니다.

```js
assert.deepEqual(findMissingNumbers([1, 3, 5]), [2, 4]);
```

`findMissingNumbers([1, 2, 3, 4, 5])`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 10])`는 `[2, 3, 4, 5, 6, 7, 8, 9]`를 반환해야 합니다.

```js
assert.deepEqual(findMissingNumbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([10, 1, 10, 1, 10, 1])`는 `[2, 3, 4, 5, 6, 7, 8, 9]`를 반환해야 합니다.

```js
assert.deepEqual(findMissingNumbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([3, 1, 4, 1, 5, 9])`는 `[2, 6, 7, 8]`를 반환해야 합니다.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])`는 `[11]`를 반환해야 합니다.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11]);
```

# --seed--

## --seed-contents--

```js
function findMissingNumbers(arr) {

  return arr;
}
```

# --solutions--

```js
function findMissingNumbers(arr) {
  if (arr.length === 0) return [];
  
  const n = Math.max(...arr);
  const set = new Set(arr);
  const result = [];

  for (let i = 1; i <= n; i++) {
    if (!set.has(i)) {
      result.push(i);
    }
  }
  return result;
}
```
