---
id: 68b06e589bf2273243814771
title: "도전 과제 36: 온도 조절기 조정기"
challengeType: 28
dashedName: challenge-36
---

# --description--

현재 방 온도와 목표 온도가 주어졌을 때, 다음 조건에 따라 방 온도를 어떻게 조절해야 하는지 나타내는 문자열을 반환하세요:

- 현재 온도가 목표 온도보다 낮으면 `"heat"`를 반환하세요.
- 현재 온도가 목표 온도보다 높으면 `"cool"`를 반환하세요.
- 현재 온도가 목표 온도와 같으면 `"hold"`를 반환하세요.

# --hints--

`adjustThermostat(68, 72)`는 `"heat"`를 반환해야 합니다.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)`는 `"cool"`를 반환해야 합니다.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)`는 `"hold"`를 반환해야 합니다.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)`는 `"heat"`를 반환해야 합니다.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)`는 `"cool"`를 반환해야 합니다.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)`는 `"hold"`를 반환해야 합니다.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
