---
id: 68b06e589bf227324381476f
title: "도전 과제 35: 단어 빈도"
challengeType: 28
dashedName: challenge-35
---

# --description--

문단이 주어지면 가장 자주 등장하는 단어 세 개를 배열로 반환하세요.

- 문단의 단어들은 공백으로 구분됩니다.
- 문단에서 대소문자를 구분하지 마세요. 예를 들어 `Hello`와 `hello`을 같은 단어로 처리하세요.
- 문단에서 구두점은 무시하세요. 구두점은 쉼표(`,`), 마침표(`.`), 느낌표(`!`)로 구성됩니다.
- 반환하는 배열의 모든 단어는 소문자여야 합니다.
- 반환하는 배열은 가장 자주 등장하는 단어가 첫 번째로 오도록 내림차순이어야 합니다.

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")`는 `["coding", "python", "in"]`를 반환해야 합니다.

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")`는 `["i", "like", "coding"]`를 반환해야 합니다.

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")`는 `["debug", "test", "deploy"]`를 반환해야 합니다.

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
