---
id: 68b06e589bf2273243814775
title: "도전 과제 38: 슬러그 생성기"
challengeType: 28
dashedName: challenge-38
---

# --description--

문자열이 주어지면 다음 조건을 사용해 URL에 적합한 버전의 문자열을 반환하세요:

- 모든 문자는 소문자여야 합니다.
- 문자, 숫자, 공백이 아닌 모든 문자는 제거해야 합니다.
- 모든 공백은 URL 인코딩된 공백 코드 `%20`로 교체해야 합니다.
- 연속된 공백은 하나의 `%20`로 교체해야 합니다.
- 반환된 문자열은 앞뒤에 `%20`가 없어야 합니다.

# --hints--

`generateSlug("helloWorld")`는 `"helloworld"`를 반환해야 합니다.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")`는 `"hello%20world"`를 반환해야 합니다.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")`는 `"helloworld"`를 반환해야 합니다.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")`는 `"hello%20world"`를 반환해야 합니다.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")`는 `"h3110%20w0r1d"`를 반환해야 합니다.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
