---
id: 68b06e589bf2273243814777
title: "도전 과제 39: 탱크 채우기"
challengeType: 28
dashedName: challenge-39
---

# --description--

연료 탱크의 크기, 현재 연료량, 갤런당 가격이 주어졌을 때, 탱크를 가득 채우는 데 드는 비용을 반환하세요.

- `tankSize`는 탱크의 총 용량(갤런 단위)입니다.
- `fuelLevel`는 탱크에 현재 남아 있는 연료량(갤런 단위)입니다.
- `pricePerGallon`는 연료 1갤런당 가격입니다.
- 반환값은 소수점 둘째 자리까지 반올림한 형식인 `"$d.dd"`이어야 합니다.

# --hints--

`costToFill(20, 0, 4.00)`는 `"$80.00"`를 반환해야 합니다.

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)`는 `"$17.50"`를 반환해야 합니다.

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)`는 `"$29.25"`를 반환해야 합니다.

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)`는 `"$0.00"`를 반환해야 합니다.

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)`는 `"$21.89"`를 반환해야 합니다.

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
