---
id: 68b06e589bf2273243814773
title: "도전 과제 37: 문장 첫 글자 대문자 변환기"
challengeType: 28
dashedName: challenge-37
---

# --description--

단락이 주어지면 각 문장의 첫 글자가 대문자인 새 단락을 반환하세요.

- 나머지 모든 문자는 그대로 유지하세요.
- 문장은 마침표(`.`), 하나 이상의 물음표(`?`), 또는 하나 이상의 느낌표(`!`)로 끝날 수 있습니다.

# --hints--

`capitalize("this is a simple sentence.")`는 `"This is a simple sentence."`를 반환해야 합니다.

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")`는 `"Hello world. How are you?"`를 반환해야 합니다.

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")`는 `"I did today's coding challenge... It was fun!!"`를 반환해야 합니다.

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")`는 `"Crazy!!!Strange???Unconventional...Sentences."`를 반환해야 합니다.

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")`는 `"There's a space before this period . Why is there a space before that period ?"`를 반환해야 합니다.

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```
