---
id: 68b1f72371a5ac895ac70a06
title: "도전 과제 42: 비디오 저장"
challengeType: 28
dashedName: challenge-42
---

# --description--

비디오 크기, 비디오 크기의 단위, 하드 드라이브 용량, 하드 드라이브 단위가 주어졌을 때, 다음 조건을 사용해 하드 드라이브에 저장할 수 있는 비디오 개수를 반환하세요:

- 비디오 크기의 단위는 바이트(`"B"`), 킬로바이트(`"KB"`), 메가바이트(`"MB"`), 기가바이트(`"GB"`)일 수 있습니다.
- 위 비디오 단위 중 하나가 아니면 `"Invalid video unit"`를 반환하세요.
- 하드 드라이브 용량 단위는 기가바이트(`"GB"`) 또는 테라바이트(`"TB"`)일 수 있습니다.
- 위 하드 드라이브 단위 중 하나가 아니면 `"Invalid drive unit"`를 반환하세요.
- 하드 드라이브에 저장할 수 있는 전체 비디오 개수를 반환하세요.
- 다음 변환을 사용하세요:

| 단위 | 대응 값   |
|:----:|:---------:|
| 1 B  |   1 B     |
| 1 KB |   1000 B  |
| 1 MB |   1000 KB |
| 1 GB |   1000 MB |
| 1 TB |   1000 GB |

예를 들어, `500`, `"MB"`, `100`, `"GB"`를 인수로 받아 100 GB 하드 드라이브에 500 MB 비디오가 몇 개 저장 가능한지 계산하세요.

# --hints--

`numberOfVideos(500, "MB", 100, "GB")`는 `200`를 반환해야 합니다.

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")`는 `"Invalid video unit"`를 반환해야 합니다.

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")`는 `"Invalid drive unit"`를 반환해야 합니다.

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")`는 `4000`를 반환해야 합니다.

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")`는 `1466`를 반환해야 합니다.

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
