---
id: 68b1f72371a5ac895ac70a0a
title: "도전 과제 44: 문자열 거울"
challengeType: 28
dashedName: challenge-44
---

# --description--

두 개의 문자열이 주어졌을 때, 두 번째 문자열이 첫 번째 문자열의 거울인지 확인하세요.

- 문자열이 거울이라고 간주되려면, 같은 문자를 역순으로 포함해야 합니다.
- 대문자와 소문자는 구분하세요.
- 알파벳이 아닌 모든 문자는 무시하세요.

# --hints--

`isMirror("helloworld", "helloworld")` 는 `false` 을 반환해야 합니다.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` 는 `false` 을 반환해야 합니다.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` 는 `false` 을 반환해야 합니다.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` 는 `true` 을 반환해야 합니다.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
