---
id: 68b1f72371a5ac895ac70a02
title: "도전 과제 40: 사진 저장"
challengeType: 28
dashedName: challenge-40
---

# --description--

사진 크기(메가바이트, MB)와 하드 드라이브 용량(기가바이트, GB)이 주어졌을 때, 다음 조건을 사용해 하드 드라이브에 저장할 수 있는 사진 수를 반환하세요:

- 1 기가바이트는 1000 메가바이트입니다.
- 드라이브에 저장할 수 있는 전체 사진 수를 반환하세요.

# --hints--

`numberOfPhotos(1, 1)`는 `1000`을 반환해야 합니다.

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)`는 `500`을 반환해야 합니다.

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)`는 `64000`을 반환해야 합니다.

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)`는 `214285`을 반환해야 합니다.

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)`는 `1571`을 반환해야 합니다.

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
