---
id: 68b1f72371a5ac895ac70a08
title: "도전 과제 43: 숫자 대 문자"
challengeType: 28
dashedName: challenge-43
---

# --description--

문자열이 주어지면, 숫자가 문자보다 많으면 `"digits"`을, 문자가 숫자보다 많으면 `"letters"`을, 숫자와 문자의 개수가 같으면 `"tie"`를 반환하세요.

- 숫자는 `0-9`로 구성됩니다.
- 문자는 대문자 또는 소문자의 `a-z`로 구성됩니다.
- 그 외의 문자는 무시하세요.

# --hints--

`digitsOrLetters("abc123")`는 `"tie"`를 반환해야 합니다.

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")`는 `"letters"`를 반환해야 합니다.

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")`는 `"digits"`를 반환해야 합니다.

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")`는 `"letters"`를 반환해야 합니다.

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")`는 `"digits"`를 반환해야 합니다.

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")`는 `"tie"`를 반환해야 합니다.

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
