---
id: 68b1f72371a5ac895ac70a04
title: "도전 과제 41: 파일 저장"
challengeType: 28
dashedName: challenge-41
---

# --description--

파일 크기, 파일 크기의 단위, 그리고 기가바이트(GB) 단위의 하드 드라이브 용량이 주어졌을 때, 다음 조건을 사용하여 하드 드라이브에 저장할 수 있는 파일 수를 반환하세요:

- 파일 크기의 단위는 바이트(`"B"`), 킬로바이트(`"KB"`), 또는 메가바이트(`"MB"`)일 수 있습니다.
- 드라이브에 저장할 수 있는 전체 파일 수를 반환하세요.
- 다음 단위 변환을 사용하세요:

| 단위 | 대응 값    |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

예를 들어, `500`, `"KB"`, `1`를 인수로 받았을 때, 1 GB 하드 드라이브에 500 KB 파일이 몇 개 저장될 수 있는지 구하세요.

# --hints--

`numberOfFiles(500, "KB", 1)`는 `2000`를 반환해야 합니다.

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)`는 `20000`를 반환해야 합니다.

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)`는 `200`를 반환해야 합니다.

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)`는 `366210`를 반환해야 합니다.

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)`는 `453514`를 반환해야 합니다.

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)`는 `166666`를 반환해야 합니다.

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
