---
id: 68b7cadffed0e75a517da675
title: "도전 과제 53: 십진수를 이진수로 변환하기"
challengeType: 28
dashedName: challenge-53
---

# --description--

음이 아닌 정수가 주어지면, 그 이진수 표현을 문자열로 반환하세요.

이진수는 숫자 `0`와 `1`만 사용하여 어떤 숫자든 나타냅니다. 십진수를 이진수로 변환하려면, 숫자를 `2`로 반복해서 나누고 나머지를 기록하세요. 숫자가 0이 될 때까지 반복합니다. 기록한 나머지를 마지막부터 처음까지 읽으면 됩니다. 예를 들어, `12`을 이진수로 변환하면:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12`의 이진수는 `1100`입니다.

# --hints--

`toBinary(5)`는 `"101"`을 반환해야 합니다.

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)`는 `"1100"`을 반환해야 합니다.

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)`는 `"110010"`을 반환해야 합니다.

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)`는 `"1100011"`을 반환해야 합니다.

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
