---
id: 68b7687dded630607aceccb3
title: "도전 과제 49: CSV 헤더 파서"
challengeType: 28
dashedName: challenge-49
---

# --description--

쉼표로 구분된 값(CSV) 파일의 첫 번째 줄이 주어지면, 제목을 포함하는 배열을 반환하세요.

- CSV 파일의 첫 번째 줄에는 쉼표로 구분된 제목들이 있습니다.
- 각 제목의 앞뒤 공백을 제거하세요.

# --hints--

`getHeadings("name,age,city")`는 `["name", "age", "city"]`을 반환해야 합니다.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")`는 `["first name", "last name", "phone"]`을 반환해야 합니다.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")`는 `["username", "email", "signup date"]`를 반환해야 합니다.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
