---
id: 68b7687dded630607aceccb1
title: "도전 과제 48: 스팸 번호 감지기"
challengeType: 28
dashedName: challenge-48
---

# --description--

다음과 같은 형식의 전화번호가 주어졌을 때 `"+A (BBB) CCC-DDDD"`, 각 문자는 다음과 같이 숫자를 나타냅니다:

- `A`은 국가 번호를 나타내며, 숫자 자릿수는 제한이 없습니다.
- `BBB`는 지역 번호를 나타내며 항상 세 자리 숫자입니다.
- `CCC`과 `DDDD`는 지역 번호의 나머지 부분으로 각각 세 자리와 네 자리 숫자입니다.

다음 기준에 따라 스팸 번호인지 판단하세요:

- 국가 번호가 2자리보다 길거나 0(`0`)으로 시작하지 않는 경우.
- 지역 번호가 900보다 크거나 200보다 작은 경우.
- 지역 번호 첫 세 자리 숫자의 합이 지역 번호 마지막 네 자리 숫자 안에 포함된 경우.
- 번호에 같은 숫자가 네 번 이상 연속으로 반복되는 경우(형식 문자는 무시).

# --hints--

`isSpam("+0 (200) 234-0182")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
