---
id: 68b7687dded630607aceccab
title: "도전 과제 45: 완전 제곱수"
challengeType: 28
dashedName: challenge-45
---

# --description--

정수가 주어졌을 때, 그것이 완전 제곱수인지 판단하세요.

- 어떤 수가 완전 제곱수라는 것은, 정수를 자기 자신과 곱해서 그 수를 만들 수 있다는 뜻입니다. 예를 들어, 9는 3을 자기 자신과 곱해서 만들 수 있으므로 완전 제곱수입니다.

# --hints--

`isPerfectSquare(9)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
