---
id: 68b7cadffed0e75a517da677
title: "도전 과제 54: P@ssw0rd 강도!"
challengeType: 28
dashedName: challenge-54
---

# --description--

비밀번호 문자열이 주어지면, 비밀번호 강도에 따라 `"weak"`, `"medium"`, 또는 `"strong"`를 반환하세요.

비밀번호는 다음 규칙에 따라 평가됩니다:

- 최소 8자 이상이어야 합니다.
- 대문자와 소문자를 모두 포함해야 합니다.
- 숫자를 최소 하나 포함해야 합니다.
- 다음 집합에서 최소 하나의 특수 문자를 포함해야 합니다: `!`, `@`, `#`, `$`, `%`, `^`, `&`, 또는 `*`.

규칙을 두 개 미만으로 충족하면 `"weak"`를 반환하세요.
규칙을 2개 또는 3개 충족하면 `"medium"`를 반환하세요.
모든 4개 규칙을 충족하면 `"strong"`를 반환하세요.

# --hints--

`checkStrength("123456")`는 `"weak"`를 반환해야 합니다.

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")`는 `"weak"`를 반환해야 합니다.

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")`는 `"weak"`를 반환해야 합니다.

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("PASSWORD")`는 `"weak"`를 반환해야 합니다.

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")`는 `"medium"`를 반환해야 합니다.

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")`는 `"medium"`를 반환해야 합니다.

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")`는 `"medium"`를 반환해야 합니다.

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")`는 `"strong"`를 반환해야 합니다.

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")`는 `"strong"`를 반환해야 합니다.

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
