---
id: 68b7687dded630607aceccaf
title: "도전 과제 47: 과속 단속"
challengeType: 28
dashedName: challenge-47
---

# --description--

차량들이 관찰된 속도를 나타내는 배열과 제한 속도를 나타내는 숫자가 주어졌을 때, 과속한 차량 수와 그 차량들의 제한 속도 초과 평균 속도를 포함하는 배열을 반환하세요.

- 과속한 차량이 없으면 `[0, 0]`를 반환하세요.

# --hints--

`speeding([50, 60, 55], 60)`는 `[0, 0]`를 반환해야 합니다.

```js
assert.deepEqual(speeding([50, 60, 55], 60), [0, 0]);
```

`speeding([58, 50, 60, 55], 55)`는 `[2, 4]`를 반환해야 합니다.

```js
assert.deepEqual(speeding([58, 50, 60, 55], 55), [2, 4]);
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)`는 `[4, 8.5]`를 반환해야 합니다.

```js
assert.deepEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5]);
```

`speeding([100, 105, 95, 102], 100)`는 `[2, 3.5]`를 반환해야 합니다.

```js
assert.deepEqual(speeding([100, 105, 95, 102], 100), [2, 3.5]);
```

`speeding([40, 45, 44, 50, 112, 39], 55)`는 `[1, 57]`를 반환해야 합니다.

```js
assert.deepEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57]);
```

# --seed--

## --seed-contents--

```js
function speeding(speeds, limit) {

  return speeds;
}
```

# --solutions--

```js
function speeding(speeds, limit) {
  let speeding = 0;
  let totalOver = 0;

  for (let speed of speeds) {
    if (speed > limit) {
      speeding++;
      totalOver += (speed - limit);
    }
  }

  if (speeding === 0) return [0, 0];

  return [speeding, totalOver / speeding];
}
```
