---
id: 68b7cadffed0e75a517da673
title: "도전 과제 52: 이진수를 십진수로 변환하기"
challengeType: 28
dashedName: challenge-52
---

# --description--

이진수를 나타내는 문자열이 주어지면, 그 십진수 값을 숫자로 반환하세요.

이진수는 숫자 `0`와 `1`만 사용하여 어떤 수든 나타냅니다. 이진수를 십진수로 변환하려면 각 숫자에 `2`의 거듭제곱을 곱한 후 모두 더하세요. 가장 오른쪽 숫자에 `2^0`을 곱하고, 그 왼쪽 숫자에는 `2^1`를 곱하는 식으로 진행합니다. 모든 숫자에 `2`의 거듭제곱을 곱한 후 결과를 모두 더하세요.

예를 들어, 이진수 `101`는 십진수로 `5`와 같습니다. 이유는 다음과 같습니다:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")`는 `5`를 반환해야 합니다.

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")`는 `10`를 반환해야 합니다.

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")`는 `18`를 반환해야 합니다.

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")`는 `85`를 반환해야 합니다.

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
