---
id: 68b7687dded630607aceccad
title: "도전 과제 46: 두 번째로 큰 수"
challengeType: 28
dashedName: challenge-46
---

# --description--

배열이 주어지면 두 번째로 큰 서로 다른 수를 반환하세요.

# --hints--

`secondLargest([1, 2, 3, 4])`는 `3`을 반환해야 합니다.

```js
assert.equal(secondLargest([1, 2, 3, 4]), 3);
```

`secondLargest([20, 139, 94, 67, 31])`는 `94`을 반환해야 합니다.

```js
assert.equal(secondLargest([20, 139, 94, 67, 31]), 94);
```

`secondLargest([2, 3, 4, 6, 6])`는 `4`을 반환해야 합니다.

```js
assert.equal(secondLargest([2, 3, 4, 6, 6]), 4);
```

`secondLargest([10, -17, 55.5, 44, 91, 0])`는 `55.5`을 반환해야 합니다.

```js
assert.equal(secondLargest([10, -17, 55.5, 44, 91, 0]), 55.5);
```

`secondLargest([1, 0, -1, 0, 1, 0, -1, 1, 0])`는 `0`을 반환해야 합니다.

```js
assert.equal(secondLargest([1, 0, -1, 0, 1, 0, -1, 1, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function secondLargest(arr) {

  return arr;
}
```

# --solutions--

```js
function secondLargest(arr) {
  const unique = [...new Set(arr)];
  unique.sort((a, b) => b - a);
  return unique[1];
}
```
