---
id: 68b7cadffed0e75a517da671
title: "도전 과제 51: 전화번호 형식 지정기"
challengeType: 28
dashedName: challenge-51
---

# --description--

11자리 숫자로 이루어진 문자열이 주어지면, 이 문자열을 `"+D (DDD) DDD-DDDD"` 형식의 전화번호로 반환하세요.

# --hints--

`formatNumber("05552340182")`은 `"+0 (555) 234-0182"`를 반환해야 합니다.

```js
assert.equal(formatNumber("05552340182"), "+0 (555) 234-0182");
```

`formatNumber("15554354792")`은 `"+1 (555) 435-4792"`를 반환해야 합니다.

```js
assert.equal(formatNumber("15554354792"), "+1 (555) 435-4792");
```

# --seed--

## --seed-contents--

```js
function formatNumber(number) {

  return number;
}
```

# --solutions--

```js
function formatNumber(number) {
  const country = number[0];
  const area = number.slice(1, 4);
  const prefix = number.slice(4, 7);
  const line = number.slice(7);

  return `+${country} (${area}) ${prefix}-${line}`;
}
```
