---
id: 68c1a929005bf54d342aa8d2
title: "도전 과제 55: 우주 주간 1일차: 별 분류"
challengeType: 28
dashedName: challenge-55
---

# --description--

10월 4일은 세계 우주 주간의 시작일입니다. 앞으로 7일 동안 천문학을 주제로 한 코딩 도전 과제가 이어집니다.

오늘의 도전 과제에서는 별의 표면 온도(켈빈, K)가 주어지고, 다음 범위에 따라 별의 분류를 결정해야 합니다:

- `"O"`: 30,000 K 이상
- `"B"`: 10,000 K - 29,999 K
- `"A"`: 7,500 K - 9,999 K
- `"F"`: 6,000 K - 7,499 K
- `"G"`: 5,200 K - 5,999 K
- `"K"`: 3,700 K - 5,199 K
- `"M"`: 0 K - 3,699 K

- 주어진 별의 분류를 반환하세요.

# --hints--

`classification(5778)`는 `"G"`을 반환해야 합니다.

```js
assert.equal(classification(5778), "G");
```

`classification(2400)`는 `"M"`을 반환해야 합니다.

```js
assert.equal(classification(2400), "M");
```

`classification(9999)`는 `"A"`을 반환해야 합니다.

```js
assert.equal(classification(9999), "A");
```

`classification(3700)`는 `"K"`을 반환해야 합니다.

```js
assert.equal(classification(3700), "K");
```

`classification(3699)`는 `"M"`을 반환해야 합니다.

```js
assert.equal(classification(3699), "M");
```

`classification(210000)`는 `"O"`을 반환해야 합니다.

```js
assert.equal(classification(210000), "O");
```

`classification(6000)`는 `"F"`을 반환해야 합니다.

```js
assert.equal(classification(6000), "F");
```

`classification(11432)`는 `"B"`을 반환해야 합니다.

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
