---
id: 68c1a929005bf54d342aa8d3
title: "도전 과제 56: 우주 주간 2일차: 외계 행성 탐색"
challengeType: 28
dashedName: challenge-56
---

# --description--

우주 주간 2일차에는 각 문자가 별의 광도 측정값을 나타내는 문자열이 주어집니다. 통과법(transit method)을 사용해 외계 행성이 탐지되었는지 판단하세요. 통과법은 행성이 별 앞을 지나가면서 관측된 광도가 줄어드는 현상입니다.

- 광도 측정값은 `0-9`와 `A-Z` 문자로만 구성되며, 각 측정값은 다음 수치에 해당합니다:
- `0-9` 문자는 `0-9` 광도 수준에 해당합니다.
- `A-Z` 문자는 `10-35` 광도 수준에 해당합니다.

별에 외계 행성이 있다고 간주하려면, 단일 측정값이 모든 측정값 평균의 80% 이하인 경우가 있어야 합니다. 예를 들어, 별의 평균 광도가 10이라면, 단일 측정값이 8 이하일 때 외계 행성이 있다고 판단합니다.

# --hints--

`hasExoplanet("665544554")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
