---
id: 68c1a929005bf54d342aa8d5
title: "도전 과제 58: 우주 주간 4일차: 착륙 지점"
challengeType: 28
dashedName: challenge-58
---

# --description--

우주 주간 4일차에는 로버의 잠재적 착륙 지점을 나타내는 숫자 행렬(배열의 배열)이 주어집니다. 다음 규칙에 따라 가장 안전한 착륙 지점을 찾으세요:

- 행렬의 각 지점은 `0-9`부터 포함하는 숫자를 가집니다.
- `0`은 잠재적 착륙 지점을 나타냅니다.
- `0`가 아닌 숫자는 착륙하기에 너무 위험합니다. 숫자가 클수록 위험도가 높습니다.
- 가장 안전한 지점은 주변 셀(대각선 제외, 최대 4개 이웃)의 위험도 합이 가장 낮은 `0` 셀로 정의됩니다.
- 경계 밖 이웃은 무시하세요(모서리와 가장자리 셀은 이웃이 적습니다).
- 가장 안전한 착륙 지점의 인덱스를 반환하세요. 항상 가장 안전한 지점은 하나뿐입니다.

예를 들어, 다음이 주어지면:

```js
[
  [1, 0],
  [2, 0]
]
```

첫 번째 배열에서 `[0, 1]`의 인덱스인 `0`를 반환하세요.

# --hints--

`findLandingSpot([[1, 0], [2, 0]])`는 `[0, 1]`를 반환해야 합니다.

```js
assert.deepEqual(findLandingSpot([[1, 0], [2, 0]]), [0, 1]);
```

`findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])`는 `[1, 1]`를 반환해야 합니다.

```js
assert.deepEqual(findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1]);
```

`findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])`는 `[2, 2]`를 반환해야 합니다.

```js
assert.deepEqual(findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2]);
```

`findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])`는 `[2, 1]`를 반환해야 합니다.

```js
assert.deepEqual(findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1]);
```

# --seed--

## --seed-contents--

```js
function findLandingSpot(matrix) {

  return matrix;
}
```

# --solutions--

```js
function findLandingSpot(matrix) {
  let bestSpot = null;
  let lowestNeighborSum = Infinity;

  for (let i = 0; i < matrix.length; i++) {
    for (let j = 0; j < matrix[i].length; j++) {

      if (matrix[i][j] === 0) {
        let currentNeighborSum = 0;

        if (i > 0) currentNeighborSum += matrix[i - 1][j];
        if (j < matrix[i].length - 1) currentNeighborSum += matrix[i][j + 1];
        if (i < matrix.length - 1) currentNeighborSum += matrix[i + 1][j];
        if (j > 0) currentNeighborSum += matrix[i][j - 1];

        if (currentNeighborSum < lowestNeighborSum) {
          lowestNeighborSum = currentNeighborSum;
          bestSpot = [i, j];
        }
      }
    }
  }

  return bestSpot;
}
```
