---
id: 68c1a929005bf54d342aa8d6
title: "도전 과제 59: 우주 주간 5일차: 골디락스 존"
challengeType: 28
dashedName: challenge-59
---

# --description--

우주 주간 5일차에는 별의 "골디락스 존"을 계산합니다. 골디락스 존은 별 주위에서 액체 상태의 물이 존재하기에 조건이 "딱 맞는" 영역입니다.

별의 질량이 주어지면, 천문 단위로 골디락스 존의 시작 거리와 끝 거리를 배열로 반환하세요.

골디락스 존을 계산하는 방법:

1. 별의 광도를 질량의 3.5제곱으로 구하세요.
2. 존의 시작 거리는 광도의 제곱근에 0.95를 곱한 값입니다.
3. 존의 끝 거리는 광도의 제곱근에 1.37을 곱한 값입니다.

- 거리는 소수점 둘째 자리까지 반올림해서 반환하세요.

예를 들어, 질량이 `1`일 때 `[0.95, 1.37]`을 반환하세요.

# --hints--

`goldilocksZone(1)`는 `[0.95, 1.37]`을 반환해야 합니다.

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)`는 `[0.28, 0.41]`을 반환해야 합니다.

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)`는 `[21.85, 31.51]`을 반환해야 합니다.

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)`는 `[9.38, 13.52]`을 반환해야 합니다.

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)`는 `[179.69, 259.13]`을 반환해야 합니다.

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
