---
id: 68c1a929005bf54d342aa8d4
title: "도전 과제 57: 우주 주간 3일차: 집으로 전화하기"
challengeType: 28
dashedName: challenge-57
---

# --description--

우주 주간 3일차에는 통신 라우트에서 자신, 위성, 그리고 고향 행성 사이의 거리(킬로미터 단위)를 나타내는 배열이 주어집니다. 다음 조건을 사용해 메시지가 라우트를 통해 목적지 행성에 도달하는 데 걸리는 시간을 계산하세요:

- 배열의 첫 번째 값은 현재 위치에서 첫 번째 위성까지의 거리입니다.
- 마지막 값을 제외한 각 값은 다음 위성까지의 거리입니다.
- 배열의 마지막 값은 이전 위성에서 고향 행성까지의 거리입니다.
- 메시지는 초당 300,000 km 속도로 이동합니다.
- 메시지가 **통과하는** 각 위성은 0.5초의 전송 지연을 추가합니다.
- 소수점 넷째 자리까지 반올림한 숫자를 반환하되, 끝의 0은 제거하세요.

# --hints--

`sendMessage([300000, 300000])`는 `2.5`을 반환해야 합니다.

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])`는 `3.0627`을 반환해야 합니다.

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])`는 `364.5`를 반환해야 합니다.

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])`는 `1674.3333`을 반환해야 합니다.

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])`는 `2.4086`를 반환해야 합니다.

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])`는 `21.1597`을 반환해야 합니다.

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
