---
id: 68c497f3aaefc9fd9f1b0e25
title: "도전 과제 61: 우주 주간 7일차: 발사 연료"
challengeType: 28
dashedName: challenge-61
---

# --description--

우주 주간 마지막 날에는 궤도에 보내려는 페이로드의 질량(킬로그램 단위)이 주어집니다. 다음 규칙을 사용해 페이로드를 궤도에 보내는 데 필요한 연료량을 계산하세요:

- 로켓은 들어 올려야 할 질량 5kg당 1kg의 연료가 필요합니다.
- 연료 자체도 질량이 있습니다. 따라서 연료를 추가하면 들어 올려야 할 질량이 증가하고, 이는 더 많은 연료를 필요로 하며, 다시 질량이 증가하는 식입니다.
- 총 필요한 연료를 계산하려면: 페이로드 질량으로 시작해 필요한 연료를 계산하고, 그 연료를 총 질량에 더한 뒤 다시 계산하세요. 추가로 필요한 연료가 1kg 미만이 될 때까지 이 과정을 반복한 후 멈추세요.
- 로켓 자체의 질량은 무시합니다. 페이로드와 그 연료를 들어 올리는 데 필요한 연료만 계산하세요.

예를 들어, 페이로드 질량이 50kg이라면, 들어 올리기 위해 10kg의 연료가 필요합니다(페이로드 / 5). 총 질량은 60kg이 되고, 이 질량을 들어 올리려면 12kg의 연료가 필요합니다(추가로 2kg). 총 질량은 62kg이 되고, 이 질량을 들어 올리려면 12.4kg의 연료가 필요합니다(추가로 0.4kg). 추가 연료가 1kg 미만이므로 여기서 멈춥니다. 총 들어 올려야 할 질량은 62.4kg이며, 이 중 50kg은 초기 페이로드, 12.4kg은 연료입니다.

- 필요한 연료량을 소수점 첫째 자리까지 반올림해 반환하세요.

# --hints--

`launchFuel(50)` should return `12.4`.

```js
assert.equal(launchFuel(50), 12.4);
```

`launchFuel(500)` should return `124.8`.

```js
assert.equal(launchFuel(500), 124.8);
```

`launchFuel(243)` should return `60.7`.

```js
assert.equal(launchFuel(243), 60.7);
```

`launchFuel(11000)` should return `2749.8`.

```js
assert.equal(launchFuel(11000), 2749.8);
```

`launchFuel(6214)` should return `1553.4`.

```js
assert.equal(launchFuel(6214), 1553.4);
```

# --seed--

## --seed-contents--

```js
function launchFuel(payload) {

  return payload;
}
```

# --solutions--

```js
function launchFuel(payload) {
  let totalMass = payload;
  let additionalFuel = totalMass / 5;
  let totalFuel = additionalFuel;

  while (additionalFuel >= 1) {
    totalMass += additionalFuel;
    additionalFuel = (totalMass / 5) - totalFuel;
    totalFuel += additionalFuel;
  }

  return Number(totalFuel.toFixed(1));
}
```
