---
id: 68c497f3aaefc9fd9f1b0e24
title: "도전 과제 60: 우주 주간 6일차: 달의 위상"
challengeType: 28
dashedName: challenge-60
---

# --description--

우주 주간 6일차에는 `"YYYY-MM-DD"` 형식의 날짜가 주어지고, 다음 규칙을 사용해 그날의 달 위상을 결정해야 합니다:

28일로 단순화한 월령 주기를 사용하며, 네 개의 같은 길이 위상으로 나눕니다:

- `"New"`: 1일차부터 7일차까지
- `"Waxing"`: 8일차부터 14일차까지
- `"Full"`: 15일차부터 21일차까지
- `"Waning"`: 22일차부터 28일차까지

28일이 지나면 주기가 1일차인 초승달부터 다시 반복됩니다.

- 주어진 날짜의 위상을 결정하기 위해 `"2000-01-06"`를 기준 초승달(주기 1일차)로 사용하세요.
- 기준 날짜 이전의 날짜는 주어지지 않습니다.
- 올바른 위상을 문자열로 반환하세요.

**참고:** 1일차는 초승달이 뜨는 날로, 마지막 초승달 이후 0일이 지난 상태를 의미합니다.

# --hints--

`moonPhase("2000-01-12")`는 `"New"`를 반환해야 합니다.

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")`는 `"Waxing"`를 반환해야 합니다.

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")`는 `"Full"`를 반환해야 합니다.

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")`는 `"Waning"`를 반환해야 합니다.

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")`는 `"New"`를 반환해야 합니다.

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
