---
id: 68c497f3aaefc9fd9f1b0e26
title: "도전 과제 62: 16진수를 10진수로 변환하기"
challengeType: 28
dashedName: challenge-62
---

# --description--

16진수(기수 16)를 나타내는 문자열이 주어지면, 10진수(기수 10) 값으로 정수를 반환하세요.

16진수는 16개의 숫자를 사용하는 숫자 체계입니다:

- `0-9`는 `0`부터 `9`까지의 값을 나타냅니다.
- `A-F`는 `10`부터 `15`까지의 값을 나타냅니다.

부분 변환 표는 다음과 같습니다:

| 16진수 | 10진수 |
|:-------:|:-------:|
| 0       | 0       |
| 1       | 1       |
| ...     | ...     |
| 9       | 9       |
| A       | 10      |
| ...     | ...     |
| F       | 15      |
| 10      | 16      |
| ...     | ...     |
| 9F      | 159     |
| A0      | 160     |
| ...     | ...     |
| FF      | 255     |
| 100     | 256     |

- 문자열에는 `0–9`와 `A–F` 문자만 포함됩니다.

# --hints--

`hexToDecimal("A")`는 `10`를 반환해야 합니다.

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")`는 `21`를 반환해야 합니다.

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")`는 `46`를 반환해야 합니다.

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")`는 `255`를 반환해야 합니다.

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")`는 `2623`를 반환해야 합니다.

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```
