---
id: 68cae5b538ff798bbd4da001
title: "도전 과제 63: 단어 대결"
challengeType: 28
dashedName: challenge-63
---

# --description--

팀과 상대 팀을 나타내는 두 문장이 주어질 때, 각 팀의 단어가 상대 팀의 해당 단어와 대결하여 다음 규칙에 따라 어느 팀이 이기는지 결정하세요:

- 주어진 문장은 항상 같은 수의 단어를 포함합니다.
- 단어는 한 칸의 공백으로 구분되며, 오직 알파벳 문자만 포함합니다.
- 각 단어의 값은 그 단어를 구성하는 글자들의 합입니다.
- `a`부터 `z`까지의 글자는 `1`부터 `26`까지의 값에 대응합니다. 예를 들어, `a`는 `1`이고, `z`는 `26`입니다.
- 대문자는 글자의 값을 두 배로 만듭니다. 예를 들어, `A`는 `2`이고, `Z`는 `52`입니다.
- 단어들은 순서대로 대결합니다: 팀의 첫 번째 단어가 상대 팀의 첫 번째 단어와 대결하고, 그 다음 단어들도 같은 방식입니다.
- 단어의 값이 상대 단어의 값보다 크면 그 단어가 이깁니다.
- 더 많은 단어가 이긴 팀이 승리합니다.

팀이 이기면 `"We win"`를, 팀이 지면 `"We lose"`를, 두 팀이 이긴 단어 수가 같으면 `"Draw"`를 반환하세요.

# --hints--

`battle("hello world", "hello word")`는 `"We win"`를 반환해야 합니다.

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")`는 `"We win"`를 반환해야 합니다.

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")`는 `"We lose"`를 반환해야 합니다.

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")`는 `"Draw"`를 반환해야 합니다.

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")`는 `"We win"`를 반환해야 합니다.

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")`는 `"We lose"`를 반환해야 합니다.

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")`는 `"Draw"`를 반환해야 합니다.

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
