---
id: 68cae5b538ff798bbd4da006
title: "도전 과제 68: 신용카드 번호 가리기"
challengeType: 28
dashedName: challenge-68
---

# --description--

신용카드 번호 문자열이 주어지면 다음 조건을 사용해 가린 버전을 반환하세요:

- 문자열은 네 자리 숫자 4세트(`0-9`)로 구성되며, 모든 세트는 공백 하나 또는 하이픈 하나(`-`)로 구분됩니다.
- 마지막 네 자리를 제외한 모든 숫자는 별표(`*`)로 바꾸세요.
- 나머지 문자는 변경하지 않고 그대로 두세요.

예를 들어, `"4012-8888-8888-1881"`가 주어지면 `"****-****-****-1881"`를 반환하세요.

# --hints--

`mask("4012-8888-8888-1881")`는 `"****-****-****-1881"`를 반환해야 합니다.

```js
assert.equal(mask("4012-8888-8888-1881"), "****-****-****-1881");
```

`mask("5105 1051 0510 5100")`는 `"**** **** **** 5100"`를 반환해야 합니다.

```js
assert.equal(mask("5105 1051 0510 5100"), "**** **** **** 5100");
```

`mask("6011 1111 1111 1117")`는 `"**** **** **** 1117"`를 반환해야 합니다.

```js
assert.equal(mask("6011 1111 1111 1117"), "**** **** **** 1117");
```

`mask("2223-0000-4845-0010")`는 `"****-****-****-0010"`를 반환해야 합니다.

```js
assert.equal(mask("2223-0000-4845-0010"), "****-****-****-0010");
```

# --seed--

## --seed-contents--

```js
function mask(card) {

  return card;
}
```

# --solutions--

```js
function mask(card) {
  const split = card.split(card.includes('-') ? '-' : ' ');
  
  return card.includes('-') ?
    `****-****-****-${split[3]}` :
    `**** **** **** ${split[3]}`;
}
```
