---
id: 68cae5b538ff798bbd4da003
title: "도전 과제 65: 문자열 개수 세기"
challengeType: 28
dashedName: challenge-65
---

# --description--

두 개의 문자열이 주어졌을 때, 두 번째 문자열이 첫 번째 문자열에 몇 번 나타나는지 알아내세요.

- 패턴 문자열은 첫 번째 문자열 안에서 겹칠 수 있습니다. 예를 들어, `"aaa"`에는 `"aa"`가 두 번 나타납니다. 첫 번째 두 `a`와 두 번째 두 `a`입니다.

# --hints--

`count('abcdefg', 'def')`는 `1`를 반환해야 합니다.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')`는 `0`를 반환해야 합니다.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')`는 `2`를 반환해야 합니다.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')`는 `3`를 반환해야 합니다.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')`는 `10`를 반환해야 합니다.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
