---
id: 68cae5b538ff798bbd4da004
title: "도전 과제 66: HTML 태그 제거기"
challengeType: 28
dashedName: challenge-66
---

# --description--

HTML 코드 문자열이 주어지면 태그를 제거하고 일반 텍스트 내용을 반환하세요.

- 입력 문자열에는 유효한 HTML만 포함됩니다.
- HTML 태그는 중첩될 수 있습니다.
- 태그와 모든 속성을 제거하세요.

예를 들어, `'<a href="#">Click here</a>'`는 `"Click here"`를 반환해야 합니다.

# --hints--

`stripTags('<a href="#">Click here</a>')`는 `"Click here"`를 반환해야 합니다.

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')`는 `"Hello World!"`를 반환해야 합니다.

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')`는 빈 문자열(`""`)을 반환해야 합니다.

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')`는 `sectionsection`를 반환해야 합니다.

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
