---
id: 68cae5b538ff798bbd4da008
title: "도전 과제 70: HTML 속성 추출기"
challengeType: 28
dashedName: challenge-70
---

# --description--

유효한 HTML 요소의 문자열이 주어지면, 다음 기준에 따라 요소의 속성을 반환하세요:

- 요소는 하나만 주어집니다.
- 속성은 `attribute="value"` 형식입니다.
- 각 속성의 이름과 값을 쉼표로 구분하여 `["attribute1, value1", "attribute2, value2"]` 형식의 문자열 배열로 반환하세요.
- 속성은 주어진 순서대로 반환하세요.
- 속성이 없으면 빈 배열을 반환하세요.

# --hints--

`extractAttributes('<span class="red"></span>')`는 `["class, red"]`를 반환해야 합니다.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')`는 `["charset, UTF-8"]`를 반환해야 합니다.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')`는 `["name, email", "type, email", "required, true"]`를 반환해야 합니다.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')`는 `["id, submit", "class, btn btn-primary"]`를 반환해야 합니다.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
