---
id: 68cae5b538ff798bbd4da009
title: "도전 과제 71: 팁 계산기"
challengeType: 28
dashedName: challenge-71
---

# --description--

식사 가격과 사용자 지정 팁 비율이 주어지면 15%, 20%, 그리고 사용자 지정 금액의 세 가지 팁 값을 배열로 반환하세요.

- 가격은 `"$N.NN"` 형식으로 제공됩니다.
- 사용자 지정 팁 비율은 `"25%"` 형식으로 제공됩니다.
- 금액은 동일한 `"$N.NN"` 형식으로 소수점 둘째 자리까지 반올림하여 반환하세요.

예를 들어, `"$10.00"` 식사 가격과 `"25%"` 사용자 지정 팁 값이 주어지면 `["$1.50", "$2.00", "$2.50"]`를 반환하세요.

# --hints--

`calculateTips("$10.00", "25%")`는 `["$1.50", "$2.00", "$2.50"]`를 반환해야 합니다.

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")`는 `["$13.45", "$17.93", "$23.31"]`를 반환해야 합니다.

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")`는 `["$2.98", "$3.97", "$1.79"]`를 반환해야 합니다.

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
