---
id: 68cae5b538ff798bbd4da002
title: "도전 과제 64: 24시간제를 12시간제로 변환하기"
challengeType: 28
dashedName: challenge-64
---

# --description--

24시간 형식의 `"HHMM"` 문자열이 주어지면, 이에 해당하는 12시간 형식의 `"H:MM AM"` 또는 `"H:MM PM"` 문자열을 반환하세요.

- 주어진 입력은 항상 `"0000"`부터 `"2359"`까지의 24시간 형식 네 자리 문자열입니다.

# --hints--

`to12("1124")`는 `"11:24 AM"`를 반환해야 합니다.

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")`는 `"9:00 AM"`를 반환해야 합니다.

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")`는 `"2:55 PM"`를 반환해야 합니다.

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")`는 `"11:46 PM"`를 반환해야 합니다.

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")`는 `"12:30 AM"`를 반환해야 합니다.

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
