---
id: 68cae5b538ff798bbd4da005
title: "도전 과제 67: 이메일 유효성 검사기"
challengeType: 28
dashedName: challenge-67
---

# --description--

문자열이 주어졌을 때, 다음 조건을 사용해 유효한 이메일 주소인지 판단하세요:

- 정확히 하나의 `@` 기호를 포함해야 합니다.
- 로컬 부분(`@` 이전):
  - 문자(`a-z`, `A-Z`), 숫자(`0-9`), 점(`.`), 밑줄(`_`), 하이픈(`-`)만 포함할 수 있습니다.
  - 점으로 시작하거나 끝나면 안 됩니다.
- 도메인 부분(`@` 이후):
  - 최소한 하나의 점을 포함해야 합니다.
  - 점 다음에 최소 두 글자의 문자로 끝나야 합니다.
- 로컬 부분과 도메인 부분 모두 연속된 두 개의 점을 포함할 수 없습니다.

# --hints--

`validate("a@b.cd")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
