---
id: 68d2ba1468508398389487cf
title: "도전 과제 74: 좋아하는 노래"
challengeType: 28
dashedName: challenge-74
---

# --description--

아이팟 기억하세요? 첫 모델은 오늘로부터 24년 전인 2001년 10월 23일에 출시되었습니다.

iPod 재생 목록을 나타내는 노래 객체 배열이 주어지면, 가장 많이 재생된 노래 두 곡의 제목을 배열로 반환하세요. 가장 많이 재생된 노래가 첫 번째에 오도록 하세요.

- 각 객체는 `"title"` 속성(문자열)과 `"plays"` 속성(정수)을 가집니다.

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])`는 `["Sync or Swim", "Earbud Blues"]`를 반환해야 합니다.

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])`는 `["Clickwheel Love", "99 Downloads"]`를 반환해야 합니다.

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])`는 `["Song B", "Song C"]`를 반환해야 합니다.

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```
