---
id: 68d2ba1468508398389487ce
title: "도전 과제 73: 현명하게 말하세요"
challengeType: 28
dashedName: challenge-73
---

# --description--

문장이 주어지면, 다음 규칙을 사용해 현명한 스승의 조언처럼 들리는 문장 버전을 반환하세요:

- 단어들은 한 칸의 공백으로 구분됩니다.
- 문장에서 다음 단어들 중 첫 번째로 나타나는 단어를 찾으세요: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- 그 단어를 포함해 그 앞에 있는 모든 단어를 문장 끝으로 옮기고:
  - 옮길 때 단어들의 순서를 유지하세요.
  - 모두 소문자로 만드세요.
  - 그리고 그 앞에 쉼표와 공백을 추가하세요.
- 새 문장의 첫 단어 첫 글자는 대문자로 만드세요.
- 모든 주어진 문장은 한 개의 구두점으로 끝납니다. 원래 문장의 구두점을 유지하고 새 문장 끝으로 옮기세요.
- 새 문장을 반환할 때, 단어 사이에 한 칸의 공백만 있고 문장 시작과 끝에는 공백이 없도록 하세요.

예를 들어, `"You must speak wisely."`가 주어지면 `"Speak wisely, you must."`를 반환하세요.

# --hints--

`wiseSpeak("You must speak wisely.")`는 `"Speak wisely, you must."`를 반환해야 합니다.

```js
assert.equal(wiseSpeak("You must speak wisely."), "Speak wisely, you must.");
```

`wiseSpeak("You can do it!")`는 `"Do it, you can!"`를 반환해야 합니다.

```js
assert.equal(wiseSpeak("You can do it!"), "Do it, you can!");
```

`wiseSpeak("Do you think you will complete this?")`는 `"Complete this, do you think you will?"`를 반환해야 합니다.

```js
assert.equal(wiseSpeak("Do you think you will complete this?"), "Complete this, do you think you will?");
```

`wiseSpeak("All your base are belong to us.")`는 `"Belong to us, all your base are."`를 반환해야 합니다.

```js
assert.equal(wiseSpeak("All your base are belong to us."), "Belong to us, all your base are.");
```

`wiseSpeak("You have much to learn.")`는 `"Much to learn, you have."`를 반환해야 합니다.

```js
assert.equal(wiseSpeak("You have much to learn."), "Much to learn, you have.");
```

# --seed--

## --seed-contents--

```js
function wiseSpeak(sentence) {

  return sentence;
}
```

# --solutions--

```js
function wiseSpeak(sentence) {
  const triggers = ["have", "must", "are", "will", "can"];
  const punctuation = sentence[sentence.length - 1];
  const words = sentence.split(" ");

  const triggerIndex = words.findIndex(w => triggers.includes(w));

  const toMove = words.slice(0, triggerIndex + 1).map(w => w.toLowerCase());
  const remaining = words.slice(triggerIndex + 1);

  const newStart = remaining.map((w, i) => {
    return i === 0 ? w[0].toUpperCase() + w.slice(1) :
      i === remaining.length - 1 ? w.slice(0, w.length - 1) : w;
  })

  return newStart.join(" ") + ", " + toMove.join(" ") + punctuation;
}
```
