---
id: 68d30845cc08266018fc46bd
title: "도전 과제 77: 기간 포맷터"
challengeType: 28
dashedName: challenge-77
---

# --description--

정수 초 단위가 주어지면, 같은 기간을 `"H:MM:SS"` 형식으로 나타내는 문자열을 반환하세요. 여기서 `"H"`은 시간 수, `"MM"`는 분 수, `"SS"`은 초 수입니다. 다음 규칙에 따라 시간을 반환하세요:

- 초: 항상 두 자리 숫자로 표시하세요.
- 분: 필요하지 않은 경우 선행 0을 생략하세요. 기간이 1분 미만이면 `"0"`를 사용하세요.
- 시간: 0보다 클 때만 포함하세요.

# --hints--

`format(500)`는 `"8:20"`를 반환해야 합니다.

```js
assert.equal(format(500), "8:20");
```

`format(4000)`는 `"1:06:40"`를 반환해야 합니다.

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)`는 `"0:01"`를 반환해야 합니다.

```js
assert.equal(format(1), "0:01");
```

`format(5555)`는 `"1:32:35"`를 반환해야 합니다.

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)`는 `"27:46:39"`를 반환해야 합니다.

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
