---
id: 68d30845cc08266018fc46be
title: "도전 과제 78: 정수 수열"
challengeType: 28
dashedName: challenge-78
---

# --description--

양의 정수가 주어지면, `1`부터 주어진 수까지 모든 정수를 숫자 순서대로 포함하는 문자열을 반환하세요.

예를 들어, `5`이 주어지면 `"12345"`를 반환하세요.

# --hints--

`sequence(5)`는 `"12345"`를 반환해야 합니다.

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)`는 `"12345678910"`를 반환해야 합니다.

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)`는 `"1"`를 반환해야 합니다.

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)`는 `"123456789101112131415161718192021222324252627"`를 반환해야 합니다.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
