---
id: 68d30845cc08266018fc46bc
title: "도전 과제 76: 보색의 DNA"
challengeType: 28
dashedName: challenge-76
---

# --description--

DNA 수열을 나타내는 문자열이 주어지면 다음 규칙을 사용해 보색 가닥을 반환하세요:

- DNA는 `"A"`, `"C"`, `"G"`, `"T"` 문자로 구성됩니다.
- `"A"`와 `"T"` 문자는 서로 보색입니다.
- `"C"`와 `"G"` 문자는 서로 보색입니다.

예를 들어, `"ACGT"`가 주어지면 `"TGCA"`를 반환하세요.

# --hints--

`complementaryDNA("ACGT")`는 `"TGCA"`를 반환해야 합니다.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")`는 `"TACGCATGCAATCG"`를 반환해야 합니다.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")`는 `"CCGAATGCTAGCTTC"`를 반환해야 합니다.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")`는 `"CTAGATCGATCCGATCGATC"`를 반환해야 합니다.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
