---
id: 68e39ed6106dac2f0a98fd62
title: "도전 과제 80: 이메일 정렬기"
challengeType: 28
dashedName: challenge-80
---

# --description--

1971년 10월 29일, 최초의 이메일이 발송되었으며, 지금도 사용하는 `username@domain` 형식을 소개했습니다. 현재는 수십억 개의 이메일 주소가 있습니다.

이 도전 과제에서는 이메일 주소 목록이 주어지고, 도메인 이름( `@` 이후 부분)을 기준으로 먼저 알파벳순으로 정렬한 다음, 사용자 이름( `@` 이전 부분)을 기준으로 정렬해야 합니다.

- 정렬은 대소문자를 구분하지 않습니다.
- 동일한 도메인을 가진 이메일이 여러 개 있으면 사용자 이름으로 정렬합니다.
- 정렬된 주소를 배열로 반환하세요.
- 반환된 주소는 원래 대소문자를 유지해야 합니다.

예를 들어, `["jill@mail.com", "john@example.com", "jane@example.com"]`가 주어지면 `["jane@example.com", "john@example.com", "jill@mail.com"]`를 반환하세요.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])`는 `["jane@example.com", "john@example.com", "jill@mail.com"]`를 반환해야 합니다.

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])`는 `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`를 반환해야 합니다.

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])`는 `["user@x.com", "user@y.com", "user@z.com"]`를 반환해야 합니다.

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])`는 `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`를 반환해야 합니다.

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])`는 `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`를 반환해야 합니다.

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```
