---
id: 68e39ed6106dac2f0a98fd64
title: "도전 과제 82: SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

변수 이름을 나타내는 문자열이 주어지면, 다음 조건에 따라 "spooky case"로 변환하세요:

- 모든 밑줄(`_`)과 하이픈(`-`)을 물결표(`~`)로 바꾸세요.
- 문자열의 첫 글자와 그 이후의 모든 홀수 번째 글자는 대문자로 만드세요. 물결표 문자는 글자 수를 셀 때 무시하세요. 나머지 글자는 모두 소문자로 만드세요.

예를 들어, `hello_world`가 주어지면 `HeLlO~wOrLd`를 반환하세요.

# --hints--

`spookify("hello_world")`는 `"HeLlO~wOrLd"`를 반환해야 합니다.

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")`는 `"SpOoKy~CaSe"`를 반환해야 합니다.

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")`는 `"TrIcK~oR~tReAt"`를 반환해야 합니다.

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")`는 `"C~a~N~d~Y~~b~O~w~L"`를 반환해야 합니다.

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")`는 `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`를 반환해야 합니다.

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
