---
id: 68e39ed6106dac2f0a98fd63
title: "도전 과제 81: N번째 소수"
challengeType: 28
dashedName: challenge-81
---

# --description--

소수는 1보다 큰 양의 정수로 1과 자기 자신만으로 나누어지는 수입니다. 처음 다섯 개 소수는 `2`, `3`, `5`, `7`, 그리고 `11`입니다.

양의 정수 `n`가 주어지면 `n`번째 소수를 반환하세요. 예를 들어, `5`가 주어지면 5번째 소수인 `11`를 반환합니다.

# --hints--

`nthPrime(5)`는 `11`를 반환해야 합니다.

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)`는 `29`를 반환해야 합니다.

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)`는 `53`를 반환해야 합니다.

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)`는 `523`를 반환해야 합니다.

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)`는 `7919`를 반환해야 합니다.

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
