---
id: 68e39ed6106dac2f0a98fd65
title: "도전 과제 83: 서명 유효성 검사"
challengeType: 28
dashedName: challenge-83
---

# --description--

문자열 메시지, 비밀 키 문자열, 그리고 서명 숫자가 주어졌을 때, 이 인코딩 방식을 사용하여 서명이 유효한지 판단하세요:

- 메시지와 비밀 키의 문자들은 다음 값을 가집니다:
  - `a`부터 `z`까지는 각각 `1`부터 `26`까지의 값을 가집니다.
  - `A`부터 `Z`까지는 각각 `27`부터 `52`까지의 값을 가집니다.
- 그 외 모든 문자는 값이 없습니다.
- 메시지의 합과 비밀 키의 합을 더해 서명을 계산하세요.

예를 들어, 메시지 `"foo"`와 비밀 키 `"bar"`가 주어지면, 서명은 `57`이 됩니다:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

계산된 서명이 주어진 서명과 일치하는지 확인하세요.

# --hints--

`verify("foo", "bar", 57)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
