---
id: 68ee9e3066cfd4eb2328e8a5
title: "도전 과제 86: 이미지 검색"
challengeType: 28
dashedName: challenge-86
---

# --description--

2001년 11월 4일, 구글은 이미지 검색을 출시하여 사람들이 검색어를 사용해 이미지를 찾을 수 있게 했습니다. 이 도전 과제에서는 이미지 검색을 모방할 것입니다.

이미지 이름 배열과 검색어가 주어지면, 검색어를 포함하는 이미지 이름 배열을 반환하세요.

- 검색어를 비교할 때 대소문자를 무시하세요.
- 입력 배열에 나타나는 순서대로 이미지를 반환하세요.

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")`는 `["dog.png"]`를 반환해야 합니다.

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")`는 `["Sunset.jpg", "sunflower.jpeg"]`를 반환해야 합니다.

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")`는 `["Moon.png", "stars.png"]`를 반환해야 합니다.

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")`는 `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`를 반환해야 합니다.

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```
