---
id: 68ee9e3066cfd4eb2328e8a8
title: "도전 과제 89: 카드 세기"
challengeType: 28
dashedName: challenge-89
---

# --description--

표준 카드 한 벌에는 각 무늬마다 13장의 고유 카드가 있습니다. 카드 덱에서 뽑을 카드 수를 나타내는 정수가 주어지면, 뽑을 수 있는 고유 카드 조합의 수를 반환하세요.

- 순서는 중요하지 않습니다. 카드 A를 먼저 뽑고 카드 B를 뽑는 것은 카드 B를 먼저 뽑고 카드 A를 뽑는 것과 같습니다.

예를 들어 `52`가 주어지면 `1`를 반환하세요. 52장 카드 덱에서 52장을 뽑는 조합은 단 하나뿐입니다. 그리고 `2`가 주어지면 `1326`를 반환하세요. 카드 덱에서 2장을 뽑을 때 가능한 카드 조합은 1326가지입니다.

# --hints--

`combinations(52)`는 `1`를 반환해야 합니다.

```js
assert.equal(combinations(52), 1);
```

`combinations(1)`는 `52`를 반환해야 합니다.

```js
assert.equal(combinations(1), 52);
```

`combinations(2)`는 `1326`를 반환해야 합니다.

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)`는 `2598960`를 반환해야 합니다.

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)`는 `15820024220`를 반환해야 합니다.

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)`는 `1326`를 반환해야 합니다.

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
