---
id: 68ee9e3066cfd4eb2328e8a7
title: "도전 과제 88: 요일 찾기"
challengeType: 28
dashedName: challenge-88
---

# --description--

`YYYY-MM-DD` 형식의 문자열 날짜가 주어지면 요일을 반환하세요.

유효한 반환 요일은 다음과 같습니다:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

시간대는 반드시 무시하세요.

# --hints--

`getWeekday("2025-11-06")`는 `Thursday`를 반환해야 합니다.

```js
assert.equal(getWeekday("2025-11-06"), "Thursday");
```

`getWeekday("1999-12-31")`는 `Friday`를 반환해야 합니다.

```js
assert.equal(getWeekday("1999-12-31"), "Friday");
```

`getWeekday("1111-11-11")`는 `Saturday`를 반환해야 합니다.

```js
assert.equal(getWeekday("1111-11-11"), "Saturday");
```

`getWeekday("2112-12-21")`는 `Wednesday`를 반환해야 합니다.

```js
assert.equal(getWeekday("2112-12-21"), "Wednesday");
```

`getWeekday("2345-10-01")`는 `Monday`를 반환해야 합니다.

```js
assert.equal(getWeekday("2345-10-01"), "Monday");
```

# --seed--

## --seed-contents--

```js
function getWeekday(dateString) {

  return dateString;
}
```

# --solutions--

```js
function getWeekday(dateString) {
  const days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"];
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  return days[date.getUTCDay()];
}
```
