---
id: 68ee9e3066cfd4eb2328e8a4
title: "도전 과제 85: 단어 개수 세기"
challengeType: 28
dashedName: challenge-85
---

# --description--

문장 문자열이 주어지면, 문장에 있는 단어의 개수를 반환하세요.

- 단어는 공백이 아닌 문자들의 수열이며, 단어들은 한 칸의 공백으로 구분됩니다.

# --hints--

`countWords("Hello world")`는 `2`을 반환해야 합니다.

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")`는 `9`을 반환해야 합니다.

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")`는 `4`를 반환해야 합니다.

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")`는 `7`을 반환해야 합니다.

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")`는 `7`를 반환해야 합니다.

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
