---
id: 68ee9e3066cfd4eb2328e8a6
title: "도전 과제 87: 행렬 생성기"
challengeType: 28
dashedName: challenge-87
---

# --description--

두 개의 정수(행 수와 열 수)가 주어지면, 주어진 크기로 0(`0`)으로 채워진 행렬(배열의 배열)을 반환하세요.

예를 들어, `2`과 `3`가 주어지면 다음을 반환하세요:

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`buildMatrix(2, 3)`는 `[[0, 0, 0], [0, 0, 0]]`를 반환해야 합니다.

```js
assert.deepEqual(buildMatrix(2, 3), [[0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(3, 2)`는 `[[0, 0], [0, 0], [0, 0]]`를 반환해야 합니다.

```js
assert.deepEqual(buildMatrix(3, 2), [[0, 0], [0, 0], [0, 0]]);
```

`buildMatrix(4, 3)`는 `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`를 반환해야 합니다.

```js
assert.deepEqual(buildMatrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(9, 1)`는 `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`를 반환해야 합니다.

```js
assert.deepEqual(buildMatrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]]);
```

# --seed--

## --seed-contents--

```js
function buildMatrix(rows, cols) {

  return rows;
}
```

# --solutions--

```js
function buildMatrix(rows, cols) {
  const matrix = [];
  for (let i = 0; i < rows; i++) {
    const row = new Array(cols).fill(0);
    matrix.push(row);
  }
  return matrix;
}
```
