---
id: 68f6587287ad1f4ad39b0c7e
title: "도전 과제 92: 확장자 추출기"
challengeType: 28
dashedName: challenge-92
---

# --description--

파일 이름을 나타내는 문자열이 주어지면, 파일의 확장자를 반환하세요.

- 확장자는 파일 이름에서 마지막 점(`.`) 이후에 나오는 부분입니다.
- 파일 이름에 점이 없거나 점으로 끝나면 `"none"`을 반환하세요.
- 확장자는 대소문자를 그대로 유지하여 반환해야 합니다.

# --hints--

`getExtension("document.txt")`는 `"txt"`을 반환해야 합니다.

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")`는 `"none"`을 반환해야 합니다.

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")`는 `"PNG"`을 반환해야 합니다.

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")`는 `"gitignore"`을 반환해야 합니다.

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")`는 `"gz"`을 반환해야 합니다.

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")`는 `"none"`을 반환해야 합니다.

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
