---
id: 68f6587287ad1f4ad39b0c81
title: "도전 과제 95: 배열 시프트"
challengeType: 28
dashedName: challenge-95
---

# --description--

배열과 배열을 몇 칸 이동할지 나타내는 정수가 주어지면, 이동된 배열을 반환하세요.

- 양수는 배열을 왼쪽으로 이동합니다.
- 음수는 배열을 오른쪽으로 이동합니다.
- 이동은 배열을 감싸면서 이루어집니다.

예를 들어, `[1, 2, 3]`과 `1`이 주어지면, 배열을 왼쪽으로 1칸 이동하여 `[2, 3, 1]`를 반환하세요.

# --hints--

`shiftArray([1, 2, 3], 1)`는 `[2, 3, 1]`를 반환해야 합니다.

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)`는 `[3, 1, 2]`를 반환해야 합니다.

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)`는 `["charlie", "alpha", "bravo"]`를 반환해야 합니다.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)`는 `["bravo", "charlie", "alpha"]`를 반환해야 합니다.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)`는 `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`를 반환해야 합니다.

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
