---
id: 68f6587287ad1f4ad39b0c85
title: "도전 과제 99: 지문 검사"
challengeType: 28
dashedName: challenge-99
---

# --description--

두 개의 문자열로 된 지문이 주어졌을 때, 다음 규칙을 사용해 두 지문이 일치하는지 판단하세요:

- 각 지문은 소문자만 포함합니다 (`a-z`).
- 두 지문이 일치한다고 간주하려면:
  - 길이가 같아야 합니다.
  - 서로 다른 문자의 수가 지문 길이의 10%를 넘지 않아야 합니다.

# --hints--

`isMatch("helloworld", "helloworld")`은 `true`를 반환해야 합니다.

```js
assert.isTrue(isMatch("helloworld", "helloworld"));
```

`isMatch("helloworld", "helloworlds")`은 `false`를 반환해야 합니다.

```js
assert.isFalse(isMatch("helloworld", "helloworlds"));
```

`isMatch("helloworld", "jelloworld")`은 `true`를 반환해야 합니다.

```js
assert.isTrue(isMatch("helloworld", "jelloworld"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")`은 `true`를 반환해야 합니다.

```js
assert.isTrue(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"));
```

`isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")`은 `true`를 반환해야 합니다.

```js
assert.isTrue(isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")`은 `false`를 반환해야 합니다.

```js
assert.isFalse(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"));
```

# --seed--

## --seed-contents--

```js
function isMatch(fingerprintA, fingerprintB) {

  return fingerprintA;
}
```

# --solutions--

```js
function isMatch(fingerprintA, fingerprintB) {
  if (fingerprintA.length !== fingerprintB.length) return false;

  const length = fingerprintA.length;
  let mismatches = 0;

  for (let i = 0; i < length; i++) {
    if (fingerprintA[i] !== fingerprintB[i]) {
      mismatches++;
      if (mismatches > length * 0.1) return false;
    }
  }

  return true;
}
```
