---
id: 68f6587287ad1f4ad39b0c82
title: "도전 과제 96: 주말인가요?"
challengeType: 28
dashedName: challenge-96
---

# --description--

`"YYYY-MM-DD"` 형식의 날짜가 주어지면 주말까지 남은 일수를 반환하세요.

- 주말은 토요일부터 시작합니다.
- 주어진 날짜가 토요일이나 일요일이면 `"It's the weekend!"`을 반환하세요.
- 그렇지 않으면 `"X days until the weekend."`를 반환하세요. 여기서 `X`은 토요일까지 남은 일수입니다.
- `X`가 `1`일 경우, 복수형인 `"day"` 대신 단수형인 `"days"`을 사용하세요.
- 계산 시 로컬 시간대는 무시해야 합니다.

# --hints--

`daysUntilWeekend("2025-11-14")`는 `"1 day until the weekend."`를 반환해야 합니다.

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")`는 `"3 days until the weekend."`를 반환해야 합니다.

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")`는 `"It's the weekend!"`를 반환해야 합니다.

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")`는 `"4 days until the weekend."`를 반환해야 합니다.

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")`는 `"5 days until the weekend."`를 반환해야 합니다.

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")`는 `"It's the weekend!"`를 반환해야 합니다.

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
