---
id: 68f6587287ad1f4ad39b0c80
title: "도전 과제 94: 이메일 서명 생성기"
challengeType: 28
dashedName: challenge-94
---

# --description--

사람의 이름, 직함, 회사에 대한 문자열이 주어지면 다음 규칙을 사용해 이메일 서명을 하나의 문자열로 반환하세요:

- 이름은 먼저 나오며, 이름의 첫 글자에 따라 접두사가 붙습니다(대소문자 구분 없음):
  - `A-I`: `>>`을 접두사로 사용하세요.
  - `J-R`: `--`을 접두사로 사용하세요.
  - `S-Z`: `::`을 접두사로 사용하세요.
- 이름 뒤에는 쉼표와 공백(`, `)이 따라야 합니다.
- 직함과 회사는 쉼표와 공백 뒤에 나오며, `" at "` (양쪽에 공백 포함)으로 구분하세요.

예를 들어, `"Quinn Waverly"`, `"Founder and CEO"`, `"TechCo"`이 주어지면 `"--Quinn Waverly, Founder and CEO at TechCo"`을 반환하세요.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")`은 `"--Quinn Waverly, Founder and CEO at TechCo"`을 반환해야 합니다.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")`는 `">>Alice Reed, Engineer at TechCo"`을 반환해야 합니다.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")`은 `"::Tina Vaughn, Developer at example.com"`을 반환해야 합니다.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")`은 `">>B. B., Product Tester at AcmeCorp"`을 반환해야 합니다.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")`은 `"::windstorm, Cloud Architect at Atmospheronics"`을 반환해야 합니다.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
