---
id: 68f6587287ad1f4ad39b0c7d
title: "도전 과제 91: 단어 찾기"
challengeType: 28
dashedName: challenge-91
---

# --description--

한 글자씩 들어있는 행렬(배열의 배열)과 찾을 단어가 주어지면, 행렬에서 단어의 시작과 끝 인덱스를 반환하세요.

- 주어진 행렬은 모두 소문자로 채워져 있습니다(`a-z`).
- 찾을 단어는 행렬에 정확히 한 번만 나타납니다.
- 찾을 단어는 다음 방향 중 하나의 직선상에 항상 있습니다:
  - 왼쪽에서 오른쪽
  - 오른쪽에서 왼쪽
  - 위에서 아래
  - 아래에서 위

예를 들어, 다음 행렬이 주어졌을 때:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

그리고 단어 `"cat"`가 주어지면, 다음을 반환하세요:

```md
[[0, 1], [2, 1]]
```

여기서 `[0, 1]`는 `"c"`(단어 시작)의 인덱스이고, `[2, 1]`는 `"t"`(단어 끝)의 인덱스입니다.

# --hints--

`findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")`는 `[[0, 1], [2, 1]]`를 반환해야 합니다.

```js
assert.deepEqual(findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]]);
```

`findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")`는 `[[0, 0], [0, 2]]`를 반환해야 합니다.

```js
assert.deepEqual(findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]]);
```

`findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")`는 `[[3, 3], [0, 3]]`를 반환해야 합니다.

```js
assert.deepEqual(findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]]);
```

`findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")`는 `[[1, 3], [1, 1]]`를 반환해야 합니다.

```js
assert.deepEqual(findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]]);
```

# --seed--

## --seed-contents--

```js
function findWord(matrix, word) {

  return matrix;
}
```

# --solutions--

```js
function findWord(matrix, word) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const len = word.length;

  const directions = [
    [0, 1],
    [0, -1],
    [1, 0],
    [-1, 0]
  ];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      for (let [dr, dc] of directions) {
        let match = true;
        for (let i = 0; i < len; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols || matrix[nr][nc] !== word[i]) {
            match = false;
            break;
          }
        }
        if (match) {
          return [
            [r, c],
            [r + dr * (len - 1), c + dc * (len - 1)]
          ];
        }
      }
    }
  }
}
```
