---
id: 68f6587287ad1f4ad39b0c7f
title: "도전 과제 93: 모음과 자음"
challengeType: 28
dashedName: challenge-93
---

# --description--

문자열이 주어지면, 문자열에 있는 모음과 자음의 개수를 배열로 반환하세요.

- 모음은 대소문자 구분 없이 `a`, `e`, `i`, `o`, `u`로 구성됩니다.
- 자음은 대소문자 구분 없이 나머지 모든 문자로 구성됩니다.
- 문자 이외의 모든 문자는 무시하세요.

예를 들어, `"Hello World"`가 주어지면 `[3, 7]`를 반환하세요.

# --hints--

`count("Hello World")`는 `[3, 7]`를 반환해야 합니다.

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")`는 `[3, 7]`를 반환해야 합니다.

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")`는 `[1, 5]`를 반환해야 합니다.

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")`는 `[5, 7]`를 반환해야 합니다.

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")`는 `[3, 7]`를 반환해야 합니다.

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")`는 `[11, 24]`를 반환해야 합니다.

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
