---
id: 68f6587287ad1f4ad39b0c7c
title: "도전 과제 90: 글자 수 제한"
challengeType: 28
dashedName: challenge-90
---

# --description--

이 도전 과제에서는 문자열이 주어지고, 그것이 소셜 미디어 게시물에 들어맞는지 판단해야 합니다. 주어진 규칙에 따라 다음 문자열을 반환하세요:

- 40자 제한 내에 들어맞으면 `"short post"`를 반환하세요.
- 40자를 초과하고 80자 제한 내에 들어맞으면 `"long post"`를 반환하세요.
- 두 제한 모두에 들어맞지 않으면 `"invalid post"`를 반환하세요.

# --hints--

`canPost("Hello world")`는 `"short post"`를 반환해야 합니다.

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")`는 `"long post"`를 반환해야 합니다.

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")`는 `"invalid post"`를 반환해야 합니다.

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```
