---
id: 68f6587287ad1f4ad39b0c84
title: "도전 과제 98: 직사각형 개수"
challengeType: 28
dashedName: challenge-98
---

# --description--

너비와 높이를 나타내는 두 개의 양의 정수가 주어졌을 때, 주어진 직사각형 안에 몇 개의 직사각형이 들어갈 수 있는지 구하세요.

- 너비와 높이가 정수인 직사각형만 계산하세요.

예를 들어, `1`와 `3`이 주어지면 `6`를 반환하세요. 1x1 직사각형 3개, 1x2 직사각형 2개, 1x3 직사각형 1개입니다.

# --hints--

`countRectangles(1, 3)`는 `6`를 반환해야 합니다.

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)`는 `18`를 반환해야 합니다.

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)`는 `3`를 반환해야 합니다.

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)`는 `150`를 반환해야 합니다.

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)`는 `12540`를 반환해야 합니다.

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
