---
id: 68f6587287ad1f4ad39b0c83
title: "도전 과제 97: 최대공약수"
challengeType: 28
dashedName: challenge-97
---

# --description--

두 개의 양의 정수가 주어지면, 그들의 최대공약수(GCD)를 반환하세요.

- 두 정수의 최대공약수는 나머지를 남기지 않고 두 수 모두를 나누는 가장 큰 수입니다.

예를 들어, `4`의 약수는 `1`, `2`, 그리고 `4`입니다. `6`의 약수는 `1`, `2`, `3`, 그리고 `6`입니다. 따라서 `4`와 `6`이 주어지면, 두 약수 집합 모두에 나타나는 가장 큰 수인 `2`을 반환하세요.

# --hints--

`gcd(4, 6)`는 `2`을 반환해야 합니다.

```js
assert.equal(gcd(4, 6), 2);
```

`gcd(20, 15)`는 `5`을 반환해야 합니다.

```js
assert.equal(gcd(20, 15), 5);
```

`gcd(13, 17)`는 `1`을 반환해야 합니다.

```js
assert.equal(gcd(13, 17), 1);
```

`gcd(654, 456)`는 `6`을 반환해야 합니다.

```js
assert.equal(gcd(654, 456), 6);
```

`gcd(3456, 4320)`는 `864`을 반환해야 합니다.

```js
assert.equal(gcd(3456, 4320), 864);
```

# --seed--

## --seed-contents--

```js
function gcd(x, y) {

  return x;
}
```

# --solutions--

```js
function gcd(x, y) {
  while (y !== 0) {
    const temp = y;
    y = x % y;
    x = temp;
  }
  return x;
}
```
