---
id: 68ffb91507a5b645769328c5
title: "도전 과제 102: 가장 긴 단어"
challengeType: 28
dashedName: challenge-102
---

# --description--

문장 문자열이 주어지면, 문장에서 가장 긴 단어를 반환하세요.

- 단어는 한 칸 띄어쓰기로 구분됩니다.
- 단어 길이에 포함되는 것은 문자(`a-z`, 대소문자 구분 없음)만 해당합니다.
- 길이가 같은 단어가 여러 개 있으면, 가장 먼저 나오는 단어를 반환하세요.
- 구두점은 제거한 상태로, 주어진 문자열에 나타난 그대로 단어를 반환하세요.

# --hints--

`longestWord("The quick red fox")`은 `"quick"`를 반환해야 합니다.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")`은 `"challenge"`를 반환해야 합니다.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")`은 `"This"`를 반환해야 합니다.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")`은 `"exclamation"`를 반환해야 합니다.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")`은 `"tie"`를 반환해야 합니다.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")`은 `"Wouldnt"`를 반환해야 합니다.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
