---
id: 68ffb91507a5b645769328c9
title: "도전 과제 106: 메시지 유효성 검사기"
challengeType: 28
dashedName: challenge-106
---

# --description--

메시지 문자열과 유효성 검사 문자열이 주어졌을 때, 메시지가 유효한지 판단하세요.

- 메시지는 메시지의 각 단어가 유효성 검사 문자열의 해당 문자로 시작할 때 유효합니다. 순서대로 검사하세요.
- 문자는 대소문자를 구분하지 않습니다.
- 메시지의 단어들은 단일 공백으로 구분됩니다.

# --hints--

`isValidMessage("hello world", "hw")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
