---
id: 68ffb91507a5b645769328c7
title: "도전 과제 104: 레시피 배율 조정기"
challengeType: 28
dashedName: challenge-104
---

# --description--

레시피 재료 배열과 레시피 배율 숫자가 주어지면, 해당 배율에 맞게 양이 조정된 배열을 반환하세요.

- 주어진 배열의 각 항목은 `"quantity unit ingredient"` 형식의 문자열입니다. 예를 들어 `"2 C Flour"`가 있습니다.
- 주어진 숫자로 양을 배율 조정하세요. 끝에 불필요한 0은 포함하지 말고 단위는 변환하지 마세요.
- 조정된 항목들은 주어진 순서대로 반환하세요.

# --hints--

`scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2)`는 `["4 C Flour", "3 T Sugar"]`를 반환해야 합니다.

```js
assert.deepEqual(scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2), ["4 C Flour", "3 T Sugar"]);
```

`scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5)`는 `["6 T Flour", "1.5 C Milk", "3 T Oil"]`를 반환해야 합니다.

```js
assert.deepEqual(scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5), ["6 T Flour", "1.5 C Milk", "3 T Oil"]);
```

`scaleRecipe(["3 C Milk", "2 C Oats"], 0.5)`는 `["1.5 C Milk", "1 C Oats"]`를 반환해야 합니다.

```js
assert.deepEqual(scaleRecipe(["3 C Milk", "2 C Oats"], 0.5), ["1.5 C Milk", "1 C Oats"]);
```

`scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5)`는 `["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]`를 반환해야 합니다.

```js
assert.deepEqual(scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5), ["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]);
```

# --seed--

## --seed-contents--

```js
function scaleRecipe(ingredients, scale) {

  return ingredients;
}
```

# --solutions--

```js
function scaleRecipe(ingredients, scale) {
  return ingredients.map(item => {
    const [quantityStr, unit, ...rest] = item.split(" ");
    const ingredientName = rest.join(" ");
    
    const quantity = parseFloat(quantityStr) * scale;
    
    return `${quantity} ${unit} ${ingredientName}`;
  });
}
```
