---
id: 68ffb91507a5b645769328ca
title: "도전 과제 107: FizzBuzz"
challengeType: 28
dashedName: challenge-107
---

# --description--

정수 (`n`)가 주어지면 `1`부터 `n`까지(포함) 정수 배열을 반환하세요. 다음 조건에 따라 숫자를 바꾸세요:

- 3의 배수는 `"Fizz"`로 바꾸세요.
- 5의 배수는 `"Buzz"`로 바꾸세요.
- 3과 5의 배수는 `"FizzBuzz"`로 바꾸세요.

# --hints--

`fizzBuzz(2)`는 `[1, 2]`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzz(2), [1, 2]);
```

`fizzBuzz(4)`는 `[1, 2, "Fizz", 4]`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzz(4), [1, 2, "Fizz", 4]);
```

`fizzBuzz(8)`는 `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8]`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzz(8), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8]);
```

`fizzBuzz(20)`는 `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzz(20), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]);
```

`fizzBuzz(50)`는 `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzz(50), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzz(n) {
  const result = [];
  for (let i = 1; i <= n; i++) {
    if (i % 15 === 0) {
      result.push("FizzBuzz");
    } else if (i % 3 === 0) {
      result.push("Fizz");
    } else if (i % 5 === 0) {
      result.push("Buzz");
    } else {
      result.push(i);
    }
  }
  return result;
}
```
