---
id: 68ffb91507a5b645769328c6
title: "도전 과제 103: 최소공배수"
challengeType: 28
dashedName: challenge-103
---

# --description--

두 정수가 주어지면, 두 수의 최소공배수(LCM)를 반환하세요.

두 수의 최소공배수는 두 수 모두의 배수인 가장 작은 양의 정수입니다. 예를 들어 `4`와 `6`이 주어지면 `12`를 반환하세요. 이유는 다음과 같습니다:

- `4`의 배수는 `4`, `8`, `12` 등입니다.
- `6`의 배수는 `6`, `12`, `18` 등입니다.
- `12`은 두 수 모두의 배수인 가장 작은 수입니다.

# --hints--

`lcm(4, 6)`는 `12`을 반환해야 합니다.

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)`는 `18`을 반환해야 합니다.

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)`는 `100`을 반환해야 합니다.

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)`는 `221`을 반환해야 합니다.

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)`는 `630`을 반환해야 합니다.

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```
