---
id: 68ffb91507a5b645769328c8
title: "도전 과제 105: 문자 개수 세기"
challengeType: 28
dashedName: challenge-105
---

# --description--

문장 문자열이 주어지면, 알파벳 순서대로 각 문자의 개수를 담은 배열을 반환하세요.

- 대문자와 소문자는 같은 문자로 취급하세요.
- 숫자, 공백, 구두점 등은 무시하세요.
- 개수와 문자는 `"letter count"` 형식으로 반환하세요. 예를 들어, `"a 3"`.
- 반환하는 모든 문자는 소문자여야 합니다.
- 주어진 문자열에 없는 문자의 개수는 반환하지 마세요.

# --hints--

`countCharacters("hello world")`는 `["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]`를 반환해야 합니다.

```js
assert.deepEqual(countCharacters("hello world"), ["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]);
```

`countCharacters("I love coding challenges!")`는 `["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]`를 반환해야 합니다.

```js
assert.deepEqual(countCharacters("I love coding challenges!"), ["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]);
```

`countCharacters("// TODO: Complete this challenge ASAP!")`는 `["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]`를 반환해야 합니다.

```js
assert.deepEqual(countCharacters("// TODO: Complete this challenge ASAP!"), ["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]);
```

# --seed--

## --seed-contents--

```js
function countCharacters(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countCharacters(sentence) {
  const counts = {};

  for (let char of sentence.toLowerCase()) {
    if (char >= "a" && char <= "z") {
      counts[char] = (counts[char] || 0) + 1;
    }
  }

  return Object.keys(counts)
               .sort()
               .map(letter => `${letter} ${counts[letter]}`);
}
```
