---
id: 68ffb91507a5b645769328cc
title: "도전 과제 109: 내 나이는 몇 살일까요?"
challengeType: 28
dashedName: challenge-109
---

# --description--

`YYYY-MM-DD` 형식으로 주어진 생일 날짜를 기준으로 2025년 11월 27일 현재 나이를 반환하세요.

- 모든 생일은 2025년 11월 27일 이전의 유효한 날짜라고 가정하세요.
- 나이는 정수로 반환하세요.
- 2025년에 이미 생일이 지났는지 여부를 반드시 고려하세요.

# --hints--

`calculateAge("2000-11-20")`은 `25`를 반환해야 합니다.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")`은 `24`를 반환해야 합니다.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")`은 `11`를 반환해야 합니다.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")`은 `31`를 반환해야 합니다.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")`은 `30`를 반환해야 합니다.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
