---
id: 69162d64f96574d9bb629efc
title: "도전 과제 111: 공의 궤적"
challengeType: 28
dashedName: challenge-111
---

# --description--

오늘 도전 과제는 1972년 11월 29일에 출시된 비디오 게임 Pong에서 영감을 받았습니다.

공의 위치(`2`)와 공의 이전 위치(`1`)가 포함된 행렬(배열의 배열)이 주어질 때, 공의 다음 위치에 해당하는 행렬 인덱스를 반환하세요.

- 공은 항상 직선으로 움직입니다.
- 이동 방향은 공이 `1`에서 `2`로 이동한 방식에 따라 결정됩니다.
- 행렬의 가장자리는 벽으로 간주합니다. 공이 다음에 부딪히면:
  - 위쪽 또는 아래쪽 벽에 부딪히면 수직 방향을 반대로 하여 튕깁니다.
  - 왼쪽 또는 오른쪽 벽에 부딪히면 수평 방향을 반대로 하여 튕깁니다.
  - 모서리에 부딪히면 두 방향 모두 반대로 하여 튕깁니다.

# --hints--

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])`는 `[2, 3]`를 반환해야 합니다.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3]);
```

`getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])`는 `[3, 0]`를 반환해야 합니다.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0]);
```

`getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])`는 `[1, 2]`를 반환해야 합니다.

```js
assert.deepEqual(getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])`는 `[1, 1]`를 반환해야 합니다.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])`는 `[2, 2]`를 반환해야 합니다.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2]);
```

# --seed--

## --seed-contents--

```js
function getNextLocation(matrix) {

  return matrix;
}
```

# --solutions--

```js
function getNextLocation(matrix) {
  let prev = null;
  let curr = null;

  for (let row = 0; row < matrix.length; row++) {
    for (let col = 0; col < matrix[0].length; col++) {
      if (matrix[row][col] === 1) prev = [row, col];
      if (matrix[row][col] === 2) curr = [row, col];
    }
  }

  const [prevRow, prevCol] = prev;
  const [currRow, currCol] = curr;

  let dirX = currCol - prevCol;
  let dirY = currRow - prevRow;

  let nextRow = currRow + dirY;
  let nextCol = currCol + dirX;

  const maxRow = matrix.length - 1;
  const maxCol = matrix[0].length - 1;


  if (nextCol < 0 || nextCol > maxCol) dirX *= -1;
  if (nextRow < 0 || nextRow > maxRow) dirY *= -1;

  nextRow = currRow + dirY;
  nextCol = currCol + dirX;

  return [nextRow, nextCol];
}
```
