---
id: 69162d64f96574d9bb629eff
title: "도전 과제 114: 카멜 표기법을 스네이크 표기법으로 변환하기"
challengeType: 28
dashedName: challenge-114
---

# --description--

카멜 표기법으로 된 문자열이 주어지면, 다음 규칙을 사용해 스네이크 표기법 버전의 문자열을 반환하세요:

- 입력 문자열은 오직 문자(`A-Z`와 `a-z`)만 포함하며 항상 소문자로 시작합니다.
- 카멜 표기법 문자열의 모든 대문자는 새로운 단어의 시작입니다.
- 모든 문자를 소문자로 변환하세요.
- 단어는 밑줄(`_`)로 구분하세요.

# --hints--

`toSnake("helloWorld")`는 `"hello_world"`를 반환해야 합니다.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")`는 `"my_variable_name"`를 반환해야 합니다.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")`는 `"freecodecamp_daily_challenges"`를 반환해야 합니다.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
