---
id: 69162d64f96574d9bb629f03
title: "도전 과제 118: 날짜 형식 변환기"
challengeType: 28
dashedName: challenge-118
---

# --description--

`"Month day, year"` 형식으로 주어진 날짜를 `"YYYY-MM-DD"` 형식으로 반환하세요.

- 주어진 월은 영어로 된 전체 월 이름입니다. 예를 들어: `"January"`, `"February"` 등입니다.
- 반환값에서는 월과 일이 두 자리 숫자가 되도록 필요하면 앞에 0을 채우세요.

예를 들어, `"December 6, 2025"`가 주어지면 `"2025-12-06"`를 반환하세요.

# --hints--

`formatDate("December 6, 2025")`는 `"2025-12-06"`를 반환해야 합니다.

```js
assert.equal(formatDate("December 6, 2025"), "2025-12-06");
```

`formatDate("January 1, 2000")`는 `"2000-01-01"`를 반환해야 합니다.

```js
assert.equal(formatDate("January 1, 2000"), "2000-01-01");
```

`formatDate("November 11, 1111")`는 `"1111-11-11"`를 반환해야 합니다.

```js
assert.equal(formatDate("November 11, 1111"), "1111-11-11");
```

`formatDate("September 7, 512")`는 `"512-09-07"`를 반환해야 합니다.

```js
assert.equal(formatDate("September 7, 512"), "512-09-07");
```

`formatDate("May 4, 1950")`는 `"1950-05-04"`를 반환해야 합니다.

```js
assert.equal(formatDate("May 4, 1950"), "1950-05-04");
```

`formatDate("February 29, 1992")`는 `"1992-02-29"`를 반환해야 합니다.

```js
assert.equal(formatDate("February 29, 1992"), "1992-02-29");
```

# --seed--

## --seed-contents--

```js
function formatDate(dateString) {

  return dateString;
}
```

# --solutions--

```js
function formatDate(dateString) {
  const months = {
    January: "01",
    February: "02",
    March: "03",
    April: "04",
    May: "05",
    June: "06",
    July: "07",
    August: "08",
    September: "09",
    October: "10",
    November: "11",
    December: "12"
  };

  const [month, day, year] = dateString.replace(",", "").split(" ");

  return `${year}-${months[month]}-${day.padStart(2, "0")}`;
}
```
