---
id: 69162d64f96574d9bb629efb
title: "도전 과제 110: 단어 추측기"
challengeType: 28
dashedName: challenge-110
---

# --description--

길이가 같은 두 문자열, 비밀 단어와 추측 단어가 주어졌을 때, 다음 규칙에 따라 추측 단어를 비밀 단어와 비교하세요:

- 비밀 단어와 추측 단어는 모두 대문자(`"A"`부터 `"Z"`까지)로만 구성됩니다;
- 추측 단어의 각 글자에 대해, 비밀 단어와 어떻게 일치하는지에 따라 숫자로 바꾸세요:
  - `"2"`: 글자가 비밀 단어에 있고 정확한 위치에 있을 때.
  - `"1"`: 글자가 비밀 단어에 있지만 위치가 다를 때.
  - `"0"`: 글자가 비밀 단어에 없을 때.
- 비밀 단어의 각 글자는 최대 한 번만 사용할 수 있습니다.
- 정확히 일치하는 글자(`"2"`)를 먼저 할당하고, 그 다음 왼쪽에서 오른쪽으로 부분 일치(`"1"`)를 할당하세요.
- 추측 단어에 같은 글자가 여러 번 나올 경우, 비밀 단어에 그 글자가 나타난 횟수만큼만 일치할 수 있습니다.

예를 들어, 비밀 단어가 `"APPLE"`이고 추측 단어가 `"POPPA"`일 때, `"10201"`를 반환하세요:

첫 번째 `"P"`는 올바른 위치에 있지 않으므로(`"1"`), `"O"`는 비밀 단어에 없으므로(`"0"`), 두 번째 `"P"`는 올바른 위치에 있으므로(`"2"`), 세 번째 `"P"`는 두 개의 `"0"`가 이미 사용되어 0(`"P"`)이고, `"A"`는 올바른 위치에 있지 않습니다(`"1"`).

# --hints--

`compare("APPLE", "POPPA")`는 `"10201"`를 반환해야 합니다.

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")`는 `"11221"`를 반환해야 합니다.

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")`는 `"000000"`를 반환해야 합니다.

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")`는 `"0000222222"`를 반환해야 합니다.

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")`는 `"110200"`를 반환해야 합니다.

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")`는 `"10021000"`를 반환해야 합니다.

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```
