---
id: 69162d64f96574d9bb629f01
title: "도전 과제 116: 순열 개수"
challengeType: 28
dashedName: challenge-116
---

# --description--

문자열이 주어지면, 그 문자열의 문자들로 만들 수 있는 서로 다른 순열의 개수를 반환하세요.

- 순열은 문자열의 문자들을 재배열한 모든 경우입니다.
- 중복된 순열은 세지 마세요.
- 문자열에 반복된 문자가 있으면, 반복된 배열은 한 번만 세야 합니다.
- 문자열은 오직 문자(`A-Z`, `a-z`)만 포함합니다.

예를 들어, `"abb"`가 주어지면 `3`을 반환하세요. 문자를 배열하는 고유한 방법이 세 가지이기 때문입니다: `"abb"`, `"bab"`, 그리고 `"bba"`.

# --hints--

`countPermutations("abb")`는 `3`을 반환해야 합니다.

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")`는 `6`을 반환해야 합니다.

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")`는 `630`를 반환해야 합니다.

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")`는 `39916800`를 반환해야 합니다.

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
