---
id: 69162d64f96574d9bb629f02
title: "도전 과제 117: 대칭 차집합"
challengeType: 28
dashedName: challenge-117
---

# --description--

두 개의 배열이 주어지면, 그들의 대칭 차집합을 포함하는 새 배열을 반환하세요.

- 두 집합의 대칭 차집합은 두 집합 중 하나에는 있지만 둘 다에는 없는 값들의 집합입니다.
- 입력 배열에서 처음 나타난 순서대로 값을 반환하세요.

# --hints--

`difference([1, 2, 3], [3, 4, 5])`는 `[1, 2, 4, 5]`을 반환해야 합니다.

```js
assert.deepEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5]);
```

`difference(["a", "b"], ["c", "b"])`는 `["a", "c"]`을 반환해야 합니다.

```js
assert.deepEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"]);
```

`difference([1, "a", 2], [2, "b", "a"])`는 `[1, "b"]`를 반환해야 합니다.

```js
assert.deepEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"]);
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])`는 `[2, 4, 6, 8]`을 반환해야 합니다.

```js
assert.deepEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8]);
```

# --seed--

## --seed-contents--

```js
function difference(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function difference(arr1, arr2) {
  const diff = [];

  for (const v of arr1) {

    if (!diff.includes(v) && !arr2.includes(v)) diff.push(v)
  }
  
  for (const v of arr2) {
    if (!diff.includes(v) && !arr1.includes(v)) diff.push(v)
  }

  return diff;
}
```
