---
id: 69162d64f96574d9bb629f00
title: "도전 과제 115: 마크다운 순서 있는 목록 항목 변환기"
challengeType: 28
dashedName: challenge-115
---

# --description--

마크다운에서 순서 있는 목록 항목을 나타내는 문자열이 주어지면, 이에 해당하는 HTML 문자열을 반환하세요.

유효한 마크다운 순서 있는 목록 항목은 다음 조건을 만족해야 합니다:

- 0개 이상의 공백으로 시작하고,
- 그다음에 숫자(1 이상)와 마침표(`.`)가 나오며,
- 최소 한 칸 이상의 공백이 있고,
- 그 뒤에 목록 항목 텍스트가 옵니다.

문자열이 위 형식과 정확히 일치하지 않으면 `"Invalid format"`을 반환하세요. 그렇지 않으면 목록 항목 텍스트를 `li` 태그로 감싸서 문자열을 반환하세요.

예를 들어, `"1. My item"`이 주어지면 `"<li>My item</li>"`를 반환하세요.

참고: 콘솔은 문자열 내 HTML 태그를 로그할 때 표시하지 않을 수 있습니다. 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`convertListItem("1. My item")`는 `"<li>My item</li>"`를 반환해야 합니다.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")`는 `"<li>Another item</li>"`를 반환해야 합니다.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")`는 `"Invalid format"`를 반환해야 합니다.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")`는 `"<li>list item text</li>"`를 반환해야 합니다.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")`는 `"Invalid format"`를 반환해야 합니다.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")`는 `"Invalid format"`를 반환해야 합니다.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
