---
id: 69162d64f96574d9bb629f04
title: "도전 과제 119: 문자열 압축"
challengeType: 28
dashedName: challenge-119
---

# --description--

문자열 sentence가 주어지면, 연속해서 중복된 단어를 해당 단어 뒤에 반복 횟수를 소괄호로 표시하여 교체한 압축된 문장을 반환하세요.

- 오직 연속된 중복만 압축합니다.
- 단어는 한 칸의 공백으로 구분됩니다.

예를 들어 `"yes yes yes please"`가 주어지면 `"yes(3) please"`를 반환하세요.

# --hints--

`compressString("yes yes yes please")`는 `"yes(3) please"`를 반환해야 합니다.

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")`는 `"I have(3) apples"`를 반환해야 합니다.

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")`는 `"one(2) three and to the(4)"`를 반환해야 합니다.

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")`는 `"route(6) tee(6)"`를 반환해야 합니다.

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
