---
id: 69162d64f96574d9bb629efd
title: "도전 과제 112: AI 감지기"
challengeType: 28
dashedName: challenge-112
---

# --description--

오늘의 도전 과제는 2022년 11월 30일 ChatGPT 출시에서 영감을 받았습니다.

하나 이상의 문장으로 이루어진 문자열이 주어졌을 때, 다음 규칙을 사용해 AI가 생성했을 가능성이 있는지 판단하세요:

- 두 개 이상의 대시(`-`)가 포함되어 있습니다.
- 두 개 이상의 괄호 세트(`()`)가 포함되어 있습니다. 괄호 안에 텍스트가 있을 수 있습니다.
- 7자 이상인 단어가 세 개 이상 포함되어 있습니다.

- 단어는 한 칸의 공백으로 구분되며 오직 문자(`A-Z`)로만 구성됩니다. 구두점이나 기타 문자가 단어의 일부로 포함되지 않도록 하세요.

주어진 문장이 위 규칙 중 하나라도 충족하면 `"AI"`를 반환하고, 그렇지 않으면 `"Human"`를 반환하세요.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")`는 `"Human"`를 반환해야 합니다.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")`는 `"Human"`를 반환해야 합니다.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")`는 `"AI"`를 반환해야 합니다.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")`는 `"AI"`를 반환해야 합니다.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")`는 `"AI"`를 반환해야 합니다.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
