---
id: 69162d64f96574d9bb629efe
title: "도전 과제 113: 마일을 킬로미터로 변환하기"
challengeType: 28
dashedName: challenge-113
---

# --description--

숫자로 주어진 마일 거리를 킬로미터로 환산하여 반환하세요.

- 입력은 항상 0 이상의 숫자입니다.
- 1마일은 1.60934킬로미터입니다.
- 결과는 소수점 둘째 자리까지 반올림하세요.
- 반올림한 결과에서 불필요한 끝의 0은 제거하세요.

# --hints--

`convertToKm(1)`는 `1.61`을 반환해야 합니다.

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)`는 `33.8`을 반환해야 합니다.

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)`는 `5.63`을 반환해야 합니다.

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)`는 `0`을 반환해야 합니다.

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)`는 `1`을 반환해야 합니다.

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
