---
id: 691b559495c5cb5a37b9b480
title: "도전 과제 120: 파운드를 킬로그램으로 변환"
challengeType: 28
dashedName: challenge-120
---

# --description--

숫자로 주어진 파운드 단위 무게를 받아 문자열 `"(lbs) pounds equals (kgs) kilograms."`을 반환하세요.

- `"(lbs)"`을 입력 숫자로 바꾸세요.
- `"(kgs)"`를 킬로그램으로 변환한 값으로 바꾸되, 소수점 둘째 자리까지 반올림하고 항상 두 자리 소수점을 포함하세요.
- 1파운드는 0.453592킬로그램입니다.
- 입력이 `1`일 경우 `"pound"` 대신 `"pounds"`를 사용하세요.
- 변환된 값이 `1`일 경우 `"kilogram"` 대신 `"kilograms"`를 사용하세요.

# --hints--

`convertToKgs(1)`는 `"1 pound equals 0.45 kilograms."`을 반환해야 합니다.

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)`는 `"0 pounds equals 0.00 kilograms."`을 반환해야 합니다.

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)`는 `"100 pounds equals 45.36 kilograms."`을 반환해야 합니다.

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)`는 `"2.5 pounds equals 1.13 kilograms."`을 반환해야 합니다.

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)`는 `"2.20462 pounds equals 1.00 kilogram."`을 반환해야 합니다.

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
