---
id: 691b559495c5cb5a37b9b486
title: "도전 과제 126: 첫 글자만 대문자로"
challengeType: 28
dashedName: challenge-126
---

# --description--

문자열 title이 주어지면 다음 규칙에 따라 제목 형식으로 새 문자열을 반환하세요:

- 각 단어의 첫 글자는 대문자로 만드세요.
- 각 단어의 나머지 글자는 모두 소문자로 만드세요.
- 단어들은 항상 한 칸의 공백으로 구분됩니다.

# --hints--

`titleCase("hello world")`는 `"Hello World"`을 반환해야 합니다.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")`는 `"The Quick Brown Fox"`을 반환해야 합니다.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")`는 `"Javascript And Python"`을 반환해야 합니다.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")`는 `"Avocado Toast For Breakfast"`을 반환해야 합니다.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
