---
id: 691b559495c5cb5a37b9b488
title: "도전 과제 128: 자음 개수 세기"
challengeType: 28
dashedName: challenge-128
---

# --description--

문자열과 목표 숫자가 주어졌을 때, 문자열에 정확히 목표 숫자만큼 자음이 포함되어 있는지 확인하세요.

- 자음은 대소문자 구분 없이 `"a"`, `"e"`, `"i"`, `"o"`, `"u"`를 제외한 모든 알파벳 문자입니다.
- 숫자, 구두점, 공백 및 기타 비문자 문자는 세지 마세요.

# --hints--

`hasConsonantCount("helloworld", 7)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
