---
id: 691b559495c5cb5a37b9b484
title: "도전 과제 124: 재고 업데이트"
challengeType: 28
dashedName: challenge-124
---

# --description--

가게의 재고를 나타내는 2차원 배열과 받은 배송을 나타내는 또 다른 2차원 배열이 주어질 때, 업데이트된 재고를 반환하세요.

- 배열의 각 요소는 `[quantity, "item"]` 형식을 가지며, 여기서 `quantity`은 정수이고 `"item"`는 문자열입니다.
- 배송에서 일치하는 항목이 있으면 그 수량을 더해 재고를 업데이트하세요.
- 받은 항목이 현재 재고에 없으면, 재고 끝에 새 항목으로 추가하세요.
- 재고는 주어진 순서대로 반환하되, 새 항목은 배송에 나타난 순서대로 끝에 추가하세요.

예를 들어, 재고가 `[[2, "apples"], [5, "bananas"]]`이고 배송이 `[[1, "apples"], [3, "bananas"]]`일 때, `[[3, "apples"], [8, "bananas"]]`를 반환하세요.

# --hints--

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])`는 `[[3, "apples"], [8, "bananas"]]`를 반환해야 합니다.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]]);
```

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])`는 `[[3, "apples"], [8, "bananas"], [4, "oranges"]]`를 반환해야 합니다.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]]);
```

`updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])`는 `[[10, "apples"], [30, "bananas"], [20, "oranges"]]`를 반환해야 합니다.

```js
assert.deepEqual(updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]]);
```

`updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])`는 `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]`를 반환해야 합니다.

```js
assert.deepEqual(updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]);
```

# --seed--

## --seed-contents--

```js
function updateInventory(inventory, shipment) {
  return inventory;
}
```

# --solutions--

```js
function updateInventory(inventory, shipment) {
  const inventoryMap = new Map();
  inventory.forEach(([qty, item], index) => {
    inventoryMap.set(item, index);
  });

  shipment.forEach(([qty, item]) => {
    if (inventoryMap.has(item)) {
      const index = inventoryMap.get(item);
      inventory[index][0] += qty;
    } else {
      inventory.push([qty, item]);
      inventoryMap.set(item, inventory.length - 1);
    }
  });

  return inventory;
}
```
