---
id: 691b559495c5cb5a37b9b483
title: "도전 과제 123: 로마 숫자 생성기"
challengeType: 28
dashedName: challenge-123
---

# --description--

정수를 입력받아 해당 값을 로마 숫자로 변환해 반환하세요.

로마 숫자는 다음 기호들을 사용합니다:

| 기호 | 값    |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

로마 숫자는 가장 큰 값부터 가장 작은 값 순서로 왼쪽에서 오른쪽으로 작성하며 다음 규칙을 따릅니다:

- 기호 다음에 같은 값이거나 더 작은 값의 기호가 오면 덧셈을 사용합니다. 예를 들어 `18`는 `XVIII`로 작성합니다 (10 + 5 + 1 + 1 + 1 = 18).
- 더 작은 기호가 더 큰 기호 앞에 오면 뺄셈을 사용하여 4나 9를 나타냅니다. 예를 들어 19는 `XIX`로 작성합니다 (10 + (10 - 1)).
- 같은 기호를 연속으로 세 번 이상 반복할 수 없습니다. 만약 세 번 이상 반복해야 한다면 뺄셈을 사용합니다. 따라서 작성할 수 있는 가장 큰 숫자는 3999입니다.

추가 예시로, `1464`가 주어지면 `"MCDLXIV"`를 반환하세요 (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`toRoman(18)`는 `"XVIII"`를 반환해야 합니다.

```js
assert.equal(toRoman(18), "XVIII");
```

`toRoman(19)`는 `"XIX"`를 반환해야 합니다.

```js
assert.equal(toRoman(19), "XIX");
```

`toRoman(1464)`는 `"MCDLXIV"`를 반환해야 합니다.

```js
assert.equal(toRoman(1464), "MCDLXIV");
```

`toRoman(2025)`는 `"MMXXV"`를 반환해야 합니다.

```js
assert.equal(toRoman(2025), "MMXXV");
```

`toRoman(3999)`는 `"MMMCMXCIX"`를 반환해야 합니다.

```js
assert.equal(toRoman(3999), "MMMCMXCIX");
```

# --seed--

## --seed-contents--

```js
function toRoman(num) {
  return num;
}
```

# --solutions--

```js
function toRoman(num) {
  const symbols = [
    ["M", 1000],
    ["CM", 900],
    ["D", 500],
    ["CD", 400],
    ["C", 100],
    ["XC", 90],
    ["L", 50],
    ["XL", 40],
    ["X", 10],
    ["IX", 9],
    ["V", 5],
    ["IV", 4],
    ["I", 1]
  ];

  let result = "";
  for (const [sym, val] of symbols) {
    while (num >= val) {
      result += sym;
      num -= val;
    }
  }
  return result;
}
```
