---
id: 691b559495c5cb5a37b9b487
title: "도전 과제 127: 속도 확인"
challengeType: 28
dashedName: challenge-127
---

# --description--

시속 마일(MPH) 단위로 이동하는 속도와 시속 킬로미터(KPH) 단위의 제한 속도가 주어졌을 때, 과속 중인지 그리고 경고를 받을지 벌금을 받을지 판단하세요.

- 1마일은 1.60934킬로미터입니다.
- 제한 속도 이하로 이동 중이면 `"Not Speeding"`를 반환하세요.
- 제한 속도보다 5 KPH 이하로 초과하면 `"Warning"`를 반환하세요.
- 제한 속도보다 5 KPH 초과하면 `"Ticket"`를 반환하세요.

# --hints--

`speedCheck(30, 70)`는 `"Not Speeding"`를 반환해야 합니다.

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)`는 `"Warning"`를 반환해야 합니다.

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)`는 `"Not Speeding"`를 반환해야 합니다.

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)`는 `"Ticket"`를 반환해야 합니다.

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)`는 `"Warning"`를 반환해야 합니다.

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)`는 `"Ticket"`를 반환해야 합니다.

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
