---
id: 691b559495c5cb5a37b9b482
title: "도전 과제 122: 마크다운 굵게 파서"
challengeType: 28
dashedName: challenge-122
---

# --description--

마크다운에서 일부 굵게 표시된 텍스트가 포함될 수 있는 문자열이 주어지면, 이에 상응하는 HTML 문자열을 반환하세요.

- 마크다운에서 굵게 표시된 텍스트는 두 개의 별표(`**`) 또는 두 개의 밑줄(`__`)로 시작하고 끝나는 모든 텍스트입니다.
- 텍스트와 별표 또는 밑줄 사이에는 공백이 있을 수 없지만, 텍스트 자체에는 공백이 있을 수 있습니다.
- 모든 굵게 표시된 부분을 HTML `b` 태그로 변환하여 문자열을 반환하세요.

예를 들어, `"**This is bold**"`가 주어지면 `"<b>This is bold</b>"`를 반환하세요.

참고: 콘솔에 문자열을 로그할 때 HTML 태그가 표시되지 않을 수 있습니다. 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parseBold("**This is bold**")`는 `"<b>This is bold</b>"`를 반환해야 합니다.

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")`는 `"<b>This is also bold</b>"`를 반환해야 합니다.

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")`는 `"**This is not bold **"`를 반환해야 합니다.

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")`는 `"__ This is also not bold__"`를 반환해야 합니다.

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")`는 `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`를 반환해야 합니다.

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
