---
id: 691b559495c5cb5a37b9b485
title: "도전 과제 125: 생명의 게임"
challengeType: 28
dashedName: challenge-125
---

# --description--

Conway의 생명의 게임에서 현재 상태를 나타내는 행렬(배열의 배열)이 주어졌을 때, 다음 규칙을 사용해 행렬의 다음 상태를 반환하세요:

- 각 셀은 `1` (살아 있음) 또는 `0` (죽음)입니다.
- 셀의 이웃은 최대 8개의 주변 셀(수직, 수평, 대각선 방향)입니다.
- 가장자리 셀은 8개 미만의 이웃을 가집니다.

각 셀을 업데이트하는 규칙:

- 살아 있는 셀 중 이웃이 2개 미만인 셀은 죽습니다(과소 번식).
- 살아 있는 셀 중 이웃이 2개 또는 3개인 셀은 계속 살아 있습니다.
- 살아 있는 셀 중 이웃이 3개 초과인 셀은 죽습니다(과잉 번식).
- 죽은 셀 중 이웃이 정확히 3개인 셀은 살아납니다(재생산).

예를 들어, 다음이 주어졌을 때:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

다음을 반환하세요:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

각 셀은 살아 있는 이웃 수에 따라 업데이트됩니다. 예를 들어, `[0][0]`는 죽은 상태를 유지합니다(이웃 2개), `[0][1]`는 살아 있는 상태를 유지합니다(이웃 2개), `[0][2]`는 죽습니다(이웃 3개) 등입니다.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])`는 `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`를 반환해야 합니다.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])`는 `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`를 반환해야 합니다.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])`는 `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`를 반환해야 합니다.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])`는 `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`를 반환해야 합니다.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
