---
id: 691b559495c5cb5a37b9b481
title: "도전 과제 121: 가장 자주 나타나는 요소"
challengeType: 28
dashedName: challenge-121
---

# --description--

요소 배열이 주어지면 가장 자주 나타나는 요소를 반환하세요.

- 항상 가장 자주 나타나는 요소가 하나만 있습니다.

# --hints--

`mostFrequent(["a", "b", "a", "c"])`는 `"a"`을 반환해야 합니다.

```js
assert.equal(mostFrequent(["a", "b", "a", "c"]), "a");
```

`mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9])`는 `2`을 반환해야 합니다.

```js
assert.equal(mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9]), 2);
```

`mostFrequent([true, false, "false", "true", false])`는 `false`을 반환해야 합니다.

```js
assert.isFalse(mostFrequent([true, false, "false", "true", false]));
```

`mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60])`는 `40`을 반환해야 합니다.

```js
assert.equal(mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60]), 40);
```

# --seed--

## --seed-contents--

```js
function mostFrequent(arr) {
  return arr;
}
```

# --solutions--

```js
function mostFrequent(arr) {
  const freq = {};
  let maxCount = 0;
  let mostElem = null;

  for (const el of arr) {
    freq[el] = (freq[el] || 0) + 1;
    if (freq[el] > maxCount) {
      maxCount = freq[el];
      mostElem = el;
    }
  }

  return mostElem;
}
```
