---
id: 691b5597f0f3e85a588a5c94
title: "도전 과제 130: 체커보드"
challengeType: 28
dashedName: challenge-130
---

# --description--

두 개의 숫자가 들어 있는 배열이 주어집니다. 첫 번째 숫자는 행 수이고 두 번째 숫자는 열 수입니다. 주어진 크기로 `"X"`와 `"O"` 문자로 채워진 행렬(배열의 배열)을 반환하세요.

- 문자는 체커보드처럼 번갈아가며 배치되어야 합니다.
- 왼쪽 상단 셀은 항상 `"X"`여야 합니다.

예를 들어, `[3, 3]`가 주어지면 다음을 반환하세요:

```json
[
  ["X", "O", "X"],
  ["O", "X", "O"],
  ["X", "O", "X"]
]
```

# --hints--

`createBoard([3, 3])`는 `[["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]`를 반환해야 합니다.

```js
assert.deepEqual(createBoard([3, 3]), [["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]);
```

`createBoard([6, 1])`는 `[["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]`를 반환해야 합니다.

```js
assert.deepEqual(createBoard([6, 1]), [["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]);
```

`createBoard([2, 10])`는 `[["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]`를 반환해야 합니다.

```js
assert.deepEqual(createBoard([2, 10]), [["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]);
```

`createBoard([5, 4])`는 `[["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]`를 반환해야 합니다.

```js
assert.deepEqual(createBoard([5, 4]), [["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]);
```

# --seed--

## --seed-contents--

```js
function createBoard(dimensions) {
  return dimensions;
}
```

# --solutions--

```js
function createBoard(dimensions) {
  const [rows, cols] = dimensions;
  const board = [];

  for (let r = 0; r < rows; r++) {
    const row = [];
    for (let c = 0; c < cols; c++) {
      row.push((r + c) % 2 === 0 ? "X" : "O");
    }
    board.push(row);
  }

  return board;
}
```
