---
id: 6925e2068081f40f549ced1a
title: "도전 과제 136: 마크다운 이미지 파서"
challengeType: 28
dashedName: challenge-136
---

# --description--

마크다운 이미지 문자열이 주어지면, 이에 상응하는 HTML 문자열을 반환하세요.

마크다운 이미지는 다음 형식을 가집니다: `"![alt text](image_url)"`. 여기서:

- `alt text` 은 이미지 설명(`alt` 속성 값)입니다.
- `image_url` 은 이미지의 소스 URL(`src` 속성 값)입니다.

이미지 URL에 `img` 을, 대체 텍스트에 `src` 을 설정한 HTML `alt` 태그 문자열을 반환하세요.

예를 들어, `"![Cute cat](cat.png)"` 이 주어지면 `'<img src="cat.png" alt="Cute cat">'` 을 반환하세요.

- 태그, 속성 순서, 공백, 따옴표 사용이 위와 동일한지 꼭 확인하세요.

참고: 콘솔에 문자열로 HTML 태그가 출력되지 않을 수 있으니, 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parseImage("![Cute cat](cat.png)")` 은 `'<img src="cat.png" alt="Cute cat">'` 을 반환해야 합니다.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` 은 `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'` 을 반환해야 합니다.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` 은 `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'` 을 반환해야 합니다.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
