---
id: 6925e2068081f40f549ced1c
title: "도전 과제 138: 약수의 합"
challengeType: 28
dashedName: challenge-138
---

# --description--

양의 정수가 주어지면, 그 수의 모든 약수의 합을 반환하세요.

- 약수는 그 수를 나누었을 때 나머지가 `0`인 모든 정수입니다.
- 각 약수는 한 번만 계산하세요.

예를 들어, `6`이 주어지면 `12`를 반환하세요. `6`의 약수는 `1`, `2`, `3`, `6`이고, 이들의 합은 `12`입니다.

# --hints--

`sumDivisors(6)`는 `12`를 반환해야 합니다.

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)`는 `14`를 반환해야 합니다.

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)`는 `56`를 반환해야 합니다.

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)`는 `224`를 반환해야 합니다.

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)`는 `806`를 반환해야 합니다.

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)`는 `23520`를 반환해야 합니다.

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
