---
id: 6925e2068081f40f549ced1b
title: "도전 과제 137: 스노우플레이크 생성기"
challengeType: 28
dashedName: challenge-137
---

# --description--

줄 바꿈을 나타내는 개행 문자(`\n`)를 사용하는 여러 줄 문자열이 주어질 때, 각 줄을 수평으로 거울처럼 뒤집어 원래 줄 끝에 붙인 새 문자열을 반환하세요.

- 줄을 거울처럼 뒤집는다는 것은 공백을 포함한 모든 문자를 역순으로 만드는 것을 의미합니다.

예를 들어, `"* \n *\n* "`이 주어지면 콘솔에 다음과 같이 출력됩니다:

```sh
* 
 *
* 
```

이때 `"*  *\n ** \n*  *"`를 반환하면 콘솔에 다음과 같이 출력됩니다:

```sh
*  *
 ** 
*  *
```

주어진 문자열과 반환할 문자열의 공백에 특히 주의하세요. 공백을 절대 잘라내지 마세요.

# --hints--

`generateSnowflake("* \n *\n* ")`는 `"*  *\n ** \n*  *"`를 반환해야 합니다.

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")`는 `"X=~~=X"`를 반환해야 합니다.

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")`는 `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`를 반환해야 합니다.

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")`는 `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`를 반환해야 합니다.

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")`는 `"*  --  *\n * -- * \n*  --  *"`를 반환해야 합니다.

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
