---
id: 6925e2068081f40f549ced1d
title: "도전 과제 139: 가위, 바위, 보"
challengeType: 28
dashedName: challenge-139
---

# --description--

두 개의 문자열이 주어집니다. 첫 번째 문자열은 플레이어 1을, 두 번째 문자열은 플레이어 2를 나타냅니다. 가위, 바위, 보 경기의 승자를 결정하세요.

- 입력 문자열은 항상 `"Rock"`, `"Paper"`, 또는 `"Scissors"`입니다.
- `"Rock"`는 `"Scissors"`를 이깁니다.
- `"Paper"`는 `"Rock"`를 이깁니다.
- `"Scissors"`는 `"Paper"`를 이깁니다.

다음 값을 반환하세요:

- 플레이어 1이 이기면 `"Player 1 wins"`.
- 플레이어 2가 이기면 `"Player 2 wins"`.
- 두 플레이어가 같은 선택을 하면 `"Tie"`.

# --hints--

`rockPaperScissors("Rock", "Rock")`는 `"Tie"`를 반환해야 합니다.

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")`는 `"Player 2 wins"`를 반환해야 합니다.

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")`는 `"Player 1 wins"`를 반환해야 합니다.

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")`는 `"Player 1 wins"`를 반환해야 합니다.

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")`는 `"Tie"`를 반환해야 합니다.

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")`는 `"Player 2 wins"`를 반환해야 합니다.

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
