---
id: 69272dcf1c24b44fd79137c5
title: "도전 과제 142: 문자열의 합 구하기"
challengeType: 28
dashedName: challenge-142
---

# --description--

숫자와 다른 문자가 섞여 있는 문자열이 주어질 때, 문자열에 포함된 모든 숫자의 합을 반환하세요.

- 연속된 숫자는 하나의 숫자로 처리하세요. 예를 들어, `"13"`는 1 + 3이 아니라 13으로 계산합니다.
- 숫자가 아닌 문자는 무시하세요.

# --hints--

`stringSum("3apples2bananas")`는 `5`를 반환해야 합니다.

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")`는 `17`를 반환해야 합니다.

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")`는 `125344`를 반환해야 합니다.

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")`는 `321`를 반환해야 합니다.

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")`는 `1653`를 반환해야 합니다.

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
