---
id: 69272dcf1c24b44fd79137c3
title: "도전 과제 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

변수 이름을 나타내는 문자열이 주어지면, 해당 변수 이름을 SCREAMING_SNAKE_CASE로 변환해 반환하세요.

주어진 변수 이름은 다음 형식 중 하나로 작성됩니다:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

위 형식에서 단어들은 밑줄(`_`), 하이픈(`-`), 또는 새 단어가 대문자로 시작하는 방식으로 구분됩니다.

SCREAMING_SNAKE_CASE로 변환하려면:

- 모든 문자를 대문자로 만드세요
- 단어 사이를 밑줄(`_`)로 구분하세요

# --hints--

`toScreamingSnakeCase("userEmail")`는 `"USER_EMAIL"`를 반환해야 합니다.

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")`는 `"USER_PASSWORD"`를 반환해야 합니다.

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")`는 `"USER_ID"`를 반환해야 합니다.

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")`는 `"USER_ADDRESS"`를 반환해야 합니다.

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")`는 `"USERNAME"`를 반환해야 합니다.

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

`toScreamingSnakeCase("my_variable_name")`는 `"MY_VARIABLE_NAME"`를 반환해야 합니다.

```js
assert.equal(toScreamingSnakeCase("my_variable_name"), "MY_VARIABLE_NAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();
}
```
