---
id: 69272dcf1c24b44fd79137c6
title: "도전 과제 143: 마크다운 이탤릭 파서"
challengeType: 28
dashedName: challenge-143
---

# --description--

마크다운에서 일부 이탤릭 텍스트가 포함될 수 있는 문자열이 주어지면, 이에 상응하는 HTML 문자열을 반환하세요.

- 마크다운에서 이탤릭 텍스트는 단일 별표(`*`) 또는 단일 밑줄(`_`)로 시작하고 끝나는 모든 텍스트입니다.
- 텍스트와 별표 또는 밑줄 사이에는 공백이 있을 수 없지만, 텍스트 자체에는 공백이 있을 수 있습니다.
- 모든 이탤릭 표현을 HTML `i` 태그로 변환하고 문자열을 반환하세요.

예를 들어, `"*This is italic*"`가 주어지면 `"<i>This is italic</i>"`를 반환하세요.

참고: 콘솔은 문자열 내 HTML 태그를 로그할 때 표시하지 않을 수 있습니다. 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parseItalics("*This is italic*")`는 `"<i>This is italic</i>"`를 반환해야 합니다.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")`는 `"<i>This is also italic</i>"`를 반환해야 합니다.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")`는 `"*This is not italic *"`를 반환해야 합니다.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")`는 `"_ This is also not italic_"`를 반환해야 합니다.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")`는 `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`를 반환해야 합니다.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
