---
id: 69272dcf1c24b44fd79137c4
title: "도전 과제 141: 이륙 연료"
challengeType: 28
dashedName: challenge-141
---

# --description--

현재 비행기에 있는 연료 갤런 수와 목적지까지 필요한 연료 리터 수가 주어졌을 때, 추가로 몇 갤런의 연료를 넣어야 하는지 계산하세요.

- 1갤런은 3.78541리터입니다.
- 비행기에 이미 충분한 연료가 있다면 `0`를 반환하세요.
- 연료는 정수 갤런 단위로만 추가할 수 있습니다.
- 반환 값에 소수점은 포함하지 마세요.

# --hints--

`fuelToAdd(0, 1)`는 `1`를 반환해야 합니다.

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)`는 `6`를 반환해야 합니다.

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)`는 `0`를 반환해야 합니다.

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)`는 `4520`를 반환해야 합니다.

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)`는 `12209`를 반환해야 합니다.

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
