---
id: 69306364df283fcaff2e1ad7
title: "도전 과제 146: 왼손잡이용 좌석"
challengeType: 28
dashedName: challenge-146
---

# --description--

4x2 행렬(배열의 배열)로 표현된 식사 자리 배치에서 왼손잡이가 앉을 수 있는 좌석 수를 구하세요.

- 왼손잡이는 바로 왼쪽에 앉은 사람이 오른손잡이인 자리에 앉을 수 없습니다.

주어진 행렬에서:

- `"R"`는 오른손잡이가 앉은 자리입니다.
- `"L"`는 왼손잡이가 앉은 자리입니다.
- `"U"`는 비어 있는 자리입니다.
- 앉을 수 있는 자리는 비어 있는 자리뿐입니다.
- 윗줄 좌석은 "아래쪽"을 향하고, 아랫줄 좌석은 "위쪽"을 향하고 있어(테이블처럼) 왼쪽과 오른쪽은 좌석 방향에 따라 상대적입니다.
- 모서리 좌석은 옆에 한 자리만 있습니다.

예를 들어, 주어진 행렬에서:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

왼쪽 위 좌석은 바로 왼쪽에 오른손잡이가 있어 앉을 수 없습니다. 나머지 두 빈 좌석은 바로 왼쪽에 오른손잡이가 없으므로 앉을 수 있습니다.

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])`는 `2`를 반환해야 합니다.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])`는 `8`를 반환해야 합니다.

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])`는 `0`를 반환해야 합니다.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])`는 `1`를 반환해야 합니다.

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])`는 `5`를 반환해야 합니다.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```
