---
id: 69306364df283fcaff2e1ad9
title: "도전 과제 148: 타이어 압력"
challengeType: 28
dashedName: challenge-148
---

# --description--

차량의 네 타이어 압력을 psi 단위로 나타내는 네 개 숫자의 배열과, 타이어의 최솟값과 최댓값 압력을 bar 단위로 나타내는 두 개 숫자의 배열이 주어질 때, 각 타이어 상태를 설명하는 네 개 문자열의 배열을 반환하세요.

- 1 bar는 14.5038 psi와 같습니다.

각 타이어에 대해 다음 값을 포함하는 배열을 반환하세요:

- 타이어 압력이 허용된 최솟값보다 낮으면 `"Low"`.
- 최솟값과 최댓값 사이면 `"Good"`.
- 최댓값보다 높으면 `"High"`.

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])`는 `["Good", "Low", "Good", "Low"]`를 반환해야 합니다.

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])`는 `["Good", "Low", "High", "Good"]`를 반환해야 합니다.

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])`는 `["Low", "Low", "Good", "Low"]`를 반환해야 합니다.

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])`는 `["High", "High", "High", "Good"]`를 반환해야 합니다.

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])`는 `["Good", "Good", "Good", "Good"]`를 반환해야 합니다.

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
