---
id: 69306364df283fcaff2e1ad8
title: "도전 과제 147: 윤년 계산기"
challengeType: 28
dashedName: challenge-147
---

# --description--

정수 연도가 주어지면 윤년인지 아닌지 판단하세요.

윤년은 다음 규칙을 만족할 때 윤년입니다:

- 연도가 4로 나누어떨어지고,
- 연도가 100으로 나누어떨어지지 않거나,
- 연도가 400으로 나누어떨어질 때입니다.

# --hints--

`isLeapYear(2024)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
