---
id: 69306364df283fcaff2e1ada
title: "도전 과제 149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

문자열이 주어지면 모든 모음을 대문자로 변환하고 나머지 알파벳 문자는 소문자로 변환한 새 문자열을 반환하세요.

- 모음은 `"a"`, `"e"`, `"i"`, `"o"`, `"u"`이며 대소문자 구분 없이 포함됩니다.
- 알파벳이 아닌 문자는 변경하지 않고 그대로 둡니다.

# --hints--

`vowelCase("vowelcase")`는 `"vOwElcAsE"`를 반환해야 합니다.

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")`는 `"cOdIng Is fUn"`를 반환해야 합니다.

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")`는 `"hEllO, wOrld!"`를 반환해야 합니다.

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")`는 `"gIt chErry-pIck"`를 반환해야 합니다.

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")`는 `"hEAd~1"`를 반환해야 합니다.

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
