---
id: 69306364df283fcaff2e1ad5
title: "도전 과제 144: 해상도 연속 성공"
challengeType: 28
dashedName: challenge-144
---

# --description--

배열의 배열이 주어집니다. 각 하위 배열은 해당 날짜의 결심 활동을 나타내며 세 개의 정수를 포함합니다: 그날 걸은 걸음 수, 화면 사용 시간(분), 읽은 페이지 수입니다. 결심을 잘 지키고 있는지 판단하세요.

- 첫 번째 하위 배열은 1일 차, 두 번째는 2일 차, 이런 식입니다.

다음 세 가지 목표를 모두 달성한 날은 성공한 날로 간주합니다:

- 최소 10,000걸음 이상 걸었습니다.
- 화면 사용 시간이 120분을 넘지 않았습니다.
- 최소 5페이지 이상 읽었습니다.

주어진 모든 날이 성공적이면 `"Resolution on track: N day streak."`를 반환하세요. 여기서 `N`은 성공한 날 수입니다.

하나 이상의 날이 성공하지 못했다면 `"Resolution failed on day X: N day streak."`를 반환하세요. 여기서 `X`은 첫 번째 실패한 날 번호이고, `N`는 첫 실패한 날 이전의 성공한 날 수입니다.

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])`는 `"Resolution on track: 3 day streak."`를 반환해야 합니다.

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])`는 `"Resolution failed on day 2: 1 day streak."`를 반환해야 합니다.

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])`는 `"Resolution failed on day 3: 2 day streak."`를 반환해야 합니다.

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])`는 `"Resolution failed on day 4: 3 day streak."`를 반환해야 합니다.

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])`는 `"Resolution on track: 10 day streak."`를 반환해야 합니다.

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```
