---
id: 69306364df283fcaff2e1ad6
title: "도전 과제 145: N번째 피보나치 수"
challengeType: 28
dashedName: challenge-145
---

# --description--

정수 `n`가 주어지면 피보나치 수열에서 `n`번째 수를 반환하세요.

피보나치 수열은 각 숫자가 바로 앞의 두 숫자의 합인 수열입니다. 수열의 처음 10개 숫자는 `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`입니다.

# --hints--

`nthFibonacci(4)`는 `2`를 반환해야 합니다.

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)`는 `34`를 반환해야 합니다.

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)`는 `377`를 반환해야 합니다.

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)`는 `63245986`를 반환해야 합니다.

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)`는 `1304969544928657`를 반환해야 합니다.

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
