---
id: 69373793f5a867f769cde135
title: "도전 과제 150: 마크다운 순서 없는 목록 파서"
challengeType: 28
dashedName: challenge-150
---

# --description--

유효한 마크다운 순서 없는 목록 문자열이 주어지면, 이에 해당하는 HTML 문자열을 반환하세요.

순서 없는 목록은 하나 이상의 목록 항목으로 구성됩니다. 유효한 목록 항목은 각자 한 줄에 나타나며 다음 조건을 만족합니다:

- 대시(`"-"`)로 시작하고,
- 그 다음에 최소 한 칸 이상의 공백이 있으며,
- 목록 항목 텍스트가 이어집니다.

목록은 줄바꿈 문자(`"\n"`)로 구분된 단일 문자열로 주어집니다. 항목 텍스트에 줄바꿈 문자를 포함하지 마세요.

각 목록 항목을 HTML `li` 태그로 감싸고, 전체 목록은 `ul` 태그로 감싸세요.

예를 들어, `"- Item A\n- Item B"`가 주어지면 `"<ul><li>Item A</li><li>Item B</li></ul>"`를 반환하세요.

참고: 콘솔에서는 문자열 내 HTML 태그가 표시되지 않을 수 있습니다. 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parseUnorderedList("- Item A\n- Item B")`는 `"<ul><li>Item A</li><li>Item B</li></ul>"`를 반환해야 합니다.

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")`는 `"<ul><li>JavaScript</li><li>Python</li></ul>"`를 반환해야 합니다.

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")`는 `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`를 반환해야 합니다.

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")`는 `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`를 반환해야 합니다.

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
