---
id: 69373793f5a867f769cde13a
title: "도전 과제 155: 작물 심기"
challengeType: 28
dashedName: challenge-155
---

# --description--

농장 밭 크기를 나타내는 정수와 농장 밭 크기의 단위를 나타내는 `"acres"` 또는 `"hectares"`, 그리고 작물 종류가 주어질 때, 해당 밭에 심을 수 있는 그 작물의 개수를 구하세요.

- 1에이커는 4046.86 제곱미터입니다.
- 1헥타르는 10,000 제곱미터입니다.

다음은 입력으로 주어질 작물과 한 그루가 차지하는 공간입니다:

|작물|한 그루당 공간|
|-|-|
|`"corn"`|1 제곱미터|
|`"wheat"`|0.1 제곱미터|
|`"soybeans"`|0.5 제곱미터|
|`"tomatoes"`|0.25 제곱미터|
|`"lettuce"`|0.2 제곱미터|

밭에 심을 수 있는 작물 개수를 소수점 이하는 버리고 정수로 반환하세요.

# --hints--

`getNumberOfPlants(1, "acres", "corn")`는 `4046`를 반환해야 합니다.

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")`는 `100000`를 반환해야 합니다.

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")`는 `161874`를 반환해야 합니다.

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")`는 `150000`를 반환해야 합니다.

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")`는 `271139`를 반환해야 합니다.

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
