---
id: 69373793f5a867f769cde137
title: "도전 과제 152: 순환 소수"
challengeType: 28
dashedName: challenge-152
---

# --description--

정수가 주어졌을 때, 그것이 순환 소수인지 판단하세요.

순환 소수는 숫자의 모든 자리 회전이 모두 소수인 정수입니다.

예를 들어, `197`는 그 자리 숫자들의 모든 회전인 `197`, `971`, 그리고 `719`가 모두 소수이기 때문에 순환 소수입니다.

# --hints--

`isCircularPrime(197)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
