---
id: 69373793f5a867f769cde139
title: "도전 과제 154: 홀의 기준 타수"
challengeType: 28
dashedName: challenge-154
---

# --description--

두 개의 정수가 주어집니다. 하나는 골프 홀의 기준 타수(par)이고, 다른 하나는 골퍼가 그 홀에서 친 획 수입니다. 골프 용어를 사용해 골퍼의 점수를 반환하세요.

반환값:

- 한 획을 쳤다면 `"Hole in one!"`을 반환하세요.
- 기준 타수보다 두 획 적게 쳤다면 `"Eagle"`을 반환하세요.
- 기준 타수보다 한 획 적게 쳤다면 `"Birdie"`를 반환하세요.
- 기준 타수와 같은 획 수를 쳤다면 `"Par"`을 반환하세요.
- 기준 타수보다 한 획 더 쳤다면 `"Bogey"`를 반환하세요.
- 기준 타수보다 두 획 더 쳤다면 `"Double bogey"`를 반환하세요.

# --hints--

`golfScore(3, 3)`는 `"Par"`을 반환해야 합니다.

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)`는 `"Birdie"`을 반환해야 합니다.

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)`는 `"Hole in one!"`을 반환해야 합니다.

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)`는 `"Double bogey"`을 반환해야 합니다.

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)`는 `"Bogey"`을 반환해야 합니다.

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)`는 `"Eagle"`을 반환해야 합니다.

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
