---
id: 69373793f5a867f769cde136
title: "도전 과제 151: 정렬된 배열인가요?"
challengeType: 28
dashedName: challenge-151
---

# --description--

숫자 배열이 주어졌을 때, 숫자가 오름차순, 내림차순으로 정렬되었는지 아니면 둘 다 아닌지 판단하세요.

주어진 배열이:

- 오름차순(낮은 값에서 높은 값)이라면 `"Ascending"`를 반환하세요.
- 내림차순(높은 값에서 낮은 값)이라면 `"Descending"`를 반환하세요.
- 오름차순이나 내림차순으로 정렬되지 않았다면 `"Not sorted"`를 반환하세요.

# --hints--

`isSorted([1, 2, 3, 4, 5])`는 `"Ascending"`를 반환해야 합니다.

```js
assert.equal(isSorted([1, 2, 3, 4, 5]), "Ascending");
```

`isSorted([10, 8, 6, 4, 2])`는 `"Descending"`를 반환해야 합니다.

```js
assert.equal(isSorted([10, 8, 6, 4, 2]), "Descending");
```

`isSorted([1, 3, 2, 4, 5])`는 `"Not sorted"`를 반환해야 합니다.

```js
assert.equal(isSorted([1, 3, 2, 4, 5]), "Not sorted");
```

`isSorted([3.14, 2.71, 1.61, 0.57])`는 `"Descending"`를 반환해야 합니다.

```js
assert.equal(isSorted([3.14, 2.71, 1.61, 0.57]), "Descending");
```

`isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])`는 `"Ascending"`를 반환해야 합니다.

```js
assert.equal(isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending");
```

`isSorted([0.4, 0.5, 0.3])`는 `"Not sorted"`를 반환해야 합니다.

```js
assert.equal(isSorted([0.4, 0.5, 0.3]), "Not sorted");
```

# --seed--

## --seed-contents--

```js
function isSorted(arr) {

  return arr;
}
```

# --solutions--

```js
function isSorted(arr) {
  let ascending = true;
  let descending = true;

  for (let i = 1; i < arr.length; i++) {
    if (arr[i] < arr[i - 1]) ascending = false;
    if (arr[i] > arr[i - 1]) descending = false;
  }

  if (ascending) return "Ascending";
  if (descending) return "Descending";
  return "Not sorted";
}
```
