---
id: 69373793f5a867f769cde138
title: "도전 과제 153: 틱택토"
challengeType: 28
dashedName: challenge-153
---

# --description--

완성된 틱택토 게임을 나타내는 3×3 행렬(배열의 배열)이 주어질 때, 승자를 결정하세요.

- 주어진 행렬의 각 요소는 `"X"` 또는 `"O"`입니다.

플레이어는 가로, 세로, 대각선 중 한 줄에 자신의 문자가 세 개 있으면 승리합니다.

반환값:

- 플레이어 X가 한 줄에 세 개를 가지면 `"X wins"`를 반환하세요.
- 플레이어 O가 한 줄에 세 개를 가지면 `"O wins"`를 반환하세요.
- 어느 플레이어도 한 줄에 세 개를 가지지 않으면 `"Draw"`를 반환하세요.

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])`는 `"X wins"`를 반환해야 합니다.

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])`는 `"X wins"`를 반환해야 합니다.

```js
assert.equal(ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])`는 `"Draw"`를 반환해야 합니다.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]])`는 `"O wins"`를 반환해야 합니다.

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])`는 `"X wins"`를 반환해야 합니다.

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])`는 `"Draw"`를 반환해야 합니다.

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```
