---
id: 6939b873185d8e00d453563c
title: "도전 과제 157: 마크다운 링크 파서"
challengeType: 28
dashedName: challenge-157
---

# --description--

마크다운 링크 문자열이 주어지면, 이에 해당하는 HTML 문자열을 반환하세요.

마크다운 이미지는 다음 형식을 가집니다: `"[link_text](link_url)"`. `a` 태그 문자열을 반환하되, `href`는 `link_url`로 설정하고 `link_text`는 태그 내용으로 사용하세요.

예를 들어, `"[freeCodeCamp](https://freecodecamp.org/)"`가 주어지면 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`를 반환하세요;

참고: 콘솔은 문자열 내 HTML 태그를 로그로 출력할 때 표시하지 않을 수 있으니, 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")`는 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`를 반환해야 합니다.

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")`는 `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`를 반환해야 합니다.

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")`는 `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`를 반환해야 합니다.

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
