---
id: 6939b873185d8e00d453563f
title: "도전 과제 160: 나이트 이동"
challengeType: 28
dashedName: challenge-160
---

# --description--

체스판 위 나이트의 위치가 주어졌을 때, 나이트가 이동할 수 있는 유효한 칸의 수를 반환하세요.

표준 체스판은 8x8 크기로, 열은 왼쪽에서 오른쪽으로 `A`부터 `H`까지 레이블이 붙어 있고, 행은 아래에서 위로 `1`부터 `8`까지 레이블이 붙어 있습니다. 체스판은 다음과 같습니다:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

나이트는 "L"자 모양으로 이동합니다: 한 방향(가로 또는 세로)으로 두 칸, 그리고 수직 방향으로 한 칸 이동합니다.

즉, 나이트는 최대 여덟 군데로 이동할 수 있지만, 체스판 가장자리 근처에서는 이동 가능한 칸이 더 적습니다. 예를 들어, 나이트가 `A1`에 있다면 `B3` 또는 `C2`로만 이동할 수 있습니다.

# --hints--

`knightMoves("A1")`는 `2`를 반환해야 합니다.

```js
assert.equal(knightMoves("A1"), 2);
```

`knightMoves("D4")`는 `8`를 반환해야 합니다.

```js
assert.equal(knightMoves("D4"), 8);
```

`knightMoves("G6")`는 `6`를 반환해야 합니다.

```js
assert.equal(knightMoves("G6"), 6);
```

`knightMoves("B8")`는 `3`를 반환해야 합니다.

```js
assert.equal(knightMoves("B8"), 3);
```

`knightMoves("H3")`는 `4`를 반환해야 합니다.

```js
assert.equal(knightMoves("H3"), 4);
```

# --seed--

## --seed-contents--

```js
function knightMoves(position) {

  return position;
}
```

# --solutions--

```js
function knightMoves(position) {
  const col = position[0].charCodeAt(0) - 65;
  const row = parseInt(position[1]) - 1;

  const moves = [
    [2, 1], [2, -1],
    [-2, 1], [-2, -1],
    [1, 2], [1, -2],
    [-1, 2], [-1, -2]
  ];

  let valid = 0;

  for (const [dx, dy] of moves) {
    const newCol = col + dx;
    const newRow = row + dy;

    if (newCol >= 0 && newCol < 8 && newRow >= 0 && newRow < 8) {
      valid++;
    }
  }

  return valid;
}
```
