---
id: 6939b873185d8e00d453563b
title: "도전 과제 156: 홀수인지 짝수인지?"
challengeType: 28
dashedName: challenge-156
---

# --description--

양의 정수가 주어지면, 홀수라면 `"Odd"`를, 짝수라면 `"Even"`를 반환하세요.

# --hints--

`oddOrEven(1)`는 `"Odd"`를 반환해야 합니다.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)`는 `"Even"`를 반환해야 합니다.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)`는 `"Odd"`를 반환해야 합니다.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)`는 `"Even"`를 반환해야 합니다.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)`는 `"Odd"`를 반환해야 합니다.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
