---
id: 6939b873185d8e00d453563d
title: "도전 과제 158: 배열 값 교환"
challengeType: 28
dashedName: challenge-158
---

# --description--

두 개의 값이 있는 배열이 주어지면, 값이 서로 바뀐 배열을 반환하세요.

예를 들어, `["A", "B"]`가 주어지면 `["B", "A"]`을 반환하세요.

# --hints--

`arraySwap(["A", "B"])`는 `["B", "A"]`을 반환해야 합니다.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])`는 `[20, 25]`을 반환해야 합니다.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])`는 `[false, true]`을 반환해야 합니다.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])`는 `[1, "1"]`을 반환해야 합니다.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
