---
id: 6939b873185d8e00d453563e
title: "도전 과제 159: 정수 빗변"
challengeType: 28
dashedName: challenge-159
---

# --description--

두 개의 양의 정수가 직각삼각형 두 변(짧은 두 변)의 길이를 나타낼 때, 빗변의 길이가 정수인지 판단하세요.

빗변의 길이는 두 변 길이의 제곱을 더한 후 그 합의 제곱근을 구해서 계산합니다 (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

`isIntegerHypotenuse(3, 4)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
