---
id: 694596b0585c11170ac7c7fa
title: "도전 과제 162: 에너지 소비"
challengeType: 28
dashedName: challenge-162
---

# --description--

운동 중 소모한 칼로리 수와 운동 중 전자 기기가 사용한 와트시 수가 주어졌을 때, 어느 쪽이 더 많은 에너지를 사용했는지 판단하세요.

두 값을 비교하려면 다음 변환을 사용해 둘 다 줄 단위로 변환하세요:

- 1 칼로리는 4184 줄입니다.
- 1 와트시는 3600 줄입니다.

다음 값을 반환하세요:

- 운동이 더 많은 에너지를 사용했다면 `"Workout"`.
- 기기가 더 많은 에너지를 사용했다면 `"Devices"`.
- 둘 다 같은 양의 에너지를 사용했다면 `"Equal"`.

# --hints--

`compareEnergy(250, 50)`는 `"Workout"`를 반환해야 합니다.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)`는 `"Devices"`를 반환해야 합니다.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)`는 `"Equal"`를 반환해야 합니다.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)`는 `"Workout"`를 반환해야 합니다.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)`는 `"Devices"`를 반환해야 합니다.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)`는 `"Equal"`를 반환해야 합니다.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
