---
id: 694596b0585c11170ac7c7fc
title: "도전 과제 164: 마크다운 인라인 코드 파서"
challengeType: 28
dashedName: challenge-164
---

# --description--

하나 이상의 인라인 코드 블록을 포함하는 마크다운 문자열이 주어지면, 이에 상응하는 HTML 문자열을 반환하세요.

마크다운에서 인라인 코드 블록은 코드 블록 텍스트의 시작과 끝에 단일 백틱(```` ` ```)을 사용합니다.

주어진 문자열에서 모든 코드 블록을 HTML `code` 태그로 변환한 문자열을 반환하세요.

예를 들어, 문자열 ```"Use `let` to declare the variable."```, return `"변수를 선언할 때 <code>let</code>을 사용하세요."``가 주어졌을 때입니다.

참고: 콘솔은 문자열 내 HTML 태그를 로그로 출력할 때 태그를 표시하지 않을 수 있습니다. 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

```parseInlineCode("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
assert.equal(parseInlineCode("Use `let` to declare the variable."), "Use <code>let</code> to declare the variable.");
```

```parseInlineCode("Use `let` or `const` to declare a variable.")``` should return `"Use <code>let</code> or <code>const</code> to declare a variable."`.

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

```parseInlineCode("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
assert.equal(parseInlineCode("Run `npm install` then `npm start`."), "Run <code>npm install</code> then <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
