---
id: 694596b0585c11170ac7c7f9
title: "도전 과제 161: 무료 배송"
challengeType: 28
dashedName: challenge-161
---

# --description--

쇼핑 카트에 담긴 항목을 나타내는 문자열 배열과 무료 배송 자격을 위한 최솟값 주문 금액이 주어졌을 때, 쇼핑 카트의 항목들이 무료 배송 자격이 되는지 판단하세요.

주어진 배열에는 아래 목록의 항목들이 포함됩니다:

| 항목 | 가격 |
| - | - |
| `"shirt"`	| 34.25 |
| `"jeans"`	| 48.50 |
| `"shoes"` |	75.00 |
| `"hat"`	| 19.95 |
| `"socks"` | 15.00 |
| `"jacket"` | 109.95 |

# --hints--

`getsFreeShipping(["shoes"], 50)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(getsFreeShipping(["shoes"], 50));
```

`getsFreeShipping(["hat", "socks"], 50)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(getsFreeShipping(["hat", "socks"], 50));
```

`getsFreeShipping(["jeans", "shirt", "jacket"], 75)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(getsFreeShipping(["jeans", "shirt", "jacket"], 75));
```

`getsFreeShipping(["socks", "socks", "hat"], 75)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(getsFreeShipping(["socks", "socks", "hat"], 75));
```

`getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100));
```

`getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200));
```

# --seed--

## --seed-contents--

```js
function getsFreeShipping(cart, minimum) {

  return cart;
}
```

# --solutions--

```js
function getsFreeShipping(cart, minimum) {
  const prices = {
    shirt: 34.25,
    jeans: 48.50,
    shoes: 75.00,
    hat: 19.95,
    socks: 15.00,
    jacket: 109.95
  };

  let total = 0;

  for (const item of cart) {
    total += prices[item];
  }

  return total >= minimum;
}
```
