---
id: 694596b0585c11170ac7c7fb
title: "도전 과제 163: 자음 케이스"
challengeType: 28
dashedName: challenge-163
---

# --description--

변수 이름을 나타내는 문자열이 주어지면 다음 규칙에 따라 자음 케이스로 변환하세요:

- 모든 자음은 대문자로 변환하세요.
- 모든 모음(`a`, `e`, `i`, `o`, `u`, 대소문자 구분 없이)은 소문자로 변환하세요.
- 모든 하이픈(`-`)은 밑줄(`_`)로 변환하세요.

# --hints--

`toConsonantCase("helloworld")`는 `"HeLLoWoRLD"`를 반환해야 합니다.

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")`는 `"HeLLoWoRLD"`를 반환해야 합니다.

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")`는 `"_HeLLo_WoRLD_"`를 반환해야 합니다.

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")`는 `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`를 반환해야 합니다.

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
