---
id: 696655d24b614176d4c9b78b
title: "도전 과제 168: 크기 조정된 이미지"
challengeType: 28
dashedName: challenge-168
---

# --description--

이미지의 너비와 높이를 나타내는 문자열과 이미지를 확대할 숫자가 주어지면, 확대된 너비와 높이를 반환하세요.

- 입력 문자열은 `"WxH"` 형식입니다. 예를 들어, `"800x600"`입니다.
- 배율은 너비와 높이에 곱할 숫자입니다.

확대된 크기를 같은 `"WxH"` 형식으로 반환하세요.

# --hints--

`scaleImage("800x600", 2)`는 `"1600x1200"`를 반환해야 합니다.

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)`는 `"1000x1000"`를 반환해야 합니다.

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)`는 `"512x384"`를 반환해야 합니다.

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)`는 `"450x300"`를 반환해야 합니다.

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
