---
id: 696655d24b614176d4c9b78c
title: "도전 과제 169: FizzBuzz 미니"
challengeType: 28
dashedName: challenge-169
---

# --description--

정수가 주어지면 다음 규칙에 따라 문자열을 반환하세요:

- 숫자가 3으로 나누어떨어지면 `"Fizz"`를 반환하세요.
- 숫자가 5로 나누어떨어지면 `"Buzz"`를 반환하세요.
- 숫자가 3과 5 모두로 나누어떨어지면 `"FizzBuzz"`를 반환하세요.
- 그렇지 않으면 주어진 숫자를 문자열로 반환하세요.

# --hints--

`fizzBuzzMini(3)`는 `"Fizz"`를 반환해야 합니다.

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)`는 `"4"`를 반환해야 합니다.

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)`는 `"Buzz"`를 반환해야 합니다.

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)`는 `"FizzBuzz"`를 반환해야 합니다.

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)`는 `"98"`를 반환해야 합니다.

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
