---
id: 696655d24b614176d4c9b78a
title: "도전 과제 167: 빙고! 글자"
challengeType: 28
dashedName: challenge-167
---

# --description--

숫자가 주어지면 해당하는 빙고 글자(대문자)를 반환하세요. 빙고 숫자는 다음과 같이 그룹화됩니다:

| 글자 | 숫자 범위 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

`getBingoLetter(75)`는 `"O"`를 반환해야 합니다.

```js
assert.equal(getBingoLetter(75), "O");
```

`getBingoLetter(54)`는 `"G"`를 반환해야 합니다.

```js
assert.equal(getBingoLetter(54), "G");
```

`getBingoLetter(25)`는 `"I"`를 반환해야 합니다.

```js
assert.equal(getBingoLetter(25), "I");
```

`getBingoLetter(38)`는 `"N"`를 반환해야 합니다.

```js
assert.equal(getBingoLetter(38), "N");
```

`getBingoLetter(11)`는 `"B"`를 반환해야 합니다.

```js
assert.equal(getBingoLetter(11), "B");
```

# --seed--

## --seed-contents--

```js
function getBingoLetter(n) {

  return n;
}
```

# --solutions--

```js
function getBingoLetter(n) {
  if (n >= 1 && n <= 15) return "B";
  if (n >= 16 && n <= 30) return "I";
  if (n >= 31 && n <= 45) return "N";
  if (n >= 46 && n <= 60) return "G";
  if (n >= 61 && n <= 75) return "O";
  return null;
}
```
