---
id: 696655d24b614176d4c9b789
title: "도전 과제 166: 16진수 유효성 검사기"
challengeType: 28
dashedName: challenge-166
---

# --description--

문자열이 주어졌을 때, 유효한 CSS 16진수 색상인지 확인하세요. 유효한 CSS 16진수 색상은 다음 조건을 만족해야 합니다:

- `#`로 시작하고,
- 그 뒤에 3개 또는 6개의 16진수 문자가 따라야 합니다.

16진수 문자는 숫자 `0`부터 `9`까지와 문자 `a`부터 `f`까지입니다(대소문자 구분 없음).

# --hints--

`isValidHex("#123")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
