---
id: 696655d24b614176d4c9b78d
title: "도전 과제 170: 홀수일 또는 짝수일"
challengeType: 28
dashedName: challenge-170
---

# --description--

타임스탬프(Unix epoch 이후 밀리초 수)가 주어지면 다음을 반환하세요:

- 해당 타임스탬프의 날짜가 홀수일 경우 `"odd"`를 반환하세요.
- 해당 타임스탬프의 날짜가 짝수일 경우 `"even"`를 반환하세요.

예를 들어, 2026년 1월 27일의 타임스탬프인 `1769472000000`가 주어지면, 날짜 숫자(`"odd"`)가 홀수이므로 `27`를 반환하세요.

참고: 타임스탬프는 밀리초 단위이며, 로컬 시간이 아닌 UTC 시간대를 기준으로 날짜를 사용해야 합니다.

# --hints--

`oddOrEvenDay(1769472000000)`는 `"odd"`를 반환해야 합니다.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)`는 `"even"`를 반환해야 합니다.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)`는 `"odd"`를 반환해야 합니다.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)`는 `"odd"`를 반환해야 합니다.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)`는 `"even"`를 반환해야 합니다.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
