---
id: 69738771fb5a7b8b24cca2a5
title: "도전 과제 179: 주머니 속 잔돈"
challengeType: 28
dashedName: challenge-179
---

# --description--

주머니에 있는 동전을 나타내는 정수 배열이 주어집니다. 각 정수는 센트 단위 동전의 가치를 나타내며, `"$D.CC"` 형식으로 총 금액을 반환하세요.

- 100센트는 1달러와 같습니다.
- 반환값에는 1달러 미만 금액에 대해 선행 0을 포함하고 센트는 항상 정확히 두 자리 숫자로 표시하세요.

# --hints--

`countChange([25, 10, 5, 1])`은 `"$0.41"`를 반환해야 합니다.

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])`은 `"$1.43"`를 반환해야 합니다.

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])`은 `"$37.55"`를 반환해야 합니다.

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])`은 `"$0.70"`를 반환해야 합니다.

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])`은 `"$0.01"`를 반환해야 합니다.

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])`은 `"$1.00"`를 반환해야 합니다.

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
