---
id: 69738771fb5a7b8b24cca29e
title: "도전 과제 172: 문자와 숫자"
challengeType: 28
dashedName: challenge-172
---

# --description--

문자와 숫자만 포함된 문자열이 주어질 때, 문자열이 문자에서 숫자로, 또는 숫자에서 문자로 바뀔 때마다 하이픈(`-`)을 삽입한 새 문자열을 반환하세요.

# --hints--

`separateLettersAndNumbers("ABC123")`은 `"ABC-123"`를 반환해야 합니다.

```js
assert.equal(separateLettersAndNumbers("ABC123"), "ABC-123");
```

`separateLettersAndNumbers("Route66")`은 `"Route-66"`를 반환해야 합니다.

```js
assert.equal(separateLettersAndNumbers("Route66"), "Route-66");
```

`separateLettersAndNumbers("H3LL0W0RLD")`는 `"H-3-LL-0-W-0-RLD"`를 반환해야 합니다.

```js
assert.equal(separateLettersAndNumbers("H3LL0W0RLD"), "H-3-LL-0-W-0-RLD");
```

`separateLettersAndNumbers("a1b2c3d4")`은 `"a-1-b-2-c-3-d-4"`를 반환해야 합니다.

```js
assert.equal(separateLettersAndNumbers("a1b2c3d4"), "a-1-b-2-c-3-d-4");
```

# --seed--

## --seed-contents--

```js
function separateLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function separateLettersAndNumbers(str) {
  if (str.length === 0) return str;

  let result = str[0];

  for (let i = 1; i < str.length; i++) {
    const prev = str[i - 1];
    const curr = str[i];

    const prevIsLetter = /[a-zA-Z]/.test(prev);
    const currIsLetter = /[a-zA-Z]/.test(curr);

    if (prevIsLetter !== currIsLetter) {
      result += "-";
    }

    result += curr;
  }

  return result;
}
```
