---
id: 69738771fb5a7b8b24cca2a2
title: "도전 과제 176: 그라운드호그 데이"
challengeType: 28
dashedName: challenge-176
---

# --description--

오늘은 그라운드호그 데이로, 그라운드호그가 자신의 그림자를 보았는지 여부에 따라 날씨를 예측합니다.

그라운드호그의 모습을 나타내는 값을 받아 올바른 예측을 반환하세요:

- 주어진 값이 불리언 `true` (그라운드호그가 그림자를 보았다면)일 경우 `"Looks like we'll have six more weeks of winter."`을 반환하세요.
- 값이 불리언 `false` (그라운드호그가 그림자를 보지 못했다면)일 경우 `"It's going to be an early spring."`을 반환하세요.
- 값이 그 외의 경우(그라운드호그가 나타나지 않았다면) `"No prediction this year."`를 반환하세요.

# --hints--

`groundhogDayPrediction(true)`는 `"Looks like we'll have six more weeks of winter."`를 반환해야 합니다.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)`는 `"It's going to be an early spring."`를 반환해야 합니다.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)`는 `"No prediction this year."`를 반환해야 합니다.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")`는 `"No prediction this year."`를 반환해야 합니다.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")`는 `"No prediction this year."`를 반환해야 합니다.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
