---
id: 69738771fb5a7b8b24cca2a1
title: "도전 과제 175: 디지털 디톡스"
challengeType: 28
dashedName: challenge-175
---

# --description--

로그인 기록 배열이 주어졌을 때, 디지털 디톡스 목표를 달성했는지 판단하세요.

각 기록은 `"YYYY-MM-DD HH:mm:ss"` 형식의 문자열입니다.

다음 두 조건이 모두 참일 때 디지털 디톡스 목표를 달성한 것입니다:

- 4시간 이내에 로그인한 횟수가 1회를 넘지 않아야 합니다.
- 하루에 로그인한 횟수가 2회를 넘지 않아야 합니다.

# --hints--

`digitalDetox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"])`은 `true`를 반환해야 합니다.

```js
assert.isTrue(digitalDetox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"]));
```

`digitalDetox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"])`은 `false`를 반환해야 합니다.

```js
assert.isFalse(digitalDetox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"]));
```

`digitalDetox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"])`은 `true`를 반환해야 합니다.

```js
assert.isTrue(digitalDetox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"]));
```

`digitalDetox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"])`은 `false`를 반환해야 합니다.

```js
assert.isFalse(digitalDetox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"]));
```

`digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])`은 `true`를 반환해야 합니다.

```js
assert.isTrue(digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 20:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]));
```

`digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])`은 `false`를 반환해야 합니다.

```js
assert.isFalse(digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]));
```

# --seed--

## --seed-contents--

```js
function digitalDetox(logs) {

  return logs;
}
```

# --solutions--

```js
function digitalDetox(logs) {
  if (logs.length === 0) return true;

  const times = logs
    .map(log => new Date(log.replace(" ", "T") + "Z").getTime())
    .sort((a, b) => a - b);

  const ONE_DAY = 24 * 60 * 60 * 1000;
  const FOUR_HOURS = 4 * 60 * 60 * 1000;

  for (let i = 1; i < times.length; i++) {
    if (times[i] - times[i - 1] < FOUR_HOURS) {
      return false;
    }
  }

  const dailyCounts = {};
  for (const time of times) {
    const day = new Date(time).toISOString().slice(0, 10);
    dailyCounts[day] = (dailyCounts[day] || 0) + 1;
    if (dailyCounts[day] > 2) {
      return false;
    }
  }

  return true;
}
```
