---
id: 69738771fb5a7b8b24cca2a4
title: "도전 과제 178: 문자열 자르기"
challengeType: 28
dashedName: challenge-178
---

# --description--

문자열이 주어지면, 20자 이하면 그대로 반환하세요. 20자보다 길면 처음 17자만 남기고 `"..."`를 끝에 붙여서(총 20자가 되도록) 결과를 반환하세요.

# --hints--

`truncateText("Hello, world!")`는 `"Hello, world!"`를 반환해야 합니다.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")`는 `"This string shoul..."`를 반환해야 합니다.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")`는 `"Exactly twenty chars"`를 반환해야 합니다.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")`는 `"...................."`를 반환해야 합니다.

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
