---
id: 69738771fb5a7b8b24cca2a0
title: "도전 과제 174: 별자리 찾기"
challengeType: 28
dashedName: challenge-174
---

# --description--

`"YYYY-MM-DD"` 형식의 날짜 문자열이 주어지면, 다음 차트를 사용해 해당 날짜의 별자리를 반환하세요.

| 날짜 범위 | 별자리 |
| - | - |
| 3월 21일 - 4월 19일 | `"Aries"` |
| 4월 20일 - 5월 20일 | `"Taurus"` |
| 5월 21일 - 6월 20일 | `"Gemini"` |
| 6월 21일 - 7월 22일 | `"Cancer"` |
| 7월 23일 - 8월 22일 | `"Leo"` |
| 8월 23일 - 9월 22일 | `"Virgo"` |
| 9월 23일 - 10월 22일 | `"Libra"` |
| 10월 23일 - 11월 21일 | `"Scorpio"` |
| 11월 22일 - 12월 21일 | `"Sagittarius"` |
| 12월 22일 - 1월 19일 | `"Capricorn"` |
| 1월 20일 - 2월 18일 | `"Aquarius"` |
| 2월 19일 - 3월 20일 | `"Pisces"` |

- 별자리는 월과 일만 기준으로 하며, 연도는 무시하세요.

# --hints--

`getSign("2026-01-31")`는 `"Aquarius"`를 반환해야 합니다.

```js
assert.equal(getSign("2026-01-31"), "Aquarius");
```

`getSign("2001-06-10")`는 `"Gemini"`를 반환해야 합니다.

```js
assert.equal(getSign("2001-06-10"), "Gemini");
```

`getSign("1985-09-07")`는 `"Virgo"`를 반환해야 합니다.

```js
assert.equal(getSign("1985-09-07"), "Virgo");
```

`getSign("2023-03-19")`는 `"Pisces"`를 반환해야 합니다.

```js
assert.equal(getSign("2023-03-19"), "Pisces");
```

`getSign("2045-11-05")`는 `"Scorpio"`를 반환해야 합니다.

```js
assert.equal(getSign("2045-11-05"), "Scorpio");
```

`getSign("1985-12-06")`는 `"Sagittarius"`를 반환해야 합니다.

```js
assert.equal(getSign("1985-12-06"), "Sagittarius");
```

`getSign("2025-12-30")`는 `"Capricorn"`를 반환해야 합니다.

```js
assert.equal(getSign("2025-12-30"), "Capricorn");
```

`getSign("2018-10-08")`는 `"Libra"`를 반환해야 합니다.

```js
assert.equal(getSign("2018-10-08"), "Libra");
```

`getSign("1958-05-04")`는 `"Taurus"`를 반환해야 합니다.

```js
assert.equal(getSign("1958-05-04"), "Taurus");
```

# --seed--

## --seed-contents--

```js
function getSign(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getSign(dateStr) {
  const [, monthStr, dayStr] = dateStr.split("-");
  const month = Number(monthStr);
  const day = Number(dayStr);

  if ((month === 3 && day >= 21) || (month === 4 && day <= 19)) return "Aries";
  if ((month === 4 && day >= 20) || (month === 5 && day <= 20)) return "Taurus";
  if ((month === 5 && day >= 21) || (month === 6 && day <= 20)) return "Gemini";
  if ((month === 6 && day >= 21) || (month === 7 && day <= 22)) return "Cancer";
  if ((month === 7 && day >= 23) || (month === 8 && day <= 22)) return "Leo";
  if ((month === 8 && day >= 23) || (month === 9 && day <= 22)) return "Virgo";
  if ((month === 9 && day >= 23) || (month === 10 && day <= 22)) return "Libra";
  if ((month === 10 && day >= 23) || (month === 11 && day <= 21)) return "Scorpio";
  if ((month === 11 && day >= 22) || (month === 12 && day <= 21)) return "Sagittarius";
  if ((month === 12 && day >= 22) || (month === 1 && day <= 19)) return "Capricorn";
  if ((month === 1 && day >= 20) || (month === 2 && day <= 18)) return "Aquarius";
  if ((month === 2 && day >= 19) || (month === 3 && day <= 20)) return "Pisces";
}
```
