---
id: 69738771fb5a7b8b24cca2a3
title: "도전 과제 177: 문자열 거울"
challengeType: 28
dashedName: challenge-177
---

# --description--

문자열이 주어지면, 주어진 문자열 뒤에 그 문자열을 거꾸로 뒤집은 복사본을 덧붙인 새 문자열을 반환하세요.

# --hints--

`mirror("freeCodeCamp")`는 `"freeCodeCamppmaCedoCeerf"`을 반환해야 합니다.

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")`는 `"RaceCarraCecaR"`을 반환해야 합니다.

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")`는 `"helloworlddlrowolleh"`을 반환해야 합니다.

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")`는 `"The quick brown fox......xof nworb kciuq ehT"`을 반환해야 합니다.

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
