---
id: 69738771fb5a7b8b24cca29d
title: "도전 과제 171: 배열 평탄화"
challengeType: 28
dashedName: challenge-171
---

# --description--

중첩된 배열을 포함하는 배열이 주어졌을 때, 모든 값을 하나의 1차원 배열로 평탄화한 새 배열을 반환하세요. 배열 내 항목들의 원래 순서를 유지하세요.

# --hints--

`flatten([1, [2, 3], 4])`는 `[1, 2, 3, 4]`를 반환해야 합니다.

```js
assert.deepEqual(flatten([1, [2, 3], 4]), [1, 2, 3, 4]);
```

`flatten([5, [4, [3, 2]], 1])`는 `[5, 4, 3, 2, 1]`를 반환해야 합니다.

```js
assert.deepEqual(flatten([5, [4, [3, 2]], 1]), [5, 4, 3, 2, 1]);
```

`flatten(["A", [[[["B"]]]], "C"])`는 `["A", "B", "C"]`를 반환해야 합니다.

```js
assert.deepEqual(flatten(["A", [[[["B"]]]], "C"]), ["A", "B", "C"]);
```

`flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]])`는 `["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]`를 반환해야 합니다.

```js
assert.deepEqual(flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]]), ["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]);
```

`flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]])`는 `["red","blue","green","yellow","purple","orange","pink","brown"]`를 반환해야 합니다.

```js
assert.deepEqual(flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]]), ["red","blue","green","yellow","purple","orange","pink","brown"]);
```

# --seed--

## --seed-contents--

```js
function flatten(arr) {

  return arr;
}
```

# --solutions--

```js
function flatten(arr) {
  let result = [];

  for (const item of arr) {
    if (Array.isArray(item)) {
      result = result.concat(flatten(item));
    } else {
      result.push(item);
    }
  }

  return result;
}
```
