---
id: 69738771fb5a7b8b24cca29f
title: "도전 과제 173: 유효한 폰 이동"
challengeType: 28
dashedName: challenge-173
---

# --description--

체스판에서 폰 하나의 위치가 주어지면, 그 폰이 이동할 수 있는 모든 유효한 칸을 오름차순으로 배열에 담아 반환하세요.

표준 체스판은 8x8 크기로, 열은 왼쪽에서 오른쪽으로 `A`부터 `H`까지 레이블이 붙어 있고, 행은 아래에서 위로 `1`부터 `8`까지 레이블이 붙어 있습니다. 다음과 같습니다:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

이 도전 과제에서:

- 당신은 체스판 아래쪽에 있는 플레이어입니다.
- 폰은 한 칸만 "위쪽"으로 이동할 수 있습니다.
- 단, 폰이 시작 행(2행)에 있으면 한 칸 또는 두 칸 위로 이동할 수 있습니다.

예를 들어, `"D4"`가 주어지면, 폰이 이동할 수 있는 유일한 칸인 `["D5"]`를 반환하세요.
`"B2"`가 주어지면, 시작 행에 있으므로 `["B3", "B4"]`를 반환하고, 오름차순으로 정렬하세요.

# --hints--

`findPawnMoves("D4")`는 `["D5"]`를 반환해야 합니다.

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")`는 `["B3", "B4"]`를 반환해야 합니다.

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")`는 `["A8"]`를 반환해야 합니다.

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")`는 `["G3", "G4"]`를 반환해야 합니다.

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")`는 `["E4"]`를 반환해야 합니다.

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
