---
id: 697a49e9860d24853adef67d
title: "도전 과제 192: 2026년 동계 경기 13일차: 노르딕 복합"
challengeType: 28
dashedName: challenge-192
---

# --description--

선수들의 점프 점수 배열이 주어졌을 때, 노르딕 복합의 크로스컨트리 구간 출발 지연 시간을 계산하세요.

가장 높은 점프 점수를 받은 선수가 가장 먼저 출발합니다(지연 시간 0초). 나머지 선수들은 최고 점수와의 차이에 따라 더 늦게 출발합니다.

각 선수의 지연 시간을 계산하려면, 최고 점수에서 해당 선수의 점프 점수를 빼고 그 결과에 1.5를 곱하세요. 지연 시간은 가장 가까운 정수로 올림하세요.

# --hints--

`calculateStartDelays([120, 110, 125])`는 `[8, 23, 0]`를 반환해야 합니다.

```js
assert.deepEqual(calculateStartDelays([120, 110, 125]), [8, 23, 0]);
```

`calculateStartDelays([118, 125, 122, 120])`는 `[11, 0, 5, 8]`를 반환해야 합니다.

```js
assert.deepEqual(calculateStartDelays([118, 125, 122, 120]), [11, 0, 5, 8]);
```

`calculateStartDelays([100, 105, 95, 110, 120, 115, 108])`는 `[30, 23, 38, 15, 0, 8, 18]`를 반환해야 합니다.

```js
assert.deepEqual(calculateStartDelays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18]);
```

`calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])`는 `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`를 반환해야 합니다.

```js
assert.deepEqual(calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12]);
```

# --seed--

## --seed-contents--

```js
function calculateStartDelays(jumpScores) {

  return jumpScores
}
```

# --solutions--

```js
function calculateStartDelays(jumpScores) {
  const bestJump = Math.max(...jumpScores);
  return jumpScores.map(score => Math.ceil((bestJump - score) * 1.5));
}
```
