---
id: 697a49e6ff50d756c9b6935d
title: "도전 과제 180: 2026년 동계 올림픽 1일차: 개막일"
challengeType: 28
dashedName: challenge-180
---

# --description--

오늘은 2026년 동계 올림픽이 시작되는 날입니다. 앞으로 17일 동안 올림픽에서 영감을 받은 코딩 도전 과제가 이어집니다.

첫 번째 도전 과제로, 두 글자의 국가 코드를 받아 해당 국가의 국기 이모지를 반환하세요.

다음 목록을 사용하세요:

| 국가 | 코드 | 국기 |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"🇦🇱"` |
| Andorra                    |            `"AD"` | `"🇦🇩"` |
| Argentina                  |            `"AR"` | `"🇦🇷"` |
| Armenia                    |            `"AM"` | `"🇦🇲"` |
| Australia                  |            `"AU"` | `"🇦🇺"` |
| Austria                    |            `"AT"` | `"🇦🇹"` |
| Azerbaijan                 |            `"AZ"` | `"🇦🇿"` |
| Belgium                    |            `"BE"` | `"🇧🇪"` |
| Benin                      |            `"BJ"` | `"🇧🇯"` |
| Bolivia                    |            `"BO"` | `"🇧🇴"` |
| Bosnia and Herzegovina     |            `"BA"` | `"🇧🇦"` |
| Brazil                     |            `"BR"` | `"🇧🇷"` |
| Bulgaria                   |            `"BG"` | `"🇧🇬"` |
| Canada                     |            `"CA"` | `"🇨🇦"` |
| Chile                      |            `"CL"` | `"🇨🇱"` |
| China                      |            `"CN"` | `"🇨🇳"` |
| Colombia                   |            `"CO"` | `"🇨🇴"` |
| Croatia                    |            `"HR"` | `"🇭🇷"` |
| Cyprus                     |            `"CY"` | `"🇨🇾"` |
| Czech Republic             |            `"CZ"` | `"🇨🇿"` |
| Denmark                    |            `"DK"` | `"🇩🇰"` |
| Ecuador                    |            `"EC"` | `"🇪🇨"` |
| Eritrea                    |            `"ER"` | `"🇪🇷"` |
| Estonia                    |            `"EE"` | `"🇪🇪"` |
| Finland                    |            `"FI"` | `"🇫🇮"` |
| France                     |            `"FR"` | `"🇫🇷"` |
| Georgia                    |            `"GE"` | `"🇬🇪"` |
| Germany                    |            `"DE"` | `"🇩🇪"` |
| Great Britain              |            `"GB"` | `"🇬🇧"` |
| Greece                     |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau              |            `"GW"` | `"🇬🇼"` |
| Haiti                      |            `"HT"` | `"🇭🇹"` |
| Hong Kong                  |            `"HK"` | `"🇭🇰"` |
| Hungary                    |            `"HU"` | `"🇭🇺"` |
| Iceland                    |            `"IS"` | `"🇮🇸"` |
| India                      |            `"IN"` | `"🇮🇳"` |
| Iran                       |            `"IR"` | `"🇮🇷"` |
| Ireland                    |            `"IE"` | `"🇮🇪"` |
| Israel                     |            `"IL"` | `"🇮🇱"` |
| Italy                      |            `"IT"` | `"🇮🇹"` |
| Jamaica                    |            `"JM"` | `"🇯🇲"` |
| Japan                      |            `"JP"` | `"🇯🇵"` |
| Kazakhstan                 |            `"KZ"` | `"🇰🇿"` |
| Kenya                      |            `"KE"` | `"🇰🇪"` |
| Kosovo                     |            `"XK"` | `"🇽🇰"` |
| Kyrgyzstan                 |            `"KG"` | `"🇰🇬"` |
| Latvia                     |            `"LV"` | `"🇱🇻"` |
| Lebanon                    |            `"LB"` | `"🇱🇧"` |
| Liechtenstein              |            `"LI"` | `"🇱🇮"` |
| Lithuania                  |            `"LT"` | `"🇱🇹"` |
| Luxembourg                 |            `"LU"` | `"🇱🇺"` |
| Madagascar                 |            `"MG"` | `"🇲🇬"` |
| Malaysia                   |            `"MY"` | `"🇲🇾"` |
| Malta                      |            `"MT"` | `"🇲🇹"` |
| Mexico                     |            `"MX"` | `"🇲🇽"` |
| Moldova                    |            `"MD"` | `"🇲🇩"` |
| Monaco                     |            `"MC"` | `"🇲🇨"` |
| Mongolia                   |            `"MN"` | `"🇲🇳"` |
| Montenegro                 |            `"ME"` | `"🇲🇪"` |
| Morocco                    |            `"MA"` | `"🇲🇦"` |
| Netherlands                |            `"NL"` | `"🇳🇱"` |
| New Zealand                |            `"NZ"` | `"🇳🇿"` |
| Nigeria                    |            `"NG"` | `"🇳🇬"` |
| North Macedonia            |            `"MK"` | `"🇲🇰"` |
| Norway                     |            `"NO"` | `"🇳🇴"` |
| Pakistan                   |            `"PK"` | `"🇵🇰"` |
| Philippines                |            `"PH"` | `"🇵🇭"` |
| Poland                     |            `"PL"` | `"🇵🇱"` |
| Portugal                   |            `"PT"` | `"🇵🇹"` |
| Puerto Rico                |            `"PR"` | `"🇵🇷"` |
| Romania                    |            `"RO"` | `"🇷🇴"` |
| San Marino                 |            `"SM"` | `"🇸🇲"` |
| Saudi Arabia               |            `"SA"` | `"🇸🇦"` |
| Serbia                     |            `"RS"` | `"🇷🇸"` |
| Singapore                  |            `"SG"` | `"🇸🇬"` |
| Slovakia                   |            `"SK"` | `"🇸🇰"` |
| Slovenia                   |            `"SI"` | `"🇸🇮"` |
| South Africa               |            `"ZA"` | `"🇿🇦"` |
| South Korea                |            `"KR"` | `"🇰🇷"` |
| Spain                      |            `"ES"` | `"🇪🇸"` |
| Sweden                     |            `"SE"` | `"🇸🇪"` |
| Switzerland                |            `"CH"` | `"🇨🇭"` |
| Thailand                   |            `"TH"` | `"🇹🇭"` |
| Trinidad & Tobago          |            `"TT"` | `"🇹🇹"` |
| Turkey                     |            `"TR"` | `"🇹🇷"` |
| Ukraine                    |            `"UA"` | `"🇺🇦"` |
| United Arab Emirates       |            `"AE"` | `"🇦🇪"` |
| United States              |            `"US"` | `"🇺🇸"` |
| Uruguay                    |            `"UY"` | `"🇺🇾"` |
| Uzbekistan                 |            `"UZ"` | `"🇺🇿"` |
| Venezuela                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")`는 `"🇦🇱"`를 반환해야 합니다.

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")`는 `"🇦🇩"`를 반환해야 합니다.

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")`는 `"🇦🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")`는 `"🇦🇲"`를 반환해야 합니다.

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")`는 `"🇦🇺"`를 반환해야 합니다.

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")`는 `"🇦🇹"`를 반환해야 합니다.

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")`는 `"🇦🇿"`를 반환해야 합니다.

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")`는 `"🇧🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")`는 `"🇧🇯"`를 반환해야 합니다.

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")`는 `"🇧🇴"`를 반환해야 합니다.

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")`는 `"🇧🇦"`를 반환해야 합니다.

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")`는 `"🇧🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")`는 `"🇧🇬"`를 반환해야 합니다.

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")`는 `"🇨🇦"`를 반환해야 합니다.

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")`는 `"🇨🇱"`를 반환해야 합니다.

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")`는 `"🇨🇳"`를 반환해야 합니다.

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")`는 `"🇨🇴"`를 반환해야 합니다.

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")`는 `"🇭🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")`는 `"🇨🇾"`를 반환해야 합니다.

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")`는 `"🇨🇿"`를 반환해야 합니다.

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")`는 `"🇩🇰"`를 반환해야 합니다.

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")`는 `"🇪🇨"`를 반환해야 합니다.

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")`는 `"🇪🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")`는 `"🇪🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")`는 `"🇫🇮"`를 반환해야 합니다.

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")`는 `"🇫🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")`는 `"🇬🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")`는 `"🇩🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")`는 `"🇬🇧"`를 반환해야 합니다.

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")`는 `"🇬🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")`는 `"🇬🇼"`를 반환해야 합니다.

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")`는 `"🇭🇹"`를 반환해야 합니다.

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")`는 `"🇭🇰"`를 반환해야 합니다.

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")`는 `"🇭🇺"`를 반환해야 합니다.

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")`는 `"🇮🇸"`를 반환해야 합니다.

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")`는 `"🇮🇳"`를 반환해야 합니다.

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")`는 `"🇮🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")`는 `"🇮🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")`는 `"🇮🇱"`를 반환해야 합니다.

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")`는 `"🇮🇹"`를 반환해야 합니다.

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")`는 `"🇯🇲"`를 반환해야 합니다.

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")`는 `"🇯🇵"`를 반환해야 합니다.

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")`는 `"🇰🇿"`를 반환해야 합니다.

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")`는 `"🇰🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")`는 `"🇽🇰"`를 반환해야 합니다.

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")`는 `"🇰🇬"`를 반환해야 합니다.

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")`는 `"🇱🇻"`를 반환해야 합니다.

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")`는 `"🇱🇧"`를 반환해야 합니다.

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")`는 `"🇱🇮"`를 반환해야 합니다.

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")`는 `"🇱🇹"`를 반환해야 합니다.

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")`는 `"🇱🇺"`를 반환해야 합니다.

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")`는 `"🇲🇬"`를 반환해야 합니다.

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")`는 `"🇲🇾"`를 반환해야 합니다.

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")`는 `"🇲🇹"`를 반환해야 합니다.

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")`는 `"🇲🇽"`를 반환해야 합니다.

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")`는 `"🇲🇩"`를 반환해야 합니다.

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")`는 `"🇲🇨"`를 반환해야 합니다.

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")`는 `"🇲🇳"`를 반환해야 합니다.

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")`는 `"🇲🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")`는 `"🇲🇦"`를 반환해야 합니다.

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")`는 `"🇳🇱"`를 반환해야 합니다.

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")`는 `"🇳🇿"`를 반환해야 합니다.

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")`는 `"🇳🇬"`를 반환해야 합니다.

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")`는 `"🇲🇰"`를 반환해야 합니다.

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")`는 `"🇳🇴"`를 반환해야 합니다.

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")`는 `"🇵🇰"`를 반환해야 합니다.

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")`는 `"🇵🇭"`를 반환해야 합니다.

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")`는 `"🇵🇱"`를 반환해야 합니다.

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")`는 `"🇵🇹"`를 반환해야 합니다.

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")`는 `"🇵🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")`는 `"🇷🇴"`를 반환해야 합니다.

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")`는 `"🇸🇲"`를 반환해야 합니다.

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")`는 `"🇸🇦"`를 반환해야 합니다.

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")`는 `"🇷🇸"`를 반환해야 합니다.

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")`는 `"🇸🇬"`를 반환해야 합니다.

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")`는 `"🇸🇰"`를 반환해야 합니다.

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")`는 `"🇸🇮"`를 반환해야 합니다.

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")`는 `"🇿🇦"`를 반환해야 합니다.

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")`는 `"🇰🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")`는 `"🇪🇸"`를 반환해야 합니다.

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")`는 `"🇸🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")`는 `"🇨🇭"`를 반환해야 합니다.

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")`는 `"🇹🇭"`를 반환해야 합니다.

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")`는 `"🇹🇹"`를 반환해야 합니다.

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")`는 `"🇹🇷"`를 반환해야 합니다.

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")`는 `"🇺🇦"`를 반환해야 합니다.

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")`는 `"🇦🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")`는 `"🇺🇸"`를 반환해야 합니다.

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")`는 `"🇺🇾"`를 반환해야 합니다.

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")`는 `"🇺🇿"`를 반환해야 합니다.

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")`는 `"🇻🇪"`를 반환해야 합니다.

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
