---
id: 697a49e6ff50d756c9b69361
title: "도전 과제 184: 2026년 동계 경기 5일차: 크로스컨트리 스키"
challengeType: 28
dashedName: challenge-184
---

# --description--

크로스컨트리 스키 경주의 완주 시간 배열이 주어질 때, 이를 우승자와의 시간 차이로 변환하세요.

- 주어진 시간은 `"H:MM:SS"` 형식의 문자열입니다.
- 주어진 시간은 가장 빠른 순서부터 느린 순서로 정렬되어 있습니다.
- 우승자의 시간(가장 빠른 시간)은 `"0"`에 해당해야 합니다.
- 나머지 시간들은 `"+M:SS"` 형식으로 우승자와의 시간 차이를 보여야 합니다.

예를 들어, `["1:25:32", "1:26:10", "1:27:05"]`가 주어지면 `["0", "+0:38", "+1:33"]`를 반환하세요.

# --hints--

`getRelativeResults(["1:25:32", "1:26:10", "1:27:05"])`는 `["0", "+0:38", "+1:33"]`를 반환해야 합니다.

```js
assert.deepEqual(getRelativeResults(["1:25:32", "1:26:10", "1:27:05"]), ["0", "+0:38", "+1:33"]);
```

`getRelativeResults(["1:00:01", "1:00:05", "1:00:10"])`는 `["0", "+0:04", "+0:09"]`를 반환해야 합니다.

```js
assert.deepEqual(getRelativeResults(["1:00:01", "1:00:05", "1:00:10"]), ["0", "+0:04", "+0:09"]);
```

`getRelativeResults(["1:10:06", "1:10:23", "1:10:48", "1:12:11"])`는 `["0", "+0:17", "+0:42", "+2:05"]`를 반환해야 합니다.

```js
assert.deepEqual(getRelativeResults(["1:10:06", "1:10:23", "1:10:48", "1:12:11"]), ["0", "+0:17", "+0:42", "+2:05"]);
```

`getRelativeResults(["0:49:13", "0:49:15", "0:50:14", "0:51:30", "0:51:58", "0:52:16", "0:53:12", "0:53:31", "0:56:19", "1:02:20"])`는 `["0", "+0:02", "+1:01", "+2:17", "+2:45", "+3:03", "+3:59", "+4:18", "+7:06", "+13:07"]`를 반환해야 합니다.

```js
assert.deepEqual(getRelativeResults(["0:49:13", "0:49:15", "0:50:14", "0:51:30", "0:51:58", "0:52:16", "0:53:12", "0:53:31", "0:56:19", "1:02:20"]), ["0", "+0:02", "+1:01", "+2:17", "+2:45", "+3:03", "+3:59", "+4:18", "+7:06", "+13:07"]);
```

`getRelativeResults(["2:01:15", "2:10:45", "2:10:53", "2:11:04", "2:11:55", "2:13:27", "2:14:30", "2:15:10"])`는 `["0", "+9:30", "+9:38", "+9:49", "+10:40", "+12:12", "+13:15", "+13:55"]`를 반환해야 합니다.

```js
assert.deepEqual(getRelativeResults(["2:01:15", "2:10:45", "2:10:53", "2:11:04", "2:11:55", "2:13:27", "2:14:30", "2:15:10"]), ["0", "+9:30", "+9:38", "+9:49", "+10:40", "+12:12", "+13:15", "+13:55"]);
```

# --seed--

## --seed-contents--

```js
function getRelativeResults(results) {

  return results;
}
```

# --solutions--

```js
function getRelativeResults(results) {
  const timeToSeconds = (timeStr) => {
    const [hours, minutes, seconds] = timeStr.split(':').map(Number);
    return hours * 3600 + minutes * 60 + seconds;
  };

  const secondsToTimeFormat = (secs) => {
    const mins = Math.floor(secs / 60);
    const secsRemainder = secs % 60;
    return `+${mins}:${String(secsRemainder).padStart(2, '0')}`;
  };

  const winnerSeconds = timeToSeconds(results[0]);

  return results.map((time, index) => {
    if (index === 0) return '0';
    const currentSeconds = timeToSeconds(time);
    const difference = currentSeconds - winnerSeconds;
    return secondsToTimeFormat(difference);
  });
}
```
