---
id: 697a49e6ff50d756c9b69365
title: "도전 과제 188: 2026년 동계 경기 9일차: 스켈레톤"
challengeType: 28
dashedName: challenge-188
---

# --description--

스켈레톤 트랙의 곡선을 나타내는 문자열이 주어질 때, 트랙의 난이도를 판단하세요.

- 주어진 문자열은 다음 문자만 포함합니다:
  - `"L"`는 왼쪽 회전을 나타냅니다
  - `"R"`는 오른쪽 회전을 나타냅니다
  - `"S"`는 직선 구간을 나타냅니다

- 방향 전환마다 15점이 추가됩니다 (`"L"` 다음에 `"R"` 또는 그 반대 경우).
- 그 외 모든 곡선은 각각 5점이 추가됩니다 (다른 모든 `"L"` 또는 `"R"` 문자).
- 직선 구간은 0점입니다.

트랙 난이도는 총 점수에 따라 결정됩니다. 다음을 반환하세요:

- 총점이 0 - 100이면 `"Easy"`
- 총점이 101 - 200이면 `"Medium"`
- 총점이 200 초과이면 `"Hard"`

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")`는 `"Easy"`을 반환해야 합니다.

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")`는 `"Hard"`을 반환해야 합니다.

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")`는 `"Medium"`을 반환해야 합니다.

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")`는 `"Hard"`을 반환해야 합니다.

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")`는 `"Medium"`을 반환해야 합니다.

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")`는 `"Easy"`을 반환해야 합니다.

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
