---
id: 697a49e6ff50d756c9b6935e
title: "도전 과제 181: 2026년 동계 경기 2일차: 스노보드"
challengeType: 28
dashedName: challenge-181
---

# --description--

스노보더의 시작 자세와 도 단위 회전이 주어졌을 때, 착지 자세를 결정하세요.

- 스노보더의 자세는 `"Regular"` 또는 `"Goofy"`입니다.
- 묘기 회전은 90도의 배수입니다. 양수는 시계 방향 회전을, 음수는 반시계 방향 회전을 나타냅니다.
- 착지 자세는 180도 회전할 때마다 뒤집힙니다.

예를 들어, `"Regular"`와 `90`가 주어지면 `"Regular"`를 반환하세요. `"Regular"`와 `180` 도가 주어지면 `"Goofy"`를 반환하세요.

# --hints--

`getLandingStance("Regular", 90)`는 `"Regular"`를 반환해야 합니다.

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)`는 `"Goofy"`를 반환해야 합니다.

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)`는 `"Regular"`를 반환해야 합니다.

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)`는 `"Goofy"`를 반환해야 합니다.

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)`는 `"Goofy"`를 반환해야 합니다.

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)`는 `"Regular"`를 반환해야 합니다.

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
