---
id: 697a49e6ff50d756c9b69366
title: "도전 과제 189: 2026년 동계 경기 10일차: 알파인 스키"
challengeType: 28
dashedName: challenge-189
---

# --description--

스키장의 수직 낙차, 수평 거리, 그리고 유형이 주어졌을 때, 경사의 난이도 등급을 결정하세요.

난이도 등급을 결정하려면:

- 낙차를 거리로 나누어 경사의 가파름을 계산하세요.
- 그런 다음, 경사 가파름에 경사 유형에 따른 조정 계수를 곱하세요:
  - `"Downhill"`: 경사 가파름에 1.2를 곱하세요
  - `"Slalom"`: 경사 가파름에 0.9를 곱하세요
  - `"Giant Slalom"`: 경사 가파름에 1.0을 곱하세요

반환값:

- 조정된 경사 가파름이 0.1 이하이면 `"Green"`을 반환하세요
- 조정된 경사 가파름이 0.1 초과 0.25 이하이면 `"Blue"`를 반환하세요
- 조정된 경사 가파름이 0.25 초과이면 `"Black"`를 반환하세요

# --hints--

`getHillRating(95, 900, "Slalom")`는 `"Green"`을 반환해야 합니다.

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")`는 `"Black"`을 반환해야 합니다.

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")`는 `"Blue"`을 반환해야 합니다.

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")`는 `"Green"`을 반환해야 합니다.

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")`는 `"Blue"`을 반환해야 합니다.

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")`는 `"Black"`을 반환해야 합니다.

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```
