---
id: 697a49e9860d24853adef67c
title: "도전 과제 191: 2026년 동계 경기 12일차: 봅슬레이"
challengeType: 28
dashedName: challenge-191
---

# --description--

봅슬레이 팀 선수들의 몸무게를 나타내는 배열과 봅슬레이 자체 무게를 나타내는 숫자가 주어질 때, 팀이 경기에 출전할 자격이 있는지 판단하세요.

- 배열의 길이는 팀 인원 수를 결정합니다: 1명, 2명 또는 4명 팀입니다.
- 모든 몸무게 값은 킬로그램(kg) 단위입니다.

봅슬레이(썰매 자체)는 다음과 같은 최솟값 무게를 가져야 합니다:

- 1인 팀의 경우 162kg
- 2인 팀의 경우 170kg
- 4인 팀의 경우 210kg

봅슬레이 총 무게(선수 몸무게 합 + 썰매 무게)는 다음을 초과해서는 안 됩니다:

- 1인 팀의 경우 247kg
- 2인 팀의 경우 390kg
- 4인 팀의 경우 630kg

팀이 모든 조건을 충족하면 `"Eligible"`를 반환하고, 하나 이상의 조건을 충족하지 못하면 `"Not Eligible"`를 반환하세요.

# --hints--

`checkEligibility([78], 165)`는 `"Eligible"`를 반환해야 합니다.

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)`는 `"Not Eligible"`를 반환해야 합니다.

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)`는 `"Not Eligible"`를 반환해야 합니다.

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)`는 `"Eligible"`를 반환해야 합니다.

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)`는 `"Not Eligible"`를 반환해야 합니다.

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)`는 `"Not Eligible"`를 반환해야 합니다.

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)`는 `"Eligible"`를 반환해야 합니다.

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)`는 `"Not Eligible"`를 반환해야 합니다.

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)`는 `"Not Eligible"`를 반환해야 합니다.

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
