---
id: 697a49e9860d24853adef67b
title: "도전 과제 190: 2026년 동계 경기 11일차: 아이스하키"
challengeType: 28
dashedName: challenge-190
---

# --description--

6개 아이스하키 팀과 라운드 로빈 경기 후 기록이 담긴 배열이 주어졌을 때, 준결승 라운드의 대진을 결정하세요.

- 각 배열 항목은 팀과 `"TEAM: W-OTW-OTL-L"` 형식의 기록을 포함합니다. 여기서:
  - `"W"`은 정규 시간 승리 횟수로, 각각 3점입니다
  - `"OTW"`은 연장전 승리 횟수로, 각각 2점입니다
  - `"OTL"`은 연장전 패배 횟수로, 각각 1점입니다
  - `"L"`은 패배 횟수로, 점수는 0점입니다
  
예를 들어, `"FIN: 2-2-1-0"`는 기록을 합산해 11점을 가집니다.
  
각 팀의 총 점수를 구하고 `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."`를 반환하세요. 예를 들어, `"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])`는 `"The semi-final games will be FIN vs SWE and CAN vs USA."`를 반환해야 합니다.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.");
```

`getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])`는 `"The semi-final games will be USA vs CZE and CAN vs FIN."`를 반환해야 합니다.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.");
```

`getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])`는 `"The semi-final games will be CAN vs ITA and USA vs CZE."`를 반환해야 합니다.

```js
assert.equal(getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.");
```

`getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])`는 `"The semi-final games will be JPN vs ITA and AUT vs KOR."`를 반환해야 합니다.

```js
assert.equal(getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.");
```

# --seed--

## --seed-contents--

```js
function getSemifinalMatchups(teams) {

  return teams;
}
```

# --solutions--

```js
function getSemifinalMatchups(teams) {
    const teamPoints = teams.map(teamStr => {
        const [name, record] = teamStr.split(": ");
        const [W, OTW, OTL, L] = record.split("-").map(Number);
        const points = W * 3 + OTW * 2 + OTL * 1 + L * 0;
        return { name, points };
    });

    teamPoints.sort((a, b) => b.points - a.points);
    const [first, second, third, fourth] = teamPoints;

    return `The semi-final games will be ${first.name} vs ${fourth.name} and ${second.name} vs ${third.name}.`;
}
```
