---
id: 697a49e6ff50d756c9b69364
title: "도전 과제 187: 2026년 동계 올림픽 8일차: 루지"
challengeType: 28
dashedName: challenge-187
---

# --description--

트랙 구간을 완주하는 데 걸린 시간(초)을 나타내는 다섯 개 숫자로 이루어진 배열이 주어졌을 때, 가장 빠른 속도를 낸 구간과 그 속도를 구하세요.

트랙은 다음 구간들로 나뉘어 있습니다:

- 구간 1: 320미터
- 구간 2: 280미터
- 구간 3: 350미터
- 구간 4: 300미터
- 구간 5: 250미터

주어진 배열의 첫 번째 값은 구간 1의 시간, 두 번째 값은 구간 2의 시간, 이런 식으로 대응합니다.

구간의 속도(초당 미터)를 계산하려면 거리를 시간으로 나누세요.

`"The luger's fastest speed was X m/s on segment Y."`를 반환하세요. 여기서 `X`은 소수점 둘째 자리까지 반올림한 가장 빠른 속도이고, `Y`는 가장 빠른 속도가 나온 구간 번호입니다.

# --hints--

`getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128])`는 `"The luger's fastest speed was 35.07 m/s on segment 5."`를 반환해야 합니다.

```js
assert.equal(getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.");
```

`getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284])`는 `"The luger's fastest speed was 37.75 m/s on segment 2."`를 반환해야 합니다.

```js
assert.equal(getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.");
```

`getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912])`는 `"The luger's fastest speed was 38.49 m/s on segment 3."`를 반환해야 합니다.

```js
assert.equal(getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.");
```

`getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840])`는 `"The luger's fastest speed was 37.69 m/s on segment 1."`를 반환해야 합니다.

```js
assert.equal(getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.");
```

`getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508])`는 `"The luger's fastest speed was 39.24 m/s on segment 4."`를 반환해야 합니다.

```js
assert.equal(getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.");
```

# --seed--

## --seed-contents--

```js
function getFastestSpeed(times) {

  return times;
}
```

# --solutions--

```js
function getFastestSpeed(times) {
  const distances = [320, 280, 350, 300, 250];
  let fastestSpeed = 0;
  let fastestSegment = 0;

  for (let i = 0; i < times.length; i++) {
    const speed = distances[i] / times[i];
    if (speed > fastestSpeed) {
      fastestSpeed = speed;
      fastestSegment = i + 1;
    }
  }

  return `The luger's fastest speed was ${fastestSpeed.toFixed(2)} m/s on segment ${fastestSegment}.`;
}
```
