---
id: 697a49e6ff50d756c9b6935f
title: "도전 과제 182: 2026년 동계 경기 3일차: 바이애슬론"
challengeType: 28
dashedName: challenge-182
---

# --description--

각 값이 바이애슬론 한 라운드에서 맞힌 표적 수를 나타내는 정수 배열이 주어질 때, 선수에게 부과되는 총 벌칙 거리(스키를 더 타야 하는 거리)를 반환하세요.

- 각 라운드는 5개의 표적으로 구성됩니다.
- 놓친 표적 하나당 150미터 벌칙 루프가 부과됩니다.

# --hints--

`calculatePenaltyDistance([4, 4])`는 `300`를 반환해야 합니다.

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])`는 `0`를 반환해야 합니다.

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])`는 `450`를 반환해야 합니다.

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])`는 `150`를 반환해야 합니다.

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])`는 `1500`를 반환해야 합니다.

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
