---
id: 697a49e6ff50d756c9b69360
title: "도전 과제 183: 2026년 동계 올림픽 4일차: 스키 점프"
challengeType: 28
dashedName: challenge-183
---

# --description--

거리 점수, 스타일 점수, 바람 보정 값, K-포인트 보너스 값을 받아 스키 점프 점수를 계산하고 메달 획득 여부를 결정하세요.

- 점수는 위 네 가지 값을 모두 더해서 계산합니다.

다른 점프 선수들의 현재 총점은 다음과 같습니다:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- 점수가 가장 높으면 `"Gold"`를 반환하세요
- 두 번째로 높으면 `"Silver"`를 반환하세요
- 세 번째로 높으면 `"Bronze"`를 반환하세요
- 그렇지 않으면 `"No Medal"`를 반환하세요

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)`는 `"Gold"`를 반환해야 합니다.

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)`는 `"Bronze"`를 반환해야 합니다.

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)`는 `"Silver"`를 반환해야 합니다.

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)`는 `"No Medal"`를 반환해야 합니다.

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)`는 `"Gold"`를 반환해야 합니다.

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)`는 `"No Medal"`를 반환해야 합니다.

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```
