---
id: 697a49e9860d24853adef680
title: "도전 과제 195: 2026년 동계 경기 16일차: 컬링"
challengeType: 28
dashedName: challenge-195
---

# --description--

컬링 라운드가 끝난 후 "하우스"를 나타내는 5x5 행렬이 주어질 때, 어느 팀이 점수를 얻고 몇 점을 얻는지 결정하세요.

배치:

- 중앙 셀(인덱스 `[2, 2]`)은 "버튼"입니다.
- 버튼을 직접 둘러싼 8개의 셀은 1번 링을 나타냅니다.
- 하우스의 바깥 가장자리에 있는 16개의 셀은 2번 링을 나타냅니다.

주어진 행렬에서:

- `"."`은 빈 공간을 나타냅니다.
- `"R"`는 빨간 돌이 있는 공간을 나타냅니다.
- `"Y"`는 노란 돌이 있는 공간을 나타냅니다.

점수 규칙:

- 한 라운드에 한 팀만 점수를 얻을 수 있습니다.
- 버튼에 가장 가까운 돌을 가진 팀이 점수를 얻습니다.
- 점수를 얻는 팀은 상대 팀의 가장 가까운 돌보다 버튼에 더 가까운 자신의 돌 하나당 1점을 얻습니다.
- 링 번호가 낮을수록 돌이 중심에 더 가깝습니다.
- 양 팀의 가장 가까운 돌이 중심에서 같은 거리에 있으면 어느 팀도 점수를 얻지 못합니다.

반환값:

- `"team: number_of_points"` 형식의 문자열을 반환하세요. 예: `"R: 2"`.
- 어느 팀도 점수를 얻지 못하면 `"No points awarded"`를 반환하세요.

예를 들어, 다음이 주어지면:

```js
[
  [".", ".", "R", ".", "."],
  [".", "R", ".", ".", "."],
  ["Y", ".", ".", ".", "."],
  [".", "R", ".", ".", "."],
  [".", ".", ".", ".", "."]
]
```

`"R: 2"`를 반환하세요. 1번 링에 있는 두 개의 빨간 돌이 가장 가깝고 노란 돌 중 가장 가까운 돌보다 더 가깝기 때문입니다.

# --hints--

`scoreCurling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]])`는 `"R: 2"`를 반환해야 합니다.

```js
assert.equal(scoreCurling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]]), "R: 2");
```

`scoreCurling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]])`는 `"Y: 3"`를 반환해야 합니다.

```js
assert.equal(scoreCurling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]]), "Y: 3");
```

`scoreCurling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]])`는 `"No points awarded"`를 반환해야 합니다.

```js
assert.equal(scoreCurling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]]), "No points awarded");
```

`scoreCurling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]])`는 `"R: 4"`를 반환해야 합니다.

```js
assert.equal(scoreCurling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]]), "R: 4");
```

`scoreCurling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]])`는 `"Y: 1"`를 반환해야 합니다.

```js
assert.equal(scoreCurling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]]), "Y: 1");
```

`scoreCurling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]])`는 `"No points awarded"`를 반환해야 합니다.

```js
assert.equal(scoreCurling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]]), "No points awarded");
```

# --seed--

## --seed-contents--

```js
function scoreCurling(house) {

  return house;
}
```

# --solutions--

```js
function scoreCurling(house) {
  const stones = { R: [], Y: [] };

  for (let i = 0; i < 5; i++) {
    for (let j = 0; j < 5; j++) {
      const cell = house[i][j];
      if (cell === 'R' || cell === 'Y') {
        const distance = Math.max(Math.abs(i - 2), Math.abs(j - 2));
        stones[cell].push(distance);
      }
    }
  }

  const redClosest = stones.R.length > 0 ? Math.min(...stones.R) : Infinity;
  const yellowClosest = stones.Y.length > 0 ? Math.min(...stones.Y) : Infinity;

  if (redClosest === Infinity && yellowClosest === Infinity) {
    return "No points awarded";
  }

  if (redClosest === yellowClosest) {
    return "No points awarded";
  }

let winner, opponentClosest, winningStones;
  if (redClosest < yellowClosest) {
    winner = 'R';
    opponentClosest = yellowClosest;
    winningStones = stones.R;
  } else {
    winner = 'Y';
    opponentClosest = redClosest;
    winningStones = stones.Y;
  }

  const points = winningStones.filter(d => d < opponentClosest).length;

  return points === 0 ? "No points awarded" : `${winner}: ${points}`;
}
```
