---
id: 697a49e9860d24853adef681
title: "도전 과제 196: 2026년 동계 올림픽 17일차: 폐막일"
challengeType: 28
dashedName: challenge-196
---

# --description--

메달 수상자 2D 배열이 주어질 때, 각 국가별 메달 수를 CSV 문자열로 반환하세요.

- 주어진 배열에서 각 하위 배열은 단일 이벤트를 나타냅니다: `[gold_winner, silver_winner, bronze_winner]`

- 반환할 CSV 문자열은 다음 형식을 가져야 하며, 첫 번째 줄은 헤더입니다:

```sh
Country,Gold,Silver,Bronze,Total
country_name,gold_count,silver_count,bronze_count,total_medals
```

- 새 줄은 줄 바꿈 문자(`"\n"`)로 구분하세요.
- 쉼표 주변이나 줄 끝에 공백을 포함하지 마세요.
- 반환된 CSV는 금메달 수를 기준으로 내림차순 정렬하세요. 두 국가가 금메달 수가 같으면 알파벳 순으로 정렬하세요.

예를 들어, 다음이 주어지면:

```js
[
  ["USA", "CAN", "NOR"],
  ["NOR", "USA", "CAN"],
  ["USA", "NOR", "SWE"]
]
```

다음을 반환하세요:

```sh
"Country,Gold,Silver,Bronze,Total\nUSA,2,1,0,3\nNOR,1,1,1,3\nCAN,0,1,1,2\nSWE,0,0,1,1"
```

출력하면 다음과 같습니다:

```sh
Country,Gold,Silver,Bronze,Total
USA,2,1,0,3
NOR,1,1,1,3
CAN,0,1,1,2
SWE,0,0,1,1
```

# --hints--

`countMedals([["USA", "CAN", "NOR"], ["NOR", "USA", "CAN"], ["USA", "NOR", "SWE"]])`는 `"Country,Gold,Silver,Bronze,Total\nUSA,2,1,0,3\nNOR,1,1,1,3\nCAN,0,1,1,2\nSWE,0,0,1,1"`를 반환해야 합니다.

```js
assert.equal(countMedals([["USA", "CAN", "NOR"], ["NOR", "USA", "CAN"], ["USA", "NOR", "SWE"]]), "Country,Gold,Silver,Bronze,Total\nUSA,2,1,0,3\nNOR,1,1,1,3\nCAN,0,1,1,2\nSWE,0,0,1,1");
```

`countMedals([["NOR","SWE","FIN"]])`는 `"Country,Gold,Silver,Bronze,Total\nNOR,1,0,0,1\nFIN,0,0,1,1\nSWE,0,1,0,1"`를 반환해야 합니다.

```js
assert.equal(countMedals([["NOR","SWE","FIN"]]), "Country,Gold,Silver,Bronze,Total\nNOR,1,0,0,1\nFIN,0,0,1,1\nSWE,0,1,0,1");
```

`countMedals([["ITA", "CHN", "CHN"], ["JPN", "ITA", "JPN"]])`는 `"Country,Gold,Silver,Bronze,Total\nITA,1,1,0,2\nJPN,1,0,1,2\nCHN,0,1,1,2"`를 반환해야 합니다.

```js
assert.equal(countMedals([["ITA", "CHN", "CHN"], ["JPN", "ITA", "JPN"]]), "Country,Gold,Silver,Bronze,Total\nITA,1,1,0,2\nJPN,1,0,1,2\nCHN,0,1,1,2");
```

`countMedals([["USA","CAN","NOR"], ["GER","FRA","ITA"], ["JPN","KOR","CHN"], ["SWE","FIN","NOR"], ["CAN","USA","SWE"], ["FRA","GER","ITA"]])`는 `"Country,Gold,Silver,Bronze,Total\nCAN,1,1,0,2\nFRA,1,1,0,2\nGER,1,1,0,2\nJPN,1,0,0,1\nSWE,1,0,1,2\nUSA,1,1,0,2\nCHN,0,0,1,1\nFIN,0,1,0,1\nITA,0,0,2,2\nKOR,0,1,0,1\nNOR,0,0,2,2"`를 반환해야 합니다.

```js
assert.equal(countMedals([["USA","CAN","NOR"], ["GER","FRA","ITA"], ["JPN","KOR","CHN"], ["SWE","FIN","NOR"], ["CAN","USA","SWE"], ["FRA","GER","ITA"]]), "Country,Gold,Silver,Bronze,Total\nCAN,1,1,0,2\nFRA,1,1,0,2\nGER,1,1,0,2\nJPN,1,0,0,1\nSWE,1,0,1,2\nUSA,1,1,0,2\nCHN,0,0,1,1\nFIN,0,1,0,1\nITA,0,0,2,2\nKOR,0,1,0,1\nNOR,0,0,2,2");
```

`countMedals([["ESP","ITA","FRA"], ["ITA","ESP","GER"], ["NOR","SWE","FIN"], ["FIN","NOR","SWE"], ["USA","CAN","MEX"], ["CAN","USA","MEX"], ["JPN","KOR","CHN"], ["CHN","JPN","KOR"]])`는 `"Country,Gold,Silver,Bronze,Total\nCAN,1,1,0,2\nCHN,1,0,1,2\nESP,1,1,0,2\nFIN,1,0,1,2\nITA,1,1,0,2\nJPN,1,1,0,2\nNOR,1,1,0,2\nUSA,1,1,0,2\nFRA,0,0,1,1\nGER,0,0,1,1\nKOR,0,1,1,2\nMEX,0,0,2,2\nSWE,0,1,1,2"`를 반환해야 합니다.

```js
assert.equal(countMedals([["ESP","ITA","FRA"], ["ITA","ESP","GER"], ["NOR","SWE","FIN"], ["FIN","NOR","SWE"], ["USA","CAN","MEX"], ["CAN","USA","MEX"], ["JPN","KOR","CHN"], ["CHN","JPN","KOR"]]), "Country,Gold,Silver,Bronze,Total\nCAN,1,1,0,2\nCHN,1,0,1,2\nESP,1,1,0,2\nFIN,1,0,1,2\nITA,1,1,0,2\nJPN,1,1,0,2\nNOR,1,1,0,2\nUSA,1,1,0,2\nFRA,0,0,1,1\nGER,0,0,1,1\nKOR,0,1,1,2\nMEX,0,0,2,2\nSWE,0,1,1,2");
```

`countMedals([["USA","CAN","GER"], ["NOR","SWE","FIN"], ["USA","NOR","SWE"], ["GER","FRA","ITA"], ["JPN","KOR","CHN"], ["USA","GER","CAN"], ["SWE","NOR","FIN"], ["CAN","USA","NOR"], ["FRA","GER","ITA"], ["JPN","CHN","KOR"], ["SWE","FIN","NOR"], ["GER","ITA","FRA"]])`는 `"Country,Gold,Silver,Bronze,Total\nUSA,3,1,0,4\nGER,2,2,1,5\nJPN,2,0,0,2\nSWE,2,1,1,4\nCAN,1,1,1,3\nFRA,1,1,1,3\nNOR,1,2,2,5\nCHN,0,1,1,2\nFIN,0,1,2,3\nITA,0,1,2,3\nKOR,0,1,1,2"`를 반환해야 합니다.

```js
assert.equal(countMedals([["USA","CAN","GER"], ["NOR","SWE","FIN"], ["USA","NOR","SWE"], ["GER","FRA","ITA"], ["JPN","KOR","CHN"], ["USA","GER","CAN"], ["SWE","NOR","FIN"], ["CAN","USA","NOR"], ["FRA","GER","ITA"], ["JPN","CHN","KOR"], ["SWE","FIN","NOR"], ["GER","ITA","FRA"]]), "Country,Gold,Silver,Bronze,Total\nUSA,3,1,0,4\nGER,2,2,1,5\nJPN,2,0,0,2\nSWE,2,1,1,4\nCAN,1,1,1,3\nFRA,1,1,1,3\nNOR,1,2,2,5\nCHN,0,1,1,2\nFIN,0,1,2,3\nITA,0,1,2,3\nKOR,0,1,1,2");
```

# --seed--

## --seed-contents--

```js
function countMedals(winners) {

  return winners;
}
```

# --solutions--

```js
function countMedals(winners) {
  const counts = {};

  for (const [gold, silver, bronze] of winners) {
    if (!counts[gold]) counts[gold] = { Gold: 0, Silver: 0, Bronze: 0 };
    if (!counts[silver]) counts[silver] = { Gold: 0, Silver: 0, Bronze: 0 };
    if (!counts[bronze]) counts[bronze] = { Gold: 0, Silver: 0, Bronze: 0 };

    counts[gold].Gold += 1;
    counts[silver].Silver += 1;
    counts[bronze].Bronze += 1;
  }

  const countries = Object.entries(counts).map(([country, medals]) => ({
    Country: country,
    Gold: medals.Gold,
    Silver: medals.Silver,
    Bronze: medals.Bronze,
    Total: medals.Gold + medals.Silver + medals.Bronze
  }));

  countries.sort((a, b) => b.Gold - a.Gold || a.Country.localeCompare(b.Country));

  const lines = ["Country,Gold,Silver,Bronze,Total"];
  for (const c of countries) {
    lines.push(`${c.Country},${c.Gold},${c.Silver},${c.Bronze},${c.Total}`);
  }

  return lines.join("\n");
}
```
