---
id: 697a49e9860d24853adef67e
title: "도전 과제 193: 2026년 동계 경기 14일차: 스키 마운티니어링"
challengeType: 28
dashedName: challenge-193
---

# --description--

산의 적설 깊이와 경사를 고려해 눈사태 위험이 있는지 판단하세요.

- 적설 깊이 값은 `"Shallow"`, `"Moderate"`, 또는 `"Deep"`입니다.
- 경사 값은 `"Gentle"`, `"Steep"`, 또는 `"Very Steep"`입니다.

다음 표를 기준으로 `"Safe"` 또는 `"Risky"`을 반환하세요:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalancheRisk("Shallow", "Gentle")`은 `"Safe"`를 반환해야 합니다.

```js
assert.equal(avalancheRisk("Shallow", "Gentle"), "Safe");
```

`avalancheRisk("Shallow", "Steep")`은 `"Safe"`를 반환해야 합니다.

```js
assert.equal(avalancheRisk("Shallow", "Steep"), "Safe");
```

`avalancheRisk("Shallow", "Very Steep")`은 `"Safe"`를 반환해야 합니다.

```js
assert.equal(avalancheRisk("Shallow", "Very Steep"), "Safe");
```

`avalancheRisk("Moderate", "Gentle")`은 `"Safe"`를 반환해야 합니다.

```js
assert.equal(avalancheRisk("Moderate", "Gentle"), "Safe");
```

`avalancheRisk("Moderate", "Steep")`은 `"Risky"`를 반환해야 합니다.

```js
assert.equal(avalancheRisk("Moderate", "Steep"), "Risky");
```

`avalancheRisk("Moderate", "Very Steep")`은 `"Risky"`를 반환해야 합니다.

```js
assert.equal(avalancheRisk("Moderate", "Very Steep"), "Risky");
```

`avalancheRisk("Deep", "Gentle")`은 `"Safe"`를 반환해야 합니다.

```js
assert.equal(avalancheRisk("Deep", "Gentle"), "Safe");
```

`avalancheRisk("Deep", "Steep")`은 `"Risky"`를 반환해야 합니다.

```js
assert.equal(avalancheRisk("Deep", "Steep"), "Risky");
```

`avalancheRisk("Deep", "Very Steep")`은 `"Risky"`를 반환해야 합니다.

```js
assert.equal(avalancheRisk("Deep", "Very Steep"), "Risky");
```

# --seed--

## --seed-contents--

```js
function avalancheRisk(snowDepth, slope) {

  return snowDepth;
}
```

# --solutions--

```js
function avalancheRisk(snowDepth, slope) {
  const riskTable = {
    Gentle: {
      Shallow: "Safe",
      Moderate: "Safe",
      Deep: "Safe",
    },
    Steep: {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
    "Very Steep": {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
  };

  return riskTable[slope][snowDepth];
}
```
