---
id: 697a49e9860d24853adef67f
title: "도전 과제 194: 2026년 동계 경기 15일차: 프리스타일 스키"
challengeType: 28
dashedName: challenge-194
---

# --description--

두 단어로 이루어진 묘기 이름이 프리스타일 스키에서 유효한 묘기 이름인지 판단하세요.

묘기는 첫 번째 단어가 유효한 첫 단어 목록에 있고, 두 번째 단어가 유효한 두 번째 단어 목록에 있을 때 유효합니다.

- 두 단어는 한 칸의 공백으로 구분됩니다.

유효한 첫 단어:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

유효한 두 번째 단어:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`isValidTrick("Polar Vortex")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidTrick("Polar Vortex"));
```

`isValidTrick("Solar Icequake")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidTrick("Solar Icequake"));
```

`isValidTrick("Thunder Blizzard")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidTrick("Thunder Blizzard"));
```

`isValidTrick("Phantom Frostbite")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidTrick("Phantom Frostbite"));
```

`isValidTrick("Ghost Avalanche")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidTrick("Ghost Avalanche"));
```

`isValidTrick("Snowstorm Shadow")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isValidTrick("Snowstorm Shadow"));
```

`isValidTrick("Solar Sky")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isValidTrick("Solar Sky"));
```

# --seed--

## --seed-contents--

```js
function isValidTrick(trickName) {

  return trickName;
}
```

# --solutions--

```js
function isValidTrick(trickName) {
  const validFirst = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"];
  const validSecond = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"];

  const words = trickName.split(" ");
  const [first, second] = words;
  return validFirst.includes(first) && validSecond.includes(second);
}
```
