---
id: 698a1a73ade5ac0e19180fa7
title: "도전 과제 204: 문자 합계"
challengeType: 28
dashedName: challenge-204
---

# --description--

문자열이 주어지면 그 문자열의 문자 합계를 반환하세요.

- 문자는 대문자 또는 소문자 A-Z입니다
- 문자 값은 `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26입니다
- 대문자와 소문자는 같은 값을 가집니다.
- 문자 이외의 모든 문자는 무시하세요.

# --hints--

`sumLetters("Hello")`는 `52`를 반환해야 합니다.

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")`는 `94`를 반환해야 합니다.

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")`는 `473`를 반환해야 합니다.

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")`는 `1681`를 반환해야 합니다.

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")`는 `0`를 반환해야 합니다.

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```
