---
id: 698a1a73ade5ac0e19180fa9
title: "도전 과제 206: 카드 값 계산"
challengeType: 28
dashedName: challenge-206
---

# --description--

주어진 카드 배열에서 각 카드의 숫자 값을 담은 새 배열을 반환하세요.

카드 값:

- 에이스(`"A"`)는 값이 1입니다.
- 숫자 카드(`"2"` - `"10"`)는 각각 2부터 10까지 해당 숫자 값을 가집니다.
- 그림 카드: 잭(`"J"`), 퀸(`"Q"`), 킹(`"K"`)은 각각 10의 값입니다.

무늬:

- 각 카드는 스페이드(`"S"`), 클로버(`"C"`), 다이아몬드(`"D"`), 하트(`"H"`) 중 하나의 무늬를 가집니다.

카드 형식:

- 각 카드는 문자열로 나타냅니다: `"valueSuit"`. 예를 들어, `"AS"`은 스페이드 에이스, `"10H"`는 하트 10, `"QC"`은 클로버 퀸입니다.

# --hints--

`cardValues(["3H", "4D", "5S"])`는 `[3, 4, 5]`를 반환해야 합니다.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])`는 `[1, 10, 10, 6, 7]`를 반환해야 합니다.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])`는 `[8, 10, 2, 10, 9]`를 반환해야 합니다.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])`는 `[1, 10]`를 반환해야 합니다.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])`는 `[10, 10, 10, 10, 1]`를 반환해야 합니다.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
