---
id: 698a1a73ade5ac0e19180fa0
title: "도전 과제 197: 혈액형 적합성"
challengeType: 28
dashedName: challenge-197
---

# --description--

기증자 혈액형과 수혈자 혈액형이 주어졌을 때, 기증자가 수혈자에게 혈액을 줄 수 있는지 판단하세요.

각 혈액형은 다음으로 구성됩니다:

- 한 글자: `"A"`, `"B"`, `"AB"`, 또는 `"O"`
- 그리고 Rh 인자: `"+"` 또는 `"-"`

혈액형은 유효한 글자 중 하나와 Rh 인자가 결합된 형태입니다. 예를 들어, `"AB+"`와 `"O-"`는 유효한 혈액형입니다.

글자 규칙:

- `"O"`는 다른 글자 유형에 기증할 수 있습니다.
- `"A"`는 `"A"`와 `"AB"`에 기증할 수 있습니다.
- `"B"`는 `"B"`와 `"AB"`에 기증할 수 있습니다.
- `"AB"`는 `"AB"`에만 기증할 수 있습니다.

Rh 규칙:

- 음성(`"-"`)은 `"-"`와 `"+"` 모두에 기증할 수 있습니다.
- 양성(`"+"`)은 `"+"`에만 기증할 수 있습니다.

기증자가 수혈자에게 기증할 수 있으려면 글자 규칙과 Rh 규칙 모두 충족해야 합니다.

# --hints--

`canDonate("B+", "B+")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
