---
id: 698a1a863194f1f4e63f645f
title: "도전 과제 208: 길 따라 이동하기"
challengeType: 28
dashedName: challenge-208
---

# --description--

길 지도를 나타내는 문자열 배열이 주어질 때, 목표 지점에 도달하기 위해 필요한 이동 경로를 문자열로 반환하세요.

주어진 문자열에는 다음 값들이 포함됩니다:

- `"C"`: 현재 위치
- `"G"`: 목표 지점
- `"T"`: 이동 가능한 길 부분
- `"-"`: 이동 불가능한 지도 부분

현재 위치에서 목표 지점까지 길을 따라 이동하는 데 필요한 경로를 나타내는 문자열을 반환하세요. 여기서:

- `"R"`는 오른쪽으로 이동
- `"D"`는 아래로 이동
- `"L"`는 왼쪽으로 이동
- `"U"`는 위로 이동

- 현재 위치에서 목표 지점까지 항상 하나의 연속된 길이 있으며, 갈래가 없습니다.
- 각 길 위치는 최대 두 개의 이동 가능한 위치와 인접해 있습니다.

예를 들어, 다음이 주어지면:

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

`"RDDRDD"`를 반환하세요.

# --hints--

`navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])`는 `"RDDRDD"`를 반환해야 합니다.

```js
assert.equal(navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD");
```

`navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"])`는 `"RRUUURR"`를 반환해야 합니다.

```js
assert.equal(navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR");
```

`navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])`는 `"DLDDRRRRD"`를 반환해야 합니다.

```js
assert.equal(navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD");
```

`navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])`는 `"RRRDDL"`를 반환해야 합니다.

```js
assert.equal(navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL");
```

`navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])`는 `"ULLLUUURRRRRRDDDR"`를 반환해야 합니다.

```js
assert.equal(navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR");
```

# --seed--

## --seed-contents--

```js
function navigateTrail(map) {

  return map;
}
```

# --solutions--

```js
function navigateTrail(map) {
  const rows = map.length;
  const cols = map[0].length;

  const directions = [
    [0, 1, "R"],
    [1, 0, "D"],
    [0, -1, "L"],
    [-1, 0, "U"],
  ];

  let row, col;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (map[r][c] === "C") {
        row = r;
        col = c;
      }
    }
  }

  let result = "";
  let prev = null;

  while (map[row][col] !== "G") {
    for (let [dr, dc, move] of directions) {
      const newRow = row + dr;
      const newCol = col + dc;

      if (
        newRow >= 0 &&
        newRow < rows &&
        newCol >= 0 &&
        newCol < cols &&
        (map[newRow][newCol] === "T" ||
          map[newRow][newCol] === "G") &&
        (!prev || newRow !== prev[0] || newCol !== prev[1])
      ) {
        result += move;
        prev = [row, col];
        row = newRow;
        col = newCol;
        break;
      }
    }
  }

  return result;
}
```
