---
id: 698a1a73ade5ac0e19180fa3
title: "도전 과제 200: 문자와 숫자 개수 세기"
challengeType: 28
dashedName: challenge-200
---

# --description--

문자열이 주어지면, 그 안에 포함된 문자와 숫자의 개수를 알려주는 메시지를 반환하세요.

- 문자는 `A-Z`와 `a-z`입니다.
- 숫자는 `0-9`입니다.
- 그 외 모든 문자는 무시하세요.

`"The string has X letters and Y numbers."`를 반환하세요. 여기서 `"X"`는 문자의 개수이고 `"Y"`는 숫자의 개수입니다. 둘 중 하나라도 개수가 1이면, 해당 항목은 단수형을 사용하세요. 예를 들어 `"1 letter"` 대신 `"1 letters"`을, `"1 number"` 대신 `"1 numbers"`를 사용합니다.

# --hints--

`countLettersAndNumbers("helloworld123")`는 `"The string has 10 letters and 3 numbers."`를 반환해야 합니다.

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")`는 `"The string has 5 letters and 2 numbers."`를 반환해야 합니다.

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")`는 `"The string has 1 letter and 1 number."`를 반환해야 합니다.

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")`는 `"The string has 0 letters and 5 numbers."`를 반환해야 합니다.

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")`는 `"The string has 8 letters and 0 numbers."`를 반환해야 합니다.

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
