---
id: 698a1a863194f1f4e63f6460
title: "도전 과제 209: 요소 크기"
challengeType: 28
dashedName: challenge-209
---

# --description--

주어진 창 크기, 뷰포트 너비 `"vw"` 단위로 표현된 요소의 너비, 뷰포트 높이 `"vh"` 단위로 표현된 요소의 높이를 기준으로 요소의 크기를 픽셀 단위로 계산하세요.

- 주어진 창 크기와 반환할 요소 크기는 `"width x height"`, `"1200 x 800"` 형식의 문자열입니다.

- `"vw"` 단위는 창 너비의 백분율입니다. 예를 들어 `"50vw"` 는 창 너비의 50%를 의미합니다.
- `"vh"` 단위는 창 높이의 백분율입니다. 예를 들어 `"50vh"` 는 창 높이의 50%를 의미합니다.

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` 는 `"600 x 400"` 를 반환해야 합니다.

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` 은 `"80 x 240"` 를 반환해야 합니다.

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` 는 `"70 x 15"` 를 반환해야 합니다.

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` 는 `"1824 x 1080"` 를 반환해야 합니다.

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` 는 `"0 x 0"` 를 반환해야 합니다.

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` 는 `"1440 x 1026"` 를 반환해야 합니다.

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
