---
id: 698a1a73ade5ac0e19180fa8
title: "도전 과제 205: 완전 세제곱수 개수"
challengeType: 28
dashedName: challenge-205
---

# --description--

두 정수가 주어졌을 때, 두 수를 포함한 범위 내에 존재하는 완전 세제곱수의 개수를 구하세요.

- 더 작은 수가 첫 번째 인수라는 보장은 없습니다.
- 어떤 수가 완전 세제곱수라는 것은 정수 `n`가 존재하여 `n * n * n = number`인 경우입니다. 예를 들어, 27은 `3 * 3 * 3 = 27`이기 때문에 완전 세제곱수입니다.

# --hints--

`countPerfectCubes(3, 30)`는 `2`를 반환해야 합니다.

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)`는 `3`를 반환해야 합니다.

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)`는 `4`를 반환해야 합니다.

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)`는 `9`를 반환해야 합니다.

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)`는 `41`를 반환해야 합니다.

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
