---
id: 698a1a863194f1f4e63f645e
title: "도전 과제 207: 가장 작은 간격"
challengeType: 28
dashedName: challenge-207
---

# --description--

문자열이 주어지면, 두 개의 동일한 문자 사이에 있는 가장 작은 문자 수(가장 작은 간격)를 가진 부분 문자열을 반환하세요.

- 항상 일치하는 문자 쌍이 최소 한 쌍 이상 있습니다.
- 반환하는 부분 문자열에는 일치하는 문자는 포함하지 않습니다.
- 두 개 이상의 간격이 같은 길이일 경우, 첫 번째 간격에 해당하는 문자를 반환하세요.

예를 들어 `"ABCDAC"`가 주어지면 `"DA"`를 반환하는데, 그 이유는:

- 문자열에서 `"A"`와 `"C"`만 반복됩니다.
- 두 `"A"` 문자 사이의 문자 수는 3이고, 두 `"C"` 문자 사이의 문자 수는 2입니다.
- 따라서 두 `"C"` 문자 사이의 문자열을 반환합니다.

# --hints--

`smallestGap("ABCDAC")`는 `"DA"`를 반환해야 합니다.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")`는 `"e"`를 반환해야 합니다.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")`는 `"#q6e&rkf(p"`를 반환해야 합니다.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")`는 `""`를 반환해야 합니다.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")`는 `"fox"`를 반환해야 합니다.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
