---
id: 698a1a73ade5ac0e19180fa2
title: "도전 과제 199: 연속 차이"
challengeType: 28
dashedName: challenge-199
---

# --description--

숫자 배열이 주어지면, 각 숫자에서 다음 숫자로 가기 위해 필요한 값을 포함하는 새 배열을 반환하세요.

- 마지막 숫자의 경우 다음 숫자가 없으므로 `0`를 사용하세요.

예를 들어, `[1, 2, 4, 7]`이 주어지면 `[1, 2, 3, 0]`를 반환하세요.

# --hints--

`findDifferences([1, 2, 4, 7])`은 `[1, 2, 3, 0]`를 반환해야 합니다.

```js
assert.deepEqual(findDifferences([1, 2, 4, 7]), [1, 2, 3, 0]);
```

`findDifferences([10, 15, 19, 22, 24, 25])`는 `[5, 4, 3, 2, 1, 0]`를 반환해야 합니다.

```js
assert.deepEqual(findDifferences([10, 15, 19, 22, 24, 25]), [5, 4, 3, 2, 1, 0]);
```

`findDifferences([25, 20, 16, 13, 11, 10])`은 `[-5, -4, -3, -2, -1, 0]`를 반환해야 합니다.

```js
assert.deepEqual(findDifferences([25, 20, 16, 13, 11, 10]), [-5, -4, -3, -2, -1, 0]);
```

`findDifferences([0, 1, 2, 2, 3, 3, 4, 5])`는 `[ 1, 1, 0, 1, 0, 1, 1, 0 ]`를 반환해야 합니다.

```js
assert.deepEqual(findDifferences([0, 1, 2, 2, 3, 3, 4, 5]), [ 1, 1, 0, 1, 0, 1, 1, 0 ]);
```

`findDifferences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1])`은 `[1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]`를 반환해야 합니다.

```js
assert.deepEqual(findDifferences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1]), [1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]);
```

# --seed--

## --seed-contents--

```js
function findDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function findDifferences(arr) {
  const result = [];

  for (let i = 0; i < arr.length; i++) {
    result.push(arr[i + 1] - arr[i] || 0);
  }

  return result;
}
```
