---
id: 698a1a73ade5ac0e19180fa6
title: "도전 과제 203: 평탄화된 배열"
challengeType: 28
dashedName: challenge-203
---

# --description--

배열이 주어졌을 때, 그것이 평탄한지 판단하세요.

- 배열의 요소 중 배열이 하나도 없으면 평탄한 배열입니다.

# --hints--

`isFlat([1, 2, 3, 4])`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
