---
id: 698a1a73ade5ac0e19180fa4
title: "도전 과제 201: 행렬 이동"
challengeType: 28
dashedName: challenge-201
---

# --description--

숫자로 이루어진 행렬(배열의 배열)과 정수가 주어졌을 때, 행렬의 모든 값을 주어진 수만큼 이동하세요.

- 양수 이동은 값을 오른쪽으로 이동합니다.
- 음수 이동은 값을 왼쪽으로 이동합니다.

값은 감싸서 이동해야 합니다:

- 행렬을 하나의 연속된 수열로 간주하세요.
- 값이 한 행의 끝을 지나면 다음 행의 처음 위치에서 계속됩니다.
- 값이 행렬의 마지막 위치를 지나면 첫 위치로 감싸서 이동합니다.
- 왼쪽으로 이동할 때도 동일하게 적용됩니다.

예를 들어, 다음이 주어졌을 때:

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

이동 값이 `1`이면 모든 숫자를 오른쪽으로 한 칸 이동하세요:

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shiftMatrix([[1, 2, 3], [4, 5, 6]], 1)`은 `[[6, 1, 2], [3, 4, 5]]`를 반환해야 합니다.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6]], -1)`은 `[[2, 3, 4], [5, 6, 1]]`를 반환해야 합니다.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)`는 `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]`를 반환해야 합니다.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)`은 `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]`를 반환해야 합니다.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)`는 `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]`를 반환해야 합니다.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)`은 `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]`를 반환해야 합니다.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]);
```

# --seed--

## --seed-contents--

```js
function shiftMatrix(matrix, shift) {

  return matrix;
}
```

# --solutions--

```js
function shiftMatrix(matrix, shift) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const flat = matrix.flat();
  const n = flat.length;
  const s = ((shift % n) + n) % n;
  const shifted = flat.slice(-s).concat(flat.slice(0, n - s));
  const result = [];
  for (let i = 0; i < rows; i++) {
    result.push(shifted.slice(i * cols, (i + 1) * cols));
  }

  return result;
}
```
