---
id: 698a1a73ade5ac0e19180fa1
title: "도전 과제 198: 영업일 수 계산"
challengeType: 28
dashedName: challenge-198
---

# --description--

시작 날짜와 종료 날짜가 주어지면, 두 날짜 사이의 영업일 수를 반환하세요.

- 주어진 날짜는 `"YYYY-MM-DD"` 형식입니다.
- 평일은 영업일입니다(월요일부터 금요일까지).
- 주말은 영업일이 아닙니다(토요일과 일요일).
- 시작 날짜와 종료 날짜 모두 포함하여 계산하세요.

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")`은 `3`를 반환해야 합니다.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")`은 `4`를 반환해야 합니다.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")`은 `5`를 반환해야 합니다.

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")`은 `7`를 반환해야 합니다.

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")`은 `262`를 반환해야 합니다.

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
