---
id: 6994cff2290543b3aec9f510
title: "도전 과제 211: 배열 합계"
challengeType: 28
dashedName: challenge-211
---

# --description--

숫자 배열이 주어지면, 모든 숫자의 합을 반환하세요.

# --hints--

`sumArray([1, 2, 3, 4, 5])`는 `15`을 반환해야 합니다.

```js
assert.equal(sumArray([1, 2, 3, 4, 5]), 15);
```

`sumArray([42])`는 `42`을 반환해야 합니다.

```js
assert.equal(sumArray([42]), 42);
```

`sumArray([5, -2, 7, -3])`는 `7`을 반환해야 합니다.

```js
assert.equal(sumArray([5, -2, 7, -3]), 7);
```

`sumArray([203, 145, -129, 6293, 523, -919, 845, 2434])`는 `9395`을 반환해야 합니다.

```js
assert.equal(sumArray([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395);
```

`sumArray([0, 0])`는 `0`을 반환해야 합니다.

```js
assert.equal(sumArray([0, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function sumArray(numbers) {

  return numbers;
}
```

# --solutions--

```js
function sumArray(numbers) {

  return numbers.reduce((a, c) => c + a, 0);
}
```
