---
id: 6994cff2290543b3aec9f511
title: "도전 과제 212: 배열 삽입"
challengeType: 28
dashedName: challenge-212
---

# --description--

배열과 배열에 삽입할 값, 그리고 값을 삽입할 인덱스가 주어지면, 지정된 인덱스에 값이 삽입된 새 배열을 반환하세요.

# --hints--

`insertIntoArray([2, 4, 8, 10], 6, 2)`는 `[2, 4, 6, 8, 10]`를 반환해야 합니다.

```js
assert.deepEqual(insertIntoArray([2, 4, 8, 10], 6, 2), [2, 4, 6, 8, 10]);
```

`insertIntoArray(["the", "quick", "fox"], "brown", 2)`는 `["the", "quick", "brown", "fox"]`를 반환해야 합니다.

```js
assert.deepEqual(insertIntoArray(["the", "quick", "fox"], "brown", 2), ["the", "quick", "brown", "fox"]);
```

`insertIntoArray([], 0, 0)`는 `[0]`를 반환해야 합니다.

```js
assert.deepEqual(insertIntoArray([], 0, 0), [0]);
```

`insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5)`는 `[0, 1, 1, 2, 3, 5, 8, 13]`를 반환해야 합니다.

```js
assert.deepEqual(insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5), [0, 1, 1, 2, 3, 5, 8, 13]);
```

# --seed--

## --seed-contents--

```js
function insertIntoArray(arr, value, index) {

  return arr;
}
```

# --solutions--

```js
function insertIntoArray(arr, value, index) {
  return [
    ...arr.slice(0, index),
    value,
    ...arr.slice(index)
  ];
}
```
