---
id: 6994cff2290543b3aec9f50f
title: "도전 과제 210: HSL 유효성 검사기"
challengeType: 28
dashedName: challenge-210
---

# --description--

문자열이 유효한 CSS `hsl()` 색상 값인지 확인하세요.

- 유효한 HSL 값은 `"hsl(h, s%, l%)"` 형식이어야 하며:
  - `h` (색상) 값은 0에서 360 사이(포함)의 숫자여야 합니다.
  - `s` (채도) 값은 0%에서 100% 사이의 백분율이어야 합니다.
  - `l` (밝기) 값은 0%에서 100% 사이의 백분율이어야 합니다.

- 공백은 다음 위치에만 허용됩니다:
  - 여는 괄호 뒤
  - 쉼표 앞뒤
  - 닫는 괄호 앞뒤

- 선택적으로, 값은 세미콜론(`";"`)으로 끝날 수 있습니다.

예를 들어, `"hsl(240, 50%, 50%)"` 은 유효한 HSL 값입니다.

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` 은 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` 은 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` 은 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` 은 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` 은 `true` 을 반환해야 합니다.

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` 은 `false` 을 반환해야 합니다.

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` 은 `false` 을 반환해야 합니다.

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` 은 `false` 을 반환해야 합니다.

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` 은 `false` 을 반환해야 합니다.

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
