---
id: 699c8e045ee7cb94ed2322d9
title: "도전 과제 218: 정확히 나누어지는지 확인하기"
challengeType: 28
dashedName: challenge-218
---

# --description--

두 개의 정수가 주어졌을 때, 첫 번째 정수를 두 번째 정수로 정확히 나눌 수 있는지 확인하세요.

# --hints--

`isEvenlyDivisible(4, 2)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
