---
id: 699c8e045ee7cb94ed2322d5
title: "도전 과제 214: 도미노 체인 검증기"
challengeType: 28
dashedName: challenge-214
---

# --description--

도미노 수열을 나타내는 2차원 배열이 주어졌을 때, 유효한 체인을 이루는지 판단하세요.

- 배열의 각 요소는 도미노를 나타내며, 1부터 6까지(포함) 두 숫자로 이루어진 배열입니다.
- 체인이 유효하려면 각 도미노의 두 번째 숫자가 다음 도미노의 첫 번째 숫자와 일치해야 합니다.
- 첫 번째 도미노의 첫 숫자와 마지막 도미노의 마지막 숫자는 아무것과도 일치할 필요가 없습니다.

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
