---
id: 699c8e045ee7cb94ed2322dd
title: "도전 과제 222: 춘분 그림자"
challengeType: 28
dashedName: challenge-222
---

# --description--

오늘은 춘분으로, 태양이 적도 바로 위에 있고 정오에 머리 위에 있습니다. 주어진 시간에 4피트 높이의 수직 기둥이 만드는 그림자를 구하세요.

- 시간은 `"HH:MM"` 24시간 형식의 문자열입니다 (예: `"15:00"`은 오후 3시입니다).
- 시간은 30분 단위로만 주어집니다.

규칙:

- 태양은 오전 6시에 `"east"`에서 떠서 오후 6시에 `"west"`에서 집니다.
- 그림자는 항상 태양과 반대 방향을 가리킵니다.
- 그림자의 길이(피트 단위)는 정오에서 떨어진 시간(시간 단위)의 세제곱입니다.
- 일출 전(오전 6시 이전), 일몰 후(오후 6시 이후), 그리고 정오에는 그림자가 없습니다.

반환:

- 그림자가 있으면 `"(length)ft (direction)"`를 반환하세요. 예를 들어 `"8ft west"`입니다.
- 그렇지 않으면 `"No shadow"`를 반환하세요.

예를 들어 `"10:00"`이 주어지면, 오전 10시는 정오에서 2시간 떨어져 있으므로 2<sup>3</sup> = 8피트이고, 태양이 동쪽에 있으므로 그림자는 서쪽을 가리켜 `"8ft west"`를 반환합니다.

# --hints--

`getShadow("10:00")`는 `"8ft west"`를 반환해야 합니다.

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")`는 `"27ft east"`를 반환해야 합니다.

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")`는 `"No shadow"`를 반환해야 합니다.

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")`는 `"166.375ft east"`를 반환해야 합니다.

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")`는 `"No shadow"`를 반환해야 합니다.

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")`는 `"216ft west"`를 반환해야 합니다.

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")`는 `"No shadow"`를 반환해야 합니다.

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")`는 `"91.125ft west"`를 반환해야 합니다.

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")`는 `"No shadow"`를 반환해야 합니다.

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
