---
id: 699c8e045ee7cb94ed2322da
title: "도전 과제 219: 기념일 이정표"
challengeType: 28
dashedName: challenge-219
---

# --description--

부부가 결혼한 연수를 나타내는 정수가 주어지면, 다음 차트에 따라 가장 최근의 기념일 이정표를 반환하세요:

| 결혼 연수 | 이정표    |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- 첫 번째 이정표에 도달하지 못했다면 `"Newlyweds"`을 반환하세요.

# --hints--

`getMilestone(0)`은 `"Newlyweds"`을 반환해야 합니다.

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)`은 `"Paper"`을 반환해야 합니다.

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)`은 `"Wood"`을 반환해야 합니다.

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)`은 `"Tin"`을 반환해야 합니다.

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)`은 `"Silver"`을 반환해야 합니다.

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)`은 `"Ruby"`을 반환해야 합니다.

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)`은 `"Gold"`을 반환해야 합니다.

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)`은 `"Diamond"`을 반환해야 합니다.

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)`은 `"Platinum"`을 반환해야 합니다.

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
