---
id: 699c8e045ee7cb94ed2322d6
title: "도전 과제 215: 주차 요금 계산기"
challengeType: 28
dashedName: challenge-215
---

# --description--

차를 주차한 시간과 차를 찾은 시간을 나타내는 두 개의 문자열이 주어질 때, 주차 요금을 계산하세요.

- 주어진 문자열은 24시간 시계를 사용하여 `"HH:MM"` 형식으로 되어 있습니다. 예를 들어 `"14:00"`은 오후 2시입니다.
- 첫 번째 문자열은 차를 주차한 시간이고, 두 번째는 차를 찾은 시간입니다.
- 찾은 시간이 주차 시간보다 이르면, 주차가 자정을 넘어 다음 날까지 이어졌다는 뜻입니다.

요금 규칙:

- 주차한 시간당 3달러입니다.
- 부분 시간은 다음 한 시간으로 올림합니다.
- 주차가 자정을 넘어 밤새 이어지면 추가로 10달러의 야간 요금이 부과됩니다.
- 총 요금이 5달러 미만일 경우 최솟값으로 5달러가 적용됩니다.

총 요금을 `"$cost"`, `"$5"` 형식으로 반환하세요.

# --hints--

`calculateParkingFee("09:00", "11:00")`는 `"$6"`를 반환해야 합니다.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")`는 `"$5"`를 반환해야 합니다.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")`는 `"$9"`를 반환해야 합니다.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")`는 `"$27"`를 반환해야 합니다.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")`는 `"$34"`를 반환해야 합니다.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")`는 `"$82"`를 반환해야 합니다.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
