---
id: 699c8e045ee7cb94ed2322d4
title: "도전 과제 213: 단어 길이 변환기"
challengeType: 28
dashedName: challenge-213
---

# --description--

문자열로 된 단어들이 주어지면, 각 단어가 그 길이로 교체된 새 문자열을 반환하세요.

- 주어진 문자열의 단어들은 한 칸의 공백으로 구분됩니다
- 반환하는 문자열에서도 공백을 유지하세요.

예를 들어, `"hello world"`가 주어지면 `"5 5"`을 반환하세요.

# --hints--

`convertWords("hello world")`는 `"5 5"`을 반환해야 합니다.

```js
assert.equal(convertWords("hello world"), "5 5");
```

`convertWords("Thanks and happy coding")`는 `"6 3 5 6"`를 반환해야 합니다.

```js
assert.equal(convertWords("Thanks and happy coding"), "6 3 5 6");
```

`convertWords("The quick brown fox jumps over the lazy dog")`는 `"3 5 5 3 5 4 3 4 3"`을 반환해야 합니다.

```js
assert.equal(convertWords("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3");
```

`convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")`는 `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`를 반환해야 합니다.

```js
assert.equal(convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4");
```

# --seed--

## --seed-contents--

```js
function convertWords(str) {

  return str;
}
```

# --solutions--

```js
function convertWords(str) {
  return str
    .split(" ")
    .map(word => word.length)
    .join(" ");
}
```
