---
id: 699c8e045ee7cb94ed2322db
title: "도전 과제 220: 가장 큰 숫자"
challengeType: 28
dashedName: challenge-220
---

# --description--

여러 구두점으로 구분된 숫자 문자열이 주어질 때, 가장 큰 숫자를 반환하세요.

- 주어진 문자열에는 숫자와 구분 기호만 포함됩니다.
- 구분 기호는 쉼표(`","`), 느낌표(`"!"`), 물음표(`"?"`), 콜론(`":"`), 세미콜론(`";"`)일 수 있습니다.

# --hints--

`largestNumber("1,2")`는 `2`를 반환해야 합니다.

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")`는 `60`를 반환해야 합니다.

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")`는 `-402`를 반환해야 합니다.

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")`는 `99.9`를 반환해야 합니다.

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
