---
id: 69a890af247de743333bd4cc
title: "도전 과제 223: QR 디코더"
challengeType: 28
dashedName: challenge-223
---

# --description--

6x6 행렬(배열의 배열)로 주어진 QR 코드를 나타내는 이진 데이터 문자열을 반환하세요.

- QR 코드는 90도 단위로 회전된 어떤 방향으로도 주어질 수 있습니다.
- 올바르게 방향이 맞춰진 코드는 왼쪽 아래, 왼쪽 위, 오른쪽 위 모서리에 2x2 크기의 `1` (방향 표시자) 그룹이 있습니다.
- 세 개의 2x2 방향 표시자는 이진 데이터에 포함되지 않습니다.
- QR 코드가 올바르게 방향이 맞춰졌을 때 이진 데이터는 왼쪽에서 오른쪽, 위에서 아래 순서로(책을 읽듯이) 읽습니다.
- 코드는 항상 정확히 하나의 유효한 방향을 가집니다.

예를 들어, 다음이 주어졌을 때:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

또는 같은 코드가 다른 방향으로 주어졌을 때:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

세 개의 2x2 방향 표시자를 제외한 모든 이진 데이터인 `"000000000000000000000001"`를 반환하세요.

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])`는 `"000000000000000000000001"`를 반환해야 합니다.

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])`는 `"000000000000000000000001"`를 반환해야 합니다.

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])`는 `"001101000011000000110100"`를 반환해야 합니다.

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])`는 `"010001000100010101010110"`를 반환해야 합니다.

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])`는 `"010000100100100101001110"`를 반환해야 합니다.

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
