---
id: 69a890af247de743333bd4d2
title: "도전 과제 229: 텍스트 자르기 2"
challengeType: 28
dashedName: challenge-229
---

# --description--

문자열이 주어지면, 문자들의 총 너비가 50 단위를 넘지 않도록 잘라낸 새 문자열을 반환하세요.

각 문자는 특정 너비를 가집니다:

| 문자 | 너비 |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

위 표는 모든 대문자와 소문자를 포함합니다. 추가로:

- 공백(`" "`)은 너비가 2입니다
- 마침표(`"."`)는 너비가 1입니다

- 주어진 문자열이 50 단위 이하라면, 문자열을 그대로 반환하세요. 그렇지 않으면
- 문자열을 자르고 끝에 마침표 세 개(`"..."`)를 추가하세요. 마침표 세 개를 포함한 총 너비가 50 단위를 넘지 않으면서 최대한 가까워야 합니다.

# --hints--

`truncateText("The quick brown fox")`는 `"The quick brown f..."`를 반환해야 합니다.

```js
assert.equal(truncateText("The quick brown fox"), "The quick brown f...");
```

`truncateText("The silky smooth sloth")`는 `"The silky smooth s..."`를 반환해야 합니다.

```js
assert.equal(truncateText("The silky smooth sloth"), "The silky smooth s...");
```

`truncateText("THE LOUD BRIGHT BIRD")`는 `"THE LOUD BRIG..."`를 반환해야 합니다.

```js
assert.equal(truncateText("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...");
```

`truncateText("The fast striped zebra")`는 `"The fast striped z..."`를 반환해야 합니다.

```js
assert.equal(truncateText("The fast striped zebra"), "The fast striped z...");
```

`truncateText("The big black bear")`는 `"The big black bear"`를 반환해야 합니다.

```js
assert.equal(truncateText("The big black bear"), "The big black bear");
```

# --seed--

## --seed-contents--

```js
function truncateText(str) {

  return str;
}
```

# --solutions--

```js
function truncateText(str) {
  const MAX_WIDTH = 50;
  const ELLIPSIS = "...";
  const ELLIPSIS_WIDTH = 3;
  const TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH;

  const charGroups = {
    'ilI.': 1,
    'fjrt ': 2,
    'abcdeghkmnopqrstuvwxyzJL': 3,
    'ABCDEFGHKMNOPQRSTUVWXYZ': 4
  };

  function getCharWidth(char) {
    for (const key in charGroups) {
      if (key.includes(char)) return charGroups[key];
    }
    return 3;
  }

  function stringWidth(str) {
    let totalWidth = 0;
    for (const char of str) {
      totalWidth += getCharWidth(char);
    }
    return totalWidth;
  }

  if (stringWidth(str) <= MAX_WIDTH) return str;

  let result = "";
  let totalWidth = 0;

  for (const char of str) {
    const charWidth = getCharWidth(char);

    if (totalWidth + charWidth > TRUNCATE_LIMIT) break;

    result += char;
    totalWidth += charWidth;
  }

  return result + ELLIPSIS;
}
```
