---
id: 69a890af247de743333bd4cd
title: "도전 과제 224: 커피 로스트 감지기"
challengeType: 28
dashedName: challenge-224
---

# --description--

커피 한 잔을 만들 때 사용된 원두를 나타내는 문자열이 주어질 때, 그 커피의 로스트 단계를 판단하세요.

- 주어진 문자열에는 각기 다른 원두 유형을 나타내는 다음 문자가 포함됩니다:

  - 작은따옴표(`'`)는 1점짜리 라이트 로스트 원두입니다.
  - 대시(`-`)는 2점짜리 중간 로스트 원두입니다.
  - 마침표(`.`)는 3점짜리 다크 로스트 원두입니다.

- 로스트 단계는 모든 원두 점수의 평균으로 결정됩니다.

다음 값을 반환하세요:

- 평균이 1.75 미만이면 `"Light"`.
- 평균이 1.75 이상 2.5 이하이면 `"Medium"`.
- 평균이 2.5 초과이면 `"Dark"`.

# --hints--

`detectRoast("''-''''''-'-''--''''")`는 `"Light"`을 반환해야 합니다.

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")`는 `"Medium"`을 반환해야 합니다.

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")`는 `"Medium"`을 반환해야 합니다.

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")`는 `"Dark"`을 반환해야 합니다.

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")`는 `"Medium"`을 반환해야 합니다.

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")`는 `"Dark"`을 반환해야 합니다.

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")`는 `"Light"`을 반환해야 합니다.

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
