---
id: 69a890af247de743333bd4cf
title: "도전 과제 226: 시험 합격자 수 계산"
challengeType: 28
dashedName: challenge-226
---

# --description--

학생 시험 점수 배열과 합격 점수가 주어졌을 때, 시험에 합격한 학생 수를 반환하세요.

# --hints--

`passingCount([90, 85, 75, 60, 50], 70)`는 `3`을 반환해야 합니다.

```js
assert.equal(passingCount([90, 85, 75, 60, 50], 70), 3);
```

`passingCount([100, 80, 75, 88, 72, 74, 79, 71, 60, 92], 75)`는 `6`을 반환해야 합니다.

```js
assert.equal(passingCount([100, 80, 75, 88, 72, 74, 79, 71, 60, 92], 75), 6);
```

`passingCount([79, 60, 88, 72, 74, 59, 75, 71, 80, 92], 60)`는 `9`을 반환해야 합니다.

```js
assert.equal(passingCount([79, 60, 88, 72, 74, 59, 75, 71, 80, 92], 60), 9);
```

`passingCount([76, 79, 80, 70, 71, 65, 79, 78, 59, 72], 85)`는 `0`을 반환해야 합니다.

```js
assert.equal(passingCount([76, 79, 80, 70, 71, 65, 79, 78, 59, 72], 85), 0);
```

`passingCount([84, 65, 98, 53, 58, 71, 91, 80, 92, 70, 73, 83, 86, 69, 84, 77, 72, 58, 69, 75, 66, 68, 72, 96, 90, 63, 88, 63, 80, 67], 60)`는 `27`을 반환해야 합니다.

```js
assert.equal(passingCount([84, 65, 98, 53, 58, 71, 91, 80, 92, 70, 73, 83, 86, 69, 84, 77, 72, 58, 69, 75, 66, 68, 72, 96, 90, 63, 88, 63, 80, 67], 60), 27);
```

# --seed--

## --seed-contents--

```js
function passingCount(scores, passingScore) {

  return scores;
}
```

# --solutions--

```js
function passingCount(scores, passingScore) {
  let count = 0;

  for (const score of scores) {
    if (score >= passingScore) {
      count++;
    }
  }

  return count;
}
```
