---
id: 69a890af247de743333bd4ce
title: "도전 과제 225: 연속 반복 없음"
challengeType: 28
dashedName: challenge-225
---

# --description--

문자열이 주어졌을 때, 연속해서 반복되는 문자가 없는지 확인하세요.

- 문자열에 같은 문자가 두 번 이상 연속으로 나오지 않으면 반복이 없는 것입니다.

# --hints--

`hasNoRepeats("hi world")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```
