---
id: 69a890af247de743333bd4d0
title: "도전 과제 227: 쿨다운 시간"
challengeType: 28
dashedName: challenge-227
---

# --description--

두 개의 타임스탬프가 주어집니다. 첫 번째는 사용자가 시험을 마친 시점을 나타내고, 두 번째는 현재 시간을 나타냅니다. 사용자가 다시 시험을 볼 수 있는지 판단하세요.

- 두 타임스탬프는 모두 `"YYYY-MM-DDTHH:MM:SS"` 형식으로 주어집니다. 예를 들어 `"2026-03-25T14:00:00"`입니다. 시간은 24시간제로 표시됩니다.
- 사용자는 시험을 다시 보기 전에 최소 48시간을 기다려야 합니다.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
