---
id: 69a890af247de743333bd4d1
title: "도전 과제 228: 영화의 밤"
challengeType: 28
dashedName: challenge-228
---

# --description--

요일을 나타내는 문자열, 상영 시간을 나타내는 또 다른 문자열, 그리고 티켓 수를 나타내는 정수가 주어질 때, 해당 상영의 영화 티켓 총 비용을 반환하세요.

주어진 요일은 다음 중 하나입니다:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

상영 시간은 `"H:MMam"` 또는 `"H:MMpm"` 형식으로 주어집니다. 예를 들어 `"10:00am"` 또는 `"10:00pm"`입니다.

다음 규칙을 사용해 `"$D.CC"` 형식으로 총 비용을 반환하세요:

- 주말(금요일~일요일): 티켓당 $12.00.
- 평일(월요일~목요일): 티켓당 $10.00.
- 매티니(오후 5시 이전): 티켓당 $2.00 할인(화요일 제외).
- 화요일: 모든 티켓이 각각 $5.00.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)`는 `"$12.00"`을 반환해야 합니다.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)`는 `"$10.00"`을 반환해야 합니다.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)`는 `"$10.00"`을 반환해야 합니다.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)`는 `"$30.00"`을 반환해야 합니다.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)`는 `"$32.00"`을 반환해야 합니다.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)`는 `"$50.00"`을 반환해야 합니다.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)`는 `"$5.00"`을 반환해야 합니다.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
