---
id: 69b1028d6e265413d0198a2b
title: "도전 과제 232: 만기일"
challengeType: 28
dashedName: challenge-232
---

# --description--

날짜 문자열이 주어지면, 9개월 후의 날짜를 반환하세요.

- 주어진 문자열과 반환 문자열은 `"YYYY-MM-DD"` 형식입니다.
- 9개월 후의 월에 원래 날짜가 없으면, 그 달의 마지막 날짜를 반환하세요.

# --hints--

`getDueDate("2025-03-30")`은 `"2025-12-30"`를 반환해야 합니다.

```js
assert.equal(getDueDate("2025-03-30"), "2025-12-30");
```

`getDueDate("2025-04-27")`은 `"2026-01-27"`를 반환해야 합니다.

```js
assert.equal(getDueDate("2025-04-27"), "2026-01-27");
```

`getDueDate("2025-05-29")`는 `"2026-02-28"`를 반환해야 합니다.

```js
assert.equal(getDueDate("2025-05-29"), "2026-02-28");
```

`getDueDate("2026-06-30")`은 `"2027-03-30"`를 반환해야 합니다.

```js
assert.equal(getDueDate("2026-06-30"), "2027-03-30");
```

`getDueDate("2026-10-11")`는 `"2027-07-11"`를 반환해야 합니다.

```js
assert.equal(getDueDate("2026-10-11"), "2027-07-11");
```

# --seed--

## --seed-contents--

```js
function getDueDate(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getDueDate(dateStr) {
  let [year, month, day] = dateStr.split("-").map(Number);

  month += 9;

  year += Math.floor((month - 1) / 12);
  month = ((month - 1) % 12) + 1;

  const daysInMonth = new Date(year, month, 0).getDate();

  if (day > daysInMonth) {
    day = daysInMonth;
  }

  const mm = String(month).padStart(2, "0");
  const dd = String(day).padStart(2, "0");

  return `${year}-${mm}-${dd}`;
}
```
