---
id: 69b1028d6e265413d0198a2c
title: "도전 과제 233: 기상 알람"
challengeType: 28
dashedName: challenge-233
---

# --description--

알람을 맞춘 시간을 나타내는 문자열과 실제로 일어난 시간을 나타내는 문자열이 주어질 때, 일찍 일어났는지, 제시간에 일어났는지, 아니면 늦게 일어났는지 판단하세요.

- 두 시간 모두 `"HH:MM"` 24시간 형식으로 주어집니다.

반환값:

- 알람 시간보다 일찍 일어났으면 `"early"`을 반환하세요.
- 알람 시간에 일어났거나 알람 시간 이후 10분 이내의 스누즈 시간 내에 일어났으면 `"on time"`를 반환하세요.
- 알람 시간보다 10분 이상 늦게 일어났으면 `"late"`을 반환하세요.

두 시간은 같은 날에 해당합니다.

# --hints--

`alarmCheck("07:00", "06:45")`는 `"early"`를 반환해야 합니다.

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")`는 `"on time"`를 반환해야 합니다.

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")`는 `"on time"`를 반환해야 합니다.

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")`는 `"late"`를 반환해야 합니다.

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")`는 `"on time"`를 반환해야 합니다.

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")`는 `"late"`를 반환해야 합니다.

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")`는 `"early"`를 반환해야 합니다.

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
