---
id: 69b1028d6e265413d0198a29
title: "도전 과제 230: 파스칼 삼각형 행"
challengeType: 28
dashedName: challenge-230
---

# --description--

정수 `n`가 주어지면, 파스칼 삼각형의 `n`번째 행을 배열로 반환하세요.

파스칼 삼각형에서 각 행은 1로 시작하고 끝나며, 내부의 각 값은 바로 위 두 값의 합입니다.

다음은 삼각형의 처음 5개 행입니다:

```js
    1
   1 1
  1 2 1
 1 3 3 1
1 4 6 4 1
```

# --hints--

`pascalRow(5)`는 `[1, 4, 6, 4, 1]`을 반환해야 합니다.

```js
assert.deepEqual(pascalRow(5), [1, 4, 6, 4, 1]);
```

`pascalRow(3)`는 `[1, 2, 1]`를 반환해야 합니다.

```js
assert.deepEqual(pascalRow(3), [1, 2, 1]);
```

`pascalRow(1)`는 `[1]`을 반환해야 합니다.

```js
assert.deepEqual(pascalRow(1), [1]);
```

`pascalRow(10)`는 `[1, 9, 36, 84, 126, 126, 84, 36, 9, 1]`를 반환해야 합니다.

```js
assert.deepEqual(pascalRow(10), [1, 9, 36, 84, 126, 126, 84, 36, 9, 1]);
```

`pascalRow(15)`는 `[1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]`을 반환해야 합니다.

```js
assert.deepEqual(pascalRow(15), [1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]);
```

# --seed--

## --seed-contents--

```js
function pascalRow(n) {

  return n;
}
```

# --solutions--

```js
function pascalRow(n) {
  if (n === 1) return [1];

  let row = [1];

  for (let i = 2; i <= n; i++) {
    const next = [1];

    for (let j = 1; j < row.length; j++) {
      next.push(row[j - 1] + row[j]);
    }

    next.push(1);
    row = next;
  }

  return row;
}
```
