---
id: 69b1028d6e265413d0198a2a
title: "도전 과제 231: ISBN-10 유효성 검사기"
challengeType: 28
dashedName: challenge-231
---

# --description--

문자열이 주어지면, 그것이 유효한 ISBN-10인지 판단하세요.

ISBN-10은 하이픈(`"-"`)과 10개의 다른 문자로 구성됩니다. 하이픈을 제거한 후(`"-"`):

- 처음 9개 문자는 숫자여야 하며,
- 마지막 문자는 숫자이거나 숫자 10을 나타내는 문자 `"X"`일 수 있습니다.

유효성을 검사하려면:

- 각 숫자(또는 값)에 그 위치를 곱하세요(첫 번째 숫자는 1을 곱하고, 두 번째는 2를 곱하는 식으로).
- 모든 결과를 더하세요.
- 합계가 11로 나누어떨어지면 유효합니다.

# --hints--

`isValidIsbn10("0-306-40615-2")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
