---
id: 69b1028d6e265413d0198a2d
title: "도전 과제 234: 장난 숫자"
challengeType: 28
dashedName: challenge-234
---

# --description--

하나를 제외한 모든 숫자가 일정한 규칙을 따르는 숫자 배열이 주어질 때, 규칙을 따르지 않는 숫자 하나를 수정한 새 배열을 반환하세요.

규칙은 다음 중 하나입니다:

- 숫자가 일정한 값만큼 증가합니다(덧셈).
- 숫자가 일정한 값만큼 감소합니다(뺄셈).

예를 들어, `[2, 4, 7, 8, 10]`가 주어지면 `[2, 4, 6, 8, 10]`을 반환하세요.

# --hints--

`fixPrankNumber([2, 4, 7, 8, 10])`는 `[2, 4, 6, 8, 10]`을 반환해야 합니다.

```js
assert.deepEqual(fixPrankNumber([2, 4, 7, 8, 10]), [2, 4, 6, 8, 10]);
```

`fixPrankNumber([10, 10, 8, 7, 6])`는 `[10, 9, 8, 7, 6]`을 반환해야 합니다.

```js
assert.deepEqual(fixPrankNumber([10, 10, 8, 7, 6]), [10, 9, 8, 7, 6]);
```

`fixPrankNumber([12, 24, 36, 48, 61, 72, 84, 96])`는 `[12, 24, 36, 48, 60, 72, 84, 96]`을 반환해야 합니다.

```js
assert.deepEqual(fixPrankNumber([12, 24, 36, 48, 61, 72, 84, 96]), [12, 24, 36, 48, 60, 72, 84, 96]);
```

`fixPrankNumber([4, 1, -2, -5, -8, -5])`는 `[4, 1, -2, -5, -8, -11]`을 반환해야 합니다.

```js
assert.deepEqual(fixPrankNumber([4, 1, -2, -5, -8, -5]), [4, 1, -2, -5, -8, -11]);
```

`fixPrankNumber([0, 100, 200, 300, 150, 500])`는 `[0, 100, 200, 300, 400, 500]`을 반환해야 합니다.

```js
assert.deepEqual(fixPrankNumber([0, 100, 200, 300, 150, 500]), [0, 100, 200, 300, 400, 500]);
```

`fixPrankNumber([400, 425, 400, 375, 350, 325, 300])`는 `[450, 425, 400, 375, 350, 325, 300]`을 반환해야 합니다.

```js
assert.deepEqual(fixPrankNumber([400, 425, 400, 375, 350, 325, 300]), [450, 425, 400, 375, 350, 325, 300]);
```

`fixPrankNumber([-5, 5, 10, 15, 20])`는 `[0, 5, 10, 15, 20]`을 반환해야 합니다.

```js
assert.deepEqual(fixPrankNumber([-5, 5, 10, 15, 20]), [0, 5, 10, 15, 20]);
```

# --seed--

## --seed-contents--

```js
function fixPrankNumber(arr) {

  return arr;
}
```

# --solutions--

```js
function fixPrankNumber(arr) {
  const diffs = arr.slice(1).map((v, i) => v - arr[i]);

  let step = diffs[0];
  for (let i = 0; i < diffs.length - 1; i++) {
    if (diffs[i] === diffs[i + 1]) {
      step = diffs[i];
      break;
    }
  }

  const result = [...arr];

  for (let i = 0; i < result.length - 1; i++) {
    if (result[i + 1] - result[i] !== step) {
      if (i === 0 && result[2] - result[1] === step) {
        result[0] = result[1] - step;
      } else {
        result[i + 1] = result[i] + step;
      }
      break;
    }
  }

  return result;
}
```
