---
id: 69b58ce40693f140c84c8559
title: "도전 과제 240: 회문 문자열의 가운데 문자"
challengeType: 28
dashedName: challenge-240
---

# --description--

문자열이 주어졌을 때, 그것이 회문인지 판단하고 가운데 문자를 반환하세요(길이가 홀수면 가운데 한 문자, 짝수면 가운데 두 문자).

- 회문은 앞에서 읽으나 뒤에서 읽으나 같은 문자열입니다.
- 회문이 아니면 `"none"`를 반환하세요.

# --hints--

`palindromeLocator("racecar")`는 `"e"`를 반환해야 합니다.

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")`는 `"v"`를 반환해야 합니다.

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")`는 `"none"`를 반환해야 합니다.

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")`는 `"oo"`를 반환해야 합니다.

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")`는 `"00"`를 반환해야 합니다.

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
