---
id: 69b559d2903b9e4afe9075f7
title: "도전 과제 237: 방정식 유효성 검사"
challengeType: 28
dashedName: challenge-237
---

# --description--

수학 방정식을 나타내는 문자열이 주어졌을 때, 그것이 올바른지 판단하세요.

- 왼쪽에는 최대 세 개의 양의 정수와 연산자 `+`, `-`, `*`, `/`가 포함될 수 있습니다.
- 방정식은 `"number operator number = number"` 형식으로 주어집니다 (왼쪽에 두 개 또는 세 개의 숫자가 있음). 예를 들어 `"2 + 2 = 4"` 또는 `"2 + 3 - 1 = 4"`가 있습니다.
- 오른쪽은 항상 단일 정수입니다.

표준 연산 순서를 따르세요: 곱셈과 나눗셈은 덧셈과 뺄셈보다 먼저 왼쪽에서 오른쪽으로 평가됩니다.

# --hints--

`isValidEquation("2 + 2 = 4")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidEquation("2 + 2 = 4"));
```

`isValidEquation("2 + 3 - 1 = 4")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidEquation("2 + 3 - 1 = 4"));
```

`isValidEquation("8 / 2 = 4")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidEquation("8 / 2 = 4"));
```

`isValidEquation("10 * 5 = 50")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidEquation("10 * 5 = 50"));
```

`isValidEquation("2 - 2 = 0")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidEquation("2 - 2 = 0"));
```

`isValidEquation("2 + 9 / 3 = 5")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidEquation("2 + 9 / 3 = 5"));
```

`isValidEquation("20 - 2 * 3 = 14")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidEquation("20 - 2 * 3 = 14"));
```

`isValidEquation("2 + 5 = 6")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidEquation("2 + 5 = 6"));
```

`isValidEquation("10 - 2 * 3 = 24")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidEquation("10 - 2 * 3 = 24"));
```

`isValidEquation("3 + 9 / 3 = 4")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidEquation("3 + 9 / 3 = 4"));
```

# --seed--

## --seed-contents--

```js
function isValidEquation(equation) {

  return equation;
}
```

# --solutions--

```js
function isValidEquation(equation) {
  const [left, right] = equation.split(" = ");
  const tokens = left.split(" ");

  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    if (operator === "*" || operator === "/") {
      const leftOperand = Number(tokens[i - 1]);
      const rightOperand = Number(tokens[i + 1]);
      const computed = operator === "*" ? leftOperand * rightOperand : leftOperand / rightOperand;
      tokens.splice(i - 1, 3, String(computed));
      i -= 2;
    }
  }

  let result = Number(tokens[0]);
  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    const operand = Number(tokens[i + 1]);
    result = operator === "+" ? result + operand : result - operand;
  }

  return result === Number(right);
}
```
