---
id: 69b58ce40693f140c84c855c
title: "도전 과제 243: 룩 공격"
challengeType: 28
dashedName: challenge-243
---

# --description--

체스판 위 두 룩의 위치를 나타내는 두 문자열이 주어졌을 때, 두 룩이 서로 공격할 수 있는지 판단하세요.

표준 체스판은 8x8 크기로, 열은 왼쪽에서 오른쪽으로 `A`부터 `H`까지 레이블이 붙어 있고, 행은 아래에서 위로 `1`부터 `8`까지 레이블이 붙어 있습니다. 체스판은 다음과 같습니다:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

룩은 가로 또는 세로 방향으로 원하는 만큼 여러 칸 이동할 수 있습니다. 따라서 두 룩이 같은 행이나 열에 있으면 서로 공격할 수 있습니다.

# --hints--

`rookAttack("A1", "A8")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(rookAttack("A1", "A8"));
```

`rookAttack("B4", "F4")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(rookAttack("B4", "F4"));
```

`rookAttack("E3", "D4")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(rookAttack("E3", "D4"));
```

`rookAttack("H7", "F6")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(rookAttack("H7", "F6"));
```

# --seed--

## --seed-contents--

```js
function rookAttack(rook1, rook2) {

  return rook1;
}
```

# --solutions--

```js
function rookAttack(rook1, rook2) {
  return rook1[0] === rook2[0] || rook1[1] === rook2[1];
}
```
