---
id: 69b559d2903b9e4afe9075f8
title: "도전 과제 238: 숫자 회전 탈출"
challengeType: 28
dashedName: challenge-238
---

# --description--

양의 정수가 주어졌을 때, 그 수 또는 그 수의 어떤 회전이 숫자 자리수로 나누어떨어지는지 확인하세요.

회전은 첫 번째 숫자를 끝으로 옮기는 것을 의미합니다. 예를 들어, 1회전 후 123은 231이 됩니다.

- 먼저 `0` (주어진 숫자)의 회전을 확인하세요.
- 주어진 숫자에는 0이 포함되지 않습니다.
- 나누어떨어지는 회전 숫자가 있으면 첫 번째 회전 숫자를 반환하고, 없으면 `"none"`을 반환하세요.

# --hints--

`getRotation(123)`는 `0`을 반환해야 합니다.

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)`는 `3`를 반환해야 합니다.

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)`는 `"none"`을 반환해야 합니다.

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)`는 `6`를 반환해야 합니다.

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
