---
id: 69b559d2903b9e4afe9075f9
title: "도전 과제 239: 오늘은 무슨 요일일까요?"
challengeType: 28
dashedName: challenge-239
---

# --description--

밀리초 단위의 Unix 타임스탬프가 주어지면 요일을 반환하세요.

유효한 반환 요일은 다음과 같습니다:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

시간대는 반드시 무시하세요.

# --hints--

`getDayOfWeek(1775492249000)`는 `"Monday"`을 반환해야 합니다.

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)`는 `"Saturday"`을 반환해야 합니다.

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)`는 `"Tuesday"`를 반환해야 합니다.

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)`는 `"Sunday"`를 반환해야 합니다.

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)`는 `"Thursday"`를 반환해야 합니다.

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
