---
id: 69b58ce40693f140c84c855b
title: "도전 과제 242: 다음 빙고 번호"
challengeType: 28
dashedName: challenge-242
---

# --description--

주어진 빙고 번호에 대해, 순서상 다음 빙고 번호를 반환하세요.

빙고 번호는 이 차트에 따라 단일 문자 다음에 해당 범위의 숫자가 옵니다:

| 문자 | 숫자 범위 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

예를 들어, `"B10"`가 주어지면 다음 빙고 번호인 `"B11"`를 반환하세요. 마지막 빙고 번호가 주어지면 `"B1"`를 반환하세요.

# --hints--

`getNextBingoNumber("B10")`는 `"B11"`를 반환해야 합니다.

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")`는 `"N34"`를 반환해야 합니다.

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")`는 `"N31"`를 반환해야 합니다.

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")`는 `"O61"`를 반환해야 합니다.

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")`는 `"B1"`를 반환해야 합니다.

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
