---
id: 69bc6cb30c1d112a2e110a08
title: "도전 과제 251: 배열 합계 찾기"
challengeType: 28
dashedName: challenge-251
---

# --description--

숫자 배열과 목표 숫자가 주어졌을 때, 목표 숫자의 합이 되는 두 개 이상의 숫자로 이루어진 첫 번째 부분집합을 반환하세요.

- "첫 번째" 부분집합은 요소들의 인덱스가 가장 낮은 부분집합이며, 가장 이른 인덱스를 우선시합니다.
- 배열의 각 숫자는 한 번만 사용할 수 있습니다.
- 유효한 부분집합이 없으면 `"Sum not found"`를 반환하세요.

일치하는 숫자들은 원래 배열에 나타난 순서대로 배열로 반환하세요.

# --hints--

`findSum([1, 3, 5, 7], 6)`는 `[1, 5]`를 반환해야 합니다.

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)`는 `[1, 4]`를 반환해야 합니다.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)`는 `[1, 2, 3]`를 반환해야 합니다.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)`는 `[-1, -2, 4]`를 반환해야 합니다.

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)`는 `[3, 1, 4, 2]`를 반환해야 합니다.

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)`는 `[1, 2, 3, 5, 9]`를 반환해야 합니다.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)`는 `"Sum not found"`를 반환해야 합니다.

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```
