---
id: 69bc6cb30c1d112a2e110a06
title: "도전 과제 249: 문자열 수학"
challengeType: 28
dashedName: challenge-249
---

# --description--

숫자와 다른 문자들이 섞인 문자열이 주어졌을 때, 숫자 사이에 있는 숫자가 아닌 문자들의 개수에 따라 숫자들에 수학 연산을 수행하세요.

- 두 숫자를 구분하는 문자의 개수가 짝수이면 덧셈을 사용하세요.
- 홀수이면 뺄셈을 사용하세요.
- 연속된 숫자들은 하나의 숫자로 간주하세요.
- 연산은 왼쪽에서 오른쪽 순서로 적용하세요.
- 숫자가 아닌 선행 및 후행 문자는 무시하세요.

예를 들어 `"3ab10c8"`가 주어지면 `5`을 반환하세요. 두 숫자 사이에 짝수 개의 문자가 있으므로 3과 10을 더해 13이 됩니다. 그다음 결과인 13과 8 사이에 홀수 개의 문자가 있으므로 8을 빼서 최종 결과를 구합니다.

# --hints--

`doMath("3ab10c8")`는 `5`을 반환해야 합니다.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")`는 `2`을 반환해야 합니다.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")`는 `12`을 반환해야 합니다.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")`는 `15`을 반환해야 합니다.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")`는 `67`을 반환해야 합니다.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
