---
id: 69bc6cb30c1d112a2e110a04
title: "도전 과제 247: 마지막 글자"
challengeType: 28
dashedName: challenge-247
---

# --description--

문자열이 주어지면, 알파벳 순서상 가장 마지막에 나오는 글자를 반환하세요.

- 두 글자 이상이 알파벳 순서상 마지막에 같다면, 그중 첫 번째 글자를 반환하세요.
- 알파벳 글자가 아닌 모든 문자는 무시하세요.

# --hints--

`getLastLetter("world")`는 `"w"`를 반환해야 합니다.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")`는 `"W"`를 반환해야 합니다.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")`는 `"z"`를 반환해야 합니다.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")`는 `"L"`를 반환해야 합니다.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")`는 `"T"`를 반환해야 합니다.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
