---
id: 69bc6cb30c1d112a2e110a09
title: "도전 과제 252: 고유한 계단 오르기"
challengeType: 28
dashedName: challenge-252
---

# --description--

계단 수가 주어지면 한 번에 1칸 또는 2칸씩 올라갈 때 고유하게 올라갈 수 있는 방법의 수를 반환하세요.

# --hints--

`getUniqueClimbs(4)`는 `5`를 반환해야 합니다.

```js
assert.equal(getUniqueClimbs(4), 5);
```

`getUniqueClimbs(5)`는 `8`를 반환해야 합니다.

```js
assert.equal(getUniqueClimbs(5), 8);
```

`getUniqueClimbs(10)`는 `89`를 반환해야 합니다.

```js
assert.equal(getUniqueClimbs(10), 89);
```

`getUniqueClimbs(18)`는 `4181`를 반환해야 합니다.

```js
assert.equal(getUniqueClimbs(18), 4181);
```

`getUniqueClimbs(29)`는 `832040`를 반환해야 합니다.

```js
assert.equal(getUniqueClimbs(29), 832040);
```

`getUniqueClimbs(50)`는 `20365011074`를 반환해야 합니다.

```js
assert.equal(getUniqueClimbs(50), 20365011074);
```

# --seed--

## --seed-contents--

```js
function getUniqueClimbs(steps) {

  return steps;
}
```

# --solutions--

```js
function getUniqueClimbs(steps) {
  if (steps <= 0) return 0;
  if (steps === 1) return 1;
  if (steps === 2) return 2;

  let prev2 = 1, prev1 = 2;
  for (let i = 3; i <= steps; i++) {
    [prev2, prev1] = [prev1, prev2 + prev1];
  }
  return prev1;
}
```
