---
id: 69bc6cb30c1d112a2e110a05
title: "도전 과제 248: 정렬된 배열 교환"
challengeType: 28
dashedName: challenge-248
---

# --description--

정수 배열이 주어지면 다음 규칙을 사용해 새 배열을 반환하세요:

1. 정수를 오름차순으로 정렬합니다
2. 그런 다음 인덱스가 3의 배수인 모든 값을 그 앞의 값과 교환합니다.

# --hints--

`sortAndSwap([3, 1, 2, 4, 6, 5])`는 `[1, 2, 4, 3, 5, 6]`를 반환해야 합니다.

```js
assert.deepEqual(sortAndSwap([3, 1, 2, 4, 6, 5]), [1, 2, 4, 3, 5, 6]);
```

`sortAndSwap([9, 7, 5, 3, 1, 2, 4, 6, 8])`는 `[1, 2, 4, 3, 5, 7, 6, 8, 9]`를 반환해야 합니다.

```js
assert.deepEqual(sortAndSwap([9, 7, 5, 3, 1, 2, 4, 6, 8]), [1, 2, 4, 3, 5, 7, 6, 8, 9]);
```

`sortAndSwap([1, 2, 3, 4, 5, 6, 7, 8, 9])`는 `[1, 2, 4, 3, 5, 7, 6, 8, 9]`를 반환해야 합니다.

```js
assert.deepEqual(sortAndSwap([1, 2, 3, 4, 5, 6, 7, 8, 9]), [1, 2, 4, 3, 5, 7, 6, 8, 9]);
```

`sortAndSwap([12, 5, 8, 1, 3, 10, 2, 7, 6, 4, 9, 11])`는 `[1, 2, 4, 3, 5, 7, 6, 8, 10, 9, 11, 12]`를 반환해야 합니다.

```js
assert.deepEqual(sortAndSwap([12, 5, 8, 1, 3, 10, 2, 7, 6, 4, 9, 11]), [1, 2, 4, 3, 5, 7, 6, 8, 10, 9, 11, 12]);
```

`sortAndSwap([100, -50, 0, 75, -25, 50, -75, 25])`는 `[-75, -50, 0, -25, 25, 75, 50, 100]`를 반환해야 합니다.

```js
assert.deepEqual(sortAndSwap([100, -50, 0, 75, -25, 50, -75, 25]), [-75, -50, 0, -25, 25, 75, 50, 100]);
```

`sortAndSwap([5, 9, 13, 77, 88, 313, -10, -65, 0, 8, 99, 101, -4, 2])`는 `[-65, -10, 0, -4, 2, 8, 5, 9, 77, 13, 88, 101, 99, 313]`를 반환해야 합니다.

```js
assert.deepEqual(sortAndSwap([5, 9, 13, 77, 88, 313, -10, -65, 0, 8, 99, 101, -4, 2]), [-65, -10, 0, -4, 2, 8, 5, 9, 77, 13, 88, 101, 99, 313]);
```

# --seed--

## --seed-contents--

```js
function sortAndSwap(arr) {

  return arr;
}
```

# --solutions--

```js
function sortAndSwap(arr) {
  const result = [...arr].sort((a, b) => a - b);

  for (let i = 3; i < result.length; i += 3) {
    [result[i], result[i - 1]] = [result[i - 1], result[i]];
  }

  return result;
}
```
