---
id: 69bc6cb30c1d112a2e110a03
title: "도전 과제 246: 이름 이니셜"
challengeType: 28
dashedName: challenge-246
---

# --description--

전체 이름이 문자열로 주어질 때, 이니셜을 반환하세요.

- 이니셜을 만들 이름들은 공백으로 구분됩니다.
- 이니셜은 대문자로 만들어야 합니다.
- 이니셜은 점으로 구분해야 합니다.

예를 들어, `"Tommy Millwood"`는 `"T.M."`를 반환합니다.

# --hints--

`getInitials("Tommy Millwood")`는 `"T.M."`를 반환해야 합니다.

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")`는 `"S.P."`를 반환해야 합니다.

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")`는 `"F.C.C."`를 반환해야 합니다.

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")`는 `"D."`를 반환해야 합니다.

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")`는 `"D.V.C.H.N."`를 반환해야 합니다.

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
