---
id: 69c5f3d787b1725d5f00c8b7
title: "도전 과제 253: 약어 찾기"
challengeType: 28
dashedName: challenge-253
---

# --description--

약어를 나타내는 문자열이 주어지면, 아래 목록에서 해당 약어가 속한 조직의 전체 이름을 반환하세요.

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

주어진 약어의 각 글자는 조직 이름의 각 단어 첫 글자와 같은 순서로 일치해야 합니다.

# --hints--

`findOrg("NASA")`는 `"National Avocado Storage Authority"`를 반환해야 합니다.

```js
assert.equal(findOrg("NASA"), "National Avocado Storage Authority");
```

`findOrg("CIA")`는 `"Cats Infiltration Agency"`를 반환해야 합니다.

```js
assert.equal(findOrg("CIA"), "Cats Infiltration Agency");
```

`findOrg("FBI")`는 `"Fluffy Beanbag Inspectors"`를 반환해야 합니다.

```js
assert.equal(findOrg("FBI"), "Fluffy Beanbag Inspectors");
```

`findOrg("DOJ")`는 `"Department Of Jelly"`를 반환해야 합니다.

```js
assert.equal(findOrg("DOJ"), "Department Of Jelly");
```

`findOrg("WHO")`는 `"Wild Honey Organization"`를 반환해야 합니다.

```js
assert.equal(findOrg("WHO"), "Wild Honey Organization");
```

`findOrg("EPA")`는 `"Eating Pancakes Administration"`를 반환해야 합니다.

```js
assert.equal(findOrg("EPA"), "Eating Pancakes Administration");
```

# --seed--

## --seed-contents--

```js
function findOrg(acronym) {

  return acronym;
}
```

# --solutions--

```js
function findOrg(acronym) {
  const orgs = [
    "National Avocado Storage Authority",
    "Cats Infiltration Agency",
    "Fluffy Beanbag Inspectors",
    "Department Of Jelly",
    "Wild Honey Organization",
    "Eating Pancakes Administration"
  ];

  for (let org of orgs) {
    const initials = org
      .split(" ")
      .map(word => word[0])
      .join("");

    if (initials.toUpperCase() === acronym.toUpperCase()) {
      return org;
    }
  }

  return "not found";
}
```
