---
id: 69c5f3d787b1725d5f00c8b9
title: "도전 과제 255: 지구의 날 청소팀"
challengeType: 28
dashedName: challenge-255
---

# --description--

오늘은 지구의 날입니다. 청소한 항목들의 배열이 주어질 때, 아래 규칙에 따라 총 청소 점수를 반환하세요.

주어진 항목은 다음 중 하나입니다:

| 항목 | 기본 점수 |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

희귀 아이템은 `["rare", value]`로 나타냅니다. 예를 들어, `["rare", 80]`입니다. 희귀 아이템은 연속 보너스를 받지 않습니다.

- 연속 보너스: 같은 항목이 연속으로 나타나면 점수가 점점 증가합니다.
  - 첫 번째 연속 출현: 기본 점수
  - 두 번째: 기본 점수 + 1
  - 세 번째: 기본 점수 + 2
  - 계속해서 증가
- 다섯 번째 아이템 배수: 다섯 번째로 수집한 아이템마다 배수가 적용됩니다.
  - 다섯 번째 아이템: *2
  - 열 번째 아이템: *3
  - 계속해서 증가

- 보너스 계산 후에 배수를 적용하세요.

# --hints--

`getCleanupScore(["bottle", "straw", "shoe", "battery"])`는 `44`를 반환해야 합니다.

```js
assert.equal(getCleanupScore(["bottle", "straw", "shoe", "battery"]), 44);
```

`getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"])`는 `58`를 반환해야 합니다.

```js
assert.equal(getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"]), 58);
```

`getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])`는 `79`를 반환해야 합니다.

```js
assert.equal(getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79);
```

`getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])`는 `358`를 반환해야 합니다.

```js
assert.equal(getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358);
```

`getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])`는 `383`를 반환해야 합니다.

```js
assert.equal(getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383);
```

# --seed--

## --seed-contents--

```js
function getCleanupScore(items) {

  return items;
}
```

# --solutions--

```js
function getCleanupScore(items) {
  const values = {
    bottle: 10,
    can: 6,
    bag: 8,
    tire: 35,
    straw: 4,
    cardboard: 3,
    newspaper: 3,
    shoe: 12,
    electronics: 25,
    battery: 18,
    mattress: 38
  };

  let total = 0;
  let prevItem = null;
  let streak = 0;

  for (let i = 0; i < items.length; i++) {
    let item = items[i];
    let base;
    let isRare = Array.isArray(item) && item[0] === "rare";

    if (isRare) {
      base = item[1];
      prevItem = null;
      streak = 0;
    } else {
      base = values[item];
      streak = item === prevItem ? streak + 1 : 0;
      prevItem = item;
    }

    let points = base;
    if (!isRare) points += streak;

    if ((i + 1) % 5 === 0) {
      const multiplier = (i + 1) / 5 + 1;
      points *= multiplier;
    }

    total += points;
  }

  return total;
}
```
