---
id: 69c5f3d787b1725d5f00c8b8
title: "도전 과제 254: 홀수 글자 단어"
challengeType: 28
dashedName: challenge-254
---

# --description--

문자열로 된 단어들이 주어지면, 글자 수가 홀수인 단어만 반환하세요.

- 주어진 문자열의 단어들은 한 칸의 공백으로 구분됩니다.
- 단어들은 한 칸의 공백으로 구분하여 반환하세요.

# --hints--

`getOddWords("This is a super good test")`는 `"a super"`을 반환해야 합니다.

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")`는 `"one two three"`을 반환해야 합니다.

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")`는 `"split rainbow sprinkles top"`를 반환해야 합니다.

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")`는 `"The quick brown fox the river"`을 반환해야 합니다.

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
