---
id: 69c5f3d787b1725d5f00c8bd
title: "도전 과제 259: FizzBuzz 폭발"
challengeType: 28
dashedName: challenge-259
---

# --description--

정수가 주어지면, 다음 규칙을 사용해 `"fizzbuzz"` 단어를 주어진 수 이상의 `"z"`를 포함하는 문자열로 바꾸는 데 걸리는 단계 수를 반환하세요.

- 문자열 `"fizzbuzz"`로 시작하세요.
- 각 단계마다 문자열 내 글자 위치(첫 번째 `"f"`는 위치 1)를 기준으로 표준 FizzBuzz 규칙을 적용하세요.
  - 글자 위치가 3으로 나누어떨어지면 글자를 `"fizz"`로 바꾸세요.
  - 5로 나누어떨어지면 글자를 `"buzz"`로 바꾸세요.
  - 3과 5 모두로 나누어떨어지면 글자를 `"fizzbuzz"`로 바꾸세요.

따라서 1단계 후 `"fizzbuzz"`는 `"fifizzzbuzzfizzzz"`로 바뀌며, 이 문자열에는 9개의 `"z"`가 있습니다.

# --hints--

`explodeFizzbuzz(9)`는 `1`를 반환해야 합니다.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)`는 `2`를 반환해야 합니다.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)`는 `3`를 반환해야 합니다.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)`는 `4`를 반환해야 합니다.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)`는 `5`를 반환해야 합니다.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)`는 `6`를 반환해야 합니다.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)`는 `11`를 반환해야 합니다.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)`는 `14`를 반환해야 합니다.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
