---
id: 69c5f3d787b1725d5f00c8ba
title: "도전 과제 256: 가장 가까운 시간 방향"
challengeType: 28
dashedName: challenge-256
---

# --description--

두 시간이 주어졌을 때, 첫 번째 시간에서 두 번째 시간으로 앞으로 가는 것이 빠른지, 뒤로 가는 것이 빠른지 판단하세요.

- 시간은 24시간 형식으로 주어집니다 (`"HH:MM"`)
- 시계는 순환합니다 (앞으로 이동할 때 23:59 다음은 00:00이고, 뒤로 이동할 때 00:00 다음은 23:59입니다)

반환값:

- 앞으로 가는 것이 더 짧으면 `"forward"`
- 뒤로 가는 것이 더 짧으면 `"backward"`
- 두 방향 모두 같은 시간이 걸리면 `"equal"` 반환하세요

# --hints--

`getDirection("10:00", "12:00")`는 `"forward"`를 반환해야 합니다.

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")`는 `"backward"`를 반환해야 합니다.

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")`는 `"equal"`를 반환해야 합니다.

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")`는 `"forward"`를 반환해야 합니다.

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")`는 `"backward"`를 반환해야 합니다.

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")`는 `"equal"`를 반환해야 합니다.

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
