---
id: 69c6ff713a52713463aa7929
title: "도전 과제 263: 이진 십자말풀이"
challengeType: 28
dashedName: challenge-263
---

# --description--

문자가 주어졌을 때, 해당 문자의 8비트 이진 표현이 다음 격자에서 가로 또는 세로 방향으로 어느 쪽이든 존재하는지 확인하세요.

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

예를 들어, `"A"`의 이진 표현은 `01000001`이며, 이는 첫 번째 행의 왼쪽에서 오른쪽으로 나타납니다.

# --hints--

`isInCrossword("I")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
