---
id: 69c6ff713a52713463aa7926
title: "도전 과제 260: 단어 점수"
challengeType: 28
dashedName: challenge-260
---

# --description--

단어가 주어지면, 표준 문자 값 표를 사용해 점수를 반환하세요:

| Letter | Value |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- 대문자와 소문자는 같은 값을 가집니다.

# --hints--

`getWordScore("hi")`는 `17`를 반환해야 합니다.

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")`는 `52`를 반환해야 합니다.

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")`는 `169`를 반환해야 합니다.

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")`는 `94`를 반환해야 합니다.

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
