---
id: 69c6ff713a52713463aa7927
title: "도전 과제 261: 숫자 단어"
challengeType: 28
dashedName: challenge-261
---

# --description--

0부터 99까지의 정수가 주어지면, 그 숫자의 영어 단어 표현을 반환하세요.

- `0`는 `"zero"`을 반환합니다.
- 1부터 19까지 숫자는 고유한 이름이 있습니다 (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- 20부터 90까지 10의 배수는 각각 고유한 이름이 있습니다 (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- 21부터 99까지 10의 배수가 아닌 숫자는 하이픈으로 연결된 두 단어로 작성합니다. 예를 들어 `"forty-two"`와 `"fifty-three"`입니다.

# --hints--

`getNumberWords(0)`는 `"zero"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)`는 `"ten"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)`는 `"nineteen"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)`는 `"thirty"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)`는 `"fifty-three"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)`는 `"seven"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)`는 `"twelve"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)`는 `"sixty"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)`는 `"sixty-seven"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)`는 `"ninety-eight"`를 반환해야 합니다.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
