---
id: 69c6ff713a52713463aa7928
title: "도전 과제 262: URL 쿼리 파서"
challengeType: 28
dashedName: challenge-262
---

# --description--

쿼리 문자열이 포함된 URL이 주어지면, 쿼리 문자열을 키-값 쌍의 객체(또는 딕셔너리)로 파싱하세요.

- 쿼리 문자열은 `"?"` 이후에 시작합니다,
- 각 매개변수는 `"&"`로 구분되고,
- 각 키/값 쌍은 `"="`로 구분됩니다.

예를 들어, `"https://example.com/search?name=Alice&age=30"`가 주어지면 다음을 반환하세요:

```json
{
  "name": "Alice",
  "age": "30"
}
```

모든 값은 문자열로 반환해야 합니다.

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")`는 `{"name": "Alice", "age": "30"}`를 반환해야 합니다.

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")`는 `{"skill": "programming", "language": "python"}`를 반환해야 합니다.

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")`는 `{"category": "books", "sort": "asc", "page": "2"}`를 반환해야 합니다.

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")`는 `{"redirect": "freecodecamp.org/learn", "when": "now"}`를 반환해야 합니다.

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
