---
id: 69cfca90e8a0a6d4d6871c52
title: "도전 과제 268: 나르시시스트 수"
challengeType: 28
dashedName: challenge-268
---

# --description--

양의 정수가 주어졌을 때, 그것이 나르시시스트 수인지 판단하세요.

- 숫자가 나르시시스트 수라는 것은 각 자리 숫자를 전체 자리수만큼 거듭제곱한 값을 모두 더했을 때 원래 숫자와 같을 때를 말합니다.

예를 들어, 153은 3자리 숫자이고 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153이므로 나르시시스트 수입니다.

# --hints--

`isNarcissistic(153)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
