---
id: 69cfca90e8a0a6d4d6871c4f
title: "도전 과제 265: 가장 깊은 괄호"
challengeType: 28
dashedName: challenge-265
---

# --description--

균형 잡힌 괄호가 포함된 문자열이 주어질 때, 가장 깊게 중첩된 괄호 안의 내용을 반환하세요.

- 괄호는 `()`, `[]`, `{}` 세 가지 유형 중 하나일 수 있습니다.
- 입력에는 항상 하나의 가장 깊은 그룹이 존재합니다.

예를 들어, `"(hello (world))"`가 주어지면 `"world"`를 반환하세요.

# --hints--

`getDeepestBrackets("(hello (world))")`는 `"world"`를 반환해야 합니다.

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")`는 `"inner"`를 반환해야 합니다.

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")`는 `"e"`를 반환해야 합니다.

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")`는 `"fox"`를 반환해야 합니다.

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")`는 `"C"`를 반환해야 합니다.

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
