---
id: 69cfca90e8a0a6d4d6871c54
title: "도전 과제 270: 가장 긴 공통 부분 문자열"
challengeType: 28
dashedName: challenge-270
---

# --description--

문자열이 주어지면 두 번 이상 나타나는 가장 긴 부분 문자열을 반환하세요.

- 부분 문자열은 겹칠 수 있습니다.

# --hints--

`getLongestSubstring("abracadabra")`는 `"abra"`를 반환해야 합니다.

```js
assert.equal(getLongestSubstring("abracadabra"), "abra");
```

`getLongestSubstring("hello world hello")`는 `"hello"`를 반환해야 합니다.

```js
assert.equal(getLongestSubstring("hello world hello"), "hello");
```

`getLongestSubstring("mississippi")`는 `"issi"`를 반환해야 합니다.

```js
assert.equal(getLongestSubstring("mississippi"), "issi");
```

`getLongestSubstring("ha ha ha ha ha ha ha")`는 `"ha ha ha ha ha ha"`를 반환해야 합니다.

```js
assert.equal(getLongestSubstring("ha ha ha ha ha ha ha"), "ha ha ha ha ha ha");
```

`getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over")`는 `"the quick brown fox jumped over"`를 반환해야 합니다.

```js
assert.equal(getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over"), "the quick brown fox jumped over");
```

# --seed--

## --seed-contents--

```js
function getLongestSubstring(str) {

  return str;
}
```

# --solutions--

```js
function getLongestSubstring(str) {
  let longest = '';

  for (let len = str.length - 1; len >= 1; len--) {
    for (let i = 0; i <= str.length - len; i++) {
      const sub = str.slice(i, i + len);
      if (str.indexOf(sub) !== str.lastIndexOf(sub)) {
        return sub;
      }
    }
  }

  return longest;
}
```
