---
id: 69cfca90e8a0a6d4d6871c51
title: "도전 과제 267: 파섹 변환기"
challengeType: 28
dashedName: challenge-267
---

# --description--

먼 은하계에서 파섹은 시간과 거리를 모두 측정하는 데 사용됩니다. 정수로 주어진 파섹 수를 시간 또는 거리로 변환해 반환하세요.

- 주어진 정수가 홀수면 시간, 짝수면 거리를 나타냅니다.

다음 변환 비율을 사용하세요:

| 파섹 | 시간/거리 |
| - | - |
| 1 | 2시간 |
| 2 | 6광년 |

변환된 값을 정수로 반환하세요.

# --hints--

`convertParsecs(1)`는 `2`을 반환해야 합니다.

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)`는 `6`을 반환해야 합니다.

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)`는 `62`을 반환해야 합니다.

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)`는 `264`을 반환해야 합니다.

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)`는 `34`을 반환해야 합니다.

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)`는 `42`을 반환해야 합니다.

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
