---
id: 69cfca90e8a0a6d4d6871c53
title: "도전 과제 269: 알레르기 유발 물질 없는 식사"
challengeType: 28
dashedName: challenge-269
---

# --description--

식사 배열과 피해야 할 알레르기 유발 물질 배열이 주어질 때, 주어진 알레르기 유발 물질이 전혀 포함되지 않은 모든 식사의 이름을 반환하세요.

- 각 식사는 `[meal, allergens]` 형식이며, 여기서 `meal`은 식사 이름이고 `allergens`는 식사에 포함된 알레르기 유발 물질 배열입니다. 예를 들어, `["pasta", ["wheat", "milk"]]`입니다.
- 피해야 할 알레르기 유발 물질은 문자열 배열입니다.

안전한 식사 이름을 주어진 순서대로 반환하세요. 안전한 식사가 없으면 빈 배열을 반환하세요.

# --hints--

`getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"])`는 `["salad"]`를 반환해야 합니다.

```js
assert.deepEqual(getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"]), ["salad"]);
```

`getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"])`는 `["fried rice", "chicken parmesan"]`를 반환해야 합니다.

```js
assert.deepEqual(getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"]), ["fried rice", "chicken parmesan"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"])`는 `["oatmeal", "granola", "toast"]`를 반환해야 합니다.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"]), ["oatmeal", "granola", "toast"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"])`는 `["granola", "yogurt", "eggs"]`를 반환해야 합니다.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"]), ["granola", "yogurt", "eggs"]);
```

# --seed--

## --seed-contents--

```js
function getAllergenFriendlyMeals(meals, allergens) {

  return meals;
}
```

# --solutions--

```js
function getAllergenFriendlyMeals(meals, allergens) {
  return meals
    .filter(([name, mealAllergens]) =>
      !allergens.some(a => mealAllergens.includes(a)))
    .map(([name]) => name);
}
```
