---
id: 69cfca90e8a0a6d4d6871c50
title: "도전 과제 266: 좋은 하루"
challengeType: 28
dashedName: challenge-266
---

# --description--

`"HH:MM"` 형식(24시간제)으로 된 시간 문자열이 주어지면 다음을 반환하세요:

- `"Good morning"`: `05:00`부터 `11:59`까지의 시간에 대해
- `"Good afternoon"`: `12:00`부터 `17:59`까지의 시간에 대해
- `"Good evening"`: `18:00`부터 `21:59`까지의 시간에 대해
- `"Good night"`: `22:00`부터 `04:59`까지의 시간에 대해

# --hints--

`getGreeting("06:30")`는 `"Good morning"`를 반환해야 합니다.

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")`는 `"Good afternoon"`를 반환해야 합니다.

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")`는 `"Good evening"`를 반환해야 합니다.

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")`는 `"Good night"`를 반환해야 합니다.

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")`는 `"Good morning"`를 반환해야 합니다.

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
