---
id: 69e2383af7832c8032603b92
title: "도전 과제 276: 잘못된 요소"
challengeType: 28
dashedName: challenge-276
---

# --description--

오름차순으로 정렬된 정수 배열이 있는데, 그중 하나의 요소만 제자리에 있지 않은 경우, 그 요소의 인덱스를 반환하세요.

- 만약 여러 요소가 제자리에 있지 않은 것으로 간주될 수 있다면, 첫 번째 요소의 인덱스를 반환하세요.

# --hints--

`findOffender([1, 6, 2, 3, 4, 5])`는 `1`를 반환해야 합니다.

```js
assert.equal(findOffender([1, 6, 2, 3, 4, 5]), 1);
```

`findOffender([1, 2, 3, 5, 4, 5])`는 `3`를 반환해야 합니다.

```js
assert.equal(findOffender([1, 2, 3, 5, 4, 5]), 3);
```

`findOffender([2, 1])`는 `0`를 반환해야 합니다.

```js
assert.equal(findOffender([2, 1]), 0);
```

`findOffender([2, 4, 1, 6, 8])`는 `2`를 반환해야 합니다.

```js
assert.equal(findOffender([2, 4, 1, 6, 8]), 2);
```

`findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])`는 `6`를 반환해야 합니다.

```js
assert.equal(findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6);
```

# --seed--

## --seed-contents--

```js
function findOffender(arr) {

  return arr;
}
```

# --solutions--

```js
function findOffender(arr) {
  for (let i = 0; i < arr.length; i++) {
    const without = [...arr.slice(0, i), ...arr.slice(i + 1)];
    if (without.every((n, j) => j === 0 || without[j - 1] <= n)) return i;
  }
}
```
