---
id: 69e2383af7832c8032603b8f
title: "도전 과제 273: ISBN-13 유효성 검사기"
challengeType: 28
dashedName: challenge-273
---

# --description--

문자열이 주어졌을 때, 그것이 유효한 ISBN-13 번호인지 확인하세요.

유효한 ISBN-13은 다음 조건을 만족합니다:

- 숫자와 하이픈만 포함합니다
- 하이픈을 제거한 후 정확히 13자리 숫자가 있습니다
- 다음 검사를 통과합니다:
  1. 각 숫자에 1과 3을 번갈아 곱합니다 (첫 번째 숫자는 1, 두 번째는 3, 세 번째는 1, 이런 식으로).
  2. 결과들의 합이 10으로 나누어떨어져야 합니다.

# --hints--

`isValidIsbn13("9780306406157")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
