---
id: 69e2383af7832c8032603b93
title: "도전 과제 277: 거울상"
challengeType: 28
dashedName: challenge-277
---

# --description--

두 개의 문자열이 주어졌을 때, 두 번째 문자열이 첫 번째 문자열의 거울상인지 확인하세요.

거울상은 문자열을 뒤집고 각 문자를 거울에 비친 대응 문자로 바꿔서 만듭니다.

- 대칭 문자는 거울에 비친 모습이 자기 자신과 같습니다:

`W`, `T`, `Y`, `U`, `I`, `O`, `H`, `A`, `X`, `V`, `M`, `w`, `o`, `x`, `v`, `0`, `8`, `=`, `+`, `:`, `|`, `-`, `_`, `*`, `^`, `!`, `.`, 그리고 공백(` `).

- 거울상 쌍은 서로 거울에 비친 모습이 서로 바뀝니다:

| 문자 | 바뀌는 문자 |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

만약 두 문자열 중 하나라도 위 목록에 없는 문자를 포함하면, 그 문자들로 거울상을 만들 수 없습니다.

예를 들어, `"[HOW]"`의 거울상은 `"[WOH]"`입니다.

# --hints--

`isMirrorImage("[HOW]", "[WOH]")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isMirrorImage("[HOW]", "[WOH]"));
```

`isMirrorImage("MOM", "MOM")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isMirrorImage("MOM", "MOM"));
```

`isMirrorImage("vow", "wov")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isMirrorImage("vow", "wov"));
```

`isMirrorImage("TIM", "TIM")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isMirrorImage("TIM", "TIM"));
```

`isMirrorImage("{WOW}", "}WOW{")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isMirrorImage("{WOW}", "}WOW{"));
```

`isMirrorImage("XXVII", "IIV%X")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isMirrorImage("XXVII", "IIV%X"));
```

`isMirrorImage("><(((*>", "<*)))><")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isMirrorImage("><(((*>", "<*)))><"));
```

`isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"));
```

# --seed--

## --seed-contents--

```js
function isMirrorImage(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirrorImage(str1, str2) {
  const symmetric = new Set('WTYUIOHA XVMwoxv08=+:|-_*^!.');
  const pairs = { '[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<', 'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '(' };
  
  const mirrored = str1.split('').reverse().map(c => {
    if (symmetric.has(c)) return c;
    if (pairs[c]) return pairs[c];
    return null;
  });
  
  if (mirrored.includes(null)) return false;
  return mirrored.join('') === str2;
}
```
