---
id: 69e2383af7832c8032603b90
title: "도전 과제 274: 가장 나이 많은 사람"
challengeType: 28
dashedName: challenge-274
---

# --description--

각각 `"name"`와 `"age"` 속성을 가진 객체들로 이루어진 배열이 주어질 때, 가장 나이 많은 사람의 이름을 담은 배열을 반환하세요.

만약 여러 사람이 가장 나이가 많다면, 입력에 나타난 순서대로 그들의 이름을 모두 반환하세요.

# --hints--

`getOldest([{ name: "Brenda", age: 40 }])`는 `["Brenda"]`를 반환해야 합니다.

```js
assert.deepEqual(getOldest([{ name: "Brenda", age: 40 }]), ["Brenda"]);
```

`getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }])`는 `["Alice"]`를 반환해야 합니다.

```js
assert.deepEqual(getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }]), ["Alice"]);
```

`getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }])`는 `["Bill", "Carol"]`를 반환해야 합니다.

```js
assert.deepEqual(getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }]), ["Bill", "Carol"]);
```

`getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }])`는 `["George", "Holly", "Zach"]`를 반환해야 합니다.

```js
assert.deepEqual(getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }]), ["George", "Holly", "Zach"]);
```

# --seed--

## --seed-contents--

```js
function getOldest(people) {

  return people;
}
```

# --solutions--

```js
function getOldest(people) {
  const maxAge = Math.max(...people.map(p => p.age));
  return people.filter(p => p.age === maxAge).map(p => p.name);
}
```
