---
id: 69e2383af7832c8032603b91
title: "도전 과제 275: 문자 빈도"
challengeType: 28
dashedName: challenge-275
---

# --description--

문자열이 주어지면, 각 문자가 나타나는 횟수를 매핑하는 객체(JavaScript) 또는 딕셔너리(Python)를 반환하세요.

# --hints--

`getFrequency("test")`는 `{t: 2, e: 1, s: 1}`을 반환해야 합니다.

```js
assert.deepEqual(getFrequency("test"), {t: 2, e: 1, s: 1});
```

`getFrequency("mississippi")`는 `{m: 1, i: 4, s: 4, p: 2}`을 반환해야 합니다.

```js
assert.deepEqual(getFrequency("mississippi"), {m: 1, i: 4, s: 4, p: 2});
```

`getFrequency("hello world")`는 `{h: 1, e: 1, l: 3, o: 2, " ": 1, w: 1, r: 1, d: 1}`을 반환해야 합니다.

```js
assert.deepEqual(getFrequency("hello world"), {h: 1, e: 1, l: 3, o: 2, " ": 1, w: 1, r: 1, d: 1});
```

`getFrequency("She sells seashells by the seashore.")`는 `{S: 1, h: 4, e: 7, " ": 5, s: 7, l: 4, a: 2, b: 1, y: 1, t: 1, o: 1, r: 1, ".": 1}`을 반환해야 합니다.

```js
assert.deepEqual(getFrequency("She sells seashells by the seashore."), {S: 1, h: 4, e: 7, " ": 5, s: 7, l: 4, a: 2, b: 1, y: 1, t: 1, o: 1, r: 1, ".": 1});
```

`getFrequency("The quick brown fox jumps over the lazy dog.")`는 `{T: 1, h: 2, e: 3, " ": 8, q: 1, u: 2, i: 1, c: 1, k: 1, b: 1, r: 2, o: 4, w: 1, n: 1, f: 1, x: 1, j: 1, m: 1, p: 1, s: 1, v: 1, t: 1, l: 1, a: 1, z: 1, y: 1, d: 1, g: 1, ".": 1}`을 반환해야 합니다.

```js
assert.deepEqual(getFrequency("The quick brown fox jumps over the lazy dog."), {T: 1, h: 2, e: 3, " ": 8, q: 1, u: 2, i: 1, c: 1, k: 1, b: 1, r: 2, o: 4, w: 1, n: 1, f: 1, x: 1, j: 1, m: 1, p: 1, s: 1, v: 1, t: 1, l: 1, a: 1, z: 1, y: 1, d: 1, g: 1, ".": 1});
```

# --seed--

## --seed-contents--

```js
function getFrequency(str) {

  return str;
}
```

# --solutions--

```js
function getFrequency(str) {
  return str.split('').reduce((acc, char) => {
    acc[char] = (acc[char] || 0) + 1;
    return acc;
  }, {});
}
```
