---
id: 69e2383af7832c8032603b8e
title: "도전 과제 272: 전치 행렬"
challengeType: 28
dashedName: challenge-272
---

# --description--

행렬(배열의 배열)이 주어지면, 그 전치 행렬을 반환하세요.

행렬을 전치하려면 행과 열을 바꾸세요. 예를 들어, 인덱스 `[0, 1]`에 있는 값은 인덱스 `[1, 0]`로 이동해야 합니다.

예를 들어, 다음이 주어졌을 때:

```json
[
  [1, 2, 3],
  [4, 5, 6]
]
```

다음과 같이 반환하세요:

```json
[
  [1, 4],
  [2, 5],
  [3, 6]
]
```

# --hints--

`transpose([[1, 2, 3], [4, 5, 6]])`는 `[[1, 4], [2, 5], [3, 6]]`를 반환해야 합니다.

```js
assert.deepEqual(transpose([[1, 2, 3], [4, 5, 6]]), [[1, 4], [2, 5], [3, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6]])`는 `[[1, 3, 5], [2, 4, 6]]`를 반환해야 합니다.

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6]]), [[1, 3, 5], [2, 4, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6], [7, 8]])`는 `[[1, 3, 5, 7], [2, 4, 6, 8]]`를 반환해야 합니다.

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6], [7, 8]]), [[1, 3, 5, 7], [2, 4, 6, 8]]);
```

`transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]])`는 `[["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]`를 반환해야 합니다.

```js
assert.deepEqual(transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]]), [["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]);
```

`transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]])`는 `[[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]`를 반환해야 합니다.

```js
assert.deepEqual(transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]]), [[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]);
```

# --seed--

## --seed-contents--

```js
function transpose(matrix) {

  return matrix;
}
```

# --solutions--

```js
function transpose(matrix) {

  return matrix[0].map((_, i) => matrix.map(row => row[i]));
}
```
