---
id: 69f35a5bb823ed620fcb7cba
title: "도전 과제 281: 빙고 범위"
challengeType: 28
dashedName: challenge-281
---

# --description--

빙고 글자가 주어지면, 그 글자에 해당하는 숫자 범위를 반환하세요.

| 글자 | 숫자 범위 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

가장 작은 수부터 가장 큰 수까지 범위 내 모든 숫자를 포함하는 배열을 반환하세요.

# --hints--

`getBingoRange("B")`는 `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]`를 반환해야 합니다.

```js
assert.deepEqual(getBingoRange("B"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]);
```

`getBingoRange("I")`는 `[16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30]`를 반환해야 합니다.

```js
assert.deepEqual(getBingoRange("I"), [16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30]);
```

`getBingoRange("N")`는 `[31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45]`를 반환해야 합니다.

```js
assert.deepEqual(getBingoRange("N"), [31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45]);
```

`getBingoRange("G")`는 `[46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60]`를 반환해야 합니다.

```js
assert.deepEqual(getBingoRange("G"), [46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60]);
```

`getBingoRange("O")`는 `[61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75]`를 반환해야 합니다.

```js
assert.deepEqual(getBingoRange("O"), [61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75]);
```

# --seed--

## --seed-contents--

```js
function getBingoRange(letter) {

  return letter;
}
```

# --solutions--

```js
function getBingoRange(letter) {
  const ranges = { B: [1, 15], I: [16, 30], N: [31, 45], G: [46, 60], O: [61, 75] };
  const [start, end] = ranges[letter];
  return Array.from({ length: end - start + 1 }, (_, i) => start + i);
}
```
