---
id: 69f35a5bb823ed620fcb7cbd
title: "도전 과제 284: I Before E 규칙"
challengeType: 28
dashedName: challenge-284
---

# --description--

단어나 문장이 주어지면, 모든 단어가 "C 다음이 아니면 I가 E보다 먼저 온다" 규칙을 따르도록 수정된 버전을 반환하세요.

- 단어에 `"ei"`가 `"c"` 앞에 오지 않고 포함되어 있으면 `"ie"`로 바꾸세요.
- 단어에 `"ie"`가 `"c"` 앞에 오면 `"ei"`로 바꾸세요.
- 나머지 단어는 변경하지 않고 그대로 두세요.

# --hints--

`iBeforeE("beleive")`는 `"believe"`를 반환해야 합니다.

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")`는 `"receive"`를 반환해야 합니다.

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")`는 `"we received a brief"`를 반환해야 합니다.

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")`는 `"she believed the friendly niece could perceive the grief"`를 반환해야 합니다.

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")`는 `"we received relief after the thief gave us a brief piece of fierce deceit"`를 반환해야 합니다.

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
