---
id: 69f35a5bb823ed620fcb7cb9
title: "도전 과제 280: Mongo ID 날짜"
challengeType: 28
dashedName: challenge-280
---

# --description--

MongoDB ID 문자열이 주어지면, 생성 시간을 ISO 8601 문자열로 반환하세요.

- MongoDB ID는 24자리 16진수 문자열입니다. 처음 8자리는 16진수 정수로 인코딩된 유닉스 타임스탬프(초 단위)를 나타냅니다.

예를 들어, `"6a094b50bcf86cd799439011"`는 16진수로 `"6a094b50"`의 타임스탬프를 가지며, 이는 10진수로 `1778994000`이고, 생성 시간은 `"2026-05-17T05:00:00.000Z"`입니다.

# --hints--

`mongoIdToDate("6a094b50bcf86cd799439011")`는 `"2026-05-17T05:00:00.000Z"`를 반환해야 합니다.

```js
assert.equal(mongoIdToDate("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z");
```

`mongoIdToDate("695344eb1f4a4c1123042128")`는 `"2025-12-30T03:20:11.000Z"`를 반환해야 합니다.

```js
assert.equal(mongoIdToDate("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z");
```

`mongoIdToDate("386da62df34123ac54617e56")`는 `"2000-01-01T07:01:01.000Z"`를 반환해야 합니다.

```js
assert.equal(mongoIdToDate("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z");
```

`mongoIdToDate("69f571c3d7711807afd3dd55")`는 `"2026-05-02T03:38:43.000Z"`를 반환해야 합니다.

```js
assert.equal(mongoIdToDate("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z");
```

`mongoIdToDate("68adce01c0e1144d0a90295a")`는 `"2025-08-26T15:08:49.000Z"`를 반환해야 합니다.

```js
assert.equal(mongoIdToDate("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z");
```

# --seed--

## --seed-contents--

```js
function mongoIdToDate(id) {

  return id;
}
```

# --solutions--

```js
function mongoIdToDate(id) {
  const timestamp = parseInt(id.slice(0, 8), 16);
  return new Date(timestamp * 1000).toISOString();
}
```
