---
id: 69f35a5bb823ed620fcb7cbb
title: "도전 과제 282: 수면 부채"
challengeType: 28
dashedName: challenge-282
---

# --description--

오늘까지 매일 밤 잔 수면 시간을 배열로 주고, 목표 수면 시간을 알려줄 때, 수면 부채를 없애기 위해 오늘 밤 몇 시간 자야 하는지 반환하세요.

- 오늘 밤 자야 할 시간도 총 필요한 시간에 포함하세요.
- 오늘 밤 목표 수면 시간 이상 잤다면 `0`를 반환하세요.

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)`는 `20`를 반환해야 합니다.

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)`는 `10`를 반환해야 합니다.

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)`는 `4`를 반환해야 합니다.

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)`는 `0`를 반환해야 합니다.

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)`는 `0`를 반환해야 합니다.

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
