---
id: 69f35a5bb823ed620fcb7cbe
title: "도전 과제 285: 만남 시간"
challengeType: 28
dashedName: challenge-285
---

# --description--

여러 사람의 가능 시간대를 나타내는 3D 배열이 주어질 때, 모두가 한 시간 동안 자유로운 가장 이른 시간을 반환하세요. 그런 시간이 없다면 `"None"`를 반환하세요.

- 각 사람의 가능 시간은 24시간 형식의 `[start, end]` 정수 쌍 배열입니다. 예를 들어, `[10, 12]`는 그 사람이 10시부터 12시까지 가능하다는 뜻입니다. 시작 시간은 0~23 사이이고, 종료 시간은 1~24 사이입니다.

예를 들어, 다음이 주어지면:

```json
[
  [[10, 12], [15, 16]], // person 1
  [[11, 14], [15, 16]]  // person 2
]
```

모두가 처음으로 공유하는 한 시간의 시작인 `11`를 반환하세요.

# --hints--

`getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]])`는 `11`를 반환해야 합니다.

```js
assert.equal(getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]]), 11);
```

`getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]])`는 `13`를 반환해야 합니다.

```js
assert.equal(getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]]), 13);
```

`getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]])`는 `9`를 반환해야 합니다.

```js
assert.equal(getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]]), 9);
```

`getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]])`는 `None`를 반환해야 합니다.

```js
assert.equal(getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]]), "None");
```

`getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]])`는 `21`를 반환해야 합니다.

```js
assert.equal(getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]]), 21);
```

# --seed--

## --seed-contents--

```js
function getMeetingTime(availability) {

  return availability;
}
```

# --solutions--

```js
function getMeetingTime(availability) {
  for (let hour = 0; hour < 24; hour++) {
    const allFree = availability.every(person =>
      person.some(([start, end]) => start <= hour && hour + 1 <= end)
    );
    if (allFree) return hour;
  }
  return "None";
}
```
