---
id: 69f35a5bb823ed620fcb7cbc
title: "도전 과제 283: 문자열 지퍼"
challengeType: 28
dashedName: challenge-283
---

# --description--

두 개의 문자열이 주어지면, 각 문자를 한 번에 하나씩 번갈아가며 합친 새 문자열을 반환하세요. 한 문자열이 더 길면 남은 문자를 끝에 덧붙이세요.

첫 번째 문자열의 첫 문자부터 시작하세요.

# --hints--

`zipStrings("abc", "123")`는 `"a1b2c3"`를 반환해야 합니다.

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")`는 `"abcdefghijklmnopqrstuvwxyz"`를 반환해야 합니다.

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")`는 `"dnaiyght"`를 반환해야 합니다.

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")`는 `"pjyatvhaosncript"`를 반환해야 합니다.

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")`는 `"freeCodeCamp"`를 반환해야 합니다.

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
