---
id: 69f8c998d78ad3171a0713be
title: "도전 과제 291: FizzBuzz 개수 세기"
challengeType: 28
dashedName: challenge-291
---

# --description--

시작 숫자와 끝 숫자가 주어지면, 범위 내(포함)에서 fizz와 buzz가 나타나는 횟수를 세세요.

- 3으로 나누어떨어지는 숫자는 fizz로 셉니다.
- 5로 나누어떨어지는 숫자는 buzz로 셉니다.
- 3과 5 모두로 나누어떨어지는 숫자는 fizz와 buzz 모두로 셉니다.

개수를 `{ fizz, buzz }` 형식으로 객체나 딕셔너리로 반환하세요.

# --hints--

`fizzBuzzCount(1, 11)`는 `{fizz: 3, buzz: 2}`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)`는 `{fizz: 9, buzz: 6}`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)`는 `{fizz: 26, buzz: 16}`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)`는 `{fizz: 207, buzz: 125}`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)`는 `{fizz: 4074, buzz: 2444}`를 반환해야 합니다.

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
