---
id: 69f8c998d78ad3171a0713bb
title: "도전 과제 288: 비밀 숫자"
challengeType: 28
dashedName: challenge-288
---

# --description--

비밀 숫자와 추측이 주어졌을 때, 추측이 맞았는지 확인하세요.

반환값:

- 비밀 숫자가 추측보다 크면 `"higher"`를 반환하세요.
- 비밀 숫자가 추측보다 작으면 `"lower"`를 반환하세요.
- 추측이 맞으면 `"you got it!"`를 반환하세요.

# --hints--

`guessNumber(50, 30)`는 `"higher"`를 반환해야 합니다.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)`는 `"lower"`를 반환해야 합니다.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)`는 `"you got it!"`를 반환해야 합니다.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)`는 `"higher"`를 반환해야 합니다.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)`는 `"lower"`를 반환해야 합니다.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)`는 `"you got it!"`를 반환해야 합니다.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
