---
id: 69f8c998d78ad3171a0713ba
title: "도전 과제 287: 로마 숫자 수정기"
challengeType: 28
dashedName: challenge-287
---

# --description--

잘못된 로마 숫자 문자열이 주어지면, 표준 로마 숫자 표기법으로 값을 반환하세요.

입력은 덧셈 표기법만 사용하므로 각 기호가 그 값을 총합에 더합니다. 참고로, 기호와 값은 다음과 같습니다:

| Symbol | Value |
| ------ | ----- |
| `"I"`  | 1     |
| `"V"`  | 5     |
| `"X"`  | 10    |
| `"L"`  | 50    |
| `"C"`  | 100   |
| `"D"`  | 500   |
| `"M"`  | 1000  |

재인코딩할 때는 각 단계에서 가능한 가장 큰 기호를 사용하고, 필요하면 뺄셈 쌍(`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`)을 사용하세요.

# --hints--

`fixNumerals("XIIIII")`는 `"XV"`를 반환해야 합니다.

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")`는 `"LXIV"`를 반환해야 합니다.

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")`는 `"XL"`를 반환해야 합니다.

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")`는 `"MDCCXCIX"`를 반환해야 합니다.

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")`는 `"MCDLXIV"`를 반환해야 합니다.

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")`는 `"MMCMLXXXIV"`를 반환해야 합니다.

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```
