---
id: 69f8c998d78ad3171a0713bf
title: "도전 과제 292: 더 넓은 가로세로비"
challengeType: 28
dashedName: challenge-292
---

# --description--

서로 다른 이미지 크기를 나타내는 두 문자열이 주어질 때, 가로세로 비율이 더 큰 이미지의 가로세로비를 반환하세요.

- 주어진 문자열은 `"WxH"` 형식이며, 예를 들어 `"1920x1080"`과 같습니다.
- 가로세로비는 가로 길이를 세로 길이로 나눈 비율이며, 가장 작은 정수 비율로 약분합니다. 예를 들어 `"1920x1080"`는 `"16:9"`으로 약분됩니다.
- `"W:H"` 형식의 문자열을 반환하세요. 예를 들어 `"16:9"`와 같습니다.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")`는 `"16:9"`을 반환해야 합니다.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")`는 `"4:3"`을 반환해야 합니다.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")`는 `"2:1"`을 반환해야 합니다.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")`는 `"9:16"`을 반환해야 합니다.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")`는 `"43:18"`을 반환해야 합니다.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")`는 `"2089:8539"`을 반환해야 합니다.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
