---
id: 69f8c998d78ad3171a0713b9
title: "도전 과제 286: 열린 이슈"
challengeType: 28
dashedName: challenge-286
---

# --description--

이슈 번호 배열과 풀 리퀘스트(PR) 번호 배열이 주어질 때, 모든 PR이 병합된 후에도 여전히 열린 이슈들의 배열을 반환하세요.

- PR은 숫자가 서로 회전한 관계일 때 이슈를 닫습니다. 예를 들어, 이슈 123은 PR 231이나 312에 의해 닫힙니다.
- PR은 정확히 같은 번호의 이슈를 닫지 않습니다. 예를 들어, PR 123은 이슈 123을 닫지 않습니다. 따라서 모든 숫자가 같은 이슈는 닫히지 않습니다.
- 어느 쪽 숫자도 앞의 0이 제거될 수 있습니다. 예를 들어, PR 201은 이슈 12(012, 201의 회전)도 닫습니다. 마찬가지로, 이슈 201은 PR 12에 의해 닫힙니다.

주어진 순서대로 남아 있는 열린 이슈를 반환하세요.

# --hints--

`getOpenIssues([123, 234], [231])`는 `[234]`을 반환해야 합니다.

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])`는 `[345, 16]`을 반환해야 합니다.

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])`는 `[456, 332, 15]`를 반환해야 합니다.

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])`는 `[115, 296, 24]`을 반환해야 합니다.

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])`는 `[95, 395, 754, 296, 709, 237, 1802]`를 반환해야 합니다.

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
