---
id: 69f8c998d78ad3171a0713bd
title: "도전 과제 290: 피자 파티"
challengeType: 28
dashedName: challenge-290
---

# --description--

오늘 각 사람이 일한 시간을 담은 배열이 주어졌을 때, 피자 파티를 위해 주문할 피자 개수를 반환하세요.

- 각 사람이 일한 시간을 3으로 나누어 조각 수를 구하세요.
- 조각을 부분적으로 먹을 수 없으니, 각 사람의 조각 수를 가장 가까운 정수로 올림하세요.
- 각 사람은 최소 두 조각을 받습니다.
- 피자 한 판은 8조각입니다. 피자 총 개수를 가장 가까운 정수로 올림하세요.

# --hints--

`getPizzasToOrder([8, 8, 8])` should return `2`.

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` should return `3`.

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` should return `2`.

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` should return `3`.

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` should return `1`.

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` should return `5`.

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
