---
id: 69f90c1329a94b37e2a2086d
title: "도전 과제 294: 소괄호 조합"
challengeType: 28
dashedName: challenge-294
---

# --description--

정수 `n`가 주어지면, `n` 쌍의 소괄호로 만들 수 있는 유효한 조합의 개수를 반환하세요.

- 유효한 조합이란 모든 여는 소괄호가 대응하는 닫는 소괄호를 가지며, 닫는 소괄호가 그에 대응하는 여는 소괄호보다 먼저 나오지 않는 문자열입니다.

예를 들어, `2`가 주어지면 유효한 조합은 2가지입니다:

```json
(())
()()
```

# --hints--

`getCombinations(2)`는 `2`를 반환해야 합니다.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)`는 `5`를 반환해야 합니다.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)`는 `42`를 반환해야 합니다.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)`는 `1430`를 반환해야 합니다.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)`는 `742900`를 반환해야 합니다.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
