---
id: 6a0dcd03ee4e68698080ef66
title: "도전 과제 302: 시차 적응"
challengeType: 28
dashedName: challenge-302
---

# --description--

출발 도시, 도착 도시, 비행 시간(시간 단위), 그리고 이동 방향이 주어질 때, 여행자가 겪는 시차 적응 시간을 반환하세요.

주어진 도시는 다음 목록에 있는 도시들이며, 각 도시의 UTC 오프셋도 포함되어 있습니다:

| 도시 | 오프셋 |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

시차 적응 시간을 계산하는 방법:

1. 두 도시 간 시간대 차이를 시간 단위로 구하세요.
2. 이동 방향에 따른 배수를 결정하세요. `"east"` 방향으로 이동하면 1.5, 그렇지 않으면 1.0입니다.
3. 시차 적응 시간은 다음 공식으로 구하세요: 시간대 차이 + (비행 시간 * 0.1) * 방향 배수

시차 적응 시간은 소수점 첫째 자리까지 반올림하여 반환하세요.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")`는 `6.5`를 반환해야 합니다.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")`는 `5.8`를 반환해야 합니다.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")`는 `1.6`를 반환해야 합니다.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")`는 `4.7`를 반환해야 합니다.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")`는 `17.5`를 반환해야 합니다.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")`는 `5.9`를 반환해야 합니다.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")`는 `9.5`를 반환해야 합니다.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
