---
id: 6a0dcd03ee4e68698080ef6c
title: "도전 과제 308: 신용카드 유효성 검사기"
challengeType: 28
dashedName: challenge-308
---

# --description--

신용카드 번호를 나타내는 숫자 문자열이 주어졌을 때, 다음 메서드를 사용해 유효한 카드 번호인지 확인하세요:

- 끝에서 두 번째 숫자부터 시작해 왼쪽으로 한 칸씩 건너뛰며 모든 숫자를 두 배로 만드세요.
- 두 배로 만든 숫자가 9보다 크면 9를 빼세요.
- 두 배로 만든 숫자와 그렇지 않은 숫자 모두 합하세요.
- 합계가 10으로 나누어떨어지면 유효한 번호입니다.

# --hints--

`isValidCard("4532015112830366")` should return `true`.

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` should return `true`.

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` should return `true`.

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` should return `true`.

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` should return `false`.

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` should return `false`.

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` should return `false`.

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
