---
id: 6a0dcd03ee4e68698080ef69
title: "도전 과제 305: 아이디어 순위 매기기"
challengeType: 28
dashedName: challenge-305
---

# --description--

각 내부 배열에 아이디어 이름, 낙관적 예상, 현실적 예상, 비관적 예상(일 단위)이 순서대로 들어 있는 2차원 배열이 주어질 때, 예상 완료 시간 기준으로 가장 짧은 것부터 정렬된 아이디어 이름 배열을 반환하세요.

다음 공식을 사용해 각 아이디어의 예상 완료 시간을 계산하세요:

- `expected = ((optimistic + 4 * realistic + pessimistic) / 6) * length of idea name`

# --hints--

`analyzeIdeas([["Add logging", 2, 5, 15], ["SEO optimization", 4, 8, 20], ["Fix bug", 1, 3, 5]])`은 `["Fix bug", "Add logging", "SEO optimization"]`를 반환해야 합니다.

```js
assert.deepEqual(analyzeIdeas([["Add logging", 2, 5, 15], ["SEO optimization", 4, 8, 20], ["Fix bug", 1, 3, 5]]), ["Fix bug", "Add logging", "SEO optimization"]);
```

`analyzeIdeas([["Dark mode", 1, 3, 8], ["Real-time collaboration feature", 6, 12, 20], ["Add tooltip", 1, 2, 4]])`은 `["Add tooltip", "Dark mode", "Real-time collaboration feature"]`를 반환해야 합니다.

```js
assert.deepEqual(analyzeIdeas([["Dark mode", 1, 3, 8], ["Real-time collaboration feature", 6, 12, 20], ["Add tooltip", 1, 2, 4]]), ["Add tooltip", "Dark mode", "Real-time collaboration feature"]);
```

`analyzeIdeas([["Update user profile page", 3, 7, 14], ["Add pagination", 2, 5, 10], ["Add tags", 2, 3, 6], ["Fix login bug", 1, 4, 8]])`는 `["Add tags", "Fix login bug", "Add pagination", "Update user profile page"]`를 반환해야 합니다.

```js
assert.deepEqual(analyzeIdeas([["Update user profile page", 3, 7, 14], ["Add pagination", 2, 5, 10], ["Add tags", 2, 3, 6], ["Fix login bug", 1, 4, 8]]), ["Add tags", "Fix login bug", "Add pagination", "Update user profile page"]);
```

`analyzeIdeas([["Migrate database", 14, 25, 40], ["Add chat assistant", 8, 15, 24], ["Redesign onboarding flow", 3, 7, 13], ["Add language support", 6, 11, 18]])`은 `["Redesign onboarding flow", "Add language support", "Add chat assistant", "Migrate database"]`를 반환해야 합니다.

```js
assert.deepEqual(analyzeIdeas([["Migrate database", 14, 25, 40], ["Add chat assistant", 8, 15, 24], ["Redesign onboarding flow", 3, 7, 13], ["Add language support", 6, 11, 18]]), ["Redesign onboarding flow", "Add language support", "Add chat assistant", "Migrate database"]);
```

`analyzeIdeas([["Add email notifications", 3, 7, 10], ["Migrate deployment flow", 6, 10, 16], ["Add push notifications", 2, 6, 10], ["Optimize continuous integration", 5, 8, 15], ["Analyze user patterns", 5, 10, 18], ["Create onboarding curriculum", 6, 15, 25]])`는 `["Add push notifications", "Add email notifications", "Analyze user patterns", "Migrate deployment flow", "Optimize continuous integration", "Create onboarding curriculum"]`를 반환해야 합니다.

```js
assert.deepEqual(analyzeIdeas([["Add email notifications", 3, 7, 10], ["Migrate deployment flow", 6, 10, 16], ["Add push notifications", 2, 6, 10], ["Optimize continuous integration", 5, 8, 15], ["Analyze user patterns", 5, 10, 18], ["Create onboarding curriculum", 6, 15, 25]]), ["Add push notifications", "Add email notifications", "Analyze user patterns", "Migrate deployment flow", "Optimize continuous integration", "Create onboarding curriculum"]);
```

# --seed--

## --seed-contents--

```js
function analyzeIdeas(ideas) {

  return ideas;
}
```

# --solutions--

```js
function analyzeIdeas(ideas) {
  return ideas
    .map(([name, optimistic, realistic, pessimistic]) => ({
      name,
      expected: ((optimistic + 4 * realistic + pessimistic) / 6) * name.length
    }))
    .sort((a, b) => a.expected - b.expected)
    .map(idea => idea.name);
}
```
