---
id: 6a0dcd03ee4e68698080ef6d
title: "도전 과제 309: 숫자 정렬"
challengeType: 28
dashedName: challenge-309
---

# --description--

쉼표로 구분된 숫자 문자열이 주어지면, 가장 작은 수부터 가장 큰 수까지 정렬된 숫자 배열을 반환하세요.

# --hints--

`sortNumbers("3,1,2")`는 `[1, 2, 3]`을 반환해야 합니다.

```js
assert.deepEqual(sortNumbers("3,1,2"), [1, 2, 3]);
```

`sortNumbers("5,3,8,1,9,2")`는 `[1, 2, 3, 5, 8, 9]`을 반환해야 합니다.

```js
assert.deepEqual(sortNumbers("5,3,8,1,9,2"), [1, 2, 3, 5, 8, 9]);
```

`sortNumbers("12,61,49,80,19,50,77,38")`는 `[12, 19, 38, 49, 50, 61, 77, 80]`을 반환해야 합니다.

```js
assert.deepEqual(sortNumbers("12,61,49,80,19,50,77,38"), [12, 19, 38, 49, 50, 61, 77, 80]);
```

`sortNumbers("0,6,-19,44,-2,7,0")`는 `[-19, -2, 0, 0, 6, 7, 44]`을 반환해야 합니다.

```js
assert.deepEqual(sortNumbers("0,6,-19,44,-2,7,0"), [-19, -2, 0, 0, 6, 7, 44]);
```

# --seed--

## --seed-contents--

```js
function sortNumbers(str) {

  return str;
}
```

# --solutions--

```js
function sortNumbers(str) {
  return str.split(",").map(Number).sort((a, b) => a - b);
}
```
