---
id: 6a0dcc730cb92a616f86f0c5
title: "도전 과제 301: 마지막 사용량"
challengeType: 28
dashedName: challenge-301
---

# --description--

남은 세탁 세제 스쿱 수와 이전 며칠 동안 사용한 스쿱 수가 담긴 배열이 주어질 때, 남은 세제로 며칠을 완전히 사용할 수 있는지 반환하세요.

사용 기록에서 하루 평균 사용량을 계산하고 앞으로 매일 그만큼 사용할 것이라고 가정하세요.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])`는 `5`을 반환해야 합니다.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])`는 `7`을 반환해야 합니다.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])`는 `11`를 반환해야 합니다.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])`는 `12`을 반환해야 합니다.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])`는 `1`을 반환해야 합니다.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
