---
id: 6a0dcc730cb92a616f86f0c4
title: "도전 과제 300: 스키마 검증기 6부"
challengeType: 28
dashedName: challenge-300
---

# --description--

객체(JavaScript) 또는 딕셔너리(Python)가 다음 스키마와 일치하는지 확인하세요:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

UserProfile = {
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean,
  badges: string[]
}

{
  users: UserProfile[]
}
```

- 파이프(`|`) 기호는 "또는"을 의미합니다. role은 나열된 `Roles` 값 중 하나여야 합니다.
- supporter 뒤의 물음표(`?`)는 해당 필드가 선택 사항이지만 존재하면 지정된 유형이어야 함을 나타냅니다.
- `UserProfile[]`는 `UserProfile` 객체의 배열을 나타냅니다. 빈 배열도 유효합니다.
- 추가 키는 허용됩니다

# --hints--

`isValidSchema({ users: [{ username: "ron", posts: 14, verified: true, role: "creator", badges: [ "early-adopter" ]}, { username: "cher", posts: 25, verified: true, role: "moderator", supporter: true, followers: 20, badges: [ "helper" ]}]})`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({users: [{ username: "ron", posts: 14, verified: true, role: "creator", badges: [ "early-adopter" ]}, { username: "cher", posts: 25, verified: true, role: "moderator", supporter: true, followers: 20, badges: [ "helper" ]}]}));
```

`isValidSchema({ users: [] })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ users: [] }));
```

`isValidSchema({ users: { username: "anne", posts: 0, verified: false, role: "user", supporter: false, badges: []}})`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ users: { username: "anne", posts: 0, verified: false, role: "user", supporter: false, badges: []}}));
```

`isValidSchema({ users: [{ username: "tony", posts: 10, verified: true, role: "creator", supporter: true, badges: ["liked", 6]}]})`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ users: [{ username: "tony", posts: 10, verified: true, role: "creator", supporter: true, badges: ["liked", 6]}]}));
```

`isValidSchema({ users: [{ username: "ursula", posts: 3, verified: false, role: "user", supporter: "false", badges: ["comeback"]}]})`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ users: [{ username: "ursula", posts: 3, verified: false, role: "user", supporter: "false", badges: ["comeback"]}]}));
```

`isValidSchema({ users: [{ username: "benny", posts: 55, verified: true, role: "superstar", supporter: true, badges: ["veteran"]}]})`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ users: [{ username: "benny", posts: 55, verified: true, role: "superstar", supporter: true, badges: ["veteran"]}]}));
```

`isValidSchema({ users: [{ username: "chase", posts: 1, verified: "yes", role: "staff", supporter: false, badges: ["superstar"]}]})`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ users: [{ username: "chase", posts: 1, verified: "yes", role: "staff", supporter: false, badges: ["superstar"]}]}));
```

`isValidSchema({ users: [{ username: "carla", posts: "10", verified: false, role: "user", supporter: false, badges: ["newbie"]}]})`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ users: [{ username: "carla", posts: "10", verified: false, role: "user", supporter: false, badges: ["newbie"]}]}));
```

`isValidSchema({ users: [{ posts: 4, verified: false, role: "admin", supporter: false, badges: ["superuser", "veteran"]}]})`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ users: [{ posts: 4, verified: false, role: "admin", supporter: false, badges: ["superuser", "veteran"]}]}));
```

`isValidSchema({ users: [{ username: "harold", posts: 80, verified: true, role: "creator", supporter: true, badges: ["liked", "hero"]}, { username: "kim", posts: 11, verified: false, role: "admin", supporter: true, badges: ["first"]}, {}]})`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ users: [{ username: "harold", posts: 80, verified: true, role: "creator", supporter: true, badges: ["liked", "hero"]}, { username: "kim", posts: 11, verified: false, role: "admin", supporter: true, badges: ["first"]}, {}]}));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];

  function isValidUser(user) {
    return (
      typeof user.username === 'string' &&
      typeof user.posts === 'number' &&
      typeof user.verified === 'boolean' &&
      roles.includes(user.role) &&
      (user.supporter === undefined || typeof user.supporter === 'boolean') &&
      Array.isArray(user.badges) && user.badges.every(b => typeof b === 'string')
    );
  }

  return Array.isArray(obj.users) && obj.users.every(isValidUser);
}
```
