---
id: 6a0dcc730cb92a616f86f0c1
title: "도전 과제 297: 스키마 검증기 3부"
challengeType: 28
dashedName: challenge-297
---

# --description--

객체(JavaScript) 또는 딕셔너리(Python)가 다음 스키마와 일치하는지 확인하세요:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles
}
```

- 파이프(`|`) 기호는 "또는"을 의미합니다. `role`은 나열된 `Roles` 값 중 하나여야 합니다.
- 추가 키는 허용됩니다

# --hints--

`isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" }));
```

`isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 }));
```

`isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" }));
```

`isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" }));
```

`isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" }));
```

`isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" }));
```

`isValidSchema({ username: "wendy", posts: 10, verified: true })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "wendy", posts: 10, verified: true }));
```

`isValidSchema({ username: "fabian", posts: 1, verified: true, role: true })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "fabian", posts: 1, verified: true, role: true }));
```

`isValidSchema({ username: 8, posts: 1, verified: true, role: "user" })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: 8, posts: 1, verified: true, role: "user" }));
```

`isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" }));
```

`isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role)
  );
}
```
