---
id: 6a0dcc730cb92a616f86f0bf
title: "도전 과제 295: 스키마 검증기 1부"
challengeType: 28
dashedName: challenge-295
---

# --description--

객체(JavaScript) 또는 딕셔너리(Python)가 다음 스키마와 일치하는지 확인하세요:

```json
{
  username: string
}
```

- 추가 키는 허용됩니다

# --hints--

`isValidSchema({ username: "bob" })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
