---
id: 6a0dcc730cb92a616f86f0c3
title: "도전 과제 299: 스키마 검증기 5부"
challengeType: 28
dashedName: challenge-299
---

# --description--

객체(JavaScript) 또는 딕셔너리(Python)가 다음 스키마와 일치하는지 확인하세요:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean,
  badges: string[]
}
```

- 파이프(`|`) 기호는 "또는"을 의미합니다. `role`은 나열된 `Roles` 값 중 하나여야 합니다.
- `?` 뒤의 물음표(`supporter`)는 해당 필드가 선택 사항이지만 존재한다면 지정된 유형이어야 함을 의미합니다.
- `[]` 뒤의 괄호 `string`는 `badges`이 문자열의 배열(또는 빈 배열)이어야 함을 의미합니다.
- 추가 키는 허용됩니다

# --hints--

`isValidSchema({ username: "gill", posts: 12, verified: false, role: "creator", supporter: false, badges: [ "early-adopter", "popular" ] })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "gill", posts: 12, verified: false, role: "creator", supporter: false, badges: [ "early-adopter", "popular" ] }));
```

`isValidSchema({ username: "tonya", posts: 299, verified: true, role: "moderator", supporter: true, badges: [ "streak-master", "veteran" ], followers: 1233 })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "tonya", posts: 299, verified: true, role: "moderator", supporter: true, badges: [ "streak-master", "veteran" ], followers: 1233 }));
```

`isValidSchema({ username: "zara", posts: 0, verified: false, role: "user", supporter: false, badges: [] })`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidSchema({ username: "zara", posts: 0, verified: false, role: "user", supporter: false, badges: [] }));
```

`isValidSchema({ username: "nicole", posts: 65, verified: true, role: "admin", supporter: false, badges: [ "first-post", 18 ] })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "nicole", posts: 65, verified: true, role: "admin", supporter: false, badges: [ "first-post", 18 ] }));
```

`isValidSchema({ username: "tim", posts: 25, verified: true, role: "staff", supporter: false })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "tim", posts: 25, verified: true, role: "staff", supporter: false }));
```

`isValidSchema({ username: "charlie", posts: 0, verified: false, role: "user", supporter: "no", badges: [ "first-post", "anniversary" ] })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "charlie", posts: 0, verified: false, role: "user", supporter: "no", badges: [ "first-post", "anniversary" ] }));
```

`isValidSchema({ username: "wanda", posts: 15, verified: true, role: "friend", supporter: true, badges: [ "popular" ] })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "wanda", posts: 15, verified: true, role: "friend", supporter: true, badges: [ "popular" ] }));
```

`isValidSchema({ username: "guy", posts: 5, verified: "false", role: "staff", supporter: true, badges: [ "helper" ] })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "guy", posts: 5, verified: "false", role: "staff", supporter: true, badges: [ "helper" ] }));
```

`isValidSchema({ username: "carrie", verified: true, role: "moderator", supporter: true, badges: [ "helper", "sharer" ] })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: "carrie", verified: true, role: "moderator", supporter: true, badges: [ "helper", "sharer" ] }));
```

`isValidSchema({ username: true, posts: 75, verified: true, role: "creator", supporter: true, badges: [ "veteran" ] })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidSchema({ username: true, posts: 75, verified: true, role: "creator", supporter: true, badges: [ "veteran" ] }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role) &&
    (obj.supporter === undefined || typeof obj.supporter === 'boolean') &&
    Array.isArray(obj.badges) && obj.badges.every(b => typeof b === 'string')
  );
}
```
