---
id: 6a0dcd03ee4e68698080ef68
title: "도전 과제 304: 여행 일정 배치"
challengeType: 28
dashedName: challenge-304
---

# --description--

하루 여행을 위한 최소 두 개의 선택적 정지가 포함된 배열이 주어질 때, 유효한 여행 일정 배치의 수를 반환하세요.

여행 일정에는 항상 `"breakfast"`, `"lunch"`, `"dinner"`가 포함되며, 이들은 인수로 전달되지 않습니다. 선택적 정지는 다음 규칙에 따라 여행 일정 어디에나 배치할 수 있습니다:

- `"breakfast"`는 항상 첫 번째이며, `"lunch"` 이전에 최소 한 개의 정지가 있어야 합니다.
- `"lunch"`는 `"dinner"`보다 먼저 나타나야 하며, 그 사이에 최소 한 개의 정지가 있어야 합니다.
- `"dinner"` 이후에는 최대 한 개의 선택적 정지만 나타날 수 있습니다.

유효한 배치의 수를 반환하세요.

# --hints--

`getItineraryCount(["library", "park"])`는 `2`를 반환해야 합니다.

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])`는 `18`를 반환해야 합니다.

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])`는 `120`를 반환해야 합니다.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])`는 `840`를 반환해야 합니다.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])`는 `55440`를 반환해야 합니다.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
