---
id: 6a0dcd03ee4e68698080ef6b
title: "도전 과제 307: 구역 규정"
challengeType: 28
dashedName: challenge-307
---

# --description--

도시 건물 배치를 나타내는 2D 그리드(배열의 배열)가 주어질 때, 구역 규정을 위반하는 모든 건물의 좌표를 반환하세요.

그리드의 각 셀에는 아래 표에 있는 레이블 중 하나가 들어 있습니다. 건물은 수평 또는 수직으로 인접한 최대 4개의 이웃 중 허용되지 않는 유형이 있으면 위반 상태입니다.

| 레이블 | 유형 | 인접할 수 없는 유형 |
| - | - | - |
| `"i"` | 산업 | `"R"`, `"I"` |
| `"A"` | 농업 | `"C"` |
| `"R"` | 주거 | `"i"`, `"C"` |
| `"I"` | 기관 | `"i"` |
| `"C"` | 상업 | `"R"`, `"A"` |
| `""` (빈 문자열) | 미개발 | 제한 없음 |

위반하는 모든 셀의 좌표를 `[row, col]` 쌍의 배열로, 순서에 상관없이 반환하세요. 위반이 없으면 빈 배열을 반환하세요.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])`는 `[[0, 0], [0, 1]]`를 반환해야 합니다.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])`는 `[[0, 1], [1, 1]]`를 반환해야 합니다.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])`는 `[[0, 1], [0, 2], [0, 3]]`를 반환해야 합니다.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])`는 `[[2, 3], [2, 4], [3, 1], [3, 2]]`를 반환해야 합니다.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
