---
id: 6a15cadf5f240d05a264955d
title: "도전 과제 319: Frontmatter 파서"
challengeType: 28
dashedName: challenge-319
---

# --description--

frontmatter 블록을 나타내는 문자열이 주어지면, 이를 파싱해서 키와 값을 가진 객체(JavaScript) 또는 딕셔너리(Python)를 반환하세요.

Frontmatter는 `---` 구분자로 감싸져 있고, 그 안에 한 줄에 하나씩 `key: value` 쌍이 들어 있습니다. 예를 들어:

```md
---
title: My Post
draft: false
views: 100
---
```

다음과 같이 반환해야 합니다:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- 숫자, 불리언, 문자열은 각각의 타입으로 반환해야 합니다.
- 주어진 문자열은 줄바꿈 문자(`"\n"`)로 줄이 구분됩니다. 위 예시는 `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`로 주어집니다.

# --hints--

`parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")`는 `{ title: "My Post", draft: false, views: 100 }`를 반환해야 합니다.

```js
assert.deepEqual(parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---"), { title: "My Post", draft: false, views: 100 });
```

`parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")`는 `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false }`를 반환해야 합니다.

```js
assert.deepEqual(parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---"), { id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false });
```

`parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")`는 `{ version: "1.0.0", url: "https://example.com", private: true }`를 반환해야 합니다.

```js
assert.deepEqual(parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---"), { version: "1.0.0", url: "https://example.com", private: true });
```

`parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---")`는 `{ rating: 4.5, price: 9.99 }`를 반환해야 합니다.

```js
assert.deepEqual(parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---"), { rating: 4.5, price: 9.99 });
```

# --seed--

## --seed-contents--

```js
function parseFrontmatter(str) {

  return str;
}
```

# --solutions--

```js
function parseFrontmatter(str) {
  const lines = str.split("\n").filter(line => line !== "---" && line !== "");

  return lines.reduce((obj, line) => {
    const [key, ...rest] = line.split(":");
    const value = rest.join(":").trim();

    if (value === "true") obj[key.trim()] = true;
    else if (value === "false") obj[key.trim()] = false;
    else if (!isNaN(value)) obj[key.trim()] = Number(value);
    else obj[key.trim()] = value;

    return obj;
  }, {});
}
```
