---
id: 6a15cadf5f240d05a2649556
title: "도전 과제 312: 스트리밍 비용"
challengeType: 28
dashedName: challenge-312
---

# --description--

스트리밍 서비스 장바구니에 있는 영화를 나타내는 배열과 구독 등급을 나타내는 문자열이 주어지면, 영화들의 총 비용을 반환하세요.

장바구니의 각 항목은 `"format"`(`"HD"` 또는 `"4K"`)과 `"type"`(`"rent"` 또는 `"buy"`)을 가진 객체입니다. 비용은 다음과 같습니다:

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

다음 구독 등급 할인율을 적용하세요:

- `"none"`: 정가
- `"basic"`: 10% 할인
- `"premium"`: 25% 할인

총 비용을 소수점 둘째 자리까지 반올림하여 `"$D.CC"` 형식으로 반환하세요.

# --hints--

`getStreamingBill([{ format: "HD", type: "rent" }], "none")`는 `"$3.99"`를 반환해야 합니다.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }], "none"), "$3.99");
```

`getStreamingBill([{ format: "HD", type: "buy" }], "premium")`는 `"$9.74"`를 반환해야 합니다.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "buy" }], "premium"), "$9.74");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic")`는 `"$18.87"`를 반환해야 합니다.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic"), "$18.87");
```

`getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium")`는 `"$44.98"`를 반환해야 합니다.

```js
assert.equal(getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium"), "$44.98");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none")`는 `"$42.96"`를 반환해야 합니다.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none"), "$42.96");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic")`는 `"$38.66"`를 반환해야 합니다.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic"), "$38.66");
```

# --seed--

## --seed-contents--

```js
function getStreamingBill(cart, subscription) {

  return cart;
}
```

# --solutions--

```js
function getStreamingBill(cart, subscription) {
  const pricing = {
    HD: { rent: 3.99, buy: 12.99 },
    "4K": { rent: 5.99, buy: 19.99 },
  };

  const discounts = {
    none: 1,
    basic: 0.9,
    premium: 0.75,
  };

  const multiplier = discounts[subscription];

  const total = cart.reduce((sum, item) => {
    const price = pricing[item.format][item.type];
    return sum + price * multiplier;
  }, 0);

  return `$${total.toFixed(2)}`;
}
```
