---
id: 6a15cadf5f240d05a264955b
title: "도전 과제 317: BMI 계산기"
challengeType: 28
dashedName: challenge-317
---

# --description--

파운드 단위의 몸무게와 인치 단위의 키가 주어지면, 소수점 첫째 자리까지 반올림한 BMI(체질량지수)를 반환하세요.

BMI를 구하려면: 몸무게를 키의 제곱으로 나눈 후, 그 결과에 703을 곱하세요.

# --hints--

`calculateBmi(180, 70)`는 `25.8`을 반환해야 합니다.

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)`는 `24.0`을 반환해야 합니다.

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)`는 `19.5`를 반환해야 합니다.

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)`는 `39.1`을 반환해야 합니다.

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)`는 `22.8`을 반환해야 합니다.

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
