---
id: 6a15cadf5f240d05a264955a
title: "도전 과제 316: 1337 스피크"
challengeType: 28
dashedName: challenge-316
---

# --description--

소문자 문자열이 주어지면, 아래 글자를 해당하는 1337 치환 문자로 바꿔서 1337 스피크로 변환한 문자열을 반환하세요:

| 글자 | 1337 |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- 치환 문자가 없는 글자는 그대로 둡니다.

# --hints--

`makeLeet("cool")`은 `"c001"`을 반환해야 합니다.

```js
assert.equal(makeLeet("cool"), "c001");
```

`makeLeet("leet")`은 `"1337"`을 반환해야 합니다.

```js
assert.equal(makeLeet("leet"), "1337");
```

`makeLeet("hacker")`은 `"h4ck3r"`을 반환해야 합니다.

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")`은 `"547311173"`을 반환해야 합니다.

```js
assert.equal(makeLeet("satellite"), "547311173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")`은 `"4bcd3f9h1jk1mn0pqr57uvwxyz"`을 반환해야 합니다.

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return str;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", g: "9", i: "1", l: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
