---
id: 6a19b1062d1b153d8ac76d73
title: "도전 과제 325: 행운의 숫자"
challengeType: 28
dashedName: challenge-325
---

# --description--

사람의 이름과 성이 포함된 문자열이 주어지면 다음 규칙을 사용해 행운의 숫자를 계산하세요:

- 이름과 성은 공백으로 구분됩니다
- 각 이름의 모음과 자음 개수를 찾으세요
- 더 작은 모음과 자음 개수를 서로 곱한 후 더 작은 이름의 길이와 곱하세요
- 더 큰 두 개의 개수와 더 큰 이름에 대해서도 동일하게 하세요
- 더 큰 값에서 더 작은 값을 빼서 행운의 숫자를 구하세요

최종 값이 `0`이면 `13`을 반환하세요.

# --hints--

`getLuckyNumber("John Doe")`는 `21`을 반환해야 합니다.

```js
assert.equal(getLuckyNumber("John Doe"), 21);
```

`getLuckyNumber("Olivia Lewis")`는 `52`를 반환해야 합니다.

```js
assert.equal(getLuckyNumber("Olivia Lewis"), 52);
```

`getLuckyNumber("James Wilson")`는 `18`을 반환해야 합니다.

```js
assert.equal(getLuckyNumber("James Wilson"), 18);
```

`getLuckyNumber("Elizabeth Hernandez")`는 `81`을 반환해야 합니다.

```js
assert.equal(getLuckyNumber("Elizabeth Hernandez"), 81);
```

`getLuckyNumber("Mike Walker")`는 `32`를 반환해야 합니다.

```js
assert.equal(getLuckyNumber("Mike Walker"), 32);
```

`getLuckyNumber("Chloe Perez")`는 `13`을 반환해야 합니다.

```js
assert.equal(getLuckyNumber("Chloe Perez"), 13);
```

# --seed--

## --seed-contents--

```js
function getLuckyNumber(name) {

  return name;
}
```

# --solutions--

```js
function getLuckyNumber(name) {
  const vowels = "aeiouAEIOU";

  function nameStats(str) {
    const vowelCount = str.split("").filter(c => vowels.includes(c)).length;
    const consonantCount = str.split("").filter(c => c.match(/[a-zA-Z]/) && !vowels.includes(c)).length;
    return { vowelCount, consonantCount, length: str.length };
  }

  const [first, last] = name.split(" ");
  const firstStats = nameStats(first);
  const lastStats = nameStats(last);

  const largerVowels = Math.max(firstStats.vowelCount, lastStats.vowelCount);
  const smallerVowels = Math.min(firstStats.vowelCount, lastStats.vowelCount);
  const largerConsonants = Math.max(firstStats.consonantCount, lastStats.consonantCount);
  const smallerConsonants = Math.min(firstStats.consonantCount, lastStats.consonantCount);
  const largerLength = Math.max(firstStats.length, lastStats.length);
  const smallerLength = Math.min(firstStats.length, lastStats.length);

  const largerValue = largerVowels * largerConsonants * largerLength;
  const smallerValue = smallerVowels * smallerConsonants * smallerLength;
  const result = largerValue - smallerValue;

  return result === 0 ? 13 : result;
}
```
