---
id: 6a19b1062d1b153d8ac76d70
title: "도전 과제 322: Connect 3"
challengeType: 28
dashedName: challenge-322
---

# --description--

게임 그리드 위 조각을 나타내는 문자열 행렬이 주어졌을 때, 어떤 플레이어가 세 개를 연속으로 놓았는지 확인하세요.

- 각 칸에는 `"R"`, `"Y"`, 또는 `""` (빈 문자열)가 들어 있습니다.
- 세 개 연속은 가로, 세로, 대각선 방향으로 같은 종류의 빈 칸이 아닌 셀이 세 개 연속된 것을 의미합니다.

반환값:

- 승자와 그 승리한 세 칸의 좌표를 `["R", [0,2], [1,3], [2,4]]` 형식의 평탄한 배열로 반환하세요. 좌표는 위에서 아래, 왼쪽에서 오른쪽 순서로 반환합니다.
- 승자가 없으면 빈 배열을 반환하세요.

# --hints--

`connectThree([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]])`는 `["R", [3, 1], [3, 2], [3, 3]]`를 반환해야 합니다.

```js
assert.deepEqual(connectThree([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]]), ["R", [3, 1], [3, 2], [3, 3]]);
```

`connectThree([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]])`는 `["Y", [1, 1], [2, 1], [3, 1]]`를 반환해야 합니다.

```js
assert.deepEqual(connectThree([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]]), ["Y", [1, 1], [2, 1], [3, 1]]);
```

`connectThree([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]])`는 `["R", [0, 3], [1, 2], [2, 1]]`를 반환해야 합니다.

```js
assert.deepEqual(connectThree([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]]), ["R", [0, 3], [1, 2], [2, 1]]);
```

`connectThree([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]])`는 `["Y", [0, 1], [1, 2], [2, 3]]`를 반환해야 합니다.

```js
assert.deepEqual(connectThree([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]]), ["Y", [0, 1], [1, 2], [2, 3]]);
```

`connectThree([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]])`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(connectThree([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]]), []);
```

# --seed--

## --seed-contents--

```js
function connectThree(matrix) {

  return matrix;
}
```

# --solutions--

```js
function connectThree(matrix) {
  const rows = matrix.length;
  const cols = matrix[0].length;

  const directions = [[0, 1], [1, 0], [1, 1], [1, -1]];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const piece = matrix[r][c];
      if (!piece) continue;

      for (const [dr, dc] of directions) {
        const cells = [[r, c]];

        for (let i = 1; i < 3; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols) break;
          if (matrix[nr][nc] !== piece) break;
          cells.push([nr, nc]);
        }

        if (cells.length === 3) {
          cells.sort(([r1, c1], [r2, c2]) => r1 !== r2 ? r1 - r2 : c1 - c2);
          return [piece, ...cells];
        }
      }
    }
  }

  return [];
}
```
