---
id: 6a19b1062d1b153d8ac76d72
title: "도전 과제 324: 중복 문자 개수 세기"
challengeType: 28
dashedName: challenge-324
---

# --description--

두 개의 문자열이 주어지면, 두 번째 문자열에서 첫 번째 문자열에서 찾을 수 있는 문자의 개수를 반환하세요.

- 두 번째 문자열에 중복된 문자는 각각 별도로 계산하세요.

# --hints--

`duplicateCharacterCount("aloha", "hei")`는 `1`을 반환해야 합니다.

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")`는 `4`를 반환해야 합니다.

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")`는 `3`을 반환해야 합니다.

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")`는 `0`을 반환해야 합니다.

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")`는 `5`를 반환해야 합니다.

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")`는 `2`를 반환해야 합니다.

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")`는 `26`을 반환해야 합니다.

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
