---
id: 6a1d9f98e819ed70a0e994dd
title: "도전 과제 331: 가장 가까운 배수"
challengeType: 28
dashedName: challenge-331
---

# --description--

두 개의 정수가 주어지면, 첫 번째 정수를 두 번째 정수의 가장 가까운 배수로 반올림하세요.

# --hints--

`roundToNearestMultiple(5, 3)`는 `6`을 반환해야 합니다.

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)`는 `16`을 반환해야 합니다.

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)`는 `45`를 반환해야 합니다.

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)`는 `33`을 반환해야 합니다.

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)`는 `96`을 반환해야 합니다.

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
