---
id: 6a1d9f98e819ed70a0e994df
title: "도전 과제 333: 이슈 분류 2"
challengeType: 28
dashedName: challenge-333
---

# --description--

이슈 제목과 현재 레이블 배열이 주어지면, 다음 규칙에 따라 업데이트된 레이블 배열을 반환하세요:

이슈에 레이블이 없으면 다음을 추가하세요:

- 제목에 `"error"` 또는 `"bug"`가 포함되어 있으면 `"bug"`와 `"needs triage"`를 추가하세요
- 제목에 `"feature"` 또는 `"add"`가 포함되어 있으면 `"enhancement"`와 `"discussing"`을 추가하세요

그렇지 않고 주어진 레이블에 다음이 포함되어 있으면:

- `"needs triage"`가 있고 제목에 `"simple"` 또는 `"easy"`가 포함되어 있으면 `"needs triage"`를 제거하고 `"good first issue"`를 추가하세요
- `"discussing"`이 있고 제목에 `"planned"` 또는 `"next"`가 포함되어 있으면 `"discussing"`을 제거하고 `"on the roadmap"`를 추가하세요
- 그렇지 않고 `"needs triage"` 또는 `"discussing"`가 있으면 이를 제거하고 `"help wanted"`를 추가하세요

제목에 다음이 포함되어 있으면:

- `"security"`가 있으면 `"critical"` 레이블을 추가하세요

# --hints--

`triageIssue("app crashes with error", [])`는 `["bug", "needs triage"]`를 반환해야 합니다.

```js
assert.deepEqual(triageIssue("app crashes with error", []).sort(), ["bug", "needs triage"].sort());
```

`triageIssue("app crashes with error", ["bug", "needs triage"])`는 `["bug", "help wanted"]`를 반환해야 합니다.

```js
assert.deepEqual(triageIssue("app crashes with error", ["bug", "needs triage"]).sort(), ["bug", "help wanted"].sort());
```

`triageIssue("add dark mode", [])`는 `["enhancement", "discussing"]`를 반환해야 합니다.

```js
assert.deepEqual(triageIssue("add dark mode", []).sort(), ["enhancement", "discussing"].sort());
```

`triageIssue("add dark mode", ["enhancement", "discussing"])`는 `["enhancement", "help wanted"]`를 반환해야 합니다.

```js
assert.deepEqual(triageIssue("add dark mode", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted"].sort());
```

`triageIssue("xss security bug", [])`는 `["bug", "needs triage", "critical"]`를 반환해야 합니다.

```js
assert.deepEqual(triageIssue("xss security bug", []).sort(), ["bug", "needs triage", "critical"].sort());
```

`triageIssue("security vulnerability in auth", [])`는 `["critical"]`를 반환해야 합니다.

```js
assert.deepEqual(triageIssue("security vulnerability in auth", []).sort(), ["critical"].sort());
```

`triageIssue("easy a11y fix", ["bug", "needs triage"])`는 `["bug", "good first issue"]`를 반환해야 합니다.

```js
assert.deepEqual(triageIssue("easy a11y fix", ["bug", "needs triage"]).sort(), ["bug", "good first issue"].sort());
```

`triageIssue("planned api migration", ["enhancement", "discussing"])`는 `["enhancement", "on the roadmap"]`를 반환해야 합니다.

```js
assert.deepEqual(triageIssue("planned api migration", ["enhancement", "discussing"]).sort(), ["enhancement", "on the roadmap"].sort());
```

`triageIssue("improve security", ["enhancement", "discussing"])`는 `["enhancement", "help wanted", "critical"]`를 반환해야 합니다.

```js
assert.deepEqual(triageIssue("improve security", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted", "critical"].sort());
```

# --seed--

## --seed-contents--

```js
function triageIssue(title, labels) {

  return title;
}
```

# --solutions--

```js
function triageIssue(title, labels) {
  const result = [...labels];
  const t = title.toLowerCase();

  const has = label => result.includes(label);
  const add = label => { if (!has(label)) result.push(label); };
  const remove = label => { const i = result.indexOf(label); if (i > -1) result.splice(i, 1); };

  if (result.length === 0) {
    if (t.includes("error") || t.includes("bug")) {
      add("bug"); add("needs triage");
    } else if (t.includes("feature") || t.includes("add")) {
      add("enhancement"); add("discussing");
    }
  } else {
    if (has("needs triage") && (t.includes("simple") || t.includes("easy"))) {
      remove("needs triage"); add("good first issue");
    } else if (has("discussing") && (t.includes("planned") || t.includes("next"))) {
      remove("discussing"); add("on the roadmap");
    } else if (has("needs triage") || has("discussing")) {
      remove("needs triage"); remove("discussing"); add("help wanted");
    }
  }

  if (t.includes("security")) {
    add("critical");
  }

  return result;
}
```
