---
id: 6a1d9f98e819ed70a0e994d9
title: "도전 과제 327: 데이터베이스 마이그레이션"
challengeType: 28
dashedName: challenge-327
---

# --description--

두 개의 데이터베이스 객체가 주어지면, 첫 번째 객체에 없는 프로프를 두 번째 객체에 채워서 반환하세요.

- 이미 레코드에 존재하는 필드는 덮어쓰지 마세요.

# --hints--

`migrateRecord({ username: "", posts: 0 }, { verified: true })`는 `{ username: "", posts: 0, verified: true }`를 반환해야 합니다.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 })`는 `{ username: "camper", posts: 5 }`를 반환해야 합니다.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 }), { username: "camper", posts: 5 });
```

`migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" })`는 `{ username: "camper", posts: 0, verified: false }`를 반환해야 합니다.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" }), { username: "camper", posts: 0, verified: false });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" })`는 `{ username: "camper", role: "admin", posts: 0 }`를 반환해야 합니다.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" }), { username: "camper", role: "admin", posts: 0 });
```

`migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" })`는 `{ username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false }`를 반환해야 합니다.

```js
assert.deepEqual(migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" }), { username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false });
```

# --seed--

## --seed-contents--

```js
function migrateRecord(schema, record) {

  return schema;
}
```

# --solutions--

```js
function migrateRecord(schema, record) {
  return { ...schema, ...record };
}
```
