---
id: 6a1d9f98e819ed70a0e994e0
title: "도전 과제 334: 정확한 거스름돈"
challengeType: 28
dashedName: challenge-334
---

# --description--

정수 센트 금액이 주어지면, 페니(1센트), 니켈(5센트), 다임(10센트), 쿼터(25센트)를 사용해 정확한 거스름돈을 만들 수 있는 서로 다른 방법의 수를 반환하세요.

# --hints--

`exactChange(3)`는 `1`을 반환해야 합니다.

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)`는 `2`를 반환해야 합니다.

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)`는 `6`을 반환해야 합니다.

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)`는 `24`를 반환해야 합니다.

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)`는 `73`을 반환해야 합니다.

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)`는 `213`을 반환해야 합니다.

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
