---
id: 6a1d9f98e819ed70a0e994dc
title: "도전 과제 330: 소문자 단어"
challengeType: 28
dashedName: challenge-330
---

# --description--

문자열이 주어지면, 완전히 소문자인 단어만 원래 순서대로 각 단어 사이에 공백을 두고 반환하세요.

# --hints--

`getLowercaseWords("hello GOOD world")`는 `"hello world"`를 반환해야 합니다.

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")`는 `"these are all lowercase"`를 반환해야 합니다.

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")`는 `"less is more"`를 반환해야 합니다.

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")`는 `"eat pizza every day"`를 반환해야 합니다.

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")`는 `"the quick brown fox jumped over the lazy dog"`를 반환해야 합니다.

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
