---
id: 6a1d9f98e819ed70a0e994e1
title: "도전 과제 335: 다섯 개 주사위"
challengeType: 28
dashedName: challenge-335
---

# --description--

값이 1-6인 다섯 개 주사위 배열이 주어지면, 가능한 최고의 족보를 반환하세요.

족보는 낮은 순서부터 높은 순서로 다음과 같습니다:

| 족보 | 설명 |
|------|-------------|
| `"no pair"` | 한 쌍도 없는 경우 또는 그 이상 |
| `"pair"` | 같은 값의 주사위 두 개 |
| `"two pair"` | 서로 다른 두 쌍 |
| `"three of a kind"` | 같은 값의 주사위 세 개 |
| `"small straight"` | 연속된 네 개 값 |
| `"large straight"` | 연속된 다섯 개 값 |
| `"full house"` | 세 개 같은 값과 한 쌍 |
| `"four of a kind"` | 같은 값의 주사위 네 개 |
| `"five of a kind"` | 같은 값의 주사위 다섯 개 |

# --hints--

`fiveDice([1, 1, 1, 1, 1])`는 `"five of a kind"`를 반환해야 합니다.

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])`는 `"four of a kind"`를 반환해야 합니다.

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])`는 `"large straight"`를 반환해야 합니다.

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])`는 `"three of a kind"`를 반환해야 합니다.

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])`는 `"pair"`를 반환해야 합니다.

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])`는 `"no pair"`를 반환해야 합니다.

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])`는 `"small straight"`를 반환해야 합니다.

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])`는 `"full house"`를 반환해야 합니다.

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])`는 `"two pair"`를 반환해야 합니다.

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
