---
id: 6a2037a68a0bc2aef0006003
title: "도전 과제 341: 생일 카운트다운"
challengeType: 28
dashedName: challenge-341
---

# --description--

오늘 날짜와 생일이 주어지면, 그 사람의 다음 생일까지 남은 일수를 반환하세요.

- 오늘 날짜는 `"YYYY-MM-DD"` 형식의 문자열로, 앞에 0이 붙어 있으며 예를 들어 `"2026-07-16"`입니다.
- 생일은 `"M/D"` 형식의 문자열로, 앞에 0이 붙지 않으며 예를 들어 `"9/7"`입니다.
- 오늘이 생일이라면 다음 생일까지 남은 일수를 반환하세요(`0`이 아닙니다).
- 윤년도 고려해야 합니다.

# --hints--

`daysUntilBirthday("2026-07-16", "9/7")`는 `53`을 반환해야 합니다.

```js
assert.equal(daysUntilBirthday("2026-07-16", "9/7"), 53);
```

`daysUntilBirthday("2026-07-16", "3/22")`는 `249`를 반환해야 합니다.

```js
assert.equal(daysUntilBirthday("2026-07-16", "3/22"), 249);
```

`daysUntilBirthday("2026-07-16", "7/16")`는 `365`를 반환해야 합니다.

```js
assert.equal(daysUntilBirthday("2026-07-16", "7/16"), 365);
```

`daysUntilBirthday("2024-02-28", "3/1")`는 `2`를 반환해야 합니다.

```js
assert.equal(daysUntilBirthday("2024-02-28", "3/1"), 2);
```

`daysUntilBirthday("2023-04-24", "12/30")`는 `250`을 반환해야 합니다.

```js
assert.equal(daysUntilBirthday("2023-04-24", "12/30"), 250);
```

`daysUntilBirthday("2024-03-01", "2/29")`는 `1460`을 반환해야 합니다.

```js
assert.equal(daysUntilBirthday("2024-03-01", "2/29"), 1460);
```

`daysUntilBirthday("2096-03-01", "2/29")`는 `2920`을 반환해야 합니다.

```js
assert.equal(daysUntilBirthday("2096-03-01", "2/29"), 2920);
```

# --seed--

## --seed-contents--

```js
function daysUntilBirthday(today, birthday) {

  return today;
}
```

# --solutions--

```js
function daysUntilBirthday(today, birthday) {
  const [year, month, day] = today.split("-").map(Number);
  const [bMonth, bDay] = birthday.split("/").map(Number);

  const todayDate = new Date(year, month - 1, day);

  function isLeapYear(y) {
    return (y % 4 === 0 && y % 100 !== 0) || y % 400 === 0;
  }

  function nextValidBirthday(fromYear) {
    for (let y = fromYear; y <= fromYear + 8; y++) {
      if (bMonth === 2 && bDay === 29 && !isLeapYear(y)) continue;
      const date = new Date(y, bMonth - 1, bDay);
      if (date > todayDate) return date;
    }
  }

  const nextBirthday = nextValidBirthday(year);
  const diff = nextBirthday - todayDate;
  return Math.round(diff / (1000 * 60 * 60 * 24));
}
```
