---
id: 6a2037a68a0bc2aef0006002
title: "도전 과제 340: 피그 라틴 변환기"
challengeType: 28
dashedName: challenge-340
---

# --description--

문자열이 주어지면 다음 규칙을 사용해 피그 라틴으로 변환하세요:

- 단어가 모음(`"a"`, `"e"`, `"i"`, `"o"`, 또는 `"u"`)으로 시작하면 끝에 `"way"`를 추가하세요. 예를 들어, `"universe"`는 `"universeway"`로 변환됩니다.
- 단어가 하나 이상의 자음으로 시작하면 그 자음들을 끝으로 옮기고 `"ay"`를 추가하세요. 예를 들어, `"hello"`는 `"ellohay"`로 변환됩니다.
- 첫 글자의 대소문자를 유지하세요. 예를 들어, `"Hello"`는 `"Ellohay"`로 변환됩니다.

# --hints--

`pigLatin("universe")`는 `"universeway"`를 반환해야 합니다.

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")`는 `"ellohay"`를 반환해야 합니다.

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")`는 `"ellohay universeway"`를 반환해야 합니다.

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")`는 `"Ellohay universeway"`를 반환해야 합니다.

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")`는 `"Igpay Atinlay isway unfay"`를 반환해야 합니다.

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")`는 `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`를 반환해야 합니다.

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
