---
id: 6a2037a68a0bc2aef0006004
title: "도전 과제 342: 주사위 확률"
challengeType: 28
dashedName: challenge-342
---

# --description--

여섯 면체 주사위 여러 개를 굴리고 목표 합계가 주어질 때, 그 합계가 나올 확률을 `"1 in X"` 형식의 문자열로 반환하세요.

- 주사위 개수는 1개에서 6개 사이입니다.
- 주어진 주사위 개수로 목표 합계는 항상 나올 수 있습니다.
- `"X"`는 가장 가까운 정수로 반올림하세요.

# --hints--

`getOdds(1, 5)`는 `"1 in 6"`을 반환해야 합니다.

```js
assert.equal(getOdds(1, 5), "1 in 6");
```

`getOdds(2, 4)`는 `"1 in 12"`를 반환해야 합니다.

```js
assert.equal(getOdds(2, 4), "1 in 12");
```

`getOdds(3, 10)`는 `"1 in 8"`을 반환해야 합니다.

```js
assert.equal(getOdds(3, 10), "1 in 8");
```

`getOdds(4, 7)`는 `"1 in 65"`를 반환해야 합니다.

```js
assert.equal(getOdds(4, 7), "1 in 65");
```

`getOdds(5, 26)`는 `"1 in 111"`를 반환해야 합니다.

```js
assert.equal(getOdds(5, 26), "1 in 111");
```

`getOdds(6, 35)`는 `"1 in 7776"`를 반환해야 합니다.

```js
assert.equal(getOdds(6, 35), "1 in 7776");
```

# --seed--

## --seed-contents--

```js
function getOdds(dice, target) {

  return dice;
}
```

# --solutions--

```js
function getOdds(dice, target) {
  let validCombos = 0;
  const totalCombos = Math.pow(6, dice);

  function roll(diceLeft, sum) {
    if (diceLeft === 0) {
      if (sum === target) validCombos++;
      return;
    }
    for (let face = 1; face <= 6; face++) {
      roll(diceLeft - 1, sum + face);
    }
  }

  roll(dice, 0);
  return `1 in ${Math.round(totalCombos / validCombos)}`;
}
```
