---
id: 6a2037a68a0bc2aef0006008
title: "도전 과제 346: 돼지 저금통"
challengeType: 28
dashedName: challenge-346
---

# --description--

돼지 저금통을 나타내는 객체가 주어지면, 총 가치를 `"$D.CC"` 형식의 문자열로 반환하세요.

객체는 다음 중 어느 것이든 포함할 수 있습니다:

| 동전 | 가치 |
|------|-------|
| `pennies` | $0.01 |
| `nickels` | $0.05 |
| `dimes` | $0.10 |
| `quarters` | $0.25 |

# --hints--

`piggyBank({ pennies: 3, nickels: 5, dimes: 2, quarters: 6 })`는 `"$1.98"`를 반환해야 합니다.

```js
assert.equal(piggyBank({ pennies: 3, nickels: 5, dimes: 2, quarters: 6 }), "$1.98");
```

`piggyBank({ pennies: 1, nickels: 1, dimes: 1, quarters: 1 })`는 `"$0.41"`를 반환해야 합니다.

```js
assert.equal(piggyBank({ pennies: 1, nickels: 1, dimes: 1, quarters: 1 }), "$0.41");
```

`piggyBank({ nickels: 8, dimes: 6, quarters: 5 })`는 `"$2.25"`를 반환해야 합니다.

```js
assert.equal(piggyBank({ nickels: 8, dimes: 6, quarters: 5 }), "$2.25");
```

`piggyBank({})`는 `"$0.00"`를 반환해야 합니다.

```js
assert.equal(piggyBank({}), "$0.00");
```

`piggyBank({ pennies: 146, nickels: 11, dimes: 0, quarters: 19 })`는 `"$6.76"`를 반환해야 합니다.

```js
assert.equal(piggyBank({ pennies: 146, nickels: 11, dimes: 0, quarters: 19 }), "$6.76");
```

# --seed--

## --seed-contents--

```js
function piggyBank(coins) {

  return coins;
}
```

# --solutions--

```js
function piggyBank(coins) {
  const values = {
    pennies: 0.01,
    nickels: 0.05,
    dimes: 0.10,
    quarters: 0.25,
  };

  const total = Object.entries(coins).reduce((sum, [coin, count]) => {
    return sum + (values[coin] * count);
  }, 0);

  return "$" + total.toFixed(2);
}
```
