---
id: 6a2037a68a0bc2aef0006007
title: "도전 과제 345: 단어 블렌더"
challengeType: 28
dashedName: challenge-345
---

# --description--

두 단어가 주어지면 첫 번째 단어의 앞부분과 두 번째 단어의 뒷부분을 합쳐 새 단어를 반환하세요.

- 홀수 길이 단어의 경우, 앞부분이 더 짧은 절반입니다.

# --hints--

`blendWords("turtle", "toucan")`는 `"turcan"`을 반환해야 합니다.

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")`는 `"chipingo"`를 반환해야 합니다.

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")`는 `"falican"`을 반환해야 합니다.

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")`는 `"hyana"`를 반환해야 합니다.

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")`는 `"scorilla"`를 반환해야 합니다.

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")`는 `"platerine"`을 반환해야 합니다.

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
