---
id: 6a2037a68a0bc2aef0005fff
title: "도전 과제 337: 집계 카운터"
challengeType: 28
dashedName: challenge-337
---

# --description--

집계 표시 문자열이 주어지면, 나타내는 총 개수를 반환하세요.

- 각 파이프 `"|"`는 하나의 개수를 나타내요.
- 다섯 번째 표시는 슬래시 `"/"`로 나타내며, 다섯 개 그룹을 완성해요 (`"||||/"`).
- 그룹들은 공백으로 구분돼요.

# --hints--

`getTallyCount("||||")`는 `4`를 반환해야 해요.

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")`는 `5`를 반환해야 해요.

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")`는 `8`을 반환해야 해요.

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")`는 `17`을 반환해야 해요.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")`는 `41`을 반환해야 해요.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
