---
id: 6a2037a68a0bc2aef0006006
title: "도전 과제 344: 황금비"
challengeType: 28
dashedName: challenge-344
---

# --description--

두 숫자가 주어졌을 때, 그 비율이 황금비에 근접하는지 판단하세요.

- `1.618`의 황금비를 사용하세요
- `0.01`의 허용 오차를 허용하세요

# --hints--

`isGoldenRatio(21, 34)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isGoldenRatio(21, 34));
```

`isGoldenRatio(15, 20)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isGoldenRatio(15, 20));
```

`isGoldenRatio(8, 13)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isGoldenRatio(8, 13));
```

`isGoldenRatio(10, 16)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isGoldenRatio(10, 16));
```

`isGoldenRatio(1618, 1000)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

`isGoldenRatio(88, 55)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
