---
id: 6a2037a68a0bc2aef0006005
title: "도전 과제 343: 엘리베이터 정지"
challengeType: 28
dashedName: challenge-343
---

# --description--

엘리베이터의 현재 층 수와 요청된 층의 배열이 주어지면, 이동하는 층 수를 최소화하도록 엘리베이터가 방문해야 할 순서의 배열을 반환하세요.

- 동률일 경우 먼저 위로 이동하세요
- 요청된 층은 엘리베이터가 처음 지나갈 때 반드시 방문해야 합니다

# --hints--

`elevatorStops(5, [2, 8, 3, 9])`는 `[3, 2, 8, 9]`를 반환해야 합니다.

```js
assert.deepEqual(elevatorStops(5, [2, 8, 3, 9]), [3, 2, 8, 9]);
```

`elevatorStops(6, [2, 10, 8, 3, 1, 9])`는 `[8, 9, 10, 3, 2, 1]`를 반환해야 합니다.

```js
assert.deepEqual(elevatorStops(6, [2, 10, 8, 3, 1, 9]), [8, 9, 10, 3, 2, 1]);
```

`elevatorStops(1, [4, 8, 3, 6, 9])`는 `[3, 4, 6, 8, 9]`를 반환해야 합니다.

```js
assert.deepEqual(elevatorStops(1, [4, 8, 3, 6, 9]), [3, 4, 6, 8, 9]);
```

`elevatorStops(12, [6, 10, 7, 3, 1, 4])`는 `[10, 7, 6, 4, 3, 1]`를 반환해야 합니다.

```js
assert.deepEqual(elevatorStops(12, [6, 10, 7, 3, 1, 4]), [10, 7, 6, 4, 3, 1]);
```

`elevatorStops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19])`는 `[10, 9, 8, 6, 5, 2, 12, 19, 23]`를 반환해야 합니다.

```js
assert.deepEqual(elevatorStops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19]), [10, 9, 8, 6, 5, 2, 12, 19, 23]);
```

# --seed--

## --seed-contents--

```js
function elevatorStops(currentFloor, stops) {

  return currentFloor;
}
```

# --solutions--

```js
function elevatorStops(currentFloor, stops) {
  stops = stops.sort((a, b) => a - b);
  const above = stops.filter(s => s > currentFloor);
  const below = stops.filter(s => s < currentFloor).reverse();
  const upDistance = above.length ? above[above.length - 1] - currentFloor : Infinity;
  const downDistance = below.length ? currentFloor - below[below.length - 1] : Infinity;
  return downDistance < upDistance ? [...below, ...above] : [...above, ...below];
}
```
