---
id: 6a22d77ddf034bc4e35b1d57
title: "도전 과제 350: 문자 거리"
challengeType: 28
dashedName: challenge-350
---

# --description--

길이가 같은 두 문자열이 주어지면 각 문자 쌍 사이의 가장 짧은 거리들의 합을 반환하세요.

- 입력은 소문자만 포함합니다
- 알파벳은 원형으로 취급하므로 `a`와 `z` 사이 거리는 1입니다.

# --hints--

`letterDistance("abc", "bcd")`는 `3`을 반환해야 합니다.

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")`는 `9`를 반환해야 합니다.

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")`는 `10`을 반환해야 합니다.

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")`는 `43`을 반환해야 합니다.

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")`는 `47`을 반환해야 합니다.

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")`는 `55`를 반환해야 합니다.

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
