---
id: 6a22d77ddf034bc4e35b1d5c
title: "도전 과제 355: 모스 부호"
challengeType: 28
dashedName: challenge-355
---

# --description--

모스 부호 문자열이 주어지면, 다음 표를 사용해 해독된 메시지를 반환하세요:

| 부호 | 문자 | 부호 | 문자 |
|------|--------|------|--------|
| `.-` | `A` | `-.` | `N` |
| `-...` | `B` | `---` | `O` |
| `-.-.` | `C` | `.--.` | `P` |
| `-..` | `D` | `--.-` | `Q` |
| `.` | `E` | `.-.` | `R` |
| `..-.` | `F` | `...` | `S` |
| `--.` | `G` | `-` | `T` |
| `....` | `H` | `..-` | `U` |
| `..` | `I` | `...-` | `V` |
| `.---` | `J` | `.--` | `W` |
| `-.-` | `K` | `-..-` | `X` |
| `.-..` | `L` | `-.--` | `Y` |
| `--` | `M` | `--..` | `Z` |

- 글자는 한 칸 띄어쓰기로 구분합니다
- 단어는 세 칸 띄어쓰기로 구분합니다

# --hints--

`decodeMorse("--..")`는 `"Z"`를 반환해야 합니다.

```js
assert.equal(decodeMorse("--.."), "Z");
```

`decodeMorse("... --- ...")`는 `"SOS"`를 반환해야 합니다.

```js
assert.equal(decodeMorse("... --- ..."), "SOS");
```

`decodeMorse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--.")`는 `"FREECODECAMP"`를 반환해야 합니다.

```js
assert.equal(decodeMorse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--."), "FREECODECAMP");
```

`decodeMorse(".... . .-.. .-.. ---   .-- --- .-. .-.. -..")`는 `"HELLO WORLD"`를 반환해야 합니다.

```js
assert.equal(decodeMorse(".... . .-.. .-.. ---   .-- --- .-. .-.. -.."), "HELLO WORLD");
```

`decodeMorse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --.")`는 `"THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG"`를 반환해야 합니다.

```js
assert.equal(decodeMorse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --."), "THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG");
```

# --seed--

## --seed-contents--

```js
function decodeMorse(code) {

  return code;
}
```

# --solutions--

```js
function decodeMorse(code) {
  const table = {
    '.-': 'A', '-...': 'B', '-.-.': 'C', '-..': 'D', '.': 'E',
    '..-.': 'F', '--.': 'G', '....': 'H', '..': 'I', '.---': 'J',
    '-.-': 'K', '.-..': 'L', '--': 'M', '-.': 'N', '---': 'O',
    '.--.': 'P', '--.-': 'Q', '.-.': 'R', '...': 'S', '-': 'T',
    '..-': 'U', '...-': 'V', '.--': 'W', '-..-': 'X', '-.--': 'Y',
    '--..': 'Z'
  };

  return code.split('   ')
    .map(word => word.split(' ').map(c => table[c]).join(''))
    .join(' ');
}
```
