---
id: 6a22d77ddf034bc4e35b1d58
title: "도전 과제 351: Pronic 수"
challengeType: 28
dashedName: challenge-351
---

# --description--

숫자가 주어지면 그것이 pronic 수인지 판단하세요.

pronic 수는 두 연속된 정수의 곱입니다. 예를 들어, 6은 2 * 3 = 6이므로 pronic 수입니다.

# --hints--

`isPronic(6)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
