---
id: 6a22d77ddf034bc4e35b1d5a
title: "도전 과제 353: 명암 대비율 2"
challengeType: 28
dashedName: challenge-353
---

# --description--

두 개의 상대 휘도 값과 텍스트가 큰지 여부를 나타내는 불리언이 주어지면, 다음 방법을 사용해 WCAG 명암 대비 등급을 반환하세요:

각 휘도 값에 0.05를 더한 후, 더 밝은 값을 더 어두운 값으로 나누어 명암 대비율을 계산하세요. 더 밝은 값은 항상 첫 번째 인수입니다.

다음 표를 사용해 명암 대비율에 따른 등급을 반환하세요:

| 등급 | 일반 텍스트 | 큰 텍스트 |
|--------|-------------|------------|
| `"AAA"` | 7.0 이상 | 4.5 이상 |
| `"AA"` | 4.5 이상 | 3.0 이상 |
| `"Fail"` | 4.5 미만 | 3.0 미만 |

# --hints--

`getContrastRating(1.0, 0.0, false)`는 `"AAA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating(1.0, 0.0, false), "AAA");
```

`getContrastRating(0.9015, 0.1364, false)`는 `"AA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating(0.9015, 0.1364, false), "AA");
```

`getContrastRating(0.8965, 0.1628, false)`는 `"Fail"`을 반환해야 합니다.

```js
assert.equal(getContrastRating(0.8965, 0.1628, false), "Fail");
```

`getContrastRating(0.7469, 0.0957, true)`는 `"AAA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating(0.7469, 0.0957, true), "AAA");
```

`getContrastRating(0.7489, 0.2018, true)`는 `"AA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating(0.7489, 0.2018, true), "AA");
```

`getContrastRating(0.6571, 0.1974, true)`는 `"Fail"`을 반환해야 합니다.

```js
assert.equal(getContrastRating(0.6571, 0.1974, true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(l1, l2, isLargeText) {

  return l1;
}
```

# --solutions--

```js
function getContrastRating(l1, l2, isLargeText) {
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
