---
id: 6a22d77ddf034bc4e35b1d5d
title: "도전 과제 356: 마방진 해결기"
challengeType: 28
dashedName: challenge-356
---

# --description--

하나의 숫자가 빠진 3x3 격자(`0`으로 나타냄)가 주어지면, 마방진을 완성하는 빠진 숫자를 반환하세요. 유효한 숫자가 없으면 `"impossible"`을 반환하세요.

마방진은 모든 행, 열, 대각선의 합이 같은 숫자인 격자입니다.

# --hints--

`solveMagicSquare([[2, 7, 6], [9, 0, 1], [4, 3, 8]])`는 `5`를 반환해야 합니다.

```js
assert.equal(solveMagicSquare([[2, 7, 6], [9, 0, 1], [4, 3, 8]]), 5);
```

`solveMagicSquare([[0, 14, 12], [18, 10, 2], [8, 6, 16]])`는 `4`를 반환해야 합니다.

```js
assert.equal(solveMagicSquare([[0, 14, 12], [18, 10, 2], [8, 6, 16]]), 4);
```

`solveMagicSquare([[12, 17, 16], [19, 0, 10], [14, 13, 18]])`는 `"impossible"`을 반환해야 합니다.

```js
assert.equal(solveMagicSquare([[12, 17, 16], [19, 0, 10], [14, 13, 18]]), "impossible");
```

`solveMagicSquare([[15, 35, 31], [43, 27, 11], [23, 19, 0]])`는 `39`를 반환해야 합니다.

```js
assert.equal(solveMagicSquare([[15, 35, 31], [43, 27, 11], [23, 19, 0]]), 39);
```

`solveMagicSquare([[26, 41, 14], [47, 35, 0], [32, 29, 44]])`는 `"impossible"`을 반환해야 합니다.

```js
assert.equal(solveMagicSquare([[26, 41, 14], [47, 35, 0], [32, 29, 44]]), "impossible");
```

# --seed--

## --seed-contents--

```js
function solveMagicSquare(grid) {

  return grid;
}
```

# --solutions--

```js
function solveMagicSquare(grid) {
  const flat = grid.flat();
  const zeroIndex = flat.indexOf(0);
  const row = Math.floor(zeroIndex / 3);
  const col = zeroIndex % 3;

  const target = grid.find(r => !r.includes(0))
    .reduce((sum, n) => sum + n, 0);

  const missing = target - grid[row].reduce((sum, n) => sum + n, 0);
  const filled = grid.map((r, ri) =>
    r.map((n, ci) => (ri === row && ci === col) ? missing : n)
  );

  const check = (arr) => arr.reduce((sum, n) => sum + n, 0) === target;

  for (let i = 0; i < 3; i++) {
    if (!check(filled[i])) return "impossible";
    if (!check(filled.map(r => r[i]))) return "impossible";
  }
  if (!check([filled[0][0], filled[1][1], filled[2][2]])) return "impossible";
  if (!check([filled[0][2], filled[1][1], filled[2][0]])) return "impossible";

  return missing;
}
```
