---
id: 6a22d77ddf034bc4e35b1d54
title: "도전 과제 347: 게임 이론"
challengeType: 28
dashedName: challenge-347
---

# --description--

두 플레이어의 전략을 나타내는 길이가 같은 두 문자열이 주어지면, 점수를 배열 `[player1, player2]`로 반환하세요.

- 주어진 문자열은 `"C"` (협력) 또는 `"D"` (배신) 두 글자 중 하나만 포함합니다.
- 각 문자는 한 라운드를 나타내며, 점수는 다음과 같이 계산합니다:
  - 두 플레이어가 모두 협력하면 각자 3점을 얻습니다.
  - 두 플레이어가 모두 배신하면 각자 1점을 얻습니다.
  - 한 플레이어가 배신하고 다른 플레이어가 협력하면, 배신자는 5점, 협력자는 0점을 얻습니다.

# --hints--

`playGame("CCCC", "CCCC")`는 `[12, 12]`를 반환해야 합니다.

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")`는 `[4, 4]`를 반환해야 합니다.

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")`는 `[5, 10]`를 반환해야 합니다.

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")`는 `[24, 34]`를 반환해야 합니다.

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")`는 `[66, 21]`를 반환해야 합니다.

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
