---
id: 6a22d77ddf034bc4e35b1d59
title: "도전 과제 352: 명암 대비율 평가 1"
challengeType: 28
dashedName: challenge-352
---

# --description--

명암 대비율과 텍스트가 큰지 여부를 나타내는 불리언이 주어지면, 다음 표를 사용해 WCAG 평가를 반환하세요:

| 평가 | 일반 텍스트 | 큰 텍스트 |
|--------|-------------|------------|
| `"AAA"` | 7.0 이상 | 4.5 이상 |
| `"AA"` | 4.5 이상 | 3.0 이상 |
| `"Fail"` | 4.5 미만 | 3.0 미만 |

# --hints--

`getContrastRating("7.5", false)`는 `"AAA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating("7.5", false), "AAA");
```

`getContrastRating("4.8", false)`는 `"AA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating("4.8", false), "AA");
```

`getContrastRating("4.2", false)`는 `"Fail"`을 반환해야 합니다.

```js
assert.equal(getContrastRating("4.2", false), "Fail");
```

`getContrastRating("4.5", true)`는 `"AAA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating("4.5", true), "AAA");
```

`getContrastRating("3.0", true)`는 `"AA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating("3.0", true), "AA");
```

`getContrastRating("2.7", false)`는 `"Fail"`을 반환해야 합니다.

```js
assert.equal(getContrastRating("2.7", false), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(ratio, isLargeText) {

  return ratio;
}
```

# --solutions--

```js
function getContrastRating(ratio, isLargeText) {
  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
