---
id: 6a22d77ddf034bc4e35b1d5b
title: "도전 과제 354: 명암 등급 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

RGB 값을 나타내는 두 개의 배열과 텍스트가 큰지 여부를 나타내는 불리언이 주어지면, 다음 방법을 사용해 WCAG 명암 등급을 반환하세요:

먼저 각 RGB 값을 상대 휘도로 변환하세요:

- 각 채널 `[R, G, B]`를 255로 나누어 `0`과 `1` 사이의 값을 얻으세요
- 각 채널에 감마 보정 공식을 적용하세요:
  - 채널 값이 `0.04045` 이하이면: `channel / 12.92`
  - 그렇지 않으면: `((channel + 0.055) / 1.055) ^ 2.4`
- 휘도 계산: `0.2126 * R + 0.7152 * G + 0.0722 * B`

그다음, 각 휘도 값에 0.05를 더한 후, 밝은 쪽 값을 어두운 쪽 값으로 나누어 명암 대비율을 계산하세요. 밝은 쪽 값이 항상 첫 번째 인수입니다.

다음 표를 사용해 명암 대비율에 따른 등급을 반환하세요:

| 등급 | 일반 텍스트 | 큰 텍스트 |
|--------|-------------|------------|
| `"AAA"` | 7.0 이상 | 4.5 이상 |
| `"AA"` | 4.5 이상 | 3.0 이상 |
| `"Fail"` | 4.5 미만 | 3.0 미만 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)`는 `"AAA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)`는 `"AA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)`는 `"Fail"`을 반환해야 합니다.

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)`는 `"AAA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)`는 `"AA"`를 반환해야 합니다.

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)`는 `"Fail"`을 반환해야 합니다.

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
