---
id: 6a26df3e2988bcdded204894
title: "도전 과제 358: 이모지 번역기"
challengeType: 28
dashedName: challenge-358
---

# --description--

이모지 문자열이 주어지면, 다음 표를 사용해 구문을 반환하세요:

| Emoji | Word |
|-------|------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

단어들은 공백으로 구분해 반환하세요.

# --hints--

`getEmojiPhrase("🪨⭐")`는 `"rock star"`를 반환해야 합니다.

```js
assert.equal(getEmojiPhrase("🪨⭐"), "rock star");
```

`getEmojiPhrase("🥵🐕")`는 `"hot dog"`를 반환해야 합니다.

```js
assert.equal(getEmojiPhrase("🥵🐕"), "hot dog");
```

`getEmojiPhrase("👶🦈")`는 `"baby shark"`를 반환해야 합니다.

```js
assert.equal(getEmojiPhrase("👶🦈"), "baby shark");
```

`getEmojiPhrase("⭐🐟")`는 `"star fish"`를 반환해야 합니다.

```js
assert.equal(getEmojiPhrase("⭐🐟"), "star fish");
```

`getEmojiPhrase("🧊🧊👶")`는 `"ice ice baby"`를 반환해야 합니다.

```js
assert.equal(getEmojiPhrase("🧊🧊👶"), "ice ice baby");
```

`getEmojiPhrase("🐱🐟🍲")`는 `"cat fish soup"`를 반환해야 합니다.

```js
assert.equal(getEmojiPhrase("🐱🐟🍲"), "cat fish soup");
```

# --seed--

## --seed-contents--

```js
function getEmojiPhrase(str) {

  return str;
}
```

# --solutions--

```js
function getEmojiPhrase(str) {
  const table = {
    '👶': 'baby',
    '🐱': 'cat',
    '🐕': 'dog',
    '🐟': 'fish',
    '🥵': 'hot',
    '🧊': 'ice',
    '🪨': 'rock',
    '🦈': 'shark',
    '🍲': 'soup',
    '⭐': 'star',
  };

  return [...str].map(emoji => table[emoji]).join(' ');
}
```
