---
id: 6a26df95efa55a2524399743
title: "도전 과제 362: 논노그램 검증기"
challengeType: 28
dashedName: challenge-362
---

# --description--

단서 숫자 배열과 셀 배열이 주어졌을 때, 셀이 논노그램 단서를 만족하는지 판단하세요.

- 단서는 연속된 채워진 셀의 길이를 순서대로 나타내는 숫자 배열입니다. 예를 들어, 단서가 `[3, 2]`라면 3개의 연속된 채워진 셀 다음에 최소 하나의 빈 셀로 구분된 2개의 연속된 채워진 셀이 있어야 합니다.
- 행은 1(채워짐)과 0(빈 셀)로 이루어진 배열입니다.

# --hints--

`isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1])`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1]));
```

`isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1])`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1]));
```

`isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]));
```

`isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])`는 `true`를 반환해야 합니다.

```js
assert.isTrue(isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])`는 `false`를 반환해야 합니다.

```js
assert.isFalse(isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]));
```

# --seed--

## --seed-contents--

```js
function isValidNonogram(clue, cells) {

  return clue;
}
```

# --solutions--

```js
function isValidNonogram(clue, cells) {
  const runs = [];
  let count = 0;

  for (let i = 0; i <= cells.length; i++) {
    if (cells[i] === 1) {
      count++;
    } else if (count > 0) {
      runs.push(count);
      count = 0;
    }
  }

  if (runs.length !== clue.length) return false;
  return runs.every((run, i) => run === clue[i]);
}
```
