---
id: 6a26df95efa55a2524399745
title: "도전 과제 364: 두 양동이 사이"
challengeType: 28
dashedName: challenge-364
---

# --description--

RGB 색상과 가득 찬 정도가 각각 있는 두 양동이의 페인트가 주어지면, 혼합된 RGB 색상을 세 개의 정수로 된 배열로 반환하세요.

- 각 양동이는 `color` 속성(세 개의 정수로 된 배열 `[r, g, b]`)과 `fullness` 속성(0–100)을 가진 객체(JavaScript) 또는 딕셔너리(Python)입니다.
- 혼합된 색상은 두 색상의 각 채널을 가득 찬 정도에 따라 가중 평균한 값이며, 각 채널은 가장 가까운 정수로 반올림합니다.

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })`는 `[125, 125, 125]`를 반환해야 합니다.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })`는 `[200, 200, 200]`를 반환해야 합니다.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })`는 `[100, 150, 200]`를 반환해야 합니다.

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })`는 `[114, 184, 170]`를 반환해야 합니다.

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })`는 `[68, 163, 123]`를 반환해야 합니다.

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
