---
id: 6a26df3e2988bcdded204893
title: "도전 과제 357: 먹이 사슬"
challengeType: 28
dashedName: challenge-357
---

# --description--

`[predator, prey]` 쌍으로 이루어진 배열이 주어지면, 최상위 포식자부터 최하위까지 먹이 사슬을 반환하세요.

- 최상위 포식자는 다른 동물의 먹이가 되지 않는 동물입니다.
- 먹이 사슬은 문자열 배열로 반환하세요.

# --hints--

`getFoodChain([["cat", "mouse"]])`는 `["cat", "mouse"]`를 반환해야 합니다.

```js
assert.deepEqual(getFoodChain([["cat", "mouse"]]), ["cat", "mouse"]);
```

`getFoodChain([["wolf", "deer"], ["deer", "grass"]])`는 `["wolf", "deer", "grass"]`를 반환해야 합니다.

```js
assert.deepEqual(getFoodChain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"]);
```

`getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])`는 `["hawk", "snake", "frog", "fly"]`를 반환해야 합니다.

```js
assert.deepEqual(getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"]);
```

`getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])`는 `["eagle", "fox", "rabbit", "grass"]`를 반환해야 합니다.

```js
assert.deepEqual(getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"]);
```

`getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])`는 `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`를 반환해야 합니다.

```js
assert.deepEqual(getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"]);
```

# --seed--

## --seed-contents--

```js
function getFoodChain(pairs) {

  return pairs;
}
```

# --solutions--

```js
function getFoodChain(pairs) {
  const prey = new Set(pairs.map(([, p]) => p));
  const map = Object.fromEntries(pairs);
  const apex = pairs.map(([p]) => p).find(p => !prey.has(p));
  const chain = [apex];
  while (map[chain.at(-1)]) chain.push(map[chain.at(-1)]);
  return chain;
}
```
