---
id: 674ef2d357676e50e469165c
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

现在，将 `Card` 组件放入你的 `App` 组件中。为它添加 `name` 属性，设置为 `"Mark"`，`title` 属性设置为 `"Frontend developer"`，`bio` 设置为 `"I like to work with different frontend technologies and play video games."`

# --hints--

你应该在你的 `App` 组件中使用 `Card` 组件。

```js
  const testElem = await __helpers.prepTestComponent(window.index.App);
  const card = testElem.querySelector('.card');
  assert.exists(card);
```

你应该将值为 `"Mark"` 的 `name` 属性传递给你的 `Card` 组件。

```js
  const testElem = await __helpers.prepTestComponent(window.index.App);
  const card = testElem.querySelector('.card');
  assert.equal(card.querySelector('h2').textContent, 'Mark');
```

你应该将值为 `"Frontend developer"` 的 `title` 属性传递给你的 `Card` 组件。

```js
  const testElem = await __helpers.prepTestComponent(window.index.App);
  const card = testElem.querySelector('.card');
  assert.equal(card.querySelector('.card-title').textContent, 'Frontend developer');
```

你应该为你的 `Card` 组件传递 `bio` 属性，值为 `"I like to work with different frontend technologies and play video games."`。

```js
  const testElem = await __helpers.prepTestComponent(window.index.App);
  const card = testElem.querySelector('.card');
  assert.equal(card.querySelector('p:last-child').textContent, 'I like to work with different frontend technologies and play video games.');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  return (
    <div className="flex-container">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  );
}
```
