---
id: 68778af56eb5d980a4e2daed
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

你应该创建一个评测数据数组，并使用该数据动态渲染 `Card` 组件，而不是为每个单独的 `Card` 组件设置属性。

首先在 `App` 组件内创建一个 `profiles` 数组。

为你的数组添加一个对象，用于 `Mark`，其值如下：

- `id`: `1`
- `name`: `"Mark"`
- `title`: `"Frontend developer"`
- `bio`: `"I like to work with different frontend technologies and play video games."`

# --hints--

你应该在 `App` 组件内定义一个 `profiles` 数组。

```js
assert.match(code, /(const|let)\s+profiles\s*=\s*\[/);
```

你的评测数组应包含一个对象，该对象包括所有必需的属性：`id`、`name`、`title` 和 `bio`。

```js
assert.match(
  code,
  /(const|let)\s+profiles\s*=\s*\[\s*{[^}]*id\s*:\s*\d+[^}]*name\s*:\s*["'`][^"'`]+["'`][^}]*title\s*:\s*["'`][^"'`]+["'`][^}]*bio\s*:\s*["'`][^"'`]+["'`][^}]*}\s*\]/s);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  --fcc-editable-region--
     
  --fcc-editable-region--
  return (
    <div className="flex-container"></div>
  );
}
```

