---
id: 587d78b1367417b2b2512b09
title: Ein Bild responsiv machen
challengeType: 0
forumTopicId: 301140
dashedName: make-an-image-responsive
---

# --description--

Bilder mit CSS responsiv zu machen ist eigentlich ganz einfach. Sie müssen nur diese Eigenschaften zu einem Bild hinzufügen:

```css
img {
  max-width: 100%;
  height: auto;
}
```

Die `max-width` von `100%` sorgt dafür, dass das Bild niemals breiter als der Container ist, in dem es sich befindet, und die `height` von `auto` sorgt dafür, dass das Bild sein ursprüngliches Seitenverhältnis beibehält.

# --instructions--

Fügen Sie die Stilregeln zur `responsive-img` Klasse hinzu, um sie responsiv zu machen. Sie sollte niemals breiter als ihr Container sein (in diesem Fall ist das das Vorschaufenster) und ihr ursprüngliches Seitenverhältnis beibehalten. Nachdem Sie Ihren Code hinzugefügt haben, ändern Sie die Größe der Vorschau, um zu sehen, wie sich Ihre Bilder verhalten.

# --hints--

Ihre `responsive-img` Klasse sollte eine `max-width` mit dem Wert `100%` haben.

```js
assert.strictEqual(getComputedStyle(document.querySelector('.responsive-img')).maxWidth, '100%');
```

Ihre `responsive-img` Klasse sollte eine `height` mit dem Wert `auto` haben.

```js
assert.match(code, /height:\s*?auto;/g);
```

# --seed--

## --seed-contents--

```html
<style>
.responsive-img {


}

img {
  width: 600px;
}
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```

# --solutions--

```html
<style>
  .responsive-img {
    max-width: 100%;
    height: auto;
  }

  img {
    width: 600px;
  }
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```
