---
id: 587d78b1367417b2b2512b0a
title: Verwenden Sie ein Retina-Bild für Displays mit höherer Auflösung
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cVZ4Rfp'
forumTopicId: 301142
dashedName: use-a-retina-image-for-higher-resolution-displays
---

# --description--

Mit der Zunahme internetfähiger Geräte variieren deren Größen und Spezifikationen, und die verwendeten Displays können sich äußerlich und innerlich unterscheiden. Die Pixeldichte ist ein Aspekt, der sich von Gerät zu Gerät unterscheiden kann, und diese Dichte ist als Pixel pro Zoll (PPI) oder Punkte pro Zoll (DPI) bekannt. Das bekannteste Display dieser Art ist das sogenannte „Retina Display“ bei den neuesten Apple MacBook Pro Notebooks und kürzlich bei iMac-Computern. Aufgrund der unterschiedlichen Pixeldichte zwischen „Retina“- und „Nicht-Retina“-Displays können einige Bilder, die nicht für ein hochauflösendes Display erstellt wurden, auf einem hochauflösenden Display „pixelig“ aussehen, wenn sie gerendert werden.

Der einfachste Weg, Ihre Bilder auf hochauflösenden Displays wie dem „Retina Display“ der MacBook Pros korrekt darzustellen, besteht darin, deren `width` und `height` Werte auf die Hälfte der Originaldatei zu setzen. Hier ist ein Beispiel für ein Bild, das nur die halbe Originalhöhe und -breite verwendet:

```html
<style>
  img {
    height: 250px;
    width: 250px;
  }
</style>
<img src="coolPic500x500" alt="A most excellent picture" />
```

# --instructions--

Setzen Sie die `width` und `height` des `img` Tags auf die Hälfte ihrer Originalwerte. In diesem Fall sind sowohl die ursprüngliche `height` als auch die ursprüngliche `width` `200px`.

# --hints--

Ihr `img` Tag sollte eine `width` von 100 Pixel haben.

```js
assert.strictEqual(document.querySelector('img').width, 100);
```

Ihr `img` Tag sollte eine `height` von 100 Pixel haben.

```js
assert.strictEqual(document.querySelector('img').height, 100);
```

# --seed--

## --seed-contents--

```html
<style></style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```

# --solutions--

```html
<style>
  img {
    height: 100px;
    width: 100px;
  }
</style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```
