---
id: 587d78b0367417b2b2512b08
title: Eine Media Query erstellen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cqwKrtm'
forumTopicId: 301139
dashedName: create-a-media-query
---

# --description--

Media Queries sind eine neue Technik, die mit CSS3 eingeführt wurde und die Darstellung von Inhalten basierend auf verschiedenen Viewport-Größen ändert. Der Viewport ist der sichtbare Bereich einer Webseite für den Benutzer und unterscheidet sich je nach verwendetem Gerät, um auf die Seite zuzugreifen.

Media Queries bestehen aus einem Medientyp, und wenn dieser Medientyp mit dem Typ des Geräts übereinstimmt, auf dem das Dokument angezeigt wird, werden die Styles angewendet. Sie können beliebig viele Selektoren und Styles innerhalb Ihrer Media Query haben.

Hier ist ein Beispiel für eine Media Query, die den Inhalt zurückgibt, wenn die Breite des Geräts kleiner oder gleich `100px` ist:

```css
@media (max-width: 100px) {
  /* CSS Rules */
}
```

Und die folgende Media Query gibt den Inhalt zurück, wenn die Höhe des Geräts größer oder gleich `350px` ist:

```css
@media (min-height: 350px) {
  /* CSS Rules */
}
```

Denken Sie daran, dass das CSS innerhalb der Media Query nur angewendet wird, wenn der Medientyp mit dem des verwendeten Geräts übereinstimmt.

# --instructions--

Fügen Sie eine Media Query hinzu, sodass der `p`-Tag eine `font-size` von `10px` hat, wenn die Höhe des Geräts kleiner oder gleich `800px` ist.

# --hints--

Sie sollten eine `@media`-Abfrage für Geräte mit einer `height` kleiner oder gleich `800px` deklarieren.

```js
const media = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(
  media.some(x => x.media?.mediaText?.includes('(max-height: 800px)'))
);
```

Ihr `p`-Element sollte eine `font-size` von `10px` haben, wenn die `height` des Geräts kleiner oder gleich `800px` ist.

```js
const rules = new __helpers.CSSHelp(document).getRuleListsWithinMedia(
  '(max-height: 800px)'
);
assert.strictEqual(
  rules?.find(x => x.selectorText === 'p')?.style?.fontSize,
  '10px'
);
```

Ihr `p`-Element sollte eine anfängliche `font-size` von `20px` haben, wenn die `height` des Geräts größer als `800px` ist.

```js
const ifPFirst = new __helpers.CSSHelp(document)
  .getCSSRules()
  ?.find(x => x?.selectorText === 'p' || x?.media);
assert.strictEqual(ifPFirst?.style?.fontSize, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  p {
    font-size: 20px;
  }

  /* Only change code below this line */

  /* Only change code above this line */
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  p {
    font-size: 20px;
  }

  @media (max-height: 800px) {
    p {
      font-size: 10px;
    }
  }
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
