---
id: 587d78b1367417b2b2512b0c
title: Typografie responsiv gestalten
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/crzN7T8'
forumTopicId: 301141
dashedName: make-typography-responsive
---

# --description--

Anstatt `em` oder `px` zur Größenanpassung von Text zu verwenden, können Sie Viewport-Einheiten für responsiven Text nutzen. Viewport-Einheiten sind wie Prozentsätze relative Einheiten, basieren jedoch auf unterschiedlichen Bezugspunkten. Viewport-Einheiten beziehen sich auf die Abmessungen des Viewports (Breite oder Höhe) eines Geräts, während Prozentsätze sich auf die Größe des übergeordneten Container-Elements beziehen.

Die vier verschiedenen Viewport-Einheiten sind:

<ul><li><code>vw</code> (Viewport-Breite): <code>10vw</code> entspricht 10 % der Breite des Viewports.</li><li><code>vh</code> (Viewport-Höhe): <code>3vh</code> entspricht 3 % der Höhe des Viewports.</li><li><code>vmin</code> (Viewport-Minimum): <code>70vmin</code> entspricht 70 % der kleineren Viewport-Dimension (Höhe oder Breite).</li><li><code>vmax</code> (Viewport-Maximum): <code>100vmax</code> entspricht 100 % der größeren Viewport-Dimension (Höhe oder Breite).</li></ul>

Hier ist ein Beispiel, das ein `body`-Tag auf 30 % der Viewport-Breite setzt.

```css
body {
  width: 30vw;
}
```

# --instructions--

Setzen Sie die `width` des `h2`-Tags auf 80 % der Viewport-Breite und die `width` des Absatzes auf 75 % der kleineren Viewport-Dimension.

# --hints--

Ihr `h2`-Tag sollte eine `width` von 80vw haben.

```js
assert.match(
  __helpers.removeCssComments(code),
  /h2\s*?{\s*?width:\s*?80vw;\s*?}/g
);
```

Ihr `p`-Tag sollte eine `width` von 75vmin haben.

```js
assert.match(
  __helpers.removeCssComments(code),
  /p\s*?{\s*?width:\s*?75vmin;\s*?}/g
);
```

# --seed--

## --seed-contents--

```html
<style>

</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  h2 {
    width: 80vw;
  }
  p {
    width: 75vmin;
  }
</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
