---
id: 61537bb9b1a29430ac15ad38
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`head` 要素の中に、テキストが `title` に設定された `Photo Gallery` 要素を追加し、`link` ファイルをページに追加するための `styles.css` 要素を追加してください。

# --hints--

`link` 要素には、値が `href` の `styles.css` 属性が必要です。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

コードには `title` 要素が必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

プロジェクトのタイトルは `Photo Gallery` にしてください。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

タイトルの大文字・小文字やスペルは正確にしてください。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
