---
id: 61537d86bdc3dd343688fceb
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

次に、各`img`に文書内の順序に応じた`src`属性を付けてください。最初の`img`は`src`の`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`を持つ必要があります。残りは同じですが、`1`を文書内での`img`の番号に置き換えてください。

# --hints--

9つすべての`img`要素に`src`属性が必要です。

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

最初の`img`要素には、`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`属性値として`src`を設定してください。

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

2番目の`img`要素には、`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg`属性値として`src`を設定してください。

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

3番目の`img`要素には、`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg`属性値として`src`を設定してください。

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

4番目の`img`要素には、`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg`属性値として`src`を設定してください。

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

5番目の`img`要素には、`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg`属性値として`src`を設定してください。

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

6番目の`img`要素には、`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg`属性値として`src`を設定してください。

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

7番目の`img`要素には、`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg`属性値として`src`を設定してください。

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

8番目の`img`要素には、`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg`属性値として`src`を設定してください。

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

9番目の`img`要素には、`https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg`属性値として`src`を設定してください。

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
