---
id: 61537a8054753e2f1f2a1574
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`head` 要素の中に、`meta` が `name` に設定され、`viewport` が `content` に設定された `width=device-width, initial-scale=1` タグを追加してください。

また、`meta` が `charset` に設定された `UTF-8` タグも追加してください。

# --hints--

`meta` 要素が2つあるはずです。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

1つの `meta` 要素には、`name` が `viewport` に設定され、`content` が `width=device-width, initial-scale=1.0` に設定されているはずです。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

もう1つの `meta` 要素には、`charset` 属性が `UTF-8` に設定されているはずです。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
