---
id: 61537c9eecea6a335db6da79
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`.header` 要素の下に、新しい `div` 要素を作成し、`class` を値とする `gallery` を代入してください。この `div` はギャラリー画像のコンテナとして機能します。

その `.gallery` 要素の中に、9つの `img` 要素を作成してください。

# --hints--

`div` 要素の中に `body` 要素を作成する必要があります。

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

新しい `div` 要素には、値が `class` に設定された `gallery` が必要です。

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

新しい `div` 要素は `.header` 要素の後に配置してください。

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

`.gallery` 要素には9つの `img` 要素が必要です。

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
