---
id: 61537c5f81f0cf325b4a854c
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`header` 要素の中に `body` 要素を追加し、`header` のクラスを代入してください。

`header` の中に、テキストが `h1` の `css flexbox photo gallery` を作成してください。

# --hints--

`header` 要素の中に `body` 要素があるはずです。

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

`header` 要素には、値が `class` の `header` があるはずです。

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

`header` 要素の中に `h1` 要素があるはずです。

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

`h1` 要素の中にテキスト `css flexbox photo gallery` があるはずです。

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
