---
id: 61537485c4f2a624f18d7794
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

標準的なHTMLボイラープレートから始めてください。`DOCTYPE` 宣言、`html` 要素、`head` 要素、`body` 要素を追加してください。

開始`lang` タグに`<html>` 属性を追加し、値として`en` を設定してください。

# --hints--

コードには`DOCTYPE` 参照が含まれている必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れてください。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

`html` を文書型として定義してください。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

`DOCTYPE` 宣言はタイプの後に`>` で閉じてください。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 要素には開始タグが必要です。

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

開始`<html>` タグには`lang` 属性があり、値は`en` である必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

`DOCTYPE` 宣言はHTMLの先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

開始`head` タグが必要です。

```js
assert(code.match(/<head\s*>/i));
```

終了`head` タグが必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

開始`body` タグが必要です。

```js
assert(code.match(/<body\s*>/i));
```

終了`body` タグが必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` と`body` 要素は兄弟要素である必要があります。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 要素は`html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は`html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
