---
id: 64dd076a1b33c86d84b4232f
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

要素のサイズがどのようになっているかをよりわかりやすくするために、ボーダーを追加すると便利です。

`.gallery` 要素に `width` を `50%` に、`border` を `5px solid red` に設定してください。

次に、`img` 要素に `width` を `100%` に、`padding` を `5px` に、`border` を `5px solid blue` に設定してください。

# --hints--

`.gallery` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

`.gallery` セレクターには `width` プロパティがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

`.gallery` セレクターには `width` プロパティが `50%` に設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

`.gallery` セレクターには `border` プロパティがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

`.gallery` セレクターには `border` プロパティが `5px solid red` に設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

`img` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

`img` セレクターには `width` プロパティがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

`img` セレクターには `width` プロパティが `100%` に設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

`img` セレクターには `padding` プロパティがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

`img` セレクターには `padding` プロパティが `5px` に設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

`img` セレクターには `border` プロパティがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

`img` セレクターには `border` プロパティが `5px solid blue` に設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
