---
id: 61437d575fb98f57fa8f7f36
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

从标准的 HTML 模板开始。 添加 `DOCTYPE` 声明、语言设置为英语的 `html` 元素、 `head` 和 `body` 元素。

在 `head` 元素内添加带有适当 `charset` 的 `<meta>` 标签和用于移动响应的 `<meta>` 标签。

# --hints--

你的代码应该包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你应在 `DOCTYPE` 后添加一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你应将文档类型定义为 `html` 。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你应该在 `DOCTYPE` 声明的 `type` 后用 `>` 关闭。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你应该有一个开始的 `<html>` 标签，并且它的 `lang` 应为 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你应该有一个开启的 `<head>` 标签。

```js
assert(code.match(/<head\s*>/i));
```

你应该有一个闭合的 `</head>` 标签。

```js
assert(code.match(/<\/head\s*>/i));
```

你应该有一个开始的 `<body>` 标签。

```js
assert(code.match(/<body\s*>/i));
```

你应该有一个闭合的 `</body>` 标签。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素应该是同级元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你应该有两个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

一个 `meta` 元素应将 `name` 设置为 `viewport`，并将 `content` 设置为 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一个 `meta` 元素应将 `charset` 属性设置为 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
