---
id: 587d7fb4367417b2b2512c00
title: npmの外部パッケージでプロジェクトを拡張する
challengeType: 2
forumTopicId: 301527
dashedName: expand-your-project-with-external-packages-from-npm
---

# --description--

パッケージマネージャーを使う最大の理由の1つは、その強力な依存関係管理です。新しいコンピューターでプロジェクトをセットアップするたびに、すべての依存関係を手動で確認する代わりに、npmが自動的にすべてをインストールしてくれます。でも、npmはどうやってあなたのプロジェクトに必要なものを正確に知るのでしょうか？それがpackage.jsonファイルの`dependencies`セクションです。

このセクションには、プロジェクトが必要とするパッケージが次の形式で保存されています。

```json
"dependencies": {
  "package-name": "version",
  "express": "4.14.0"
}

```

# --instructions--

`1.1.0`ファイルの`@freecodecamp/example`フィールドに、`dependencies`パッケージのバージョン`package.json`を追加してください。

**注意:** `@freecodecamp/example`は学習用の架空のパッケージです。

# --hints--

`"dependencies"`には`"@freecodecamp/example"`を含める必要があります。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`"@freecodecamp/example"`のバージョンは`"1.1.0"`であるべきです。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^[\^\~]?1\.1\.0/,
        'Wrong version of "@freecodecamp/example" installed. It should be 1.1.0'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
