---
id: 587d7fb5367417b2b2512c04
title: 依存関係からパッケージを削除する
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

これまでに、package.json の dependencies セクションを使ってプロジェクトの依存関係を管理するいくつかの方法を試してきました。外部パッケージをファイルに追加し、チルダやキャレットなどの特殊文字を使って npm に希望するバージョンの種類を伝えることもしました。

では、もう必要なくなった外部パッケージを削除したい場合はどうすればよいでしょうか？おそらくもうお分かりかと思いますが、そのパッケージに対応するキーと値のペアを dependencies から削除すればよいのです。

この方法は、package.json の他のフィールドを削除する場合にも同様に適用されます。

# --instructions--

依存関係から `@freecodecamp/example` パッケージを削除してください。

**注意:** 削除後にカンマの数が正しいことを確認してください。

# --hints--

`"dependencies"` に `"@freecodecamp/example"` が含まれていないことを確認してください。

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
