---
id: 587d7fb4367417b2b2512bfe
title: package.json にライセンスを追加する
challengeType: 2
forumTopicId: 301523
dashedName: add-a-license-to-your-package-json
---

# --description--

`license` フィールドは、ユーザーにプロジェクトで何が許可されているかを知らせる場所です。

オープンソースプロジェクトでよく使われるライセンスには MIT や BSD があります。ライセンス情報は必須ではなく、ほとんどの国の著作権法によりデフォルトでは作成したものの所有権が認められます。しかし、ユーザーが何をできて何をできないかを明確に示すことは常に良い習慣です。以下は license フィールドの例です。

```json
"license": "MIT",
```

# --instructions--

プロジェクトの package.json ファイルの `license` フィールドを適切に記入してください。

# --hints--

package.json には有効な "license" キーが必要です

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.license, '"license" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
