---
id: 587d7fb4367417b2b2512bff
title: package.json にバージョンを追加する
challengeType: 2
forumTopicId: 301525
dashedName: add-a-version-to-your-package-json
---

# --description--

`version` は package.json ファイルの必須フィールドの1つです。このフィールドはプロジェクトの現在のバージョンを表します。例を示します。

```json
"version": "1.2.0",
```

# --instructions--

プロジェクトの package.json ファイルに `version` を追加してください。

# --hints--

package.json に有効な "version" キーが含まれている必要があります

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.version, '"version" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
