---
id: 587d7fb3367417b2b2512bfc
title: package.json に説明を追加する
challengeType: 2
forumTopicId: 301522
dashedName: add-a-description-to-your-package-json
---

# --description--

良い package.json ファイルの次の部分は `description` フィールドです。ここにはプロジェクトについての短くても有益な説明を記述します。

将来的に npm にパッケージを公開する予定がある場合、この文字列はユーザーがパッケージをインストールするかどうかを決める際にアイデアを伝える役割を果たします。しかし、説明の使い道はそれだけではありません。プロジェクトが何をするものかを要約するのに最適な方法です。Node.js のプロジェクトでは、他の開発者や将来のメンテナ、さらには将来の自分自身がプロジェクトを素早く理解するために非常に重要です。

プロジェクトの目的が何であれ、説明は必ず記述することをお勧めします。以下は例です。

```json
"description": "A project that does something awesome",
```

# --instructions--

プロジェクトの package.json ファイルに `description` を追加してください。

**注意:** フィールド名には必ずダブルクォーテーション（"）を使い、フィールド同士はカンマ（,）で区切ることを忘れないでください。

# --hints--

package.json に有効な "description" キーがあること

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.description, '"description" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
