---
id: 686b9720ee1d032bd77a480a
title: Налагодження валідатора ISBN
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

ISBN (Міжнародний стандартний книжковий номер) — це унікальний ідентифікатор, призначений комерційним книгам. Він може містити 10 або 13 цифр, а остання цифра є контрольним числом, яке обчислюється з інших цифр.

Camperbot намагався створити власний валідатор ISBN. Однак він припустився кількох помилок.

У цій лабораторній роботі ви виправите існуючий код і зробите його правильним.

Очікувана поведінка:
Коли користувач запускає програму, вона покаже запрошення `Enter ISBN and length: `. Користувач може ввести код ISBN, який хоче перевірити, у форматі `ISBN,length`. Код ISBN не повинен містити дефісів, а потім через кому вказується його довжина (`10` або `13`).

Приклади вхідних даних:
`1530051126,10` для кодів ISBN-10.
`9781530051120,13` для кодів ISBN-13.

Як знайти контрольне число:

- Вам не потрібно знати детальну логіку обчислення в цій лабораторній роботі. Функції `calculate_check_digit_10` і `calculate_check_digit_13` подбають про обчислення і повернуть очікуване контрольне число у вигляді рядка.
- Контрольне число для кодів ISBN-10 може бути числом від `0` до `9` або великою літерою `X`.
- Контрольне число для кодів ISBN-13 може бути числом від `0` до `9`.

**Мета:** Виконайте наведені нижче історії користувача і пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні виправити `IndentationError` у поточному коді.
1. Навіть якщо користувач не введе значення, розділене комою, програма повинна обробити `IndexError` без аварійного завершення.
1. Коли користувач не вводить значення, розділене комою, він повинен побачити повідомлення `Enter comma-separated values.` у консолі, і програма має завершити роботу.
1. Навіть якщо користувач введе нечислове значення для довжини, програма повинна обробити `ValueError` без аварійного завершення.
1. Коли користувач вводить нечислове значення для довжини, він повинен побачити повідомлення `Length must be a number.` у консолі, і програма має завершити роботу.
1. Ви повинні виправити помилку зміщення на одиницю у функції `validate_isbn`.
1. Ви повинні виправити `TypeError` у поточному коді, що виникає, коли користувач вводить дійсний код ISBN.
1. Ви повинні виправити `IndexError` у поточному коді, коли користувач вводить дійсний код ISBN.
1. Навіть якщо користувач вводить неправильний код ISBN з символами, відмінними від цифр, програма повинна обробити `ValueError`, що виникає, без аварійного завершення.
1. Коли користувач вводить неправильний код ISBN з символами, відмінними від цифр, він повинен побачити повідомлення `Invalid character was found.` у консолі.
1. Коли користувач вводить `1530051126,10`, він повинен побачити повідомлення `Valid ISBN Code.`.
1. Коли користувач вводить `9781530051120,13`, він повинен побачити повідомлення `Valid ISBN Code.`.

**Важливо**: вам потрібно буде перетворити на коментар виклик `main()` у глобальній області, щоб тести працювали правильно.

Коли ви завершите проєкт, користувач повинен бачити такі повідомлення залежно від введених значень.

| Код ISBN | Довжина | Повідомлення | Приклад введення |
|-----------|--------|-------------|------------------|
| Дійсний | Дійсний | `Valid ISBN Code.` | `1530051126,10` |
| Невірне число | Дійсний | `Invalid ISBN Code.` | `1530051125,10` |
| Не відповідає вказаній довжині або порожній | Дійсний | `ISBN-10 code should be 10 digits long.` або `ISBN-13 code should be 13 digits long.`, залежно від введеної довжини. | `9781530051120,10` або `1530051126,13` |
| Містить нечислові символи (крім контрольного числа) | Дійсний | `Invalid character was found.` | `15-0051126,10` |
| Будь-який | Невірне число | `Length should be 10 or 13.` | `1530051126,9` |
| Будь-який | Містить нечислові символи або порожній | `Length must be a number.` | `1530051125,A` |
| Не розділений комою | Не розділений комою | `Enter comma-separated values.` | `1530051125` |

Ви можете використати наведені значення для ручного тестування коду, якщо хочете.

Приклади вхідних даних для дійсних кодів ISBN-10:
`1530051126,10`
`9971502100,10`
`080442957X,10`

Приклади вхідних даних для дійсних кодів ISBN-13:
`9781530051120,13`
`9781947172104,13`

# --hints--

Ви повинні перетворити на коментар виклик функції `main`, щоб решта тестів працювала правильно.

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

У вас має бути функція `validate_isbn`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

У вас має бути функція `calculate_check_digit_10`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

У вас має бути функція `calculate_check_digit_13`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

У вас має бути функція `main`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

Коли користувач вводить значення, яке не є значенням, розділеним комою, ви повинні побачити повідомлення `Enter comma-separated values.` у консолі.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Коли користувач вводить нечислове значення для довжини, ви повинні побачити повідомлення `Length must be a number.` у консолі.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Коли користувач вводить неправильний код ISBN з символами, відмінними від цифр, ви повинні побачити повідомлення `Invalid character was found.` у консолі.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Коли користувач вводить `1530051126,10`, ви повинні побачити повідомлення `Valid ISBN Code.` у консолі.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Коли користувач вводить `9781530051120,13`, ви повинні побачити повідомлення `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Коли користувач вводить `1530051125,10`, ви повинні побачити повідомлення `Invalid ISBN Code.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

Коли користувач вводить `9781530051120,10`, ви повинні побачити повідомлення `ISBN-10 code should be 10 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

Коли користувач вводить `1530051126,13`, ви повинні побачити повідомлення `ISBN-13 code should be 13 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

Коли користувач вводить `15-0051126,10`, ви повинні побачити повідомлення `Invalid character was found.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Коли користувач вводить `1530051126,9`, ви повинні побачити повідомлення `Length should be 10 or 13.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

Коли користувач вводить `1530051125,A`, ви повинні побачити повідомлення `Length must be a number.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Коли користувач вводить `1530051125`, ви повинні побачити повідомлення `Enter comma-separated values.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Коли користувач вводить `9971502100,10`, ви повинні побачити повідомлення `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Коли користувач вводить `080442957X,10`, ви повинні побачити повідомлення `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Коли користувач вводить `9781947172104,13`, ви повинні побачити повідомлення `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
