---
id: 615f3b091162165948e1cb33
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Ikiwa utachunguza kipengele chako cha `.label` kwa kutumia zana za msanidi programu za kivinjari chako, unaweza kugundua kuwa kina upana wa pikseli 288 badala ya 270. Hii ni kwa sababu, kwa chaguo-msingi, kivinjari hujumuisha mpaka na nafasi ya ndani wakati wa kubaini ukubwa wa kipengele.

Ili kutatua hili, rejesha mfano wa kisanduku kwa kuunda kichaguzi cha `*` na kukipa sifa ya `box-sizing` yenye thamani ya `border-box`.

# --hints--

Unapaswa kuunda kichaguzi cha `*`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

Kichaguzi chako cha `*` kinapaswa kuwa na sifa ya `box-sizing` iliyowekwa kuwa `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
