---
id: 65821c1c40ccac44adf38e6b
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

Create an empty list named `expenses`. Ви використовуватимете його, щоб зберігати витрати.

# --hints--

Оголосіть нову змінну під назвою `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

Призначте порожній список до змінної `expenses`.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass

--fcc-editable-region--

--fcc-editable-region--
```
