---
id: 65821fcc010c3245718f2a06
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Параметром `expenses` функції `add_expense` буде список витрат. Важливо мати можливість додавати елементи в кінець списку. Для цього можна використати метод списку `.append()`.

Додайте виклик методу `.append()` на списку `expenses`. Наразі не передавайте жодних аргументів до `.append()`.

# --hints--

Функція повинна містити `expenses.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

You should not have `pass` in your function.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py    
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
