---
id: 66694bfa435cef241f6f9a68
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Перш ніж почати проєкт, ви дізнаєтеся про списки.

На відміну від чисел та рядків, <dfn>список</dfn> є типом даних, який працює як вмістилище для інших значень:

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

Список характеризується квадратними дужками навколо усіх значень та комою між значеннями, наприклад: `["A", "happy", "list"]`. Якщо список не містить жодних значень, це порожній список: `[]`.

Список може містити різні типи даних: `[1, "Up", ["Down", "Twice"]]`. Це включає усі можливі типи даних. Він також може містити інший список!

Створіть змінну під назвою `my_list` та призначте до неї порожній список.

# --hints--

Ви повинні мати змінну під назвою `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

Змінна `my_list` повинна мати значення порожнього списку.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
