---
id: 666ab9013cd5563f2b39947a
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Тепер, коли ви дослідили списки, видаліть увесь код, який ви написали для огляду, щоб почати проєкт.

# --hints--

Видаліть усе.

```js
({
    test: () => runPython(`
    assert not _Node(_code)
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)

my_list.pop()
print(my_list)

--fcc-editable-region--
```
