---
id: 69bac067cd15c4685d07906d
title: Passaggio 15
challengeType: 1
dashedName: step-15
---

# --description--

Ora che hai funzioni per elaborare i singoli segni di graduazione e gestire l’overflow, è il momento di creare una funzione controller per un singolo cancello. Questa funzione gestirà tutte le azioni per un cancello durante un singolo segno di graduazione della simulazione.

Crea una funzione vuota chiamata `handleGateAtTick` con i parametri:

- `gates`: L’array completo degli oggetti cancello.
- `gate`: Il cancello corrente in elaborazione.
- `tickIndex`: L’indice del segno di graduazione corrente per la simulazione.
- `throughputSummary`: Un oggetto che tiene traccia del totale dei partecipanti elaborati per ogni cancello.

# --hints--

Dovresti creare una funzione chiamata `handleGateAtTick`.

```js
assert.isFunction(handleGateAtTick);
```

La tua funzione `handleGateAtTick` dovrebbe avere i parametri `gates`, `gate`, `tickIndex` e `throughputSummary`.

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
