---
id: 69babfd65c4c4632a0c76f63
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Quando un cancello non riesce a processare tutti i partecipanti durante un tick, alcuni partecipanti rimangono in coda. Per gestire questo overflow, costruirai una funzione che reindirizza i partecipanti in overflow a un altro cancello.

Crea una funzione vuota chiamata `rerouteOverflow` con i seguenti parametri:

- `gates`: L'array completo degli oggetti cancello.
- `currentGate`: Il cancello dove si è verificato l'overflow.
- `tickIndex`: La posizione corrente del tick nella coda.
- `overflowAmount`: Il numero di partecipanti che non sono stati processati.

# --hints--

Dovresti creare una funzione chiamata `rerouteOverflow`.

```js
assert.isFunction(rerouteOverflow);
```

La tua funzione `rerouteOverflow` dovrebbe avere i parametri `gates`, `currentGate`, `tickIndex` e `overflowAmount`.

```js
const regex = __helpers.functionRegex('rerouteOverflow', ['gates', 'currentGate', 'tickIndex', 'overflowAmount']);
assert.match(__helpers.removeJSComments(rerouteOverflow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
