---
id: 69bac027df64d4433c708d44
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Quando devi gestire l'overflow, dovresti indirizzare i partecipanti al cancello successivo nell'array `gates`. Per farlo, devi trovare l'indice del cancello successivo. Normalmente, potresti aggiungere 1 all'indice corrente. Tuttavia, se il cancello corrente è l'ultimo nell'array, devi tornare al primo cancello.

Puoi usare l'operatore modulo (`%`) per gestire questa situazione:

```js
const nextIndex = (currentIndex + 1) % array.length;
```

Questo funziona perché quando `currentIndex + 1` è uguale a `array.length`, il risultato diventa 0.

Usando questo metodo, crea una variabile chiamata `nextGateIndex` che archivia l'indice del cancello successivo nell'array `gates`.

# --hints--

Dovresti usare la tecnica del modulo per consultare e assegnare il prossimo indice da `currentIndex` nell'array `gates` a una variabile chiamata `nextGateIndex`. Segui lo stesso pattern in una riga dell'esempio.

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--

--fcc-editable-region--
}
```
