---
id: 69baafd6d683dd9921b75db3
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop userai JavaScript per simulare il flusso dei partecipanti a un festival musicale.

Lavorerai con due dataset che descrivono i cancelli del festival durante i turni di mattina e di sera:

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

Ogni oggetto cancello contiene le seguenti proprietà:

- `id`: Una stringa che identifica il cancello.
- `capacity`: Il numero di partecipanti che il cancello può processare per tick.
- `queue`: Un array di numeri che rappresentano quanti partecipanti arrivano al cancello durante un tick specifico.

Per cominciare, crea una funzione vuota chiamata `initializeThroughput` che accetta un parametro chiamato `gates`.

# --hints--

Dovresti creare una funzione chiamata `initializeThroughput`.

```js
assert.isFunction(initializeThroughput);
```

La tua funzione `initializeThroughput` dovrebbe avere un parametro `gates`.

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
