---
id: 69babdf3f896711c11516373
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Durante ogni tick della simulazione:

- Un certo numero di partecipanti arriva al cancello (dal suo `queue`).
- Il cancello elabora i partecipanti in base al suo `capacity`.
- Se arrivano più partecipanti di quanti il cancello possa gestire, alcuni rimarranno (overflow).

Ora costruirai una funzione che gestisce questa logica per un singolo cancello in un singolo tick.

Crea una funzione vuota chiamata `processGateFlow` con i seguenti parametri:

- `gate`: L'oggetto gate che viene processato.
- `tickIndex`: La posizione corrente del tick nell'array `queue`.

# --hints--

Dovresti creare una funzione chiamata `processGateFlow`.

```js
assert.isFunction(processGateFlow);
```

La tua funzione `processGateFlow` dovrebbe avere i parametri `gate` e `tickIndex`.

```js
const explorer = await __helpers.Explorer(code);
const { processGateFlow } = explorer.allFunctions;
const parameters = processGateFlow?.parameters;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), "gate");
assert.equal(parameters[1].toString(), "tickIndex");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
