---
id: 69cfdc80443b715a8aa452a7
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Per completare la tua funzione `initializeThroughput`:

- Usa un `for` loop per iterare su ogni gate nell'array di input `gates`. Per ogni gate, aggiungi una proprietà a `summary` dove:
  - La chiave è l'`id` del gate.
  - Il valore è `0`.
- Dopo il loop, restituisci l'oggetto `summary`.

La tua funzione `initializeThroughput` non dovrebbe modificare l'array di input sorgente `gates`.

# --hints--

La tua funzione `initializeThroughput` dovrebbe restituire `{}` quando viene chiamata con `[]`, senza modificare l'array di input originale.

```js
const input = [];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  {},
  "Calling `initializeThroughput([])` should return an empty object."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given an empty array."
);
```

La tua funzione `initializeThroughput` dovrebbe restituire `{ North: 0 }` quando viene chiamata con `[{ id: "North", capacity: 5, queue: [1, 2, 3] }]`, senza modificare l'array di input originale.

```js
const input = [{ id: "North", capacity: 5, queue: [1, 2, 3] }];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0 },
  "Calling `initializeThroughput([{ id: 'North', ... }])` should return `{ North: 0 }`."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given one gate."
);
```

La tua funzione `initializeThroughput` dovrebbe restituire `{ North: 0, East: 0, South: 0 }` quando viene chiamata con `[{ id: "North", capacity: 5, queue: [] }, { id: "East", capacity: 3, queue: [] }, { id: "South", capacity: 4, queue: [] }]`, senza modificare l'array di input originale.

```js
const input = [
  { id: "North", capacity: 5, queue: [] },
  { id: "East", capacity: 3, queue: [] },
  { id: "South", capacity: 4, queue: [] }
];

const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0, East: 0, South: 0 },
  "Calling `initializeThroughput` with multiple gates should return an object with an entry for each gate id set to 0."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given multiple gates."
);
```

Dovresti usare un loop `for` per iterare su ogni oggetto gate dentro l'array di input `gates`.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));
assert.match(cleaned, /for\s*\(/);
```

La tua funzione `initializeThroughput` dovrebbe restituire `summary` dopo il loop, come ultima riga della funzione.

```js
const funcStr = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));

assert.match(
  funcStr,
  /returnsummary;?}$/,
  "Your `initializeThroughput` function should end with `return summary`."
);
```

La tua funzione `initializeThroughput` non dovrebbe modificare l'array di input sorgente `gates`.

```js
const morningCopy = JSON.parse(JSON.stringify(morningGates));
const nightCopy = JSON.parse(JSON.stringify(nightGates));

initializeThroughput(morningGates);
initializeThroughput(nightGates);

assert.deepEqual(
  morningGates,
  morningCopy,
  "`initializeThroughput` should not mutate the `morningGates` seed array."
);

assert.deepEqual(
  nightGates,
  nightCopy,
  "`initializeThroughput` should not mutate the `nightGates` seed array."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
--fcc-editable-region--
  
--fcc-editable-region--
}
```
