---
id: 698b42492a7f5f737895892e
title: 인접 리스트를 인접 행렬로 변환기 만들기
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

이 실습에서는 그래프의 인접 리스트 표현을 인접 행렬로 변환하는 함수를 만드세요. 인접 리스트는 각 키가 노드를 나타내고, 해당 키 노드와 연결된 노드들의 배열을 값으로 가지는 객체입니다. 인접 행렬은 2차원 배열로, 위치 `[i][j]`의 값이 노드 `1`에서 노드 `i`로 간선이 있으면 `j`이고, 그렇지 않으면 `0`입니다.

예를 들어, 다음 인접 리스트가 주어지면:

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

해당하는 인접 행렬은 다음과 같습니다:

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 인접 리스트를 인접 행렬로 변환하는 `adjacencyListToMatrix`라는 이름의 함수를 정의해야 합니다.
2. 함수는 가중치 없는 (무방향 또는 방향) 그래프의 인접 리스트를 나타내는 객체를 인수로 받아야 합니다.
3. 함수는 다음을 수행해야 합니다:
   - 인접 리스트를 인접 행렬로 변환합니다.
   - 인접 행렬의 각 행을 출력합니다.
   - 인접 행렬을 반환합니다.

예를 들어, `adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})`는 다음을 출력해야 합니다:

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

그리고 `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`를 반환해야 합니다.

# --hints--

`adjacencyListToMatrix`라는 이름의 함수를 정의하세요.

```js
assert.isFunction(adjacencyListToMatrix);
```

`adjacencyListToMatrix` 함수는 매개변수 하나를 가져야 합니다.

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

함수는 인접 리스트에서 노드 수를 정확히 파악해야 합니다.

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

함수는 존재하는 간선에 대해 행렬 값을 `1`으로 정확히 설정해야 합니다.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

함수는 행렬의 각 행을 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

함수는 인접 행렬을 반환해야 합니다.

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

인접 리스트 `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}`가 주어지면 함수는 `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`를 반환해야 합니다.

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

인접 리스트 `{0: [1], 1: [0]}`가 주어지면 함수는 `[[0, 1], [1, 0]]`를 반환해야 합니다.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

인접 리스트 `{0: [], 1: [], 2: []}`가 주어지면 함수는 `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`를 반환해야 합니다.

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
