---
id: 655cd899f8de09431eabb40c
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

合併排序是一種使用分而治之原則來排序資料群集的排序演算法。也就是說，它將一個群集「分割」成較小的子零件，並透過獨立排序這些子零件來「征服」它們，然後合併已排序的子零件。

首先建立一個名為 `merge_sort` 的函式，該函式接受參數 `array`。這個函式將控制代碼來排序一個數字列表。在函式主體中使用關鍵字 `pass`。

# --hints--

你應該宣告一個名為 `merge_sort` 的函式（程式），並帶有一個參數：`array`。別忘了在函式（程式）主體中使用 `pass` 關鍵字。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
