---
id: 656639771fed09b5c0e8fe71
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

現在你已經將 `array` 列表分成兩個獨立的列表，你會持續將每個列表分割，直到每個元素都獨立存在於自己的列表中。只有一個數字的列表永遠是已排序的。

為了達成此目的，在你的函式中遞迴呼叫 `merge_sort`，並將 `left_part` 作為引數傳遞給該呼叫。

# --hints--

你應該在函式（程式）主體底部以引數 `left_part` 呼叫 `merge_sort`。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
