---
id: 656680b0fc79f2c38a34d90e
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

現在是時候將 `left_part` 和 `right_part` 這兩個列表排序並合併到原始列表中。 

你可以透過比較兩個列表中的元素，並將較小的元素合併到主列表中來完成這個操作。你將對 `left_part` 和 `right_part` 中所有的索引進行這個比較。

創建三個變數：`left_array_index`、`right_array_index` 和 `sorted_index`，並將它們的值設為 `0`。這些變數將幫助你在排序過程中追蹤每個索引。

# --hints--

你應該在你的 `merge_sort` 函式中有一個名為 `left_array_index` 的變數

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

`left_array_index` 的值應該是 `0`。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

你應該在你的 `merge_sort` 函式中有一個名為 `right_array_index` 的變數。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

`right_array_index` 的值應該是 `0`。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

你應該在你的 `merge_sort` 函式（程式）內有一個名為 `sorted_index` 的變數。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

`sorted_index` 的值應該是 `0`。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--
    
--fcc-editable-region--
```
