---
id: 656680b0fc79f2c38a34d90e
title: 步驟7
challengeType: 20
dashedName: step-7
---

# --description--

Now it's time to sort and merge the lists (`left_part` and `right_part`) into the original list.

你可以通過比較兩個列表中的元素並將較小的元素合併到主列表中來實現。 你將對 `left_part` 和 `right_part` 中的所有索引進行此比較。

創建三個變量：`left_array_index`、`right_array_index` 和 `sorted_index`，並將它們的值設置爲 `0`。 這些變量將幫助您在排序過程中跟蹤每個索引。

# --hints--

你應該在 `merge_sort` 函數中有一個名爲 `left_array_index` 的變量

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

The value of `left_array_index` should be `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

你應該在 `merge_sort` 函數中有一個名爲 `right_array_index` 的變量

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

The value of `right_array_index` should be `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

你應該在 `merge_sort` 函數中有一個名爲 `sorted_index` 的變量

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

The value of `sorted_index` should be `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--

--fcc-editable-region--
```
