---
id: 656627b47bd2d2a4afbc945d
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

在前一步驟中，你取得了中點。你可以使用它將 `array` 分成兩部分，並將每個部分指定給新的變數。 

使用切片語法擷取 `array` 的左半部，並將其指定到名為 `left_part` 的變數。

# --hints--

你應該在你的 `merge_sort` 函式中有一個名為 `left_part` 的變數。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

你應該使用 slice 語法來擷取陣列的左半部並將它指定給 `left_part` 變數。

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
