---
id: 656710d1e0ec62253426db24
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

你創建的 `while` 迴圈比較 `left_part` 中的一個元素與 `right_part` 中的另一個元素，然後將較小的元素添加到主要的 `array` 列表中。

它將持續此操作，直到沒有元素可供比較。但 `left_part` 可能仍有元素剩餘，而 `right_part` 則沒有，反之亦然。

創建另一個 `while` 迴圈，當 `left_array_index` 小於 `len(left_part)` 時執行。在接下來的步驟中，你將使用它將 `left_part` 中剩餘的元素複製到 `array` 列表中。

# --hints--

你的 `while` 迴圈應該有這個條件：`left_array_index < len(left_part)`。別忘了在迴圈主體中使用 `pass` 關鍵字。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

你應該在你的 `while` 迴圈主體中使用 `pass` 關鍵字

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--
    
--fcc-editable-region--
```
