---
id: 656706afd65547d22bee0b68
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

在 `while` 迴圈中，將 `pass` 替換為一個 `if` 述語，用來檢查 `left_part` 中的元素是否小於 `right_part` 中的元素。

在 `if` 述語的主體中使用 `pass` 關鍵字。

# --hints--

你應該將 `pass` 關鍵字替換為帶有此條件的 `if` 述語：`left_part[left_array_index] < right_part[right_array_index]`。為你的 `if` 述語主體添加 `pass` 關鍵字。

```js
({
 test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_body().is_equivalent("if left_part[left_array_index] < right_part[right_array_index]:\\n  pass")`))
 })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        pass
--fcc-editable-region--
```
