---
id: 65670d1ef177e7e2b76d9528
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

你在前面步驟中創建的 `if` 和 `else` 述語會將元素指定到已排序的陣列。

每個指派給已排序陣列的元素都佔據列表中的一個索引。因此，在每次指派後，你必須移動到已排序陣列中的下一個索引。

在 `if`/`else` 區塊下方，但仍在 `while` 迴圈內，將 `sorted_index` 增加 `1`。

# --hints--

你應該在你的 `else` 語句之後將 `sorted_index` 的當前值增加 `1`。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_bodies()[0].find_body()[1]
    cond.is_equivalent("sorted_index += 1") or cond.is_equivalent("sorted_index = sorted_index + 1") or cond.is_equivalent("sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
--fcc-editable-region--
```
