---
id: 656702f8b4cbd8cbf0a433c6
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

接著，你需要創建一個迴圈，只要 `left_part` 和 `right_part` 中都有剩餘元素就繼續執行。

為此，創建一個具有兩個條件的 `while` 迴圈：一個檢查 `left_array_index` 是否小於 `left_part` 的長度，另一個條件檢查 `right_array_index` 是否小於 `right_part` 的長度。

# --hints--

你的 `while` 迴圈應該檢查 `left_array_index` 是否小於 `left_part` 的長度，以及 `right_array_index` 是否小於 `right_part` 的長度。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

在迴圈主體中使用 `pass` 關鍵字。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
