---
id: 656710d1e0ec62253426db24
title: 步驟15
challengeType: 20
dashedName: step-15
---

# --description--

你創建的 `while` 循環將 `left_part` 中的一個元素與 `right_part` 中的另一個元素進行比較，然後將較小的元素添加到主要 `array` 列表中。

它將繼續此操作，直到沒有剩餘元素可供比較。 但是 `left_part` 可能仍有元素，而 `right_part` 沒有元素，反之亦然。

Create another `while` loop that runs when `left_array_index` is less than `len(left_part)`. In the next steps, you'll use it to copy the remaining elements in `left_part` into the `array` list.

# --hints--

你的 `while` 循環應具有以下條件：`left_array_index < len(left_part)`。 不要忘記在循環主體中使用 `pass` 關鍵字。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

你應該在 `while` 循環主體中添加 `pass` 關鍵字

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
