---
id: 656707e11e671ed54c96f7ec
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

當 `if` 條件評估為 `True` 時，表示 `left_part` 列表中的元素小於它與之比較的 `right_part` 列表中的元素。

在那種情況下，你可以將在排序陣列中 `left_part` 內的 `left_array_index` 所找到的元素放入。

在 `if` 區塊內，移除 `pass` 並將 `left_part[left_array_index]` 指定給 `array[sorted_index]`。

# --hints--

你應該在 `if` 述語的主體中將 `pass` 替換為 `array[sorted_index] = left_part[left_array_index]`。

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
