---
id: 6569b743630ee592a65a7e2f
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

在測試 `merge_sort()` 函式（程式）之前，你需要創建一個基底案例，當 `array` 的長度小於或等於 `1` 時停止函式（程式）執行。

這個基底案例會停止遞迴呼叫。若沒有它，合併排序操作將會在列表已排序或沒有元素時仍繼續執行。

在函式（程式）宣告式之後，創建一個帶有此條件的 `if` 述語：`len(array) <= 1`。在 `if` 述語的主體中添加一個 `return` 述語以停止函式（程式）的執行。

# --hints--

你應該創建一個 `if` 述語來檢查 `len(array) <= 1`。

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

你應該在 `if` 述語的主體中有一個 `return` 述語。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_ifs()[0].find_body().has_stmt("return")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
