---
id: 6569b4e0bd29d17d4b53b16c
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

現在，你將為 `right_part` 複製相同的 `while` 迴圈邏輯。

創建一個 `while` 迴圈，當 `right_array_index` 小於 `len(right_part)` 時執行。在 `while` 迴圈內，將 `right_part[right_array_index]` 指定給 `array[sorted_index]`。最後，將 `right_array_index` 和 `sorted_index` 各自增加 1。

# --hints--

你的 `while` 迴圈應該檢查 `right_array_index` 是否小於 `len(right_part)`。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

你應該在迴圈主體中將 `right_part[right_array_index]` 指定到 `array[sorted_index]`。 

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

你應該在 `while` 迴圈中將 `right_array_index` 增加 `1`。請注意，增加動作應該在將右邊零件元素指定到排序後的陣列之後進行。

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

你應該在 `while` 迴圈中將 `sorted_index` 的當前值增加 `1`。請注意，增加應該在將右側零件元素指定到已排序陣列之後進行。

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
