---
id: 6577254891048c8f2c19e961
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

在你的函式（程式）主體底部，再次呼叫 `merge_sort()` 函式（程式）。這次，將 `right_part` 作為引數傳入該函式（程式）呼叫。

# --hints--

你應該在函式（程式）主體底部使用引數 `right_part` 呼叫 `merge_sort()`。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
