---
id: 657f3a661730891aa64f3568
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

使用切片語法擷取 `array` 的右半部，並將其指定到名為 `right_part` 的變數。

# --hints--

你應該有一個名為 `right_part` 的變數

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

你應該使用切片語法來擷取陣列的右半部並將其指定到 `right_part` 變數。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
