---
id: 612e813b3ba67633222cbe54
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Kumbuka kuwa sifa ya kipengele `class` inaweza kuwa na maadili mengi. Ili kutofautisha funguo zako nyeupe na funguo zako za rangi nyeusi, utaongeza thamani ya pili ya `class` ya `black--key`. Ongeza hii kwa vipengele vyako vya pili, vya tatu, vya tano, vya sita, na vya saba vya `.key`.

# --hints--

Kipengele chako cha pili cha `.key` kinapaswa pia kuwa na `class` ya `black--key`.

```js
const key = document.querySelectorAll('.key')?.[1];
assert(key?.className?.includes('black--key'));
```

Kipengele chako cha tatu cha `.key` kinapaswa kuwa na `black--key` katika `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert(third?.classList?.contains('black--key'));
```

Kipengele chako cha tano cha `.key` kinapaswa kuwa na `black--key` katika `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert(fifth?.classList?.contains('black--key'));
```

Kipengele chako cha sita cha `.key` kinapaswa kuwa na `black--key` katika `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert(sixth?.classList?.contains('black--key'));
```

Kipengele chako cha saba cha `.key` kinapaswa kuwa na `black--key` katika `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert(seventh?.classList?.contains('black--key'));
```

Unapaswa kuwa na vipengele vitano vya `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 5);
```

Unapaswa kuwa na vipengele saba vya `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
