---
id: 612e89d254fe5d3df7d6693d
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Sasa umeanzisha upya mfano wa sanduku la `html`, unahitaji pia kuupitisha huo kwa vipengele vilivyo ndani yake. Ili kufanya hivyo, unaweza kuweka sifa ya `box-sizing` kuwa `inherit`, ambayo itawaambia vipengele vilivyolengwa kutumia thamani ile ile kama kipengele cha mzazi.

Pia utahitaji kulenga vipengele vya bandia, ambavyo ni maneno muhimu maalum yanayofuata kichaguzi. Vipengele viwili vya bandia utakavyotumia ni vipengele vya bandia vya `::before` na `::after`.

Kichaguzi cha `::before` huunda kipengele cha bandia ambacho ni mtoto wa kwanza wa kipengele kilichochaguliwa, wakati kichaguzi cha `::after` huunda kipengele cha bandia ambacho ni mtoto wa mwisho wa kipengele kilichochaguliwa. Vipengele hivi vya bandia mara nyingi hutumika kuunda maudhui ya urembo, ambayo utaona baadaye katika mradi huu.

Kwa sasa, tengeneza kichaguzi cha CSS kulenga vipengele vyote vyenye `*`, na jumuisha vipengele vya bandia vya `::before` na `::after`. Weka sifa ya `box-sizing` kuwa `inherit`.

# --hints--

Unapaswa kuwa na kichaguzi cha `*, ::before, ::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Kichaguzi chako cha `*, ::before, ::after` kinapaswa kuwa na sifa ya `box-sizing` ikiwa imewekwa kuwa `inherit`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
