---
id: 612e83ec2eca1e370f830511
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Ongeza kipengele cha `link` ndani ya kipengele chako cha `head`. Mpe sifa ya kipengele ya `rel` iliyowekwa kuwa `stylesheet` na sifa ya kipengele ya `href` iliyowekwa kuwa `styles.css`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Hupaswi kubadilisha lebo zako za `head` zilizopo. Hakikisha hujafuta lebo ya kufunga.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Unapaswa kuwa na kipengele kimoja cha `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `rel` yenye thamani `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `href` yenye thamani `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
