---
id: 66731cd027ef3acb155669f5
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

In this workshop, you will learn how to work with basic HTML elements like headings, paragraphs, and lists by building a cat photo app.

O primeiro elemento que você vai aprender é o `h1`. O elemento `h1` é um título, usado para o cabeçalho de uma página web.

```html
<h1>This is a main heading</h1>
```

Add an `h1` element with the text of `CatPhotoApp` and watch the change in the browser preview.

Quando estiver pronto, pressione o botão "Verifique seu código" para ver se está correto.

# --hints--

O texto `CatPhotoApp` deve estar presente no código. Não se esqueça de verificar a ortografia.

```js
assert(code.match(/catphotoapp/i));
```

O elemento `h1` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

O elemento `h1` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/h1\>/));
```

O elemento `h1` deve conter o texto `CatPhotoApp`. Você omitiu o texto, digitou o texto incorretamente ou ele não está entre as tags de abertura e fechamento do elemento `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Parece que você está usando uma extensão de navegador que está modificando a página. Não se esqueça de desativar todas as extensões do navegador.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
