---
id: 66731cd027ef3acb155669f5
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você aprenderá como trabalhar com elementos básicos do HTML, como títulos, parágrafos e listas, construindo uma aplicação de fotos de gatos.

O primeiro elemento que você vai aprender é o `h1`. O elemento `h1` é um título, usado para o cabeçalho de uma página web.

```html
<h1>This is a main heading</h1>
```

Adicione um elemento `h1` com o texto de `CatPhotoApp` e veja a mudança na pré-visualização do navegador.

Quando estiver pronto, pressione o botão "Verifique seu código" para ver se está correto.

# --hints--

O texto `CatPhotoApp` deve estar presente no código. Não se esqueça de verificar a ortografia.

```js
assert.match(code, /catphotoapp/i);
```

O elemento `h1` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

O elemento `h1` deve ter uma tag de fechamento. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

O elemento `h1` deve conter o texto `CatPhotoApp`. Você omitiu o texto, digitou o texto incorretamente ou ele não está entre as tags de abertura e fechamento do elemento `h1`.

```js
assert.equal(document.querySelector('h1').innerText.toLowerCase(), 'catphotoapp');
```

Parece que você está usando uma extensão de navegador que está modificando a página. Não se esqueça de desativar todas as extensões do navegador.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
