---
id: 6690b9a24f5b0f1300040c76
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

O elemento `p` é usado para criar um parágrafo de texto nos sites. Crie um elemento `p` abaixo do elemento `h2` e dê a ele o seguinte texto:

`See more cat photos in our gallery.`

# --hints--

O elemento `p` deve ter uma tag de abertura. As tags de abertura têm a seguinte sintaxe: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

O elemento `p` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert.match(code, /<\/p\>/);
```

O texto do elemento `p` deve ser `See more cat photos in our gallery.` Você omitiu o texto ou tem um erro de digitação.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /see more cat photos in our gallery\.?$/i);
```

O elemento `p` deve estar abaixo do elemento `h2`. Eles estão na ordem errada.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```
