---
id: 66984eef66cc1752d93b53f5
title: Passo 36
challengeType: 0
dashedName: step-36
---

# --description--

Observe que tudo o que você adicionou à página até agora está dentro do elemento `body`. Todos os elementos de conteúdo da página que devem ser renderizados na página ficam dentro do elemento `body`. No entanto, outras informações importantes vão dentro do elemento `head`.

O elemento `head` é usado para conter metadados sobre o documento, como seu título, links para folhas de estilo e scripts. Os metadados são informações sobre a página que não são exibidas diretamente na página.

Adicione um elemento `head` acima do elemento `body`.

# --hints--

Você excluiu o elemento `body` ou está faltando nele uma tag de abertura ou de fechamento.

```js
assert(document.querySelector('body') && code.match(/<\/body>/));
```

O elemento `head` deve ter uma tag de abertura. As tags de abertura têm a seguinte sintaxe: `<elementName>`.

```js
assert(code.match(/\<head\>/));
```

O elemento `head` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/\<\/head\>/));
```

O elemento `head` deve estar acima do elemento `body`. Você o colocou em um lugar diferente.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/\<\/head\>\<body\>/));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>

  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```


