---
id: 68a23004e96a105df3a21fc4
title: Costruisci un assistente per il conteggio delle carte
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

Nel gioco da casinò Blackjack, un giocatore può capire se ha un vantaggio sulla mano successiva rispetto al banco tenendo traccia del numero relativo di carte alte e basse rimaste nel mazzo. Questo si chiama conteggio delle carte.

Avere più carte alte rimaste nel mazzo favorisce il giocatore. Quando il conteggio è positivo, il giocatore dovrebbe puntare alto. Quando il conteggio è zero o negativo, il giocatore dovrebbe puntare basso.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti usare `let` per dichiarare una variabile globale chiamata `count` e impostarla su `0`.
1. Dovresti avere una funzione chiamata `cardCounter`.
1. La funzione `cardCounter` dovrebbe ricevere un parametro `card` che può essere un numero o una stringa.
    - Per valori tra `2` e `10`, il parametro `card` sarà un numero.
    - Per tutti gli altri valori, il parametro `card` sarà una stringa.
1. La funzione `cardCounter` dovrebbe modificare la variabile globale `count` in base a certi criteri.
1. La variabile globale `count` dovrebbe aumentare di `1` per le carte `2`, `3`, `4`, `5` o `6`.
1. La variabile globale `count` dovrebbe rimanere invariata per le carte `7`, `8`, `9`.
1. La variabile globale `count` dovrebbe diminuire di `1` per le carte `10`, `"J"`, `"Q"`, `"K"`, `"A"`.
1. La funzione `cardCounter` dovrebbe restituire una stringa con il conteggio attuale e la stringa `Bet` se il conteggio è positivo.
1. La funzione `cardCounter` dovrebbe restituire una stringa con il conteggio attuale e la stringa `Hold` se il conteggio è minore o uguale a `0`.
1. Nell’output della funzione, il conteggio attuale e la decisione del giocatore (`Bet` o `Hold`) dovrebbero essere separati da uno spazio. Per esempio, `-3 Hold`.

# --hints--

Dovresti usare `let` per dichiarare una variabile globale chiamata `count` e impostarla su `0`. 

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

Dovresti avere una funzione chiamata `cardCounter`. 

```js
assert.isFunction(cardCounter);
```

La tua funzione dovrebbe restituire il valore di `count` e il testo (`Bet` o `Hold`) con un carattere spazio tra loro.

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

Dopo le carte `2`, `3`, `4`, `5`, chiamare `cardCounter(6)` dovrebbe restituire la stringa `5 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

Dopo le carte `7`, `8`, chiamare `cardCounter(9)` dovrebbe restituire la stringa `0 Hold`.

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

Dopo le carte `10`, `"J"`, `"Q"`, `"K"`, chiamare `cardCounter("A")` dovrebbe restituire la stringa `-5 Hold`.

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

Dopo le carte `3`, `7`, `"Q"`, `8`, chiamare `cardCounter("A")` dovrebbe restituire la stringa `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

Dopo le carte `2`, `"J"`, `9`, `2`, chiamare `cardCounter(7)` dovrebbe restituire la stringa `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

Dopo le carte `2`, `2`, chiamare `cardCounter(10)` dovrebbe restituire la stringa `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

Dopo le carte `3`, `2`, `"A"`, `10`, chiamare `cardCounter("K")` dovrebbe restituire la stringa `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
