---
id: 694bf6eb530e19a9c48c59f2
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Agora complete a frase concatenando a string ` years old` ao final de `employee_info`. Lembre-se de incluir um espaço no início da sua string.

Finalmente, imprima `employee_info`.

# --hints--

Você deve concatenar `' years old'` ao final da variável `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

Você deve imprimir `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
