---
id: 694bf6eb530e19a9c48c59f5
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

Ao trabalhar com strings, você frequentemente precisará extrair uma porção específica de uma string. Isso é chamado de **slicing**.

A sintaxe é `string[start:stop]`, onde:

* `start` é o índice onde o slice começa (**inclusivo**).
* `stop` é o índice onde o slice termina (**exclusivo**).

Por exemplo, se `text = 'Python'`, então `text[0:2]` retorna `'Py'`.

Defina `employee_code` como `'DEV-2026-JD-001'`. Depois disso, crie uma variável `department` e atribua a ela o slice de `employee_code` do índice `0` até `3`. Em seguida, imprima `department` no terminal.

# --hints--
Você deve ter uma variável chamada `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

Você deve atribuir a string `DEV-2026-JD-001` à sua variável `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

Você deve ter uma variável chamada `department`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

Você deve fatiar os três primeiros caracteres de `employee_code` e atribuí-los à sua variável `department`.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

Você deve imprimir `department` no terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
