---
id: 694bf6eb530e19a9c48c59f1
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Agora, seu endereço parece incompleto. Você também quer adicionar o número do apartamento onde o funcionário mora, então você deve modificar a variável.  

Quando você quiser adicionar conteúdo ao final de uma variável string existente, pode usar o operador de **atribuição aumentada** ,`+=`. Isso é mais curto do que escrever `var = var + 'new text'`. Por exemplo:  

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

Lembre-se de que strings são imutáveis, portanto essa operação não altera a string original. Em vez disso, ela cria uma nova string e a reatribui à variável.

Use o operador `+=` para adicionar a string `, Apartment 4B` à sua variável `address`.

# --hints--

Você deve usar o operador `+=` para adicionar a string `, Apartment 4B` à variável `address`.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
