---
id: 694bf6eb530e19a9c48c59f4
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Concatenar muitas strings usando `+` e converter números usando `str()` pode ficar confuso e difícil de ler.

Python 3.6 introduziu as **f-strings** para resolver isso. Ao adicionar a letra `f` antes da aspas de abertura, você pode colocar variáveis e expressões dentro de campos de substituição representados por chaves `{}`. Por exemplo:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

Crie uma variável `employee_card` e atribua a ela uma f-string que exibe `Employee:` seguido de um espaço e do valor da variável `full_name`.

# --hints--

Você deve ter uma variável chamada `employee_card`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

Sua variável `employee_card` deve ter o valor `f'Employee: {full_name}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
