---
id: 694bf3af05d10fd89dab05cc
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Em Python, você pode combinar strings usando o operador `+`. Isso é chamado de concatenação. Aqui está um exemplo:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

Crie uma variável `full_name` concatenando `first_name` e `last_name`. Depois imprima `full_name`.

# --hints--

Você deve ter uma variável chamada `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

Você deve atribuir uma string formada pela concatenação de `first_name` e `last_name` à sua variável `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

Você deve imprimir a variável `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
