---
id: 694bf6ea530e19a9c48c59f0
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Se você concatenar duas strings como `'John' + 'Doe'`, o resultado é `'JohnDoe'` sem espaço. Para corrigir isso, você precisa concatenar uma string contendo um espaço (`' '`) entre elas.


Atualize sua variável `full_name` para que ela concatene `first_name`, um espaço e `last_name`.

# --hints--

Você deve adicionar um espaço (`' '`) entre `first_name` e `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
