---
id: 694bf6eb530e19a9c48c59f6
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Você pode fatiar de qualquer parte da string, não apenas do início. E isso é útil em muitos casos.

Crie uma variável `year_code` e atribua a ela o slice de `employee_code` do índice `4` ao `8`. Isso irá extrair `2026`.

Então crie uma variável `initials` e atribua a ela o slice de `employee_code` do índice `9` até `11`. Isso irá extrair `JD`.

Finalmente, imprima ambas as variáveis no terminal.

# --hints--

Você deve ter uma variável chamada `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

Você deve fatiar `employee_code` do índice `4` ao `8` e atribuir o resultado à sua variável `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

Você deve ter uma variável chamada `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

Você deve fatiar `employee_code` do índice `9` ao `11` e atribuir o resultado à sua variável `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

Você deve imprimir `year_code` no terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

Você deve imprimir `initials` no terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
