---
id: 694bf6eb530e19a9c48c59f7
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Você também pode usar números negativos para fatiar a partir do final de uma string. Por exemplo, `-1` se refere ao último caractere, `-2` se refere ao penúltimo caractere e assim por diante.

Para obter os últimos três caracteres de uma string, você pode usar `string[-3:]`. Note como o parâmetro stop é omitido após os dois pontos. Isso significa que o fatiamento deve prosseguir até o limite da string.

Crie uma variável chamada `last_three`. Use indexação negativa para extrair os últimos três caracteres de `employee_code` (que representam o número de identificação). Por fim, imprima `last_three` no terminal.

Com isso, a oficina do cartão de funcionário está completa.

# --hints--

Você deve ter uma variável chamada `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

Você deve fatiar os últimos três caracteres de `employee_code` usando indexação negativa e atribuí-los à sua variável `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

Você deve imprimir `last_three` no terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2024-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
