---
id: 694bf6eb530e19a9c48c59f3
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Agora você vai usar a função `str()` mais uma vez. Assim como com age, você deve converter qualquer variável numérica em uma string antes de concatená-la com outro texto.

Crie uma variável chamada `experience_years` e atribua a ela o inteiro `5`.

Então, crie uma variável `experience_info`. Atribua a ela uma string formada pela concatenação de `'Experience: '`, a variável `experience_years` (convertida para string) e `' years'`. Imprima o resultado no terminal.

# --hints--

Você deve ter uma variável chamada `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

Você deve atribuir o inteiro `5` à sua variável `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

Você deve ter uma variável chamada `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

Você deve atribuir uma string formada pela concatenação de `'Experience: '`, `str(experience_years)` e `' years'` à sua variável `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

Você deve imprimir `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
