---
id: 695c96f82af4fb78e59fa9ed
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Strings são sequências de caracteres usadas para armazenar dados de texto. Como você deve se lembrar das lições anteriores, você pode criar uma string colocando o texto entre aspas simples (`'`) ou duplas (`"`). Por exemplo:

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

Crie uma variável `first_name` que armazena a string `John` e uma variável `last_name` que armazena a string `Doe`. Então imprima `first_name` e `last_name`.

# --hints--

Você deve ter uma variável chamada `first_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

Você deve atribuir a string `John` à sua variável `first_name`. Lembre-se de colocar o texto entre aspas simples ou duplas.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

Você deve ter uma variável chamada `last_name`.

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

Você deve atribuir a string `Doe` à sua variável `last_name`. Lembre-se de colocar o texto entre aspas simples ou duplas.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

Você deve imprimir sua variável `first_name` no terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

Você deve imprimir sua variável `last_name` no terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
