---
id: 695c9f21a8b1d53cf1b97841
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Agora crie uma variável chamada `employee_age` e atribua a ela o inteiro `28`. 


# --hints--

Você deve ter uma variável chamada `employee_age`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

Você deve atribuir o inteiro `28` à sua variável `employee_age`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
