---
id: 695c9c7c61d446f1451e0ae7
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Em seguida, crie uma variável `address` para armazenar o endereço do funcionário. Atribua a ela a string `123 Main Street` e finalmente imprima `address`.

# --hints--

Você deve ter uma variável chamada `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

Você deve atribuir a string `123 Main Street` à sua variável `address`. Lembre-se de colocar o texto entre aspas simples ou duplas.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

Você deve imprimir a variável `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```

