---
id: 695df6480de23758af5a29cd
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Agora, você quer criar uma string que exiba a idade do funcionário.

Comece criando uma variável `employee_info` e atribua a ela o resultado da concatenação:

- a variável `full_name`.
- uma string consistindo dos caracteres `is` precedida e seguida por um espaço.

# --hints--

Você deve ter uma variável chamada `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

Você deve atribuir `full_name + ' is '` à sua variável `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
