---
id: 695ea1e55c0e619e597fc5c9
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Como você pode ver, o Python gerou um `TypeError: can only concatenate str (not "int") to str`. Isso acontece porque o Python não permite que você concatene texto (strings) e números (inteiros) diretamente.

Para corrigir isso, você deve converter o número para uma string primeiro usando a função `str()`, que retorna a versão em string de um objeto:

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

Atualize sua atribuição `employee_info` para converter `employee_age` em uma string usando `str(employee_age)`.

# --hints--

Você deve modificar a atribuição da sua variável `employee_info` convertendo `employee_age` em uma string com `str(employee_age)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
