---
id: 695fd3dde03273875e5d6709
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Atualize a atribuição `employee_info` para também concatenar `employee_age` no final.

Depois de fazer isso, você verá um `TypeError` no terminal. Na próxima etapa, você trabalhará para corrigi-lo.

# --hints--

Você deve atualizar a atribuição `employee_info` para concatenar `employee_age` no final.

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
