---
id: 69727ea5c9606bedfe36d3d2
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Agora é hora de adicionar os detalhes finais ao cartão.

Crie uma variável chamada `position` com o valor da string `Data Analyst` e uma variável chamada `salary` com o valor do inteiro `75000`.

Em seguida, atualize sua f-string `employee_card` para incluir o cargo e o salário. Deve seguir este formato exato: `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. Substitua os espaços reservados pelas variáveis correspondentes.

Finalmente, imprima `employee_card` para ver o resultado.

# --hints--

Você deve ter uma variável chamada `position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

Você deve atribuir a string `Data Analyst` à sua variável `position`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

Você deve ter uma variável chamada `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Você deve atribuir o inteiro `75000` à sua variável `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Você deve atualizar `employee_card` para usar a f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

Você deve imprimir `employee_card` no terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
