---
id: 6980d3ecfe800d8a54130f30
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Atualmente, `employee_card` só mostra o nome do funcionário. Agora você vai adicionar mais informações a ele. 

Atualize a atribuição de `employee_card` para incluir a idade do funcionário. A string final deve ser assim: `Employee: [name] | Age: [age]` com `[name]` substituído pelo nome do funcionário e `[age]` substituído pela idade do funcionário.

# --hints--

Você deve atualizar `employee_card` para usar a f-string `f'Employee: {full_name} | Age: {employee_age}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
