---
id: 699d20623e57e8ffee885459
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

`parseCard` було протестовано і воно працює правильно. Змінна `cardResult` більше не потрібна.

Почніть із видалення рядка `const cardResult` та рядка `console.log(cardResult)`. Потім створіть функцію з ім’ям `parseCatalog`, яка приймає `rawCards` як параметр. Поки що оголосіть порожній масив з ім’ям `catalog` всередині функції та поверніть його.

Під функцією викличте `parseCatalog(rawCatalogCards)`, призначте результат змінній з ім’ям `catalog` і виведіть у консоль `catalog.length`.

# --hints--

Ви повинні оголосити функцію з назвою `parseCatalog` з параметром `rawCards`.

```js
const explorer = await __helpers.Explorer(code);
const { parseCatalog: parseCatalogFn } = explorer.allFunctions;
assert.exists(parseCatalogFn);
assert.equal(parseCatalogFn.parameters[0]?.toString(), 'rawCards');
```

Ви повинні оголосити змінну з назвою `catalog`, ініціалізувати її як порожній масив, а потім повернути її.

```js
const explorer = await __helpers.Explorer(code);
const { parseCatalog: parseCatalogFn } = explorer.allFunctions;
assert.exists(parseCatalogFn);
const { catalog } = parseCatalogFn.variables;
assert.exists(catalog);
const catalogValue = eval(catalog.value?.toString());
assert.isArray(catalogValue);
assert.isEmpty(catalogValue);
assert.isTrue(parseCatalogFn.hasReturn('catalog'));
```

`catalog` має бути оголошеною і призначеною `parseCatalog(rawCatalogCards)`.

```js
const explorer = await __helpers.Explorer(code);
const { catalog } = explorer.variables;
assert.exists(catalog);
assert.isTrue(catalog.value.matches('parseCatalog(rawCatalogCards)'));
```

Ви маєте видалити змінну `cardResult` та пов’язаний виклик `console.log()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bcardResult\b/);
```

Ви повинні вивести `catalog.length` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*catalog\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

--fcc-editable-region--
const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
--fcc-editable-region--
```
