---
id: 699d1f8e51fcdc8230f779fb
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите систему каталогу бібліотеки спадщини. Ви будете аналізувати сирі дані про книги, шукати та групувати записи, перевіряти якість даних і експортувати результати у різні формати.

Кожна картка каталогу зберігається як один рядок у форматі `"Title | Author | Year | Location"`, з полями, розділеними символом вертикальної риски (`|`).

Протягом усього практичного заняття ви будете використовувати масив цих рядків, щоб створювати та тестувати свої функції каталогу. Створіть масив з назвою `rawCatalogCards` і додайте до нього наступні два рядки:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

Ви повинні оголосити змінну з назвою `rawCatalogCards`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` має бути масивом.

```js
assert.isArray(rawCatalogCards);
```

Перший елемент `rawCatalogCards` має бути `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

Другий елемент `rawCatalogCards` має бути `"The Shining | King, Stephen | 1977 | Shelf K1"`.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
