---
id: 699d20623e57e8ffee885450
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Залишок карток каталогу додано. Зверніть увагу на три записи ближче до кінця з відсутніми полями; ви опрацюєте їх на наступному кроці.

`parseCard` прийматиме сирий рядок каталогу і повертатиме структурований об’єкт із чотирма пропами: `title`, `author`, `year` і `location`. Наразі створіть функцію `parseCard` з параметром на ім’я `rawString`.

# --hints--

У вас має бути функція з ім’ям `parseCard`.

```js
assert.isFunction(parseCard);
```

`parseCard` має бути функцією з параметром на ім’я `rawString`.

```js
const regex = __helpers.functionRegex('parseCard', ['rawString']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

--fcc-editable-region--

--fcc-editable-region--
```
