---
id: 697fc395f3d85f14b05cd9ca
title: Відновлення зв’язного наративу з масиву фрагментів історії
challengeType: 26
dashedName: lab-story-fragment-restoration
---

# --description--

У цій лабораторній роботі ви відновите зв’язний наратив із пошкодженого масиву фрагментів історії.

Ви попрактикуєтеся у використанні циклів, реалізуючи базові алгоритми роботи з масивами з нуля.

Ви працюватимете з масивами об’єктів-фрагментів історії. Кожен об’єкт-фрагмент має такі властивості:

| Властивість | Опис | Приклад значення |
|-|-|-|
| `id` | Позитивне ціле число, що вказує позицію фрагмента в історії | `3` |
| `text` | Безпосередній вміст історії | `"and I use Arch btw.\""` |

У цій лабораторній роботі вам надано попередньо заповнений масив під назвою `shuffledFragments`.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача**

1. Ви не повинні змінювати попередньо заповнений масив `shuffledFragments`.
2. Ви повинні створити функцію з назвою `compactFragments`, яка приймає масив фрагментів і повертає новий масив з видаленими всіма елементами undefined. Якщо функція видаляє будь-які елементи undefined, вона має вивести повідомлення в консоль. Повідомлення має починатися з префікса `[COMPACTED]`.
3. Ви повинні оголосити змінну з назвою `compactedShuffledFragments` і присвоїти їй результат виклику `compactFragments` з масивом `shuffledFragments`.
4. Ви повинні створити функцію з назвою `sortFragments`, яка приймає масив фрагментів без елементів `undefined` і повертає новий масив, відсортований за властивістю `id` за зростанням, зберігаючи фрагменти з однаковим `id` у їхньому початковому порядку. Ви не повинні використовувати вбудований метод `sort` мови JavaScript.
5. Ви повинні оголосити змінну з назвою `sortedFragments` і присвоїти їй результат виклику `sortFragments` з масивом `compactedShuffledFragments`.
6. Ви повинні створити функцію з назвою `dedupeFragments`, яка приймає відсортований масив фрагментів і повертає новий масив з видаленими дублікатами, зберігаючи лише перше входження. Ви повинні визначати дублікати як два або більше фрагментів з однаковим `id`. Для кожного `id`, який дедуплікується, функція має вивести повідомлення в консоль. Повідомлення має починатися з префікса `[DEDUPED]`.
7. Ви повинні оголосити змінну з назвою `dedupedFragments` і присвоїти їй результат виклику `dedupeFragments` з масивом `sortedFragments`.
8. Ви повинні створити функцію з назвою `fillMissingFragments`, яка приймає відсортований масив фрагментів і повертає новий масив, у якому відсутні фрагменти заповнені об’єктами-заповнювачами. Ви повинні визначати відсутні фрагменти як розриви в послідовності між найменшим і найбільшим `id`. Об’єкти-заповнювачі повинні мати формат `{ id: missingId, text: "[...]" }`. Для кожного доданого заповнювача функція має вивести повідомлення в консоль. Повідомлення має починатися з префікса `[FILLED]`.
9. Ви повинні оголосити змінну з назвою `filledFragments` і присвоїти їй результат виклику `fillMissingFragments` з масивом `dedupedFragments`.
10. Ви повинні створити функцію з назвою `assembleStory`, яка приймає відсортований масив фрагментів і повертає один рядок, що містить тексти всіх фрагментів, розділені символами нового рядка.
11. Ви повинні використати `assembleStory` з вашими `filledFragments`, щоб показати повну історію в консолі.
12. Ваші функції `compactFragments`, `sortFragments`, `dedupeFragments`, `fillMissingFragments` та `assembleStory` не повинні змінювати масив, з яким їх викликають.

**Приклад**

Ось приклад масиву, що містить фрагменти історії:

```js
const exampleArray = [
  { id: 3, text: "and I use Arch btw.\"" },
  ,
  { id: 1, text: "Naomi said:" },
  { id: 3, text: "and I use Arch btw.\"" },
];
```

Після відновлення історії з `exampleArray` вона виглядала б так:

```md
Naomi said:
[...]
and I use Arch btw."
```

# --hints--

Ви не повинні змінювати попередньо заповнений масив `shuffledFragments`.

```js
const shuffledFragments__unchanged = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

assert.deepEqual(shuffledFragments, shuffledFragments__unchanged);
```

У вас має бути функція `compactFragments`.

```js
assert.isFunction(compactFragments);
```

`compactFragments([{ id: 1, text: "Hello" }, undefined])` має вивести повідомлення, що починається з `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Коли функція `compactFragments` викликається з масивом, що містить елементи `undefined`, вона має вивести повідомлення в консоль, що починається з `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" }, undefined, { id: 11, text: "bar" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" },,])` має вивести повідомлення, що починається з `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" },,]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Коли функція `compactFragments` викликається з масивом, що містить порожні слоти, вона має вивести повідомлення в консоль, що починається з `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" },,{ id: 11, text: "bar" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Коли функція `compactFragments` викликається з масивом `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`, вона не має виводити жодних повідомлень у консоль, що починаються з `[COMPACTED]`.

```js
const shuffledFragments__withoutUndefined = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const spy = __helpers.spyOn(console, 'log');
compactFragments(shuffledFragments__withoutUndefined);
assert.isFalse(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }])` має повернути `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`.

```js
assert.deepEqual(
  compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }]),
  [{ id: 1, text: "Hello" }, { id: 2, text: "World" }]
);
```

`compactFragments` не має змінювати масив, який їй передають.

```js
const input = [{ id: 1, text: "a" }, undefined, { id: 2, text: "b" }];
const copy = structuredClone(input);
compactFragments(input);
assert.deepEqual(input, copy);
```

`compactedShuffledFragments` має бути масивом без елементів `undefined`.

```js
assert.isArray(compactedShuffledFragments);
assert.isFalse(compactedShuffledFragments.some(f => f === undefined));
```

У вас має бути функція `sortFragments`.

```js
assert.isFunction(sortFragments);
```

`sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }])` має повернути `[{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]`.

```js
assert.deepEqual(
  sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }]),
  [{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]
);
```

`sortFragments` не має використовувати вбудований метод `.sort()`.

```js
function isBuiltInSortUsed() {
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => { sortUsed = true; };
  try {
    sortFragments([{ id: 2, text: "b" }, { id: 1, text: "a" }]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

`sortFragments` не має змінювати масив, який їй передають.

```js
const input = [{ id: 3, text: "c" }, { id: 1, text: "a" }];
const copy = structuredClone(input);
sortFragments(input);
assert.deepEqual(input, copy);
```

`sortedFragments` має мати кожен `id` елемента меншим або рівним `id` наступного елемента.

```js
for (let i = 0; i < sortedFragments.length - 1; i++) {
  assert.isAtMost(sortedFragments[i].id, sortedFragments[i + 1].id);
}
```

У вас має бути функція `dedupeFragments`.

```js
assert.isFunction(dedupeFragments);
```

`dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }])` має повернути `[{ id: 1, text: "first" }, { id: 2, text: "second" }]`.

```js
assert.deepEqual(
  dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]),
  [{ id: 1, text: "first" }, { id: 2, text: "second" }]
);
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }])` має вивести повідомлення `[DEDUPED]` для кожного дубльованого `id`.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 2);
spy.restore();
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }])` має вивести повідомлення `[DEDUPED]`, коли останній `id` дубльований.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 1);
spy.restore();
```

Коли `dedupeFragments` викликається з масивом, що містить дублікати, вона має вивести повідомлення, що починається з `[DEDUPED]` для кожного дубльованого `id`.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[DEDUPED\]/)));
spy.restore();
```

`dedupeFragments` не має змінювати масив, який їй передають.

```js
const input = [{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }];
const copy = structuredClone(input);
dedupeFragments(input);
assert.deepEqual(input, copy);
```

`dedupedFragments` не має містити дубльованих значень `id`.

```js
const ids = dedupedFragments.map(f => f.id);
assert.strictEqual(ids.length, new Set(ids).size);
```

У вас має бути функція `fillMissingFragments`.

```js
assert.isFunction(fillMissingFragments);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` має повернути `[{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]`.

```js
assert.deepEqual(
  fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]),
  [{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]
);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` має вивести повідомлення, що починається з `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

Коли `fillMissingFragments` додає заповнювач, вона має вивести повідомлення, що починається з `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 5, text: "e" }, { id: 8, text: "h" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

`fillMissingFragments` не має змінювати масив, який їй передають.

```js
const input = [{ id: 1, text: "a" }, { id: 3, text: "c" }];
const copy = structuredClone(input);
fillMissingFragments(input);
assert.deepEqual(input, copy);
```

`filledFragments` не має містити розривів у послідовності `id`.

```js
for (let i = 0; i < filledFragments.length - 1; i++) {
  assert.strictEqual(filledFragments[i + 1].id, filledFragments[i].id + 1);
}
```

У вас має бути функція `assembleStory`.

```js
assert.isFunction(assembleStory);
```

`assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }])` має повернути рядок `"Hello\nWorld"`.

```js
assert.strictEqual(
  assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }]),
  "Hello\nWorld"
);
```

`assembleStory` не має змінювати масив, який їй передають.

```js
const input = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const copy = structuredClone(input);
assembleStory(input);
assert.deepEqual(input, copy);
```

Ви повинні викликати `console.log` з результатом `assembleStory(filledFragments)`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(.*assembleStory\s*\(\s*filledFragments\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];
```

# --solutions--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

const compactFragments = (fragments) => {
  let compacted = [];
  let numberOfUndefined = 0;

  for (const fragment of fragments) {
    if (fragment === undefined) {
      numberOfUndefined++;
    } else {
      compacted.push(fragment);
    }
  }

  if (numberOfUndefined > 0) {
    console.log(`[COMPACTED] Removed ${numberOfUndefined} undefined fragment(s).`);
  }

  return compacted;
};

const compactedShuffledFragments = compactFragments(shuffledFragments);

const sortFragments = (fragments) => {
  let sorted = [...fragments];

  for (let i = 0; i < sorted.length; i++) {
    for (let j = i; j > 0; j--) {
      if (sorted[j - 1].id > sorted[j].id) {
        swap(sorted, j - 1, j);
      }
    }
  }

  function swap(array, index1, index2) {
    [sorted[index1], sorted[index2]] = [sorted[index2], sorted[index1]];
  }

  return sorted;
};

const sortedFragments = sortFragments(compactedShuffledFragments);

const dedupeFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let deduped = [shallowCopy.shift()];
  let duplicates = [];
  let currentIdOccurrence = 1;

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = deduped.at(-1);
    if (fragmentToInsert.id === previousFragment.id) {
      currentIdOccurrence++;
    } else {
      deduped.push(fragmentToInsert);

      if (currentIdOccurrence > 1) {
        duplicates.push({ id: previousFragment.id, occurrence: currentIdOccurrence });
      }

      currentIdOccurrence = 1;
    }
  }

  if (currentIdOccurrence > 1) {
    duplicates.push({ id: deduped.at(-1).id, occurrence: currentIdOccurrence });
  }

  for (const duplicate of duplicates) {
    console.log(`[DEDUPED] Fragment #${duplicate.id} appeared ${duplicate.occurrence} times, kept first occurrence.`);
  }

  return deduped;
};

const dedupedFragments = dedupeFragments(sortedFragments);

const fillMissingFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let filled = [shallowCopy.shift()];

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = filled.at(-1);
    const numberOfMissingIDs = fragmentToInsert.id - 1 - previousFragment.id;
    for (let i = 1; i <= numberOfMissingIDs; i++) {
      const missingId = previousFragment.id + i;
      filled.push({ id: missingId, text: "[...]" });
      console.log(`[FILLED] Fragment #${missingId} missing, placeholder inserted.`)
    }
    filled.push(fragmentToInsert);
  }

  return filled;
};

const filledFragments = fillMissingFragments(dedupedFragments);

const assembleStory = (sortedFragments) => {
  const texts = [];
  for (const fragment of sortedFragments) {
    texts.push(fragment.text);
  }
  return texts.join("\n");
};

console.log("\nHere is the assembled story:\n");
console.log(assembleStory(filledFragments));
```
