---
id: 66b62d0ad68488dd76228d6c
title: 랜덤 배경색 변경기 디버그하기
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot이 JavaScript를 배우면서 자신만의 랜덤 배경색 변경기를 만들려고 했습니다. 하지만 몇 가지 실수를 했습니다.   

**목표:** 아래 사용자 스토리를 충족시키고 모든 테스트를 통과시켜 실습이 제대로 작동하도록 하세요.

**사용자 스토리:**

1. `darkColorsArr` 배열의 구문 오류를 고치세요.
1. `ReferenceError` 함수 안에서 `math` 객체를 참조할 때 발생하는 `getRandomIndex`를 고치세요.
1. `randomIndex`를 가장 가까운 정수로 내림하세요.
1. `body` 변수의 값을 올바른 메서드 이름을 사용하도록 업데이트하고 관련된 `TypeError`를 고치세요.
1. `bgHexCodeSpanElement`의 `id`을 가진 요소를 올바르게 선택하도록 `bg-hex-code`를 고치세요.
1. `color` 안의 `changeBackgroundColor` 변수를 `darkColorsArr` 배열에서 무작위 색상으로 설정하도록 수정하세요.
1. `btn` 변수를 버튼 요소를 가리키도록 업데이트하세요. 올바른 `index.html` 이름은 `id`를 열어 확인하세요.
1. `click`에 `btn` 이벤트 리스너를 추가하는 줄에서 `changeBackgroundColor` 함수 참조를 메서드에 전달하도록 고치세요.
1. `darkColorsArr`를 클릭할 때마다 배경이 `#btn`에 있는 색상 중 하나로 무작위로 변경되어야 합니다.

# --hints--

`darkColorsArr` 배열의 구문 오류를 해결하세요. 배열 요소는 쉼표로 구분해야 합니다.

```js
assert.isArray(darkColorsArr);
```

`math.random()` 줄에서 대소문자 오류를 고치세요.

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

`darkColorsArr.length * Math.random()`를 가장 가까운 정수로 내림하세요.

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

`body` 변수의 값을 업데이트하여 `queryselector`가 카멜 케이스를 사용하도록 하세요.

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

`"bg-hex-code"`가 `id`의 `bg-hex-code`을 올바르게 참조하도록 고치세요. CSS에서 `id` 선택자는 해시 기호(`#`)로 시작합니다.

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

`getRandomIndex` 줄 안에서 `darkColorsArr[getRandomIndex]` 함수를 호출하세요.

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

`id` 줄의 `"#click-btn"` 이름을 `id` 파일에 있는 올바른 `index.html` 이름과 일치하도록 고치세요.

```js
assert.match(btn?.id, /btn/);
```

`changeBackgroundColor` 함수를 호출하지 마세요. 대신 `addEventListener` 메서드에 함수 참조를 전달해야 합니다.

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
