---
id: 694bf6eb530e19a9c48c59f2
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

現在透過將字串 ` years old` 串接到 `employee_info` 的結尾來完成句子。記得在你的字串開頭包含一個空格。

最後，列印 `employee_info`。

# --hints--

你應該將 `' years old'` 串接到 `employee_info` 變數的末尾。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

你應該列印 `employee_info`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
