---
id: 694bf6eb530e19a9c48c59f1
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

現在，你的位址似乎不完整。你還想為你的變數添加員工居住的公寓號碼，所以你應該修改該變數。  

當你想要將內容添加到現有字串變數的結尾時，你可以使用 **增強指派** 運算子 `+=`。這比寫 `var = var + 'new text'` 更簡短。<dfn>例如</dfn>：  

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

請記住，字串是不可變的，因此此操作不會改變原始字串。相反地，它會創建一個新的字串並將其重新指定給變數。

使用 `+=` 運算子為你的 `address` 變數添加字串 `, Apartment 4B`。

# --hints--

你應該使用 `+=` 運算子將字串 `, Apartment 4B` 添加到 `address` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
