---
id: 694bf6ea530e19a9c48c59f0
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

如果你將兩個字串串接，例如 `'John' + 'Doe'`，結果會是 `'JohnDoe'`，中間沒有空格。若要修正這個問題，你需要在它們之間串接一個包含空格的字串（`' '`）。


更新你的 `full_name` 變數，使其串接 `first_name`、一個空格和 `last_name`。

# --hints--

你應該在 `first_name` 和 `last_name` 之間添加一個空格（`' '`）。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
