---
id: 694bf6eb530e19a9c48c59f3
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

現在你將再使用一次 `str()` 函式。就像處理 age 一樣，你必須先將任何數值變數轉換成字串，才能與其他文字串接。

創建一個名為 `experience_years` 的變數並指定其整數值為 `5`。

然後，創建一個變數 `experience_info`。指定它為由串接 `'Experience: '`、`experience_years` 變數（轉換為字串）和 `' years'` 所形成的字串。將結果列印到終端機。

# --hints--

你應該有一個名為 `experience_years` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

你應該將整數 `5` 指定到你的 `experience_years` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

你應該有一個名為 `experience_info` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

你應該將由串接 `'Experience: '`、`str(experience_years)` 和 `' years'` 所形成的字串指定給你的 `experience_info` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

你應該列印 `experience_info`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
