---
id: 694bf6eb530e19a9c48c59f5
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

當處理字串時，你經常需要擷取字串的特定部分。這稱為**切片**。

語法是 `string[start:stop]`，其中：

* `start` 是切片開始的索引（**包含**）。
* `stop` 是切片結束的索引（**不包含**）。

例如，如果 `text = 'Python'`，那麼 `text[0:2]` 會得到 `'Py'`。

將 `employee_code` 指定為 `'DEV-2026-JD-001'`。之後，創建一個變數 `department` 並將 `employee_code` 從索引 `0` 到 `3` 的切片設值給它。然後將 `department` 列印到終端機。

# --hints--
你應該有一個名為 `employee_code` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

你應該將字串 `DEV-2026-JD-001` 指定到你的 `employee_code` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

你應該有一個名為 `department` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

你應該擷取 `employee_code` 的前三個字元，並將它們指定到你的 `department` 變數。

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

你應該將 `department` 列印到終端機。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
