---
id: 694bf6eb530e19a9c48c59f7
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

你也可以使用負數從字串的結尾切片。例如，`-1` 取用最後一個字元，`-2` 取用倒數第二個字元，依此類推。

要取得字串的最後 3 個字元，你可以使用 `string[-3:]`。請注意冒號後省略了 stop 參數。這表示切片應該進行到字串邊界。

創建一個名為 `last_three` 的變數。使用負索引從 `employee_code` 中擷取最後三個字元（表現 ID 編號）。最後，將 `last_three` 列印到終端機。

至此，員工卡工作坊已完成。

# --hints--

你應該有一個名為 `last_three` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

你應該使用負索引切割 `employee_code` 的最後三個字元，並將它們指定給你的 `last_three` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

你應該將 `last_three` 列印到終端機。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2024-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
