---
id: 694bf6eb530e19a9c48c59f6
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

你可以從字串的任何零件、部分切割，而不只是從開頭切割。這在許多情況下都很有幫助。

創建一個變數 `year_code` 並將 `employee_code` 中從索引 `4` 到 `8` 的切片指定給它。這將提取出 `2026`。

然後創建一個變數 `initials` 並將 `employee_code` 中從索引 `9` 到 `11` 的切片指定給它。這將擷取出 `JD`。

最後，將兩個變數都列印到終端機。

# --hints--

你應該有一個名為 `year_code` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

你應該從索引 `4` 到 `8` 擷取 `employee_code`，並將結果指定到你的 `year_code` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

你應該有一個名為 `initials` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

你應該從索引 `9` 到 `11` 擷取 `employee_code`，並將結果指定到你的 `initials` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

你應該將 `year_code` 列印到終端機。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

你應該將 `initials` 列印到終端機。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
