---
id: 694bf6eb530e19a9c48c59f4
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

使用 `+` 串接許多字串以及使用 `str()` 轉換數字會變得混亂且難以閱讀。

Python 3.6 介紹了 **f-strings** 來解決這個問題。只要在開頭引號前加上字母 `f`，你就可以將變數和表達式放在由大括號 `{}` 表述的取代欄位中。例如：

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

創建一個變數 `employee_card` 並指定一個 f-string，該字串顯示 `Employee:`，後面接一個空格和 `full_name` 變數的值。

# --hints--

你應該有一個名為 `employee_card` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

你的 `employee_card` 變數應該有值 `f'Employee: {full_name}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
