---
id: 694bf3af05d10fd89dab05cc
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

在 Python 中，你可以使用 `+` 運算子來合併字串。這稱為串接。以下是一個範例：

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

透過串接 `first_name` 和 `last_name` 來創建 `full_name` 變數。然後列印 `full_name`。

# --hints--

你應該有一個名為 `full_name` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

你應該將由串接 `first_name` 和 `last_name` 所形成的字串指定給你的 `full_name` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

你應該列印 `full_name` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
