---
id: 695c9f21a8b1d53cf1b97841
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

現在創建一個名為 `employee_age` 的變數，並將整數 `28` 指定給它。 


# --hints--

你應該有一個名為 `employee_age` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

你應該將整數 `28` 指定到你的 `employee_age` 變數。不要用引號包圍文字。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
