---
id: 695c9c7c61d446f1451e0ae7
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

接著，創建一個變數 `address` 來儲存員工的位址。將字串 `123 Main Street` 指定給它，最後列印 `address`。

# --hints--

你應該有一個名為 `address` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

你應該將字串 `123 Main Street` 指定到你的 `address` 變數。請記得將文字用單引號或雙引號括起來。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

你應該列印 `address` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```

