---
id: 695c96f82af4fb78e59fa9ed
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

字串是用來儲存文字資料的字元序列。正如你可能從先前的課程中回想起的，你可以透過將文字包裹在單引號（`'`）或雙引號（`"`）中來創建字串。例如：

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

創建一個變數 `first_name`，其儲存字串 `John`，以及一個變數 `last_name`，其儲存字串 `Doe`。然後列印 `first_name` 和 `last_name`。

# --hints--

你應該有一個名為 `first_name` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

你應該將字串 `John` 指定到你的 `first_name` 變數。請記得將文字用單引號或雙引號括起來。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

你應該有一個名為 `last_name` 的變數。

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

你應該將字串 `Doe` 指定到你的 `last_name` 變數。請記得將文字用單引號或雙引號括起來。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

你應該將你的 `first_name` 變數列印到終端機。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

你應該將你的 `last_name` 變數列印到終端機。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
