---
id: 695df6480de23758af5a29cd
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

現在，你想要創建一個顯示員工年齡的字串。

首先建立一個變數 `employee_info` 並將串接結果指定給它：

- `full_name` 變數。
- 一個由字元 `is` 組成，前後各有一個空格的字串。

# --hints--

你應該有一個名為 `employee_info` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

你應該將 `full_name + ' is '` 指定到你的 `employee_info` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
