---
id: 695ea1e55c0e619e597fc5c9
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

如你所見，Python 引發了 `TypeError: can only concatenate str (not "int") to str`。這是因為 Python 不允許你直接串接文字（字串）和數字（整數）。

為了解決此問題，你必須先使用 `str()` 函式將數字轉換為字串，該函式會傳回物件的字串版本：

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

更新你的 `employee_info` 指派，使用 `str(employee_age)` 將 `employee_age` 轉換為字串。

# --hints--

你應該修改你的 `employee_info` 變數的指派，將 `employee_age` 轉換成字串，使用 `str(employee_age)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
