---
id: 695fd3dde03273875e5d6709
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

現在嘗試將 `employee_age` 串接到你的 `employee_info` 字串末端。

完成後，你會在終端機中看到 `TypeError`。在下一步，你將著手修正它。

# --hints--

你應該將 `employee_age` 串接到 `employee_info` 的末尾。

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
