---
id: 695fcfbba648b3a96a9495a4
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

終端機中的輸出變得擁擠。在你繼續之前，是時候清理它了。

從你的程式碼中移除所有的 `print()` 述語。

# --hints--

你的程式碼中不應該有任何 `print` 呼叫。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_calls("print")) == 0`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'  
last_name = 'Doe'  
--fcc-editable-region--  
print(first_name)  
print(last_name)  
full_name = first_name + ' ' + last_name  
print(full_name)  
address = '123 Main Street'  
address += ', Apartment 4B'  
print(address)  
--fcc-editable-region--
```
