---
id: 69727ea5c9606bedfe36d3d2
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

現在是為卡片添加最後細節的時候了。

創建一個名為 `position` 的變數，其值為字串 `Data Analyst`，以及一個名為 `salary` 的變數，其值為整數 `75000`。

然後，更新你的 `employee_card` f-string 以包含職位和薪資。它應該遵循這個精確格式：`Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`。將佔位符取代為相應的變數。

最後，列印 `employee_card` 以查看結果。

# --hints--

你應該有一個名為 `position` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

你應該將字串 `Data Analyst` 指定到你的 `position` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

你應該有一個名為 `salary` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

你應該將整數 `75000` 指定到你的 `salary` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

你應該更新 `employee_card` 以使用 f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

你應該將 `employee_card` 列印到終端機。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
