---
id: 6980d3ecfe800d8a54130f30
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

目前，`employee_card` 只顯示員工的姓名。現在你將為它添加更多資訊。 

更新 `employee_card` 指派以包含員工的年齡。最終的字串應該看起來像這樣：`Employee: [name] | Age: [age]`，其中 `[name]` 取代為員工的姓名，`[age]` 取代為員工的年齡。

# --hints--

你應該更新 `employee_card` 以使用 f-string `f'Employee: {full_name} | Age: {employee_age}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
