---
id: 6909ed80e1c195f49564edac
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 조건문과 비교 연산자를 다루는 방법을 배우세요.

먼저 `truthyOrFalsy`라는 변수를 선언하고 불리언 값 `true`을 할당하세요.

# --hints--

`truthyOrFalsy`라는 변수가 있어야 합니다.

```js
assert.exists(truthyOrFalsy);
```

`true`을 `truthyOrFalsy` 변수에 할당해야 합니다.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
