---
id: 690a3f884e083a6a33b81fcf
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

마지막으로 `timmyAge` 선언을 변경하여 `16`보다 작은 숫자로 설정하세요. 그다음 `"Timmy is not old enough to drive."`가 콘솔에 출력되는 것을 볼 수 있습니다.

# --hints--

`timmyAge` 변수는 `16`보다 작은 값을 가져야 합니다.

```js
assert.isAtMost(timmyAge, 15);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--
const timmyAge = 18;
--fcc-editable-region--

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```

# --solutions--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 15;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```
