---
id: 690a0f7a5c22c15222a6aa5c
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

이제 `if`을 조건으로 사용하는 또 다른 `isTimmyAGamer` 문을 만드세요. `if` 문 본문 안에서 `"Timmy loves to play World of Warcraft."`을 콘솔에 출력하세요.

새 `if` 문의 조건이 false이므로 콘솔에 아무것도 출력되지 않는다는 점을 참고하세요.

# --hints--

코드에 `if` 문이 두 개 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

두 번째 `if` 문은 `isTimmyAGamer`을 조건으로 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

두 번째 `"Timmy loves to play World of Warcraft."` 문 안에서 `if`을 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
