---
id: 690a110a76d0496d9e7c28eb
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

`timmyAge`라는 또 다른 변수를 선언하고 `18`로 설정하세요.

# --hints--

`timmyAge`라는 변수가 있어야 합니다.

```js
assert.exists(timmyAge);
```

`18`를 `timmyAge` 변수에 할당해야 합니다.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
