---
id: 690a57fef66aac39dffc61a2
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

이제 truthy 값과 falsy 값이 무엇인지 알았으니, 변수 선언과 `console.log` 문장을 코드에서 모두 제거하세요.

# --hints--

`truthyOrFalsy` 변수가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

코드에 `console.log(Boolean(truthyOrFalsy))`가 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
