---
id: 690a06833d3e83b5226ade78
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

조건문은 조건에 따라 코드를 실행할 때 사용합니다. 조건문은 `if` 키워드와, 소괄호로 감싼 조건, 그리고 중괄호로 감싼 본문으로 구성됩니다. 다음은 `if` 문장의 예입니다:

```js
if (condition) {
  console.log("condition is truthy");
}
```

`condition`가 참 같은 값(빈 문자열이 아니거나 0이 아닌 숫자, `true` 등)일 때 `if` 문장의 본문 안의 코드가 실행됩니다.

변수 아래에 `if` 문장을 만들어 `hasDeveloperJob`를 조건으로 사용하세요. `if` 문장 본문 안에서 `"Timmy is employed as a developer."`을 콘솔에 출력하세요.

# --hints--

`if` 문장이 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

`if` 문장은 `hasDeveloperJob`를 조건으로 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

`"Timmy is employed as a developer."`을 `if` 문장 안에서 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
