---
id: 690a58cc4f828a49eab1b53a
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

`hasDeveloperJob`라는 변수를 선언하고 불리언 값 `true`을 할당하세요.

# --hints--

`hasDeveloperJob`라는 변수가 있어야 합니다.

```js
assert.exists(hasDeveloperJob);
```

`true`을 `hasDeveloperJob` 변수에 할당해야 합니다.

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
