---
id: 690a3ce03013be44d116136e
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

조건문은 `else` 절을 가질 수 있으며, 이 절은 `if` 조건이 거짓일 때 코드를 실행합니다. 다음은 `if...else` 문장의 예입니다:

```js
if (condition) {
  console.log("condition is truthy");
} else {
  console.log("condition is falsy");
}
```

기존 `else` 문장에 `if` 절을 추가하세요. `else` 절 본문 안에서 `"Timmy is not old enough to drive."`를 콘솔에 출력하세요.

# --hints--

세 번째 `if` 문장에는 `else` 절이 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{/s);
```

`"Timmy is not old enough to drive."` 절 본문 안에서 `else`를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is not old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
--fcc-editable-region--
}
--fcc-editable-region--
```
