---
id: 690a0df27ee9ee36d935d530
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

`isTimmyAGamer`라는 변수를 하나 더 선언하고 불리언 `false`을 할당하세요.

# --hints--

`isTimmyAGamer`라는 변수를 선언해야 합니다.

```js
assert.exists(isTimmyAGamer);
```

불리언 값 `false`을 `isTimmyAGamer` 변수에 할당해야 합니다.

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
