---
id: 690a3ad44874a928a2fe3077
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

비교 연산자는 값을 비교하고 더 복잡한 조건문을 작성할 수 있게 해줍니다.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

예를 들어, 위 코드는 변수 `"Happy Birthday!"`과 `today`의 값이 같을 때만 `birthday`을 콘솔에 출력합니다.

세 번째 `if` 문을 만드세요. 조건에는 `>=` 연산자를 사용해 `timmyAge`가 `16`보다 크거나 같은지 확인하세요.

새로운 `if` 문 본문 안에서 `"Timmy is old enough to drive."`을 콘솔에 출력하세요.

# --hints--

코드에 세 번째 `if` 문이 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

세 번째 `if` 문은 조건으로 `timmyAge >= 16`을 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

세 번째 `"Timmy is old enough to drive."` 문 안에서 `if`을 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
