---
id: 690a57095c4afb27bbe751b3
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

거짓 같은 값은 불리언 문맥에서 `false`으로 평가되는 값입니다. 거짓 같은 값의 예는 다음과 같습니다:

- `""` (빈 문자열)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

이 도전 과제를 진행하면서 이 모든 값을 만나게 됩니다.

지금은 `truthyOrFalsy`의 할당을 빈 문자열로 변경하세요. 그다음 `false`이 콘솔에 출력되는 것을 볼 수 있습니다.

# --hints--

`truthyOrFalsy`에 빈 문자열을 할당하세요.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
