---
id: 690a51806fb743bc896c3af7
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이전 수업에서 배운 것처럼, 불리언은 `true`와 `false` 값만 가지는 데이터 유형입니다.

참 같은 값은 불리언 문맥에서 `true`로 평가되는 값입니다. 참 같은 값의 예는 다음과 같습니다:

- 비어 있지 않은 문자열
- 0이 아닌 숫자
- 불리언 `true`

그리고 나중에 만나게 될 다른 값들입니다.

값의 참 같은지를 확인하려면 `Boolean()` 함수를 사용할 수 있습니다. 예를 들어, `Boolean("Hello World!")`는 `true`이 참 같은 값이기 때문에 `"Hello World!"`를 반환합니다.

이제 `Boolean(truthyOrFalsy)`를 호출하고 결과를 콘솔에 출력하세요.

# --hints--

`Boolean(truthyOrFalsy)`를 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
