---
id: 690a562fd611e8155599ce98
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

`truthyOrFalsy`의 할당을 `"freeCodeCamp"`로 값을 설정하여 변경하세요. 빈 문자열이 아닌 문자열은 참 같은 값이기 때문에 `true`가 콘솔에 계속 출력됩니다.

# --hints--

`truthyOrFalsy` 변수는 `"freeCodeCamp"` 값을 가져야 합니다.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
