---
id: 658246c28575d653d1f89b59
title: Schritt 33
challengeType: 20
dashedName: step-33
---

# --description--

Eine `while`-Schleife ist eine weitere Art von Schleife, die einen Teil des Codes so lange ausführt, wie eine bestimmte Bedingung `True` ist. Die Schleife wird beendet, wenn die Bedingung `False` wird:

```py
while condition:
    <code>
```

Erstellen Sie unterhalb der `expenses`-Liste eine `while`-Schleife. Verwenden Sie `True` als Bedingung und geben Sie den String `'\nExpense Tracker'` im Schleifenkörper aus, um den Titel des Programms anzuzeigen.

# --hints--

Sie sollten eine `while`-Schleife mit `True` als Bedingung innerhalb Ihrer `main()`-Funktion erstellen.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Sie sollten `'\nExpense Tracker'` in Ihrer `while`-Schleife ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
