---
id: 65823cfc74aa564ffc460489
title: Schritt 26
challengeType: 20
dashedName: step-26
---

# --description--

Rufen Sie jetzt `map()` auf und übergeben Sie Ihre `lambda`-Funktion als erstes Argument und die `expenses`-Liste als zweites Argument.

# --hints--

Sie sollten die `map()`-Funktion innerhalb der `total_expenses`-Funktion aufrufen.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

Sie sollten Ihre `lambda`-Funktion als erstes Argument an den `map()`-Aufruf übergeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

Sie sollten die `expenses`-Liste als zweites Argument an den `map()`-Aufruf übergeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
