---
id: 65824dfdb6815d563b2d3256
title: Schritt 37
challengeType: 20
dashedName: step-37
---

# --description--

Die `input()`-Funktion nimmt eine Benutzereingabe entgegen und gibt diese als String zurück.

Rufen Sie innerhalb Ihrer `while`-Schleife die `input()`-Funktion auf, übergeben Sie den String `'Enter your choice: '` als Argument, und weisen Sie das Ergebnis einer Variablen namens `choice` zu.

# --hints--

Sie sollten innerhalb Ihrer `choice`-Schleife eine Variable namens `while` deklarieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].has_variable("choice")
`)) })
```

Sie sollten das Ergebnis der `input()`-Funktion mit dem String `'Enter your choice: '` als Argument Ihrer Variablen `choice` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_variable("choice").is_equivalent("choice = input('Enter your choice: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
--fcc-editable-region--
```
