---
id: 65823040bdb32949d4c1a96f
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Sie benötigen eine Funktion, um die Gesamtsumme der Ausgaben zu berechnen.

Definieren Sie eine Funktion namens `total_expenses`, die einen Parameter `expenses` entgegennimmt. Füllen Sie den Funktionskörper vorerst mit einer `pass`-Anweisung.

# --hints--

Sie sollten eine Funktion namens `total_expenses` definieren.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

Ihre `total_expenses`-Funktion sollte einen einzelnen `expenses`-Parameter annehmen.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
