---
id: 6582352cbdbcde4d0c4f7b0b
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Um eine Lambda-Funktion aufzurufen, können Sie die übliche Funktionssyntax mit einem Paar runder Klammern nach dem Variablennamen verwenden.

Rufen Sie Ihre `test` Lambda-Funktion auf und übergeben Sie `3` als Argument. Geben Sie dann das Ergebnis aus.

# --hints--

Sie sollten `print(test(3))` in Ihrem Code haben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
