---
id: 6582507654b3ed5712341382
title: Schritt 38
challengeType: 20
dashedName: step-38
---

# --description--

Sie werden bedingte Anweisungen verwenden, um die Wahl des Benutzers zu überprüfen. Wenn die Wahl `'1'` ist, bedeutet das, dass der Benutzer eine Ausgabe hinzufügen möchte.

Bleiben Sie in der `while`-Schleife und schreiben Sie unter der `choice`-Variablen eine `if`-Anweisung, um zu prüfen, ob `choice` dem String `'1'` entspricht. Wenn das zutrifft, ist dies der Ausgangspunkt, um eine neue Ausgabe hinzuzufügen.

Innerhalb des Körpers der `if`-Anweisung deklarieren Sie eine Variable `amount` und weisen ihr einen leeren `input()`-Aufruf zu.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, die überprüft, ob `choice` gleich `'1'` in Ihrer `while`-Schleife ist.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
cond.is_equivalent("choice == '1'") or cond.is_equivalent("'1' == choice")
`)) })
```

Sie sollten eine Variable namens `amount` im Körper Ihrer `if`-Anweisung deklarieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("amount")
`)) })
```

Sie sollten Ihrer `input()`-Variablen einen leeren `amount`-Aufruf zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input()")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

--fcc-editable-region--
```
