---
id: 658212ba98182f3e855e85f9
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

In diesem Projekt erstellen Sie einen einfachen, aber funktionalen Ausgaben-Tracker in Python.

Beginnen Sie damit, eine Funktion namens `add_expense` zu deklarieren, die drei Parameter entgegennimmt: `expenses`, `amount` und `category`. Verwenden Sie das Schlüsselwort `pass`, um den Funktionskörper zu füllen.

# --hints--

Sie sollten das Schlüsselwort `def` verwenden, um eine Funktion namens `add_expense` zu deklarieren.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

Ihre `add_expense`-Funktion sollte drei Parameter annehmen: `expenses`, `amount` und `category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
