---
id: 65822bd82d708c4895080c35
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

In Python ist es wichtig zu wissen, dass dieselbe Art von Anführungszeichen, mit der ein String definiert wird, nicht innerhalb dieses Strings verwendet werden kann. Zum Beispiel ist der String `'I'm a string!'` nicht gültig. Um das einfache Anführungszeichen innerhalb dieses Strings zu verwenden, sollten Sie entweder:

- Das Anführungszeichen durch Voranstellen eines Backslash maskieren: `'I\'m a string!'`
- Oder doppelte Anführungszeichen zur Definition des Strings verwenden: `"I'm a string!"` (bevorzugt).

Sie können auf Werte in einem Dictionary über dessen Schlüssel zugreifen. Sie müssen die Klammernotation verwenden und den Schlüssel zwischen den eckigen Klammern angeben:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

Sie interpolieren gerade das `expense`-Dictionary in Ihrem f-String. Ändern Sie den f-String-Ausdruck so, dass Sie auf den Wert des `'amount'`-Schlüssels und des `'category'`-Schlüssels im `expense`-Dictionary zugreifen.

# --hints--

Sie sollten `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` an Ihren `print()`-Aufruf übergeben. Denken Sie daran, innerhalb Ihres einfach in Anführungszeichen gesetzten f-Strings doppelte Anführungszeichen zu verwenden und umgekehrt.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
