---
id: 65822308ded1ba4632f66aa3
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Ein Dictionary ist ein weiterer eingebauter Datentyp in Python. Ein Dictionary ist eine Sammlung von Daten in Form von *Schlüssel*-*Wert*-Paaren. Dictionaries werden mit geschweiften Klammern (`{}`) definiert und enthalten Schlüssel-Wert-Paare, die durch Kommas getrennt sind. Jedem Schlüssel folgt ein Doppelpunkt (`:`) und der Wert:

```py
{'amount': 50.0, 'category': 'Food'}
```

Im obigen Beispiel sind `'amount'` und `'category'` Schlüssel, und `50.0` und `'Food'` sind die entsprechenden Werte.

Erstellen Sie ein Dictionary mit dem Schlüssel `'amount'` und dem Wert des Parameters `amount` und übergeben Sie Ihr neues Dictionary an den Aufruf `.append()`.

# --hints--

Sie sollten ein Dictionary als Argument an die Methode `.append()` übergeben.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

Sie sollten `{'amount': amount}` als Argument an die Methode `.append()` übergeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
