---
id: 658252f6b1526d57b103d48a
title: Schritt 39
challengeType: 20
dashedName: step-39
---

# --description--

Innerhalb der `if`-Anweisung sollten Sie den Benutzer auffordern, den Betrag für die Ausgabe einzugeben, und diesen in einer Variablen speichern.

Übergeben Sie den String `'Enter amount: '` an Ihren leeren `input()`-Aufruf, damit Sie die Ausgabe speichern können.

# --hints--

Sie sollten `'Enter amount: '` an den `input()`-Aufruf übergeben, der der Variablen `amount` zugewiesen ist.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input('Enter amount: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = input()
--fcc-editable-region--
```
