---
id: 65821fcc010c3245718f2a06
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Der `expenses`-Parameter Ihrer `add_expense`-Funktion wird eine Liste von Ausgaben sein. Sie möchten in der Lage sein, Elemente am Ende Ihrer Liste hinzuzufügen. Dafür verwenden Sie die `.append()`-Listenmethode.

Fügen Sie einen Aufruf der `.append()`-Methode an der `expenses`-Liste hinzu. Übergeben Sie vorerst keine Argumente an `.append()`.

# --hints--

Sie sollten `expenses.append()` in Ihrer Funktion haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

Sie sollten `pass` nicht in Ihrer Funktion haben.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
