---
id: 658225d560369446ece5518b
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Fügen Sie dem Wörterbuch, das Sie an die `expense`-Liste anhängen, ein weiteres Schlüssel-Wert-Paar hinzu. Verwenden Sie den String `'category'` als Schlüssel und den Parameter `category` als Wert.

# --hints--

Sie sollten dem Wörterbuch, das Sie als Argument an die `.append()`-Methode übergeben, ein zweites Schlüssel-Wert-Paar hinzufügen.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

Sie sollten `{'amount': amount, 'category': category}` als Argument an die `.append()`-Methode übergeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
