---
id: 6582267ef39f5047411c1743
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Definieren Sie als Nächstes eine Funktion namens `print_expenses`, die einen Parameter `expenses` entgegennimmt. Diese Funktion wird später verwendet, um jede Ausgabe in Ihrer Liste anzuzeigen.

Füllen Sie den Körper Ihrer neuen Funktion mit einer `pass`-Anweisung.

# --hints--

Sie sollten eine neue Funktion namens `print_expenses` definieren.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

Ihre `print_expenses`-Funktion sollte einen einzelnen Parameter `expenses` annehmen.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
