---
id: 6582575b8089f85b8b92d7c8
title: Schritt 42
challengeType: 20
dashedName: step-42
---

# --description--

Sobald Sie die Ausgabendetails haben, müssen Sie die `add_expense`-Funktion aufrufen, um die neue Ausgabe zur `expenses`-Liste hinzuzufügen.

Nachdem Sie `amount` und `category` mit `input()` erhalten haben, rufen Sie die `add_expense`-Funktion auf und übergeben drei Argumente: `expenses`, `amount` und `category`.

- `expenses` ist die leere Liste, die zuvor in der Hauptfunktion dieses Projekts erstellt wurde.
- `amount` ist der Betrag der Ausgabe.
- `category` ist die Kategorie der Ausgabe.

# --hints--

Sie sollten die `add_expense()`-Funktion innerhalb Ihrer `if`-Anweisung aufrufen.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

Sie sollten drei Argumente an Ihren `add_expense()`-Aufruf übergeben: `expenses`, `amount` und `category`. Die Reihenfolge ist wichtig.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

Ihr `add_expense()`-Aufruf sollte nach der Zuweisung von `amount` und `category` erfolgen.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
