---
id: 65824561f3478e5371a33ae5
title: Schritt 32
challengeType: 20
dashedName: step-32
---

# --description--

Der nächste Schritt besteht darin, die Hauptfunktion zu definieren, die der Einstiegspunkt des interaktiven Ausgaben-Tracker-Programms sein wird.

Definieren Sie eine Funktion namens `main` ohne Parameter. Füllen Sie den Funktionskörper mit der `expenses`-Liste, die Sie zu Beginn dieses Projekts erstellt haben. Sie werden diese Liste verwenden, um die Ausgabenaufzeichnungen zu speichern.

# --hints--

Sie sollten eine Funktion namens `main()` ohne Parameter definieren.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

Sie sollten die `expenses`-Liste in den Funktionskörper der `main()`-Funktion verschieben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
