---
id: 65823ff0d4b991510fade1a8
title: Schritt 29
challengeType: 20
dashedName: step-29
---

# --description--

Ersetzen Sie innerhalb der `filter_expenses_by_category`-Funktion `pass` durch eine `lambda`-Funktion. Verwenden Sie `expense` als Parameter und bewerten Sie den Vergleich zwischen dem Wert des `'category'`-Schlüssels des `expense`-Dictionaries und `category` mit dem Gleichheitsoperator.

# --hints--

Sie sollten eine `lambda`-Funktion erstellen, die einen Parameter namens `expense` verwendet und den Ausdruck `expense['category'] == category` innerhalb der `filter_expenses_by_category`-Funktion auswertet.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

Sie sollten kein `pass` innerhalb der `filter_expenses_by_category`-Funktion haben.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
