---
id: 658238f7604f154ea9a23e1e
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Die `sum()`-Funktion gibt die Summe der Elemente in dem iterierbaren Objekt zurück, das ihr als Argument übergeben wird. Sie werden `sum()` zusammen mit den Funktionen `map()` und `lambda` verwenden, um die Gesamtsumme der Ausgaben zu ermitteln.

Machen Sie zunächst einen kleinen Test und ändern Sie Ihren aktuellen `print()`-Aufruf, indem Sie den `list()`-Aufruf durch einen Aufruf der `sum()`-Funktion ersetzen, der den aktuellen `map()`-Aufruf als Argument übergeben bekommt.

# --hints--

Sie sollten `sum(map(test, [2, 3, 5, 8]))` ausgeben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
