---
id: 65822773f4c60a479ec89f66
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Erstellen Sie innerhalb der `print_expenses`-Funktion eine `for`-Schleife, die über jedes Element in der `expenses`-Liste iteriert. Verwenden Sie `expense` als Schleifenvariable und verschieben Sie `pass` in den Schleifenkörper.

# --hints--

Sie sollten eine `for`-Schleife erstellen, die über `expenses` in Ihrer `print_expenses`-Funktion iteriert. Denken Sie daran, `pass` innerhalb des Schleifenkörpers zu verwenden.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

Sie sollten `expense` als Schleifenvariable verwenden.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

`pass` sollte nur innerhalb Ihres Schleifenkörpers stehen.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
