---
id: 6582324b5d37fe4b622bc3c0
title: Schritt 19
challengeType: 20
dashedName: step-19
---

# --description--

Lambda-Funktionen sind kurze, anonyme Funktionen in Python, ideal für einfache, einmalige Aufgaben. Sie werden mit dem Schlüsselwort `lambda` definiert und verwenden folgende Syntax:

```py
lambda x: expr
```

Im obigen Beispiel stellt `x` einen Parameter dar, der im Ausdruck `expr` verwendet wird, und verhält sich wie jeder Parameter in einer herkömmlichen Funktion. `expr` ist der Ausdruck, der ausgewertet und zurückgegeben wird, wenn die Lambda-Funktion aufgerufen wird.

Erstellen Sie eine Variable namens `test` und weisen Sie ihr eine Lambda-Funktion zu, die einen `x`-Parameter entgegennimmt und `x * 2` zurückgibt.

# --hints--

Sie sollten eine Variable namens `test` deklarieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

Sie sollten `lambda x: x * 2` Ihrer `test`-Variable zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
