---
id: 65823e87c9741750a22085a7
title: Schritt 28
challengeType: 20
dashedName: step-28
---

# --description--

Definieren Sie als Nächstes eine Funktion namens `filter_expenses_by_category`, die zwei Parameter entgegennimmt: `expenses` und `category`. Verwenden Sie `pass`, um den Funktionskörper zu füllen.

# --hints--

Sie sollten eine Funktion namens `filter_expenses_by_category` definieren.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

Ihre `filter_expenses_by_category`-Funktion sollte `expenses` und `category` als Parameter verwenden.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
