---
id: 65823dde36392f505a39f7c7
title: Schritt 27
challengeType: 20
dashedName: step-27
---

# --description--

Rufen Sie schließlich Ihre `map()`-Funktion in der `sum()`-Funktion auf, um die Gesamtausgaben zu erhalten, und geben Sie das Ergebnis zurück.

# --hints--

Sie sollten die `sum()`-Funktion aufrufen, indem Sie ihr Ihren aktuellen `map()`-Aufruf als Argument übergeben, und das Ergebnis aus Ihrer `total_expenses`-Funktion zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
