---
id: 65823634a308cb4d6bcc5fa6
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Lambda-Funktionen lassen sich wertvoll mit der `map()`-Funktion kombinieren, die eine angegebene Funktion für jedes Element in einer Sammlung von Objekten, wie einer Liste, ausführt:

```py
map(lambda x: x * 2, [1, 2, 3])
```

Die auszuführende Funktion wird als erstes Argument übergeben, und das iterable als zweites Argument.

Das Ergebnis des obigen Beispiels wäre `[2, 4, 6]`, wobei jedes Element in der an `map()` übergebenen Liste durch die Aktion der Lambda-Funktion verdoppelt wurde.

Ändern Sie Ihren `print()`-Aufruf so, dass das Ergebnis des Aufrufs von `map()` mit `test` als erstem Argument und `[2, 3, 5, 8]` als zweitem Argument ausgegeben wird. Sie werden noch keine lesbare Ausgabe sehen können.

# --hints--

Sie sollten `print(map(test, [2, 3, 5, 8]))` in Ihrem Code haben.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
