---
id: 65824111a09164518320088d
title: Schritt 30
challengeType: 20
dashedName: step-30
---

# --description--

Die `filter()`-Funktion ermöglicht es Ihnen, Elemente aus einem iterierbaren Objekt, wie einer Liste, basierend auf dem Ergebnis einer Funktion auszuwählen:

```py
filter(my_function, my_list)
```

`filter()` nimmt eine Funktion als erstes Argument und ein iterierbares Objekt als zweites Argument. Sie gibt einen Iterator zurück, ein spezielles Objekt, mit dem Sie die Elemente einer Sammlung, wie einer Liste, durchlaufen können.

Das Ergebnis des obigen Beispiels ist ein Iterator, der die Elemente von `my_list` enthält, für die `my_function` `True` zurückgibt.

Rufen Sie innerhalb der `filter_expenses_by_category`-Funktion `filter()` auf und übergeben Sie die `lambda`-Funktion, die Sie im vorherigen Schritt geschrieben haben, als erstes Argument und die `expenses`-Liste als zweites Argument.

# --hints--

Sie sollten `filter()` innerhalb der `filter_expenses_by_category`-Funktion aufrufen.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

Sie sollten `lambda expense: expense['category'] == category` als erstes Argument an den `filter()`-Aufruf übergeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

Sie sollten `expenses` als zweites Argument an den `filter()`-Aufruf übergeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
