---
id: 65822934c8a0454803e3ef12
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Als Nächstes geben Sie die Details für jede Ausgabe aus.

Ersetzen Sie innerhalb der `for`-Schleife `pass` durch einen `print()`-Aufruf und übergeben Sie ihm den folgenden f-String: `f'Amount: {expense}, Category: {expense}'`.

# --hints--

Sie sollten `f'Amount: {expense}, Category: {expense}'` in Ihrer `for`-Schleife ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

Sie sollten `pass` nicht innerhalb des Schleifenkörpers haben.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
