---
id: 66314f2add0a8b57da5d6d91
title: Schritt 22
challengeType: 20
dashedName: step-22
---

# --description--

Sie sollten etwas wie `<map object at 0xd273a8>` in der Konsole ausgegeben sehen, was die String-Darstellung des Map-Objekts ist, das von `map()` zurückgegeben wird.

Um eine lesbare Ausgabe zu erhalten, müssen Sie das Map-Objekt in eine Liste umwandeln. Machen Sie das, indem Sie den `map()`-Aufruf als Argument an die `list()`-Funktion übergeben.

# --hints--

Sie sollten den `map()`-Aufruf als Argument an die `list()`-Funktion übergeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
