---
id: 6669539c1379793f9cb8917c
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Python-Listen sind veränderbar, was bedeutet, dass der Wert der Listenelemente geändert werden kann. Sie können den Wert eines Elements mit der Klammernotation ändern.

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

Dadurch erhält `example_list` den Wert von `[4, 'oh', 6, 7]`.

Ändern Sie das erste Element von `my_list` in `0` und geben Sie dann die Liste aus, um den Wert zu überprüfen.

# --hints--

Ändern Sie nicht die Deklaration der Liste.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Das erste Element von `my_list` sollte 0 sein.

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

Sie sollten das erste Element mit `my_list[0] = 0` ändern.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

Sie sollten die Liste ausgeben.

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

Sie sollten die Liste nach der Änderung des ersten Elements ausgeben.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
