---
id: 66694eaaeb7d253049fa7883
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Um Elemente zu einer Liste hinzuzufügen, können Sie sie zwischen die eckigen Klammern schreiben und die Werte mit Kommas trennen.

Fügen Sie die Zahlen `1` und `2` in die Liste ein.

# --hints--

Das erste Element von `my_list` sollte `1` sein.

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

Das zweite Element von `my_list` sollte `2` sein.

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` sollte zwei Elemente enthalten.

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

Sie sollten `[1, 2]` direkt `my_list` zuweisen.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
