---
id: 66694fc4bba24f33ca01fa5b
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Python hat einige wenige Listenmethoden. Zum Beispiel Methoden zum Hinzufügen oder Entfernen von Listenelementen.

Sie können ein Element mit der `append()`-Methode am Ende einer Liste hinzufügen. Zum Beispiel:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` ist jetzt `[4, 5, 6, 7]`.

Versuchen Sie, die `append()`-Methode zu verwenden, um `3` zu `my_list` hinzuzufügen. Geben Sie dann die Liste aus.

# --hints--

Das letzte Element von `my_list` sollte `3` sein.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` sollte eine Länge von 3 haben.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

Sie sollten die Deklaration der Liste nicht ändern.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Sie sollten die Liste ausgeben.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

Sie sollten die Liste nach der Verwendung von `append()` ausgeben.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
