---
id: 666951d2e4b27e3a3c72ca5f
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Sie können auch auf ein einzelnes Element zugreifen, um dessen Wert zu erhalten.

Listen sind wie Strings nullbasiert indiziert. Das bedeutet, dass das erste Element den Index `0` hat, das zweite Element den Index `1` und so weiter.

Um auf ein Element zuzugreifen, verwenden Sie die Klammernotation. Zum Beispiel ist `example_list[1]` der Zugriff auf das Element mit dem Index `1`, also das zweite Element, von `example_list`.

Geben Sie das erste Element von `my_list` aus.

# --hints--

Sie sollten das erste Element der Liste mit `print(my_list[0])` ausgeben.

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
