---
id: 66694bfa435cef241f6f9a68
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Bevor Sie mit dem Projekt beginnen, lernen Sie etwas über Listen.

Im Gegensatz zu Zahlen und Strings ist eine <dfn>Liste</dfn> ein Datentyp, der als Container für andere Werte fungiert:

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

Die Liste zeichnet sich durch die eckigen Klammern um alle Werte aus und durch ein Komma zwischen den Werten, zum Beispiel: `["A", "happy", "list"]`. Wenn die Liste keine Werte enthält, ist sie eine leere Liste: `[]`.

Eine Liste kann verschiedene Datentypen enthalten: `[1, "Up", ["Down", "Twice"]]`. Das schließt alle möglichen Datentypen ein. Sie kann auch eine weitere Liste enthalten!

Erstellen Sie eine Variable namens `my_list` und weisen Sie ihr eine leere Liste zu.

# --hints--

Sie sollten eine Variable namens `my_list` haben.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

Die Variable `my_list` sollte den Wert einer leeren Liste haben.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
