---
id: 666aae6a5d30a71f1fd7749f
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Die `insert`-Methode kann ein Element an jeder beliebigen Position in einer Liste hinzufügen. Das erste Argument ist die Position, an der das Element eingefügt werden soll, und das zweite Argument ist das hinzuzufügende Element. Zum Beispiel so fügen Sie ein neues Element an der dritten Position von `example_list` hinzu:

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

Verwenden Sie `insert()`, um `1` an der richtigen Position in `my_list` hinzuzufügen, sodass die Liste aufsteigend zählt, und geben Sie dann die Liste aus.

# --hints--

Sie sollten `1` an Index `my_list` zu `1` mit `insert()` hinzufügen.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list` sollte 4 Elemente enthalten.

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` sollte gleich `[0, 1, 2, 3]` sein.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

Sie sollten die Liste nach dem Hinzufügen des neuen Elements ausgeben.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
