---
id: 666ab82c874add3c95145213
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Die `pop()`-Methode kann verwendet werden, um ein Element aus einer Liste zu entfernen. Standardmäßig entfernt sie das letzte Element der Liste. Sie können der Methode einen Index als Argument übergeben, und sie entfernt das Element an diesem Index.

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

In diesem Fall entfernt `fruits_list.pop(2)` das Element am Index `2` der Liste.

Verwenden Sie `pop()`, um das letzte Element aus `my_list` zu entfernen, und geben Sie dann `my_list` aus.

# --hints--

Sie sollten `pop()` verwenden.

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

`my_list` sollte `[0, 1, 2]` sein.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

Sie sollten `my_list` ausgeben, nachdem Sie `pop()` verwendet haben.

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
