---
id: 68e7e729f6d14f69a23786ef
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

接下来，在你的 `main` 元素内，添加一个带有文本 `Parent Teacher Conference Form` 和类 `title` 及 `center` 的 `h1` 元素。 

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector("h1"));
```

你的 `h1` 元素应包含文本 `Parent Teacher Conference Form`。

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

你的 `h1` 元素应具有 `title` 和 `center` 类。

```js
assert.exists(document.querySelector("h1.title.center"));
```

你的 `h1` 元素应位于 `main` 元素内。

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>

</html>
```
