---
id: 68e707257ea91c2e12d489fb
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将通过创建一个家长教师会议表单来练习如何为单选按钮添加自定义样式。已为你提供了超文本标记语言公式化。

首先为你的页面添加一个名为 `container` 的 `main` 元素。

# --hints--

你应该有一个 `main` 元素。

```js
assert.exists(document.querySelector("main"));
```

你的 `main` 元素应有一个名为 `container` 的类。

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
