---
id: 68ea8c330ed07b60879f3d38
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

现在是为你的表单添加元素以收集学生成绩信息的时候了。

首先为你的页面添加另一个带有文本 `Student Grade: ` 且 `for` 属性设置为 `"grade"` 的 `label` 元素。

然后，在你的 `label` 元素下方，添加一个 `input` 元素，其 `type` 属性设置为 `"number"`。`name` 和 `id` 属性应设置为 `"grade"`。`placeholder` 属性应设置为 `"E.g., 4"`。最后，你的数字 `input` 应该是必填的。 

# --hints--

页面上应该总共有两个 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

你的第二个 `label` 元素应包含文本 `Student Grade: `。别忘了冒号后的空格。

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

你的第二个 `label` 元素应具有 `for` 属性。

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

你的第二个 `label` 元素的 `for` 属性应具有值 `"grade"`。

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

页面上应该有两个 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

你的第二个 `input` 元素应位于你的第二个 `label` 元素下方。

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

你的第二个 `input` 元素应具有设置为 `"number"` 的 `type` 属性。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

你的第二个 `input` 元素应具有设置为 `"grade"` 的 `name` 属性。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

你的第二个 `input` 元素应具有设置为 `"grade"` 的 `id` 属性。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

你的第二个 `input` 元素应具有设置为 `"E.g., 4"` 的 `placeholder` 属性。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

你的第二个 `input` 元素应具有 `required` 属性。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
