---
id: 68ea8666b8b473455f5b33ef
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

表单的第一部分将侧重于学生的信息，例如他们的姓名和年级。

在 `form` 元素内，添加一个 `fieldset` 元素。在该 `fieldset` 元素内，添加一个带有文本 `Student Information` 的 `legend` 元素。 

在 `legend` 元素下，添加一个带有文本 `Full Name: ` 和 `for` 属性值为 `student-name` 的 `label` 元素。


# --hints--

你应该有一个 `fieldset` 元素。

```js
assert.exists(document.querySelector("fieldset"));
```

你的 `fieldset` 元素应该在 `form` 元素中。

```js
assert.exists(document.querySelector("form fieldset"));
```

你应该有一个 `legend` 元素。

```js
assert.exists(document.querySelector("legend"));
```

你的 `legend` 元素应该在 `fieldset` 元素中。

```js
assert.exists(document.querySelector("fieldset legend"));
```

你的 `legend` 元素应包含文本 `Student Information`。

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

你应该有一个 `label` 元素。

```js
assert.exists(document.querySelector("label"));
```

你的 `label` 元素应包含文本 `Full Name: `。别忘了冒号后的空格。

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

你的 `label` 元素应具有 `for` 属性。

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

你的 `label` 元素的 `for` 属性应具有值 `student-name`。

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

你的 `label` 元素应位于 `legend` 元素下方。

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
