---
id: 68ea8e6a87a2816b0bb2eb7d
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

下一步是表单中用于父母信息的部分。 

首先为你的 `fieldset` 元素添加另一个 `fieldset` 元素。在该 `fieldset` 元素内，添加一个带有文本 `Parent/Guardian Information` 的 `legend` 元素。 

# --hints--

页面上共有两个 `fieldset` 元素。

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 2);
```

你应该在你的第二个 `fieldset` 元素内有一个 `legend` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

你的 `legend` 元素应包含文本 `Parent/Guardian Information`。

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Parent/Guardian Information");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
