---
id: 68ea8fdf2c9b167307e7f671
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

现在是为收集父母信息添加表单元素的时候了。

首先为你的页面添加一个带有文本 `Parent/Guardian Name: ` 且 `for` 属性设置为 `"parent-name"` 的 `label` 元素。

然后，在你的 `label` 元素下方，添加一个 `input` 元素，`type` 属性设置为 `"text"`。`name` 和 `id` 属性应设置为 `"parent-name"`。`placeholder` 属性应设置为 `"E.g., Nancy Doe"`。最后，你的 `input` 应该是必填的。 

# --hints--

页面上应该有总共三个 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

你的第三个 `label` 元素应包含文本 `Parent/Guardian Name: `。别忘了冒号后的空格。

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

你的第三个 `label` 元素应有一个 `for` 属性。

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

你的第三个 `label` 元素的 `for` 属性应具有值 `"parent-name"`。

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

页面上应该有三个 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

你的第三个 `input` 元素应位于你的第三个 `label` 元素下方。

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

你的第三个 `input` 元素应具有设置为 `"text"` 的 `type` 属性。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

你的第三个 `input` 元素应具有设置为 `"parent-name"` 的 `name` 属性。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

你的第三个 `input` 元素应具有设置为 `"parent-name"` 的 `id` 属性。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

你的第三个 `input` 元素应具有设置为 `"E.g., Nancy Doe"` 的 `placeholder` 属性。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

你的第三个 `input` 元素应具有 `required` 属性。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
