---
id: 68ea95c833f0518ae6e366cd
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

在接下来的几步中，你将添加负责收集用户首选联系方法的表单元素。 

首先为你的表单添加另一个带有嵌套 `legend` 元素的 `fieldset` 元素。你的 `legend` 元素应包含文本 `Preferred Contact Method`。

# --hints--

页面上共有三个 `fieldset` 元素。

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 3);
```

你的第三个 `fieldset` 元素内应包含一个 `legend` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

你的 `legend` 元素应包含文本 `Preferred Contact Method`。

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Preferred Contact Method");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
