---
id: 68ea975c7cbbb894a386e402
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

下一步是为电子邮件联系方式添加 `label` 和 `input` 元素。

首先为你的标签添加一个类为 `"contact-method"` 且 `for` 属性设置为 `"email"`。标签文本应为 `Email: `。

在你的 `label` 元素下方，添加一个单选按钮，其 `id` 和 `value` 属性设置为 `"email"`。`name` 属性应设置为 `"contact-method"`，类应设置为 `"contact-method-radio-btn"`。 

最后，确保此单选按钮默认被选中。

# --hints--


你的最后一个字段集应包含一个 `label` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

这个新建的 `label` 元素（用于电子邮件）应具有 `"contact-method"` 类。

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

新建的 `label` 元素的 `for` 属性应设置为 `"email"`。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

`label` 元素应包含文本 `Email: `。确保冒号后有空格。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

你应该在最后一个 `fieldset` 内有一个 `input` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

你的 `input` 元素应具有设置为 `radio` 的 `type` 属性。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

`input` 元素应具有设置为 `"email"` 的 `id` 属性。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

`input` 元素应具有设置为 `"email"` 的 `value` 属性。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

`input` 元素应具有设置为 `"contact-method"` 的 `name` 属性。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

`input` 元素应具有 `"contact-method-radio-btn"` 类。

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

你的 `radio` 按钮应默认被 `checked`。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
