---
id: 68ea8a5db73b155708992687
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

下一步是为学生的信息添加关联的 `input` 元素。

首先为你的页面添加一个 `input` 元素，其 `type` 属性设置为 `"text"`。然后添加 `name` 和 `id` 属性，二者均设置为 `student-name`。 

接下来，添加一个 `placeholder` 属性，设置为 `E.g., Jane Doe`。最后，添加一个 `required` 属性。 

# --hints--

你应该有一个 `input` 元素。

```js
assert.exists(document.querySelector("fieldset input"));
```

你的 `input` 元素应具有设置为 `"text"` 的 `type` 属性。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

你的 `input` 元素应具有设置为 `"student-name"` 的 `name` 属性。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

你的 `input` 元素应具有设置为 `"student-name"` 的 `id` 属性。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

你的 `input` 元素应具有设置为 `E.g., Jane Doe` 的 `placeholder` 属性。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

你的 `input` 元素应该具有 `required` 属性。

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
