---
id: 68ea9f01c51b66b937793611
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

接下来，为你的页面添加另一个类为 `"contact-method"` 且 `for` 属性设置为 `"phone"` 的标签。标签文本应为 `Phone: `。

在你的 `label` 元素下方，添加一个单选按钮，其 `id` 和 `value` 属性设置为 `"phone"`。`name` 属性应设置为 `"contact-method"`，类应设置为 `"contact-method-radio-btn"`。

# --hints--

你应该在最后的 `fieldset` 中有两个 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type label"), 2);
```

`label` 元素应具有 `"contact-method"` 的类。

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

`label` 元素的 `for` 属性应设置为 `"phone"`。

```js
assert.exists(document.querySelector("fieldset:last-of-type label[for='phone']"));
```

`label` 元素应包含文本 `Phone: `。

```js
const label = document.querySelector("fieldset:last-of-type label[for='phone']");
assert.match(label?.innerText, /Phone:\s?/);
```

你应该在最后一个 `fieldset` 内有两个 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type input"), 2);
```

你的 `input` 元素应具有设置为 `radio` 的 `type` 属性。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

`input` 元素应具有设置为 `"phone"` 的 `id` 属性。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("id"), "phone");
```

`input` 元素应具有设置为 `"phone"` 的 `value` 属性。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("value"), "phone");
```

`input` 元素应具有设置为 `"contact-method"` 的 `name` 属性。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

`input` 元素应具有 `"contact-method-radio-btn"` 类。

```js
assert.exists(document.querySelector("fieldset:last-of-type input[id='phone'].contact-method-radio-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
