---
id: 68ea81da9f687732368a5322
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

现在是为你的 `form` 和 `input` 元素添加时间了，它们将表现家长和学生的信息。

首先在 `p` 元素下方添加一个 `form` 元素。

# --hints--

你应该有一个 `form` 元素。

```js
assert.exists(document.querySelector(".container form"));
```

表单元素应位于 `p` 元素下方。

```js
assert.exists(document.querySelector(".container > p+form"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
