---
id: 68eaa668d4d9fbe529a0e16b
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

现在，为你的文档添加一个带有文本 `Any specific concerns or topics you'd like to discuss?` 且 `for` 属性设置为 `notes` 的 `label` 元素。

在此 `label` 下，添加一个 `textarea` 元素。将其 `id` 设置为 `notes`，`name` 设置为 `notes`，`rows` 设置为 `4`，`cols` 设置为 `50`。

# --hints--

你最后的 `fieldset` 应该包含一个 `label` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

`label` 元素应包含文本 `Any specific concerns or topics you'd like to discuss?`。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

`label` 元素的 `for` 属性应设置为 `"notes"`。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

你的最后一个 `fieldset` 应包含一个 `textarea` 元素。

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

你的 `textarea` 元素应位于 `label` 元素下方。

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

`textarea` 元素应设置 `id` 属性为 `"notes"`。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

`textarea` 元素应具有设置为 `"notes"` 的 `name` 属性。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

`textarea` 元素应具有设置为 `4` 的 `rows` 属性。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

`textarea` 元素应具有设置为 `50` 的 `cols` 属性。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
