---
id: 68eb372f9b470238bf60577d
title: 步骤 34
challengeType: 0
dashedName: step-34
---

# --description--

现在，是时候定义当 `Radio button` 被 `Check` 时会发生什么。

你只在单选按钮被选中时为内圈设置样式，所以你的选择器应该是 `.contact-method-radio-btn:checked::before`。

在规则中，将 `transform` 设置为 `translate(3px, 3px) scale(1)`，将 `background-color` 设置为 `lightgreen`。

`transform` 稍微移动内圈并使用 `scale(1)` 将其缩放到全尺寸，浅绿色背景填充圆圈以显示该选项已被选中，使其可见。

# --hints--

你应该将目标设置为 `.contact-method-radio-btn:checked::before` 伪元素。

```js
assert(new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn:checked::before'));
```

将 `transform` 属性设置为 `translate(3px, 3px) scale(1)`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn:checked::before');
assert.strictEqual(style?.transform, 'translate(3px, 3px) scale(1)');
```

将 `background-color` 属性设置为 `lightgreen`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn:checked::before');
assert.strictEqual(style?.backgroundColor, 'lightgreen');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;  
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

fieldset {
  border: 1px solid gray;
  border-radius: 5px;
  margin: 20px 0;
  padding: 20px;
}

fieldset legend {
  font-size: 1.3rem;
  font-weight: 600;
}

label {
  font-size: 1.2rem;
}

label:not(.contact-method) {
  display: block;
  margin: 10px 0;
}

input:not(.contact-method-radio-btn),
textarea {
  background-color: #ffffff1a;
  width: 95%;
  border: 1px solid gray;
  border-radius: 5px;
  padding: 10px;
}

input,
input::placeholder,
textarea {
  color: whitesmoke;
}


.contact-method-radio-btn {
  appearance: none;
  width: 20px;
  height: 20px;
  border-radius: 50%;
  border: 2px solid gray;
  vertical-align: bottom;
}

.contact-method-radio-btn::before {
  display: block;
  content: " ";
  width: 10px;
  height: 10px;
  border-radius: 50%;
  transform: translate(3px, 3px) scale(0);
  transition: all 0.3s ease-in;
}

--fcc-editable-region--

--fcc-editable-region--
```
