---
id: 68eb323ff5a6dc1cc19fa300
title: 步骤 32
challengeType: 0
dashedName: step-32
---

# --description--

现在，你应该创建当单选按钮被选中时出现的内圈。你将使用 `::before` 伪元素来实现这一点。

伪元素如 `::before` 允许你在实际元素之前插入额外内容。这通常用于装饰目的。

定位 `.contact-method-radio-btn::before` 伪元素，设置 `display` 为 `block`，`content` 为 `" "`，`width` 为 `10px`，`height` 为 `10px`，以及 `border-radius` 为 `50%`。

# --hints--

你应该有一个 `.contact-method-radio-btn::before` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before'));
```

将 `display` 属性设置为 `block`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.display, 'block');
```

将 `content` 属性设置为 `" "`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.content, '" "');
```

将 `width` 属性设置为 `10px`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.width, '10px');
```

将 `height` 属性设置为 `10px`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.height, '10px');
```

将 `border-radius` 属性设置为 `50%`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.borderRadius, '50%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;  
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

fieldset {
  border: 1px solid gray;
  border-radius: 5px;
  margin: 20px 0;
  padding: 20px;
}

fieldset legend {
  font-size: 1.3rem;
  font-weight: 600;
}

label {
  font-size: 1.2rem;
}

label:not(.contact-method) {
  display: block;
  margin: 10px 0;
}

input:not(.contact-method-radio-btn),
textarea {
  background-color: #ffffff1a;
  width: 95%;
  border: 1px solid gray;
  border-radius: 5px;
  padding: 10px;
}

input,
input::placeholder,
textarea {
  color: whitesmoke;
}


.contact-method-radio-btn {
  appearance: none;
  width: 20px;
  height: 20px;
  border-radius: 50%;
  border: 2px solid gray;
  vertical-align: bottom;
}

--fcc-editable-region--

--fcc-editable-region--
```

