---
id: 68eb058081504963447bfa60
title: 第 17 步
challengeType: 0
dashedName: step-17
---

# --description--

接下来，你应该为主容器设置样式。选择具有类 `container` 的元素，并将其 `background-color` 设置为十六进制代码 `#ffffff1a`。

十六进制代码是由数字和字母组成的六位组合，用于在超文本标记语言和 CSS 中表现颜色。它以 `#` 开头，后跟三个点对：

- 前两位数字表现红色
- 接下来的两个表现绿色
- 最后两个表现蓝色

例如，`#ffffff` 是纯白色，因为它的红色、绿色和蓝色值都是最大值。

你还可以在末尾添加两个额外的数字来控件不透明度（称为 alpha 通道）。在 `#ffffff1a` 中，`1a` 使白色半透明。alpha 值越低，颜色看起来越透明。你稍后会学习更多关于十六进制代码的内容！

# --hints--

你应该定位具有类名 `container` 的元素。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

将 `background-color` 属性设置为 `#ffffff1a`。

```js
const style = new __helpers.CSSHelp(document).getStyle('.container')
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```


```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

--fcc-editable-region--

--fcc-editable-region--
```
