---
id: 69aaafeaa4037197c75587cd
title: implementare una coda
challengeType: 26
dashedName: implement-a-queue
---

# --description--

In questo laboratorio implementerai una struttura dati coda usando funzioni. Una coda è una struttura dati First-In-First-Out (FIFO) dove gli elementi vengono aggiunti alla fine della coda e rimossi dalla parte anteriore.

**Obiettivo:** Soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione `enqueue` che aggiunge un elemento alla fine della coda.
1. Dovresti avere una funzione `dequeue` che rimuove e restituisce l'elemento anteriore della coda.
1. Dovresti avere una funzione `front` che restituisce l'elemento anteriore della coda senza rimuoverlo.
1. Dovresti avere una funzione `size` che restituisce il numero di elementi nella coda.
1. Dovresti avere una funzione `isEmpty` che restituisce `true` se la coda è vuota e `false` altrimenti.

**Nota:** La maggior parte dei test dipende dalla funzione `enqueue`. Implementala per prima, perché i test per `dequeue`, `front`, `size` e `isEmpty` richiedono di aggiungere elementi alla coda.

# --hints--

Dovresti avere una funzione `enqueue`.

```js
assert.isFunction(enqueue);
```

La funzione `enqueue` dovrebbe aggiungere un elemento alla fine della coda.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

Dovresti avere una funzione `dequeue`.

```js
assert.isFunction(dequeue);
```

La funzione `dequeue` dovrebbe rimuovere e restituire l'elemento anteriore della coda.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

`dequeue` dovrebbe restituire `undefined` quando chiamata su una coda vuota.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` dovrebbe continuare ad aggiungere elementi alla fine della coda dopo le rimozioni.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

Diverse istanze di coda non dovrebbero condividere lo stato.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

Dovresti avere una funzione `front`.

```js
assert.isFunction(front);
```

La funzione `front` dovrebbe restituire l'elemento anteriore della coda senza rimuoverlo.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`front` dovrebbe restituire `undefined` quando chiamata su una coda vuota.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

Dovresti avere una funzione `size`.

```js
assert.isFunction(size);
```

La funzione `size` dovrebbe restituire il numero corretto di elementi nella coda.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

Dovresti avere una funzione `isEmpty`.

```js
assert.isFunction(isEmpty);
```

La funzione `isEmpty` dovrebbe restituire `true` per una coda vuota.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

La funzione `isEmpty` dovrebbe restituire `false` per una coda non vuota.

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

La funzione `isEmpty` dovrebbe restituire `true` dopo che tutti gli elementi sono stati rimossi dalla coda.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
