---
id: 68c34f4aeb2b1a4170506751
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Em seguida, adicione uma quinta linha à tabela. Comece adicionando outro elemento `tr`. Dentro desse elemento `tr`, adicione um elemento `th` com um atributo `scope` definido como `"row"` e o conteúdo de texto `12:30 PM`. Então, abaixo desse elemento `th`, adicione um elemento `td` com o conteúdo de texto `Lunch Break`. Seu elemento `td` também deve ter um atributo `colspan` definido como `3` para que ele abranja as três trilhas.

# --hints--

Você deve ter um quinto elemento `tr` dentro do elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

Você deve ter um elemento `th` dentro do quinto elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

Seu elemento `th` deve ter o conteúdo de texto `12:30 PM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

Seu elemento `th` deve ter um atributo `scope` definido como `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

Você deve ter um elemento `td` dentro do quinto elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

Seu elemento `td` deve ter o conteúdo de texto `Lunch Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

Seu elemento `td` deve ter um atributo `colspan` definido como `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
