---
id: 68c3420cb0176439a95af1b8
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

O próximo passo é adicionar o cabeçalho da tabela. 

Comece adicionando um elemento `thead` abaixo do elemento `caption`. Dentro do `thead`, adicione um elemento `tr`. 

# --hints--

Você deve ter um elemento `thead`.

```js
assert.exists(document.querySelector("table thead"));
```

Você deve ter um elemento `tr` dentro do `thead`.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
