---
id: 68c34a35c91b053dbae2a8c6
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Outro valor para o atributo `scope` é `row`, que indica que uma célula de cabeçalho é um cabeçalho para toda a sua linha.

Dentro do seu elemento `th`, adicione um atributo `scope` com o valor `row`.

# --hints--

Seu elemento `th` deve ter um atributo `scope` com um valor de `row`.

```js
assert.equal(document.querySelector("tbody tr th")?.getAttribute("scope"), "row");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        <th>9:00 AM</th>
        --fcc-editable-region--
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
