---
id: 68c341332d98fe38fae733af
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Agora é hora de começar a construir a tabela. 

Comece adicionando um elemento `table` abaixo do seu elemento `h1`. Dentro do elemento `table`, adicione um elemento `caption` com o texto `Schedule by Track and Time`.

# --hints--

Você deve ter um elemento `table`.

```js
assert.exists(document.querySelector("table"));
```

Seu elemento `table` deve ter um elemento `caption`.

```js
assert.exists(document.querySelector("table caption"));
```

Seu elemento `caption` deve conter o texto `Schedule by Track and Time`.

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
