---
id: 68c33e502cd97f2fd52cb63b
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você aprenderá como criar elementos de tabela acessíveis construindo uma programação para uma conferência de tecnologia.

Comece adicionando um elemento `h1` com o texto `Tech Conference 2025 Schedule`.

# --hints--

Você deve ter um elemento `h1`.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

Seu elemento `h1` deve conter o texto `Tech Conference 2025 Schedule`.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
