---
id: 68c34ace730ca43e5fa7b8d6
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Agora é hora de adicionar outra linha à tabela. 

Comece adicionando outro elemento `tr`. Dentro desse elemento `tr`, adicione um elemento `th` com um atributo `scope` definido como `"row"` e o conteúdo de texto `10:00 AM`. 

Então, adicione três elementos `td` com o seguinte conteúdo de texto:

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

Você deve ter um segundo elemento `tr` dentro do elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

Você deve ter um elemento `th` dentro do segundo elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

Seu elemento `th` deve ter o conteúdo de texto `10:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

Seu elemento `th` deve ter um atributo `scope` definido como `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

Seu segundo elemento `tr` deve ter três elementos `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

Seu primeiro elemento `td` deve ter o conteúdo de texto `Accessibility Deep Dive`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

Seu segundo elemento `td` deve ter o conteúdo de texto `CSS for Beginners`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

Seu terceiro elemento `td` deve ter o conteúdo de texto `Inclusive Design Principles`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
