---
id: 68c3482278a14e3cf225e897
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Dentro do seu elemento `tr`, adicione um elemento `th` com o texto `9:00 AM`. Então, abaixo desse elemento `th`, adicione três elementos `td` com os textos `Keynote: Tech Future`, `Intro to Web Dev` e `UX for All`. 

# --hints--

Você deve ter um elemento `th` dentro do seu elemento `tr`.

```js
assert.exists(document.querySelector("tbody tr th"));
```

Seu elemento `th` deve ter o conteúdo de texto `9:00 AM`.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

Você deve ter um total de três elementos `td` abaixo do seu elemento `th`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

Seu primeiro elemento `td` deve ter o conteúdo de texto `Keynote: Tech Future`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

Seu segundo elemento `td` deve ter o conteúdo de texto `Intro to Web Dev`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

Seu terceiro elemento `td` deve ter o conteúdo de texto `UX for All`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
