---
id: 68c34c608cf9e53f5d36b1d5
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Em seguida, adicione uma terceira linha à tabela. Comece adicionando outro elemento `tr`. Dentro desse elemento `tr`, adicione um elemento `th` com um atributo `scope` definido como `"row"` e o conteúdo de texto `11:00 AM`. Então, abaixo desse elemento `th`, adicione um elemento `td` com o conteúdo de texto `Break`. 

# --hints--

Você deve ter um terceiro elemento `tr` dentro do elemento `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

Você deve ter um elemento `th` dentro do terceiro elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

Seu elemento `th` deve ter o conteúdo de texto `11:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

Seu elemento `th` deve ter um atributo `scope` definido como `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

Você deve ter um elemento `td` dentro do terceiro elemento `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

Seu elemento `td` deve ter o conteúdo de texto `Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
