---
id: 68c343692285523aef0fc6bf
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

O atributo `scope` é usado para especificar se uma célula de cabeçalho é um cabeçalho para uma linha, coluna ou grupo de linhas ou colunas. Aqui está um exemplo:

```html
<th scope="col">Example Header</th>
```

Isso ajuda leitores de tela a entender a relação entre células de cabeçalho e dados.

Para todos os elementos `th`, adicione um atributo `scope` com o valor `col`.

# --hints--

Cada elemento `th` deve ter um atributo `scope` com o valor `"col"`.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
