---
id: 68c34329d430583a6cd5f7f1
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Dentro do seu elemento `tr`, adicione quatro elementos `th`. O primeiro elemento `th` deve ter o texto `Time`, o segundo deve ter o texto `Track A`, o terceiro deve ter o texto `Track B` e o quarto deve ter o texto `Track C`.

# --hints--

Seu elemento `tr` deve ter quatro elementos `th`.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

Seu primeiro elemento `th` deve conter o texto `Time`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

Seu segundo elemento `th` deve conter o texto `Track A`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

Seu terceiro elemento `th` deve ter o texto `Track B`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

Seu quarto elemento `th` deve conter o texto `Track C`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
