---
id: 68c34d57ed874e4004f7b5d4
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

No momento, o elemento `td` com o conteúdo de texto `Break` abrange apenas uma coluna. Mas seria bom se ele abrangesse as três colunas. 

Como você se lembra de workshops e aulas anteriores, você pode usar o atributo `colspan` para fazer uma célula da tabela abranger várias colunas.

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

Adicione um atributo `colspan` ao elemento `td` e defina seu valor como `3`.

# --hints--

Seu elemento `td` com o conteúdo de texto `Break` deve ter um atributo `colspan` definido como `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
