---
id: 68e5293bd00d2fe134f58987
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

الآن ستنشئ استراتيجية خصم أخرى. أنشئ صنفًا باسم `FixedAmountDiscount` يرث من `DiscountStrategy`.

أضف له طريقة `__init__` تأخذ `self` و `amount` كمعلمات. أعطِ معلمة `amount` تلميح نوع `int` وتلميح نوع إرجاع `None` للطريقة. داخل طريقة `__init__`، خزّن `amount` كخاصية في المثيل.

# --hints--

يجب أن تنشئ صنفًا باسم `FixedAmountDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

يجب أن يرث صنف `FixedAmountDiscount` من `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

يجب أن يحتوي صنف `FixedAmountDiscount` على طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

يجب أن تحتوي طريقة `__init__` على المعاملات `self` و `amount: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

يجب أن يكون نوع الإرجاع لطريقة `__init__` هو `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

يجب أن تعيّن `amount` إلى `self.amount` داخل طريقة `__init__`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
