---
id: 68e5293bd00d2fe134f58984
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

حان الوقت الآن لإنشاء تنفيذ ملموس للصنف المجرد الخاص بك. أنشئ صنفًا باسم `PercentageDiscount` يرث من `DiscountStrategy`.

أضف له طريقة `__init__` تأخذ `self` و `percent` كمعلمات. أضف تلميح نوع `int` إلى `percent` وحدد نوع الإرجاع للطريقة كـ `None`. داخل طريقة `__init__`، خزّن `percent` كخاصية في المثيل.

# --hints--

يجب أن تنشئ صنفًا باسم `PercentageDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

يجب أن يرث صنف `PercentageDiscount` من `DiscountStrategy`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

يجب أن يحتوي صنف `PercentageDiscount` على طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

يجب أن تحتوي طريقة `__init__` على المعاملات `self` و `percent: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

يجب أن تحتوي طريقتك `__init__` على تلميح نوع إرجاع `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

يجب أن تعيّن `percent` إلى `self.percent` داخل طريقة `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

--fcc-editable-region--

--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
print(product)
```
