---
id: 68e52994bc7ea2e3dec530a2
title: الخطوة 24
challengeType: 20
dashedName: step-24
---

# --description--

أنشئ الآن طريقة باسم `calculate_best_price` في صنف `DiscountEngine` الخاص بك. يجب أن تأخذ هذه الطريقة المعلمات `self` و `product` و `user_tier`، وأن تُرجع قيمة من نوع `float`.

في الوقت الحالي، استخدم فقط `pass` في جسم الطريقة. ستنفذ المنطق في الخطوات التالية.

# --hints--

يجب أن يحتوي صنف `DiscountEngine` الخاص بك على طريقة باسم `calculate_best_price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("calculate_best_price")`)
})
```

يجب أن تحتوي طريقة `calculate_best_price` على المعلمات `self` و `product: Product` و `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").has_args("self, product:Product, user_tier:str")`)
})
```

يجب أن تحتوي طريقة `calculate_best_price` على تلميح نوع الإرجاع `-> float`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").has_returns("float")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

class DiscountEngine:
    def __init__(self, strategies: list[DiscountStrategy]) -> None:
        self.strategies = strategies
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
