---
id: 68e529226b0defdff26c4a8a
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

الآن، أنشئ مثيلًا من `Product` بالاسم `Wireless Mouse` والسعر `50.0`. قم بتعيينه إلى متغير اسمه `product`. ثم، اطبع `product` في وحدة التحكم لترى شكله.

# --hints--

يجب أن يكون لديك متغير اسمه `product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

يجب تعيين مثيل من `Product` إلى المتغير `product`.

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

يجب أن يحتوي مثيل `Product` على الاسم `Wireless Mouse` والسعر `50.0`.

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

يجب طباعة `product` في وحدة التحكم.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

عند طباعة `Product` بالاسم `'Wireless Mouse'` والسعر `50.0`، يجب أن يعرض `Wireless Mouse - $50.0`.

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
