---
id: 68e524c94e2e4ab7b4bfc7e2
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستبني حاسبة خصم يمكنها تطبيق استراتيجيات خصم مختلفة على المنتجات. سيحدد النظام أفضل سعر للعميل بناءً على قواعد خصم متعددة.

ابدأ بتعريف صنف باسم `Product`. أضف إلى الصنف طريقة `__init__` بالمعلمات التالية: `self`، و`name`، و`price`. داخل طريقة `__init__`، عيّن معلمات `name` و`price` إلى خصائص الكائن بنفس الاسم.

# --hints--

يجب أن تنشئ صنفًا باسم `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

يجب أن يحتوي صنف `Product` على طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

يجب أن تحتوي طريقة `__init__` على المعلمات التالية: `self`، و`name`، و`price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

يجب أن تعيّن `name` إلى `self.name` داخل طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

يجب أن تعيّن `price` إلى `self.price` داخل طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
