---
id: 68e5293bd00d2fe134f58983
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

أضف الآن طريقة مجردة ثانية باسم `apply_discount` إلى صنف `DiscountStrategy`. استخدم المزخرف `@abstractmethod` وأعطها المعلمات `self` و `product`.

أضف تلميحات النوع: يجب أن يكون `product` من نوع `Product`، ويجب أن تُرجع الطريقة قيمة من نوع `float`. استخدم `pass` في جسم الطريقة.

# --hints--

يجب أن يحتوي صنف `DiscountStrategy` على طريقة باسم `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

يجب أن تحتوي طريقة `apply_discount` على المزخرف `@abstractmethod`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

يجب أن تحتوي طريقة `apply_discount` على المعاملات `self` و `product: Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

يجب أن تحتوي طريقة `apply_discount` على تلميح نوع الإرجاع `-> float`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
        
--fcc-editable-region--
    
--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
print(product)
```
