---
id: 68e529226b0defdff26c4a89
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

في الوقت الحالي، إذا أنشأت نسخة من `Product` وطبعته، سترى تمثيلًا افتراضيًا للكائن لا يحتوي على معلومات كثيرة. على سبيل المثال، قد يبدو مثل `<__main__.Product object at 0x7f9b8c2d1d30>`.

لتحويل التمثيل الافتراضي للكائن إلى سلسلة نصية قابلة للقراءة، أنشئ طريقة `__str__` في صنف `Product` تُرجع سلسلة نصية بالتنسيق `name - $price` (حيث يجب استبدال `name` و `price` بقيم الخاصيتين المقابلتين).

# --hints--

يجب أن يحتوي صنف `Product` على طريقة `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

يجب أن تُرجع طريقة `__str__` سلسلة نصية بالتنسيق `name - $price`.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
        
--fcc-editable-region--
    
--fcc-editable-region--
```
