---
id: 68e5293bd00d2fe134f58985
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

بما أن `PercentageDiscount` يرث من صنف مجرد، يجب تنفيذ جميع الطرق المجردة. ابدأ بتنفيذ الطريقة `is_applicable`. يجب أن تأخذ الطريقة `is_applicable` المعلمات `self`، و`product: Product`، و`user_tier: str`. يجب أن تحتوي على تلميح نوع الإرجاع `bool`.

يجب أن تحدد هذه الطريقة ما إذا كان الخصم النسبي صالحًا. أرجع `True` إذا كان `self.percent` أقل من أو يساوي `70`، وإلا أرجع `False`. هذا يمنع الخصومات غير الواقعية مثل خصم 90%.

# --hints--

يجب أن يحتوي صنف `PercentageDiscount` على طريقة `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

يجب أن تحتوي طريقة `is_applicable` على المعلمات `self`، و`product: Product`، و`user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

يجب أن تُرجع طريقة `is_applicable` القيمة `True` عندما تكون قيمة `self.percent` أقل من أو تساوي `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

يجب أن تُرجع طريقة `is_applicable` القيمة `False` عندما تكون قيمة `self.percent` أكبر من `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

يجب أن تحتوي طريقة `is_applicable` على تلميح نوع الإرجاع `bool`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
