---
id: 68e529226b0defdff26c4a8b
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

الآن ستبدأ ببناء واجهة نظام الخصم. يوفر Python طريقة لتعريف الأصناف المجردة التي تعمل كخرائط تصميم لأصناف أخرى. يتم ذلك باستخدام الصنف `ABC` (الصنف الأساسي المجرد) من وحدة `abc`.

في أعلى شفرتك، أضف جملة الاستيراد التالية:

```py
from abc import ABC, abstractmethod
```

هذا يستورد الصنف `ABC` وزخرفة `abstractmethod`، التي ستستخدمها لتعريف الطرق التي يجب تنفيذها من قبل أي صنف يرث من الصنف المجرد الخاص بك.

# --hints--

يجب أن تحتوي شفرتك على جملة الاستيراد المذكورة في الأعلى.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
