---
id: 68e5293bd00d2fe134f58982
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

كما نوقش في الدرس السابق، الطرق المجردة هي طرق يجب تنفيذها من قبل أي صنف يرث من الصنف الأساسي المجرد. تعمل كعقد يضمن أن جميع استراتيجيات الخصم تحتوي على طرق مطلوبة معينة.

داخل صنف `DiscountStrategy`، عرّف طريقة باسم `is_applicable`. استخدم المزخرف `@abstractmethod` فوق تعريف الطريقة. يجب أن تحتوي الطريقة على المعلمات التالية: `self`، و`product`، و`user_tier`. استخدم جملة `pass` في جسم الطريقة.

أضف تلميحات النوع لتوضيح توقيع الطريقة: يجب أن يكون `product` من نوع `Product`، و`user_tier` من نوع `str`.


# --hints--

يجب أن يحتوي صنف `DiscountStrategy` على طريقة باسم `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("is_applicable")`)
})
```

يجب أن تحتوي طريقة `is_applicable` على المزخرف `@abstractmethod`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_decorators("abstractmethod")`)
})
```

يجب أن تحتوي طريقة `is_applicable` على المعلمات `self`، و`product: Product`، و`user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
--fcc-editable-region--
    pass
--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
print(product)
```
