---
id: 68e5293bd00d2fe134f58981
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

أنشئ صنفًا باسم `DiscountStrategy` يرث من `ABC`. سيكون هذا صنفًا أساسيًا مجردًا يحدد واجهة لجميع استراتيجيات الخصم.

في الوقت الحالي، استخدم فقط `pass` داخل جسم الصنف لإنشاء صنف فارغ.

# --hints--

يجب أن تنشئ صنفًا باسم `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

يجب أن يرث صنف `DiscountStrategy` الخاص بك من `ABC`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
