---
id: 68e79d91daf8fb4808f93c47
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

بعد تنفيذ طريقة `apply_discount`، اختبر استراتيجية الخصم الخاصة بك. أنشئ نسخة من `PercentageDiscount` بنسبة `10` بالمئة وعيّنها في متغير باسم `discount`. ثم طبّق الخصم باستخدام `discount.apply_discount(product)`، واطبع النتيجة لترى الخصم أثناء العمل.

# --hints--

يجب أن تنشئ متغيرًا باسم `discount` يحتوي على نسخة من `PercentageDiscount`.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

يجب تهيئة نسخة `PercentageDiscount` الخاصة بك بنسبة `10` بالمئة.

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

يجب استدعاء `discount.apply_discount(product)` وطباعة النتيجة.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--

--fcc-editable-region--
```
