---
id: 68e9207ddcba053c5cb24413
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

في بايثون، تلميحات النوع هي إشارات اختيارية تخبر المطوِّرين الآخرين بنوع البيانات المتوقع لمتغير أو دالة. إليك مثالًا يضيف تلميح النوع `str` إلى المعلمة `name`.

```py
def greet(name: str):
    return 'Hello, ' + name
```

بالنسبة لمعلمات طريقة `__init__`، حدّث المعلمة `name` لتحتوي على تلميح النوع `str` والمعلمة `price` لتحتوي على تلميح النوع `float`.

# --hints--

يجب أن تضيف تلميح النوع `str` إلى المعلمة `name` وتلميح النوع `float` إلى المعلمة `price` في طريقة `__init__` لصنف `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
