---
id: 68f119ad905b29ff59d858c8
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

تمامًا مثل الخطوة السابقة، أضف في `__str__` طريقة تلميح نوع الإرجاع `str` لأن هذه الطريقة تُرجع تمثيلًا كسلسلة نصية للكائن.

# --hints--

يجب أن تحتوي طريقة `__str__` على تلميح نوع إرجاع `str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
