---
id: 564944c91be2204b269d51e3
title: تغيير النص داخل عنصر باستخدام jQuery
challengeType: 6
forumTopicId: 16773
dashedName: change-text-inside-an-element-using-jquery
---

# --description--

باستخدام jQuery، يمكنك تغيير النص بين وسم البداية ووسم النهاية لعنصر. يمكنك أيضًا تغيير ترميز HTML.

لدى jQuery دالة تسمى `.html()` تتيح لك إضافة وسوم HTML ونص داخل عنصر. أي محتوى كان موجودًا سابقًا داخل العنصر سيُستبدَل تمامًا بالمحتوى الذي تقدمه باستخدام هذه الدالة.

إليك كيف تعيد كتابة النص وتؤكده في عنواننا:

```js
$("h3").html("<em>jQuery Playground</em>");
```

لدى jQuery أيضًا دالة مشابهة تسمى `.text()` تغير النص فقط دون إضافة وسوم. بمعنى آخر، هذه الدالة لا تُقيّم أي وسوم HTML تُمرر إليها، بل تعاملها كنص تريد استبدال المحتوى الحالي به.

غيّر الزر الذي يحمل المعرف `target4` بتأكيد نصه.

<a href="https://www.freecodecamp.org/news/html-elements-explained-what-are-html-tags/#em-element" target="_blank" rel="noopener noreferrer nofollow">اطلع على مقالنا الإخباري عن &lt;em&gt;</a> لتتعرف على الفرق بين `<i>` و`<em>` واستخداماتهما.

لاحظ أنه بينما كان وسم `<i>` يُستخدم تقليديًا لتأكيد النص، فقد تم تبنيه لاحقًا كوسم للأيقونات. أما وسم `<em>` فهو الآن مقبول على نطاق واسع كوسم للتأكيد. كلاهما سيعمل لهذا التحدٍّ.

# --hints--

يجب أن تؤكد النص في زر `target4` بإضافة وسوم HTML.

```js
assert.isTrue(
  /<em>|<i>\s*#target4\s*<\/em>|<\/i>/gi.test($('#target4').html())
);
```

يجب أن يبقى النص كما هو بخلاف ذلك.

```js
assert($('#target4') && $('#target4').text().trim() === '#target4');
```

يجب ألا تغير أي نص آخر.

```js
assert.isFalse(/<em>|<i>/gi.test($('h3').html()));
```

يجب أن تستخدم `.html()` وليس `.text()`.

```js
assert(code.match(/\.html\(/g));
```

يجب أن تختار `button id="target4"` باستخدام jQuery.

```js
assert(code.match(/\$\(\s*?(\"|\')#target4(\"|\')\s*?\)\.html\(/));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target4").html('<em>#target4</em>');
  });
</script>

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
