---
id: bad87fee1348bd9aed908626
title: استهداف نفس العنصر باستخدام محددات jQuery متعددة
challengeType: 6
forumTopicId: 18322
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-the-same-element-with-multiple-jquery-selectors
---

# --description--

الآن تعرف ثلاث طرق لاستهداف العناصر: حسب النوع: `$("button")`، حسب الصنف: `$(".btn")`، وحسب المعرف `$("#target1")`.

رغم إمكانية إضافة عدة أصناف في استدعاء واحد لـ `.addClass()`، دعنا نضيفها إلى نفس العنصر *بثلاث طرق منفصلة*.

باستخدام `.addClass()`، أضف صنفًا واحدًا فقط في كل مرة إلى نفس العنصر، بثلاث طرق مختلفة:

أضف الصنف `animated` إلى كل العناصر من النوع `button`.

أضف الصنف `shake` إلى كل الأزرار التي تحمل الصنف `.btn`.

أضف الصنف `btn-primary` إلى الزر الذي يحمل المعرف `#target1`.

**ملاحظة:** يجب أن تستهدف عنصرًا واحدًا فقط وتضيف صنفًا واحدًا فقط في كل مرة. في النهاية، ستضيف محدداتك الثلاثة الأصناف `shake` و`animated` و`btn-primary` إلى `#target1`.

# --hints--

يجب أن يستخدم شفرتك المحدد `$("button")`.

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?button\s*?(?:'|")/gi));
```

يجب أن يستخدم شفرتك المحدد `$(".btn")`.

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?\.btn\s*?(?:'|")/gi));
```

يجب أن يستخدم شفرتك المحدد `$("#target1")`.

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?#target1\s*?(?:'|")/gi));
```

يجب أن تضيف صنفًا واحدًا فقط مع كل من محدداتك الثلاثة.

```js
assert(
  code.match(/addClass/g) &&
    code.match(/addClass\s*?\(\s*?('|")\s*?[\w-]+\s*?\1\s*?\)/g).length > 2
);
```

يجب أن يحتوي عنصر `#target1` على الأصناف `animated` و`shake` و`btn-primary`.

```js
assert(
  $('#target1').hasClass('animated') &&
    $('#target1').hasClass('shake') &&
    $('#target1').hasClass('btn-primary')
);
```

يجب أن تستخدم jQuery فقط لإضافة هذه الأصناف إلى العنصر.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated");
    $(".btn").addClass("shake");
    $("#target1").addClass("btn-primary");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
