---
id: bad87fee1348bd9bedc08826
title: استهداف عناصر HTML باستخدام محددات jQuery
challengeType: 6
forumTopicId: 18319
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-html-elements-with-selectors-using-jquery
---

# --description--

الآن لدينا دالة `document ready`.

لنكتب الآن أول جملة jQuery لدينا. تبدأ جميع دوال jQuery بعلامة `$`، التي تُعرف عادةً بمشغل علامة الدولار، أو bling.

غالبًا ما تختار jQuery عنصر HTML باستخدام <dfn>محدِّد</dfn>، ثم تنفذ شيئًا على ذلك العنصر.

على سبيل المثال، دعنا نجعل كل عناصر `button` الخاصة بك تقفز. فقط أضف هذه الشفرة داخل دالة document ready:

```js
$("button").addClass("animated bounce");
```

لاحظ أننا أدرجنا بالفعل مكتبة jQuery ومكتبة Animate.css في الخلفية لتتمكن من استخدامهما في المحرر. لذا أنت تستخدم jQuery لتطبيق صنف Animate.css `bounce` على عناصر `button` الخاصة بك.

# --hints--

يجب أن تستخدم دالة jQuery `addClass()` لإضافة الأصناف `animated` و `bounce` إلى عناصر `button` الخاصة بك.

```js
assert($('button').hasClass('animated') && $('button').hasClass('bounce'));
```

يجب أن تستخدم jQuery فقط لإضافة هذه الأصناف إلى العنصر.

```js
assert(!code.match(/class.*animated/g));
```

يجب أن يكون شفرة jQuery الخاصة بك داخل دالة `$(document).ready();`.

```js
assert(
  code.replace(/\s/g, '').match(/\$\(document\)\.ready\(function\(\)\{\$/g)
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
