---
id: bad87fee1348bd9aedc08826
title: استهداف العناصر حسب الصنف باستخدام jQuery
challengeType: 6
forumTopicId: 18316
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-class-using-jquery
---

# --description--

هل ترى كيف جعلنا كل عناصر `button` تقفز؟ اخترناها باستخدام `$("button")`، ثم أضفنا لها بعض أصناف CSS باستخدام `.addClass("animated bounce");`.

لقد استخدمت للتو دالة `.addClass()` من jQuery، التي تتيح لك إضافة أصناف إلى العناصر.

أولًا، لنستهدف عناصر `div` التي تحمل الصنف `well` باستخدام المحدد `$(".well")`.

لاحظ أنه، تمامًا كما في إعلانات CSS، تكتب `.` قبل اسم الصنف.

ثم استخدم دالة `.addClass()` من jQuery لإضافة الأصناف `animated` و `shake`.

على سبيل المثال، يمكنك جعل كل العناصر التي تحمل الصنف `text-primary` تهتز بإضافة التالي إلى `document ready function` الخاص بك:

```js
$(".text-primary").addClass("animated shake");
```

# --hints--

يجب أن تستخدم دالة `addClass()` من jQuery لإعطاء الأصناف `animated` و `shake` لكل العناصر التي تحمل الصنف `well`.

```js
assert($('.well').hasClass('animated') && $('.well').hasClass('shake'));
```

يجب أن تستخدم jQuery فقط لإضافة هذه الأصناف إلى العنصر.

```js
assert(!code.match(/class\.\*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
