---
id: bad87fee1348bd9aeda08826
title: استهداف العناصر بواسطة id باستخدام jQuery
challengeType: 6
forumTopicId: 18317
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-id-using-jquery
---

# --description--

يمكنك أيضًا استهداف العناصر بواسطة خاصية id الخاصة بها.

ابدأ باستهداف عنصر `button` الذي يحمل id `target3` باستخدام المحدد `$("#target3")`.

لاحظ أنه، كما هو الحال مع إعلانات CSS، تكتب `#` قبل اسم الـ id.

ثم استخدم دالة `.addClass()` في jQuery لإضافة الأصناف `animated` و `fadeOut`.

إليك كيف تجعل عنصر `button` الذي يحمل id `target6` يتلاشى:

```js
$("#target6").addClass("animated fadeOut");
```

# --hints--

يجب أن تختار عنصر `button` الذي يحمل `id` بقيمة `target3` وتستخدم دالة `addClass()` في jQuery لإعطائه الصنف `animated`.

```js
assert($('#target3').hasClass('animated'));
```

يجب أن تستهدف العنصر الذي يحمل id `target3` وتستخدم دالة `addClass()` في jQuery لإعطائه الصنف `fadeOut`.

```js
assert(
  ($('#target3').hasClass('fadeOut') || $('#target3').hasClass('fadeout')) &&
    code.match(/\$\(\s*.#target3.\s*\)/g)
);
```

يجب أن تستخدم jQuery فقط لإضافة هذه الأصناف إلى العنصر.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
    $("#target3").addClass("animated fadeOut");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
