---
id: bad87fee1348bd9acdd08826
title: تعلّم كيف تعمل وسوم السكريبت ودالة جاهزية المستند
challengeType: 6
forumTopicId: 18224
dashedName: learn-how-script-tags-and-document-ready-work
---

# --description--

الآن أنت جاهز لتعلّم jQuery، أشهر أداة JavaScript على الإطلاق.

قبل أن تبدأ باستخدام jQuery، عليك إضافة بعض الأشياء إلى HTML الخاص بك.

أولًا، أضف عنصر `script` في أعلى صفحتك. تأكد من إغلاقه في السطر التالي.

المتصفح يشغّل أي شفرة JavaScript داخل عنصر `script`، بما في ذلك jQuery.

داخل عنصر `script`، أضف هذه الشفرة: `$(document).ready(function() {` إلى `script` الخاص بك. ثم أغلقها في السطر التالي (لا تزال داخل عنصر `script`) بـ: `});`

سنتعلم المزيد عن `functions` لاحقًا. المهم أن تعرف أن الشفرة التي تضعها داخل هذه الـ `function` ستُنفذ بمجرد تحميل المتصفح لصفحتك.

هذا مهم لأن بدون `document ready function`، قد تُنفذ الشفرة قبل عرض HTML، مما يسبب أخطاء برمجية.

# --hints--

يجب أن تنشئ عنصر `script` مع التأكد من صحته ووجود وسم الإغلاق الخاص به.

```js
assert(
  code.match(/<\/script\s*>/g) &&
    code.match(
      /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
    ) &&
    code.match(/<\/script\s*>/g).length ===
      code.match(
        /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
      ).length
);
```

يجب أن تضيف `$(document).ready(function() {` في بداية عنصر `script` الخاص بك.

```js
assert(
  code.match(
    /\$\s*?\(\s*?document\s*?\)\.ready\s*?\(\s*?function\s*?\(\s*?\)\s*?\{/g
  )
);
```

يجب أن تغلق دالة `$(document).ready(function() {` باستخدام `});`

```js
assert(code.match(/\n*?\s*?\}\s*?\);/g));
```

# --seed--

## --seed-contents--

```html
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
  });
</script>
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
