---
id: 587d7dbc367417b2b2512bb1
title: إنشاء عنصر JSX بسيط
challengeType: 6
forumTopicId: 301390
dashedName: create-a-simple-jsx-element
---

# --description--

React هي مكتبة عرض مفتوحة المصدر أنشأتها وتُصان بواسطة Facebook. إنها أداة رائعة لعرض واجهة المستخدم (UI) لتطبيقات الويب الحديثة.

يستخدم React امتداد صيغة نحوية للغة JavaScript يُسمى JSX يسمح لك بكتابة HTML مباشرة داخل JavaScript. لهذا عدة فوائد. يتيح لك استخدام القوة البرمجية الكاملة لـ JavaScript داخل HTML، ويساعد على الحفاظ على قابلية قراءة الكود. في الغالب، يشبه JSX HTML الذي تعلمته بالفعل، لكن هناك بعض الاختلافات الرئيسية التي ستُشرح خلال هذه التحديات.

على سبيل المثال، بما أن JSX هو امتداد صيغة نحوية لـ JavaScript، يمكنك فعليًا كتابة JavaScript مباشرة داخل JSX. للقيام بذلك، ببساطة تضمّن الكود الذي تريد اعتباره JavaScript داخل أقواس متعرجة: `{ 'this is treated as JavaScript code' }`. احتفظ بهذا في ذهنك، لأنه يُستخدم في عدة تحديات لاحقة.

ومع ذلك، لأن JSX ليس JavaScript صالحًا، يجب تحويل كود JSX إلى JavaScript. أداة التحويل Babel هي أداة شائعة لهذه العملية. لملاءمتك، تم إضافتها بالفعل خلف الكواليس لهذه التحديات. إذا كتبت JSX بصيغة نحوية غير صحيحة، سترى فشل الاختبار الأول في هذه التحديات.

يجدر بالذكر أنه تحت الكواليس، التحديات تستدعي `ReactDOM.render(JSX, document.getElementById('root'))`. هذا الاستدعاء هو الذي يضع JSX الخاص بك في تمثيل React الخفيف الخاص بـ DOM. ثم تستخدم React لقطات من DOM الخاص بها لتحسين تحديث أجزاء محددة فقط من DOM الفعلي.

# --instructions--

الكود الحالي يستخدم JSX لتعيين عنصر `div` إلى الثابت `JSX`. استبدل عنصر `div` بعنصر `h1` وأضف النص `Hello JSX!` بداخله.

# --hints--

يجب أن يُرجع الثابت `JSX` عنصر `h1`.

```js
assert(JSX.type === 'h1');
```

يجب أن يتضمن وسم `h1` النص `Hello JSX!`

```js
assert(Enzyme.shallow(JSX).contains('Hello JSX!'));
```

# --seed--

## --seed-contents--

```jsx
const JSX = <div></div>;
```

# --solutions--

```jsx
const JSX = <h1>Hello JSX!</h1>;
```
