---
id: 5a24c314108439a4d4036162
title: إنشاء مكوِّن دال عديم الحالة
challengeType: 6
forumTopicId: 301392
dashedName: create-a-stateless-functional-component
---

# --description--

المكوِّنات هي جوهر React. كل شيء في React هو مكوِّن وهنا ستتعلم كيفية إنشاء واحد.

هناك طريقتان لإنشاء مكوِّن React. الطريقة الأولى هي استخدام دالة JavaScript. تعريف مكوِّن بهذه الطريقة يُنشئ *مكوِّن دال عديم الحالة*. سيتم تناول مفهوم الحالة في التطبيق في تحديات لاحقة. الآن، فكر في المكوِّن عديم الحالة على أنه يستقبل بيانات ويعرضها، لكنه لا يدير أو يتتبع تغييرات تلك البيانات. (سنغطي الطريقة الثانية لإنشاء مكوِّن React في التحدي التالي.)

لإنشاء مكوِّن باستخدام دالة، ببساطة اكتب دالة JavaScript تُرجع إما JSX أو `null`. من المهم أن تعرف أن React يتطلب أن يبدأ اسم الدالة بحرف كبير. إليك مثالًا على مكوِّن دال عديم الحالة يعين صنف HTML في JSX:

```jsx
const DemoComponent = function() {
  return (
    <div className='customClass' />
  );
};
```

بعد التحويل، سيكون لـ `<div>` صنف CSS باسم `customClass`.

بما أن مكوِّن JSX يمثل HTML، يمكنك تجميع عدة مكوِّنات معًا لإنشاء صفحة HTML أكثر تعقيدًا. هذه واحدة من المزايا الرئيسية لهندسة المكوِّنات التي يوفرها React. تسمح لك بتكوين واجهة المستخدم من مكوِّنات منفصلة ومعزولة. هذا يجعل بناء وصيانة واجهات المستخدم المعقدة أسهل.

# --instructions--

محرر الشفرة يحتوي على دالة تسمى `MyComponent`. أكمل هذه الدالة بحيث تُرجع عنصر `div` واحد يحتوي على سلسلة نصية.

**ملاحظة:** النص يُعتبر طفلًا لعنصر `div`، لذلك لن تتمكن من استخدام وسم مغلق ذاتيًا.

# --hints--

`MyComponent` يجب أن تُرجع JSX.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.length === 1;
  })()
);
```

`MyComponent` يجب أن تُرجع عنصر `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.children().type() === 'div';
  })()
);
```

عنصر `div` يجب أن يحتوي على سلسلة نصية.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').text() !== '';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const MyComponent = function() {
  // Change code below this line



  // Change code above this line
}
```

# --solutions--

```jsx
const MyComponent = function() {
  // Change code below this line
  return (
    <div>
      Demo Solution
    </div>
  );
  // Change code above this line
}
```
