---
id: 5a24c314108439a4d4036178
title: إنشاء مُدخل مُتحكم به
challengeType: 6
forumTopicId: 301385
dashedName: create-a-controlled-input
---

# --description--

قد يحتوي تطبيقك على تفاعلات أكثر تعقيدًا بين `state` وواجهة المستخدم المعروضة. على سبيل المثال، عناصر التحكم في النموذج الخاصة بإدخال النص، مثل `input` و`textarea`، تحتفظ بحالتها الخاصة في DOM أثناء كتابة المستخدم. مع React، يمكنك نقل هذه الحالة المتغيرة إلى `state` الخاص بمكوِّن React. يصبح إدخال المستخدم جزءًا من `state` الخاص بالتطبيق، لذا يتحكم React في قيمة حقل الإدخال هذا. عادةً، إذا كان لديك مكونات React تحتوي على حقول إدخال يمكن للمستخدم الكتابة فيها، فسيكون ذلك نموذج إدخال مُتحكم به.

# --instructions--

يحتوي محرر الشفرة على هيكل مكوِّن يُسمى `ControlledInput` لإنشاء عنصر `input` مُتحكم به. تم تهيئة `state` الخاص بالمكوِّن مسبقًا بخاصية `input` تحمل سلسلة نصية فارغة. تمثل هذه القيمة النص الذي يكتبه المستخدم في حقل الـ `input`.

أولًا، أنشئ طريقة تسمى `handleChange()` تحتوي على معلمة تسمى `event`. عند استدعاء الطريقة، تستقبل كائن `event` يحتوي على سلسلة نصية من عنصر `input`. يمكنك الوصول إلى هذه السلسلة باستخدام `event.target.value` داخل الطريقة. حدّث خاصية `input` في `state` الخاص بالمكوِّن بهذه السلسلة الجديدة.

في طريقة `render`، أنشئ عنصر `input` فوق وسم `h4`. أضف خاصية `value` تساوي خاصية `input` في `state` الخاص بالمكوِّن. ثم أضف خاصية `onChange` مضبوطة على طريقة معالِج الحدث `handleChange()`.

عند الكتابة في مربع الإدخال، يُعالج هذا النص بواسطة طريقة `handleChange()`، ثم يُحفظ كخاصية `input` في الـ `state` المحلي، ويُعرض كقيمة في مربع الـ `input` على الصفحة. الـ `state` في المكوِّن هو المصدر الوحيد الموثوق لبيانات الإدخال.

أخيرًا وليس آخرًا، لا تنسَ إضافة الربط اللازم في دالة البناء.

# --hints--

يجب أن تُعيد `ControlledInput` عنصر `div` يحتوي على عنصر `input` ووسم `p`.

```js
assert(
  Enzyme.mount(React.createElement(ControlledInput))
    .find('div')
    .children()
    .find('input').length === 1 &&
    Enzyme.mount(React.createElement(ControlledInput))
      .find('div')
      .children()
      .find('p').length === 1
);
```

يجب أن يبدأ `ControlledInput` بحالة تحتوي على خاصية `input` مُهيأة كسلسلة نصية فارغة.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(ControlledInput)).state('input'),
  ''
);
```

يجب أن يؤدي الكتابة في عنصر الإدخال إلى تحديث الحالة وقيمة الإدخال، ويجب أن يعرض عنصر `p` هذه الحالة أثناء الكتابة.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(ControlledInput));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return waitForIt(() => mockedComponent.state('input'));
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => ({
      state: mockedComponent.state('input'),
      text: mockedComponent.find('p').text(),
      inputVal: mockedComponent.find('input').props().value
    }));
  };
  const before = await _1();
  const after = await _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.text === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
```

# --seed--

## --seed-contents--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        { /* Change code below this line */}

        { /* Change code above this line */}
        <h4>Controlled Input:</h4>
        <p>{this.state.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  render() {
    return (
      <div>
        <input
          value={this.state.input}
          onChange={this.handleChange} />
        <h4>Controlled Input:</h4>

        <p>{this.state.input}</p>
      </div>
    );
  }
};
```
