---
id: 5a24c314108439a4d4036169
title: تمرير الخصائص إلى مكوِّن دالي عديم الحالة
challengeType: 6
forumTopicId: 301402
dashedName: pass-props-to-a-stateless-functional-component
---

# --description--

غطت التحديات السابقة الكثير حول إنشاء وتركيب عناصر JSX، والمكوِّنات الدالية، ومكوِّنات الصنف بأسلوب ES6 في React. مع هذا الأساس، حان الوقت للنظر في ميزة أخرى شائعة جدًا في React: **الخصائص**. في React، يمكنك تمرير الخصائص إلى المكوِّنات الفرعية. لنفترض أن لديك مكوِّن `App` يعرض مكوِّنًا فرعيًا يسمى `Welcome` وهو مكوِّن دالي عديم الحالة. يمكنك تمرير خاصية `user` إلى `Welcome` بكتابة:

```jsx
<App>
  <Welcome user='Mark' />
</App>
```

تستخدم **خصائص HTML المخصصة** التي تنشئها وتدعمها React ليتم تمريرها إلى المكوِّن. في هذه الحالة، الخاصية المنشأة `user` تُمرر إلى المكوِّن `Welcome`. وبما أن `Welcome` مكوِّن دالي عديم الحالة، فإنه يمكنه الوصول إلى هذه القيمة كما يلي:

```jsx
const Welcome = (props) => <h1>Hello, {props.user}!</h1>
```

من المعتاد تسمية هذه القيمة `props` وعند التعامل مع المكوِّنات الدالية عديمة الحالة، تعتبرها ببساطة معلمة لدالة تُرجع JSX. يمكنك الوصول إلى قيمة المعلمة داخل جسم الدالة. مع مكوِّنات الصنف، سترى أن الأمر مختلف قليلًا.

# --instructions--

يوجد مكوِّنان هما `Calendar` و`CurrentDate` في محرر الشفرة. عند عرض مكوِّن `CurrentDate` من داخل مكوِّن `Calendar`، مرِّر خاصية باسم `date` وعيِّن لها التاريخ الحالي باستخدام كائن `Date` من جافاسكريبت. بعد ذلك، وصِل إلى هذه الـ `prop` داخل مكوِّن `CurrentDate`، واعرض قيمتها داخل وسم `p`. لاحظ أنه حتى تُقيَّم قيمة الـ `prop` كجافاسكريبت، يجب وضعها بين أقواس معقوفة، مثل `date={Date()}`.

# --hints--

يجب أن يُرجع مكوِّن `Calendar` عنصر `div` واحدًا فقط.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().type() === 'div';
  })()
);
```

يجب أن يكون الطفل الثاني لمكوِّن `Calendar` هو مكوِّن `CurrentDate`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).name() === 'CurrentDate';
  })()
);
```

يجب أن يحتوي مكوِّن `CurrentDate` على خاصية تسمى `date`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).props().date;
  })()
);
```

يجب أن تحتوي خاصية `date` في مكوِّن `CurrentDate` على سلسلة نصية.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    const prop = mockedComponent.children().childAt(1).props().date;
    return typeof prop === 'string' && prop.length > 0;
  })()
);
```

يجب أن تُولد خاصية `date` باستدعاء `Date()`

```js
assert(/<CurrentDatedate={Date\(\)}\/>/.test(__helpers.removeWhiteSpace(code)));
```

يجب أن يعرض مكوِّن `CurrentDate` القيمة من خاصية `date` داخل وسم `p`.

```js
let date = 'dummy date';
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(CurrentDate, { date })
    );
    return mockedComponent.find('p').html().includes(date);
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: </p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: {props.date}</p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate date={Date()} />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
