---
id: 5a24c314108439a4d403617b
title: تمرير استدعاء راجع كخصائص
challengeType: 6
forumTopicId: 301400
dashedName: pass-a-callback-as-props
---

# --description--

يمكنك تمرير `state` كخصائص إلى المكونات الفرعية، لكنك لست محدودًا بتمرير البيانات فقط. يمكنك أيضًا تمرير دوال المعالِج أو أي طريقة معرفة على مكوِّن React إلى مكوِّن فرعي. هذه هي الطريقة التي تسمح بها للمكونات الفرعية بالتفاعل مع مكونات الوالد. تمرر الطرق إلى المكوِّن الفرعي تمامًا مثل الخاصية العادية. تُعيَّن لها اسم ويمكنك الوصول إلى هذا الاسم تحت `this.props` في المكوِّن الفرعي.

# --instructions--

هناك ثلاثة مكونات موضحة في محرر الشفرة. مكوِّن `MyApp` هو المكوِّن الوالد الذي سيعرض المكونات الفرعية `GetInput` و `RenderInput`. أضف مكوِّن `GetInput` إلى طريقة العرض في `MyApp`، ثم مرر له خاصية اسمها `input` مُعيَّنة إلى `inputValue` من `state` الخاص بـ `MyApp`. أنشئ أيضًا خاصية اسمها `handleChange` ومرر لها معالِج الإدخال `handleChange`.

بعد ذلك، أضف `RenderInput` إلى طريقة العرض في `MyApp`، ثم أنشئ خاصية اسمها `input` ومرر لها `inputValue` من `state`. بمجرد الانتهاء، ستتمكن من الكتابة في حقل `input` في مكوِّن `GetInput`، الذي يستدعي بعد ذلك طريقة المعالِج في المكوِّن الوالد عبر الخصائص. هذا يُحدِّث الإدخال في `state` الخاص بالوالد، والذي يُمرَّر كخصائص إلى كلا المكونين الفرعيين. راقب كيف تتدفق البيانات بين المكونات وكيف يبقى المصدر الوحيد للحقيقة هو `state` الخاص بالمكوِّن الوالد. بصراحة، هذا المثال مبالغ فيه قليلًا، لكنه يوضح كيف يمكن تمرير البيانات والاستدعاءات الراجعة بين مكونات React.

# --hints--

يجب أن يعرض مكوِّن `MyApp`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('MyApp').length === 1;
  })()
);
```

يجب أن يعرض مكوِّن `GetInput`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('GetInput').length === 1;
  })()
);
```

يجب أن يعرض مكوِّن `RenderInput`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('RenderInput').length === 1;
  })()
);
```

يجب أن يستقبل مكوِّن `GetInput` الخاصية `inputValue` من حالة `MyApp` كخاصية، ويحتوي على عنصر `input` يُعدِّل حالة `MyApp`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: '' });
    return waitForIt(() => mockedComponent.state());
  };
  const state_2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => mockedComponent.state());
  };
  const updated_1 = await state_1();
  const updated_2 = await state_2();
  assert(updated_1.inputValue === '' && updated_2.inputValue === 'TestInput');
```

يجب أن يستقبل مكوِّن `RenderInput` الخاصية `inputValue` من حالة `MyApp` كخاصية.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: 'TestName' });
    return waitForIt(() => mockedComponent);
  };
  const updated_1 = await state_1();
  assert(updated_1.find('p').text().includes('TestName'));
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
  this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
         <GetInput
           input={this.state.inputValue}
           handleChange={this.handleChange}/>
         <RenderInput
           input={this.state.inputValue}/>
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```
