---
id: 5a24c314108439a4d4036164
title: إنشاء مكوِّن باستخدام التكوين
challengeType: 6
forumTopicId: 301383
dashedName: create-a-component-with-composition
---

# --description--

سننظر الآن في كيفية تكوين عدة مكوِّنات React معًا. تخيل أنك تبني تطبيقًا وأنشأت ثلاثة مكوِّنات: `Navbar` و`Dashboard` و`Footer`.

لتكوين هذه المكوِّنات معًا، يمكنك إنشاء مكوِّن أب `App` يعرض كلًا من هذه المكوِّنات الثلاثة كأبناء. لعرض مكوِّن كطفل داخل مكوِّن React، تضمّن اسم المكوِّن مكتوبًا كوسم HTML مخصص في JSX. على سبيل المثال، في طريقة `render` يمكنك كتابة:

```jsx
return (
 <App>
  <Navbar />
  <Dashboard />
  <Footer />
 </App>
)
```

عندما يواجه React وسم HTML مخصص يشير إلى مكوِّن آخر (اسم المكوِّن ملفوف بـ `< />` كما في هذا المثال)، يعرض الترميز الخاص بذلك المكوِّن في موقع الوسم. هذا يوضح علاقة الأب/الطفل بين مكوِّن `App` و`Navbar` و`Dashboard` و`Footer`.

# --instructions--

في محرر الشفرة، هناك مكوِّن دالي بسيط يسمى `ChildComponent` ومكوِّن صنف يسمى `ParentComponent`. كوّن الاثنين معًا عن طريق عرض `ChildComponent` داخل `ParentComponent`. تأكد من إغلاق وسم `ChildComponent` بشرطة مائلة للأمام.

**ملاحظة:** تم تعريف `ChildComponent` باستخدام دالة سهمية ES6 لأن هذه ممارسة شائعة جدًا عند استخدام React.

# --hints--

يجب أن يعيد مكوِّن React عنصر `div` واحدًا فقط.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.type() === 'div';
  })()
);
```

يجب أن يعيد المكوِّن عنصرين مسكنين.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.children().length === 2;
  })()
);
```

يجب أن يعيد المكوِّن `ChildComponent` كطفله الثاني.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ParentComponent));
    return (
      mockedComponent.find('ParentComponent').find('ChildComponent').length ===
      1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }


        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }
        <ChildComponent />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
