---
id: 5a24c314108439a4d403618c
title: استخدام Array.filter() لتصفية مصفوفة بشكل ديناميكي
challengeType: 6
forumTopicId: 301416
dashedName: use-array-filter-to-dynamically-filter-an-array
---

# --description--

طريقة المصفوفة `map` أداة قوية ستستخدمها كثيرًا عند العمل مع React. طريقة أخرى مرتبطة بـ `map` هي `filter`، التي تُصفّي محتويات المصفوفة بناءً على شرط، ثم تُعيد مصفوفة جديدة. على سبيل المثال، إذا كان لديك مصفوفة من المستخدمين وكل منهم لديه خاصية `online` يمكن تعيينها إلى `true` أو `false`، يمكنك تصفية المستخدمين المتصلين فقط بكتابة:

```js
let onlineUsers = users.filter(user => user.online);
```

# --instructions--

في محرر الشفرة، يتم تهيئة `state` في `MyComponent` بمصفوفة من المستخدمين. بعض المستخدمين متصلون بالإنترنت وبعضهم غير متصل. صفِّ المصفوفة بحيث تظهر فقط المستخدمين المتصلين بالإنترنت. للقيام بذلك، استخدم أولًا `filter` لإرجاع مصفوفة جديدة تحتوي فقط على المستخدمين الذين تكون خاصية `online` لديهم تساوي `true`. بعد ذلك، في المتغير `renderOnline`، مرّ على المصفوفة المصفاة وأرجع عنصر `li` لكل مستخدم يحتوي على نص `username` الخاص به. تأكد أيضًا من تضمين `key` فريد، كما في التحديات السابقة.

# --hints--

يجب أن يوجد `MyComponent` ويُعرض على الصفحة.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('MyComponent').length,
  1
);
```

يجب تهيئة حالة `MyComponent` إلى مصفوفة تحتوي على ستة مستخدمين.

```js
assert(
  Array.isArray(
    Enzyme.mount(React.createElement(MyComponent)).state('users')
  ) === true &&
    Enzyme.mount(React.createElement(MyComponent)).state('users').length === 6
);
```

يجب أن يُرجع `MyComponent` عنصر `div` وعنصر `h1` ثم قائمة غير مرتبة تحتوي على عناصر `li` لكل مستخدم تم تعيين حالة الاتصال الخاصة به إلى `true`.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MyComponent));
  const users = (bool) => ({
    users: [
      { username: 'Jeff', online: bool },
      { username: 'Alan', online: bool },
      { username: 'Mary', online: bool },
      { username: 'Jim', online: bool },
      { username: 'Laura', online: bool }
    ]
  });
  const result = () => comp.find('li').length;
  const _1 = result();
  const _2 = () => {
    comp.setState(users(true));
    return result();
  };
  const _3 = () => {
    comp.setState(users(false));
    return result();
  };
  const _4 = () => {
    comp.setState({ users: [] });
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  assert(
    comp.find('div').length === 1 &&
      comp.find('h1').length === 1 &&
      comp.find('ul').length === 1 &&
      _1 === 4 &&
      _2_val === 5 &&
      _3_val === 0 &&
      _4_val === 0
  );
})();
```

يجب أن يعرض `MyComponent` عناصر `li` التي تحتوي على اسم المستخدم `username` لكل مستخدم متصل.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MyComponent));
  const users = (bool) => ({
    users: [
      { username: 'Jeff', online: bool },
      { username: 'Alan', online: bool },
      { username: 'Mary', online: bool },
      { username: 'Jim', online: bool },
      { username: 'Laura', online: bool }
    ]
  });
  const ul = () => {
    comp.setState(users(true));
    return comp.find('ul').html();
  };
  const html = ul();
  assert(
    html ===
      '<ul><li>Jeff</li><li>Alan</li><li>Mary</li><li>Jim</li><li>Laura</li></ul>'
  );
})();
```

يجب أن يحتوي كل عنصر في القائمة على خاصية `key` فريدة.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(MyComponent)).find('ul');
    console.log(ul.debug());
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key()
    ]);
    return keys.size === 4;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      users: [
        {
          username: 'Jeff',
          online: true
        },
        {
          username: 'Alan',
          online: false
        },
        {
          username: 'Mary',
          online: true
        },
        {
          username: 'Jim',
          online: false
        },
        {
          username: 'Sara',
          online: true
        },
        {
          username: 'Laura',
          online: true
        }
      ]
    };
  }
  render() {
    const usersOnline = null; // Change this line
    const renderOnline = null; // Change this line
    return (
      <div>
        <h1>Current Online Users:</h1>
        <ul>{renderOnline}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      users: [
        {
          username: 'Jeff',
          online: true
        },
        {
          username: 'Alan',
          online: false
        },
        {
          username: 'Mary',
          online: true
        },
        {
          username: 'Jim',
          online: false
        },
        {
          username: 'Sara',
          online: true
        },
        {
          username: 'Laura',
          online: true
        }
      ]
    };
  }
  render() {
    const usersOnline = this.state.users.filter(user => {
      return user.online;
    });
    const renderOnline = usersOnline.map(user => {
      return <li key={user.username}>{user.username}</li>;
    });
    return (
      <div>
        <h1>Current Online Users:</h1>
        <ul>{renderOnline}</ul>
      </div>
    );
  }
}
```
