---
id: 5a24c314108439a4d403617d
title: استخدام طريقة دورة الحياة componentDidMount
challengeType: 6
forumTopicId: 301422
dashedName: use-the-lifecycle-method-componentdidmount
---

# --description--

يحتاج معظم مطوّري الويب، في مرحلة ما، إلى استدعاء نقطة نهاية API لاسترجاع البيانات. إذا كنت تعمل مع React، من المهم أن تعرف أين تنفذ هذا الإجراء.

أفضل الممارسات مع React هي وضع استدعاءات API أو أي استدعاءات إلى الخادم في طريقة دورة الحياة `componentDidMount()`. تُستدعى هذه الطريقة بعد تركيب المكوِّن في DOM. أي استدعاءات إلى `setState()` هنا ستُحفّز إعادة عرض المكوِّن. عندما تستدعي API في هذه الطريقة وتضبط الحالة بالبيانات التي تُرجعها API، ستُحدث تلقائيًا تحديثًا بمجرد استلام البيانات.

# --instructions--

هناك استدعاء API وهمي في `componentDidMount()`. يضبط الحالة بعد 2.5 ثانية لمحاكاة استدعاء خادم لاسترجاع البيانات. يطلب هذا المثال العدد الإجمالي الحالي للمستخدمين النشطين لموقع. في طريقة العرض، اعرض قيمة `activeUsers` داخل وسم `h1` بعد النص `Active Users:`. راقب ما يحدث في المعاينة، ولا تتردد في تغيير مهلة الانتظار لرؤية التأثيرات المختلفة.

# --hints--

يجب أن يعرض `MyComponent` عنصر `div` يحتوي على وسم `h1` بداخله.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

يجب تحديث حالة المكوِّن باستخدام دالة مهلة في `componentDidMount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return new RegExp('setTimeout(.|\n)+setState(.|\n)+activeUsers').test(
      String(mockedComponent.instance().componentDidMount)
    );
  })()
);
```

يجب أن يعرض وسم `h1` قيمة `activeUsers` من حالة `MyComponent`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ activeUsers: 1237 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ activeUsers: 1000 });
    return mockedComponent.find('h1').text();
  };
  assert(new RegExp('1237').test(first()) && new RegExp('1000').test(second()));
})();
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        {/* Change code below this line */}
        <h1>Active Users: </h1>
        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        <h1>Active Users: {this.state.activeUsers}</h1>
      </div>
    );
  }
}
```
