---
id: 5a24c314108439a4d4036172
title: عرض الحالة في واجهة المستخدم بطريقة أخرى
challengeType: 6
forumTopicId: 301408
dashedName: render-state-in-the-user-interface-another-way
---

# --description--

هناك طريقة أخرى للوصول إلى `state` في المكوِّن. في طريقة `render()`، قبل جملة `return`، يمكنك كتابة جافاسكريبت مباشرة. على سبيل المثال، يمكنك إعلان دوال، الوصول إلى بيانات من `state` أو `props`، إجراء حسابات على هذه البيانات، وهكذا. ثم يمكنك تعيين أي بيانات إلى متغيرات، والتي يمكنك الوصول إليها في جملة `return`.

# --instructions--

في طريقة العرض لمكوِّن `MyComponent`، عرّف ثابتًا باستخدام `const` باسم `name` واجعله مساويًا لقيمة الاسم الموجودة في `state` للمكوِّن. بما أنك تستطيع كتابة JavaScript مباشرة في هذا الجزء من الشفرة، لا تحتاج إلى وضع هذا المرجع داخل أقواس متعرجة.

بعد ذلك، في جملة الإرجاع، اعرض هذه القيمة داخل وسم `h1` باستخدام المتغير `name`. تذكر، يجب استخدام صيغة JSX (الأقواس المتعرجة لجافاسكريبت) في جملة الإرجاع.

# --hints--

يجب أن يحتوي `MyComponent` على مفتاح `name` بقيمة `freeCodeCamp` مخزنة في حالته.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

يجب أن يعرض `MyComponent` عنصر عنوان `h1` محاطًا بعنصر `div` واحد.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

يجب أن يتضمن وسم `h1` المعروض مرجعًا إلى `{name}`.

```js
assert(/<h1>\n*\s*\{\s*name\s*\}\s*\n*<\/h1>/.test(code));
```

يجب أن يحتوي عنصر العنوان `h1` المعروض على نص يتم عرضه من حالة المكوِّن.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(firstValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line

    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line
    const name = this.state.name;
    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
