---
id: 5a24c314108439a4d4036171
title: عرض الحالة في واجهة المستخدم
challengeType: 6
forumTopicId: 301409
dashedName: render-state-in-the-user-interface
---

# --description--

بمجرد إعلان الحالة الأولية لمكوِّن، يمكنك عرض أي جزء منها في واجهة المستخدم التي تُعرض. إذا كان المكوِّن ذو حالة، فسيكون لديه دائمًا وصول إلى البيانات في `state` داخل طريقة `render()`. يمكنك الوصول إلى البيانات باستخدام `this.state`.

إذا أردت الوصول إلى قيمة من الحالة داخل `return` لطريقة العرض، يجب أن تحيط القيمة بأقواس متعرجة.

`state` هي واحدة من أقوى ميزات المكوِّنات في React. تتيح لك تتبع البيانات المهمة في تطبيقك وعرض واجهة مستخدم استجابةً للتغييرات في هذه البيانات. إذا تغيرت بياناتك، ستتغير واجهة المستخدم. يستخدم React ما يسمى DOM الافتراضي، لتتبع التغييرات خلف الكواليس. عندما تتحدث بيانات الحالة، يُحفز ذلك إعادة عرض للمكوِّنات التي تستخدم تلك البيانات - بما في ذلك المكوِّنات الفرعية التي استلمت البيانات كخاصية. يقوم React بتحديث DOM الفعلي، ولكن فقط حيثما يكون ذلك ضروريًا. هذا يعني أنك لا تحتاج للقلق بشأن تغيير DOM. فقط أعلن كيف يجب أن تبدو واجهة المستخدم.

لاحظ أنه إذا جعلت المكوِّن ذو حالة، فلن تكون المكوِّنات الأخرى على علم بـ `state` الخاص به. فـ `state` خاصته مغلق تمامًا، أو محلي لذلك المكوِّن، ما لم تمرر بيانات الحالة إلى مكوِّن فرعي كـ `props`. هذه الفكرة الخاصة بـ `state` المغلق مهمة جدًا لأنها تتيح لك كتابة منطق معين، ثم احتواء هذا المنطق وعزله في مكان واحد في شفرتك.

# --instructions--

في محرر الشفرة، `MyComponent` هو بالفعل ذو حالة. أعلن وسم `h1` في طريقة العرض للمكوِّن ليعرض قيمة `name` من حالة المكوِّن.

**ملاحظة:** يجب أن يعرض وسم `h1` فقط القيمة من `state` ولا شيء غيرها. في JSX، أي شفرة تكتبها داخل الأقواس المتعرجة `{ }` ستُعامل كجافاسكريبت. لذا للوصول إلى القيمة من `state` فقط أحط المرجع بالأقواس المتعرجة.

# --hints--

يجب أن يحتوي `MyComponent` على مفتاح `name` بقيمة `freeCodeCamp` مخزنة في حالته.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

يجب أن يعرض `MyComponent` عنصر عنوان `h1` محاطًا داخل `div` واحد.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

يجب أن يحتوي عنصر العنوان `h1` المعروض على نص يُعرض فقط من حالة المكوِّن.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  const getValue = firstValue.replace(/\s/g, '');
  assert(getValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{this.state.name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
