---
id: 5a24c314108439a4d4036174
title: "اربط 'this' بطريقة الصنف"
challengeType: 6
forumTopicId: 301379
dashedName: bind-this-to-a-class-method
---

# --description--

بالإضافة إلى تعيين وتحديث `state`، يمكنك أيضًا تعريف طرق لصنف المكوِّن الخاص بك. عادةً تحتاج طريقة الصنف إلى استخدام الكلمة المفتاحية `this` لكي تصل إلى الخصائص في الصنف (مثل `state` و`props`) داخل نطاق الطريقة. هناك عدة طرق تسمح لطرق الصنف بالوصول إلى `this`.

طريقة شائعة هي ربط `this` صراحةً في دالة البناء حتى يرتبط `this` بطرق الصنف عند تهيئة المكوِّن. ربما لاحظت أن التحدٍّ الأخير استخدم `this.handleClick = this.handleClick.bind(this)` لطريقة `handleClick` في دالة البناء. حين تستدعي دالة مثل `this.setState()` داخل طريقة الصنف، تشير `this` إلى الصنف ولن تكون `undefined`.

**ملاحظة:** الكلمة المفتاحية `this` من أكثر الجوانب إرباكًا في JavaScript لكنها تلعب دورًا مهمًا في React. رغم أن سلوكها هنا طبيعي تمامًا، هذه الدروس ليست المكان المناسب لمراجعة معمقة لـ `this` لذا يرجى الرجوع إلى دروس أخرى إذا كان ما سبق مربكًا!

# --instructions--

محرر الشفرة يحتوي على مكوِّن به `state` يتتبع النص. كما يحتوي على طريقة تسمح لك بتعيين النص إلى `You clicked!`. مع ذلك، لا تعمل الطريقة لأنها تستخدم الكلمة المفتاحية `this` التي هي undefined. أصلح ذلك بربط `this` صراحةً بطريقة `handleClick()` في دالة البناء للمكوِّن.

بعد ذلك، أضف معالِج نقر إلى عنصر `button` في طريقة العرض. يجب أن يُفعِّل طريقة `handleClick()` عندما يتلقى الزر حدث نقر. تذكر أن الطريقة التي تمررها إلى معالِج `onClick` تحتاج إلى أقواس متعرجة لأنها يجب أن تُفسر مباشرة كـ JavaScript.

بمجرد إكمال الخطوات السابقة، يجب أن تتمكن من النقر على الزر ورؤية `You clicked!`.

# --hints--

يجب أن تُرجع `MyComponent` عنصر `div` يحتوي على عنصرين، زر وعنصر `h1`، بهذا الترتيب.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).find('div').length === 1 &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(0)
      .type() === 'button' &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(1)
      .type() === 'h1'
);
```

يجب أن يبدأ `MyComponent` بحالة (state) تحتوي على زوج المفتاح والقيمة `{ text: "Hello" }`.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('text') === 'Hello'
);
```

النقر على عنصر `button` يجب أن يشغّل طريقة `handleClick` ويعيّن حالة `text` إلى `You clicked!`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ text: 'Hello' });
    return waitForIt(() => mockedComponent.state('text'));
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return waitForIt(() => mockedComponent.state('text'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Hello' && secondValue === 'You clicked!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    // Change code below this line

    // Change code above this line
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }
        <button>Click Me</button>
        { /* Change code above this line */ }
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```
