---
id: 5a24c314108439a4d4036160
title: تعريف صنف HTML في JSX
challengeType: 6
forumTopicId: 301393
dashedName: define-an-html-class-in-jsx
---

# --description--

الآن بعد أن أصبحت مرتاحًا في كتابة JSX، قد تتساءل كيف يختلف عن HTML.

حتى الآن، قد يبدو أن HTML وJSX متشابهان تمامًا.

فرق رئيسي في JSX هو أنك لم تعد تستطيع استخدام الكلمة `class` لتعريف أصناف HTML. ذلك لأن `class` كلمة محجوزة في JavaScript. بدلًا من ذلك، يستخدم JSX `className`.

في الواقع، يصبح اصطلاح التسمية لجميع خصائص HTML ومرجعيات الأحداث في JSX بتنسيق سنام الجمل. على سبيل المثال، حدث النقر في JSX هو `onClick` بدلًا من `onclick`. وبالمثل، `onchange` تصبح `onChange`. رغم أن هذا فرق بسيط، إلا أنه مهم أن تضعه في اعتبارك للمضي قدمًا.

# --instructions--

طبّق صنفًا اسمه `myDiv` على العنصر `div` الموجود في شفرة JSX.

# --hints--

يجب أن تُرجع الثابتة `JSX` عنصر `div`.

```js
assert.strictEqual(JSX.type, 'div');
```

يجب أن يحتوي العنصر `div` على صنف `myDiv`.

```js
assert.strictEqual(JSX.props.className, 'myDiv');
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h1>Add a class to this div</h1>
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div className = 'myDiv'>
  <h1>Add a class to this div</h1>
</div>);
```
