---
id: 5a24c314108439a4d4036167
title: عرض مكوِّن صنف إلى DOM
challengeType: 6
forumTopicId: 301404
dashedName: render-a-class-component-to-the-dom
---

# --description--

قد تتذكر استخدام واجهة برمجة التطبيقات ReactDOM في تحدٍّ سابق لعرض عناصر JSX إلى DOM. ستكون عملية عرض مكوِّنات React مشابهة جدًا. ركزت التحديات السابقة على المكوِّنات والتكوين، لذا تم العرض لك ضمن الخلفية. مع ذلك، لا تُعرض أي شفرة React تكتبها إلى DOM بدون استدعاء واجهة برمجة التطبيقات ReactDOM.

إليك تذكيرًا بالصّيغة النحوية: `ReactDOM.render(componentToRender, targetNode)`. المعلمة الأولى هي مكوِّن React الذي تريد عرضه. المعلمة الثانية هي عقدة DOM التي تريد عرض ذلك المكوِّن داخلها.

تُمرر مكوِّنات React إلى `ReactDOM.render()` بطريقة مختلفة قليلاً عن عناصر JSX. لعناصر JSX، تمرر اسم العنصر الذي تريد عرضه. لكن بالنسبة لمكوِّنات React، تحتاج إلى استخدام نفس الصيغة كما لو كنت تعرض مكوِّنًا مسكنًا، مثلًا `ReactDOM.render(<ComponentToRender />, targetNode)`. تستخدم هذه الصيغة لكل من مكوِّنات الصنف ES6 والمكوِّنات الدالية.

# --instructions--

تم تعريف كل من المكوِّنات `Fruits` و `Vegetables` لك ضمن الخلفية. اعرض كلا المكوِّنَين كأبناء لمكوِّن `TypesOfFood`، ثم اعرض `TypesOfFood` إلى DOM. يوجد `div` بمعرف `id='challenge-node'` متاح للاستخدام.

# --hints--

يجب أن يُرجع مكوِّن `TypesOfFood` عنصر `div` واحدًا فقط.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

يجب أن يعرض مكوِّن `TypesOfFood` المكوِّن `Fruits` بعد عنصر `h1`.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

يجب أن يعرض مكوِّن `TypesOfFood` المكوِّن `Vegetables` بعد `Fruits`.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

يجب أن يعرض مكوِّن `TypesOfFood` إلى DOM داخل `div` الذي يحمل المعرف `challenge-node`.

```js
assert(
  (function () {
    const hasRenderCall = /ReactDOM\.render\(<TypesOfFood\/>,document\.getElementById\((['"`])challenge-node\1\)\);?/.test(
      __helpers.removeWhiteSpace(code)
    );
    if (!hasRenderCall) {
      return false;
    }

    const e = React.createElement;
    globalThis.Fruits = function Fruits() {
      return e(
        'div',
        null,
        e('h2', null, 'Fruits:'),
        e('h4', null, 'Non-Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Apples'),
          e('li', null, 'Blueberries'),
          e('li', null, 'Strawberries'),
          e('li', null, 'Bananas')
        ),
        e('h4', null, 'Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Lemon'),
          e('li', null, 'Lime'),
          e('li', null, 'Orange'),
          e('li', null, 'Grapefruit')
        )
      );
    };
    globalThis.Vegetables = function Vegetables() {
      return e(
        'div',
        null,
        e('h2', null, 'Vegetables:'),
        e(
          'ul',
          null,
          e('li', null, 'Brussel Sprouts'),
          e('li', null, 'Broccoli'),
          e('li', null, 'Squash')
        )
      );
    };

    ReactDOM.render(
      React.createElement(TypesOfFood),
      document.getElementById('challenge-node')
    );

    const root = document.getElementById('challenge-node').childNodes[0];
    if (!root) {
      return false;
    }
    const html = root.innerHTML;
    return (
      html.includes(
        '<div><h2>Fruits:</h2><h4>Non-Citrus:</h4><ul><li>Apples</li><li>Blueberries</li><li>Strawberries</li><li>Bananas</li></ul><h4>Citrus:</h4><ul><li>Lemon</li><li>Lime</li><li>Orange</li><li>Grapefruit</li></ul></div>'
      ) &&
      html.includes(
        '<div><h2>Vegetables:</h2><ul><li>Brussel Sprouts</li><li>Broccoli</li><li>Squash</li></ul></div>'
      )
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
```

# --solutions--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}
          <Fruits />
           <Vegetables />
         {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
ReactDOM.render(<TypesOfFood />, document.getElementById('challenge-node'));
```
