---
id: 5a24c314108439a4d4036165
title: استخدام React لعرض مكوِّنات مسكنة
challengeType: 6
forumTopicId: 301420
dashedName: use-react-to-render-nested-components
---

# --description--

أظهر التحدٍّ الأخير طريقة بسيطة لتركيب مكوِّنَين، لكن هناك طرقًا عديدة لتركيب المكوِّنات باستخدام React.

تركيب المكوِّنات من الميزات القوية في React. عند العمل مع React، من المهم أن تبدأ بالتفكير في واجهة المستخدم من حيث المكوِّنات مثل مثال التطبيق في التحدٍّ الأخير. تفكك واجهة المستخدم إلى وحداتها الأساسية، وتصبح تلك الوحدات مكوِّنات. هذا يساعد على فصل الشفرة المسؤولة عن واجهة المستخدم عن الشفرة المسؤولة عن منطق التطبيق. يمكن أن يبسط بشكل كبير تطوير وصيانة المشاريع المعقدة.

# --instructions--

هناك مكوِّنان وظيفيان معرفان في محرر الشفرة، اسمهما `TypesOfFruit` و `Fruits`. خذ مكوِّن `TypesOfFruit` وركّبه، أو *اسكنه*، داخل مكوِّن `Fruits`. ثم خذ مكوِّن `Fruits` واسكنه داخل مكوِّن `TypesOfFood`. النتيجة يجب أن تكون مكوِّنًا فرعيًا مسكنًا داخل مكوِّن أب، وهو بدوره مسكن داخل مكوِّن أب خاص به!

# --hints--

يجب أن يُرجع مكوِّن `TypesOfFood` عنصر `div` واحدًا.

```js
assert(Enzyme.shallow(React.createElement(TypesOfFood)).type() === 'div');
```

يجب أن يُرجع مكوِّن `TypesOfFood` مكوِّن `Fruits`.

```js
assert(
  Enzyme.shallow(React.createElement(TypesOfFood)).props().children[1].type
    .name === 'Fruits'
);
```

يجب أن يُرجع مكوِّن `Fruits` مكوِّن `TypesOfFruit`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('h2').html() ===
    '<h2>Fruits:</h2>'
);
```

يجب أن يُرجع مكوِّن `TypesOfFruit` عنصري `h2` و `ul`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('ul').text() ===
    'ApplesBlueberriesStrawberriesBananas'
);
```

# --seed--

## --seed-contents--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }

      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }
        <TypesOfFruit />
      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }
        <Fruits />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
