---
id: 5a24c314108439a4d4036170
title: إنشاء مكوِّن ذو حالة
challengeType: 6
forumTopicId: 301391
dashedName: create-a-stateful-component
---

# --description--

أحد أهم المواضيع في React هو `state`. تتكون الحالة من أي بيانات يحتاج تطبيقك إلى معرفتها، والتي يمكن أن تتغير مع مرور الوقت. تريد أن تستجيب تطبيقاتك لتغييرات الحالة وتعرض واجهة مستخدم محدثة عند الحاجة. يقدم React حلاً جيدًا لإدارة الحالة في تطبيقات الويب الحديثة.

تُنشئ الحالة في مكوّن React بإعلان خاصية `state` على صنف المكوّن داخل `constructor` الخاص به. وهذا يهيّئ المكوّن بـ `state` عند إنشائه. ويجب ضبط الخاصية `state` على `object` في JavaScript. ويبدو إعلانها هكذا:

```jsx
this.state = {

}
```

لديك وصول إلى كائن `state` طوال عمر المكوِّن. يمكنك تحديثه، عرضه في واجهة المستخدم، وتمريره كـ props إلى المكوِّنات الفرعية. يمكن أن يكون كائن `state` معقدًا أو بسيطًا حسب حاجتك. لاحظ أنه يجب إنشاء مكوِّن صنفي بتمديد `React.Component` لكي تنشئ `state` بهذه الطريقة.

# --instructions--

هناك مكوِّن في محرر الشفرة يحاول عرض خاصية `firstName` من `state` الخاص به. مع ذلك، لا يوجد `state` معرف. ابدأ تهيئة المكوِّن بـ `state` داخل دالة البناء `constructor` وعيّن اسمك إلى خاصية `firstName`.

# --hints--

يجب أن يوجد `StatefulComponent` ويعرض محتوى.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return mockedComponent.find('StatefulComponent').length === 1;
  })()
);
```

يجب أن يعرض `StatefulComponent` عنصر `div` وعنصر `h1`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

يجب تهيئة حالة `StatefulComponent` بخاصية `firstName` مُعيّنة إلى سلسلة نصية.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return (
      typeof initialState === 'object' && typeof initialState.firstName === 'string'
    );
  })()
);
```

يجب أن تعرض الخاصية `firstName` في حالة `StatefulComponent` داخل عنصر `h1`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return mockedComponent.find('h1').text() === initialState.firstName;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    // Only change code below this line

    // Only change code above this line
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      firstName: 'freeCodeCamp!'
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```
