---
id: 5a24c314108439a4d4036176
title: استخدام الحالة لتبديل عنصر
challengeType: 6
forumTopicId: 301421
dashedName: use-state-to-toggle-an-element
---

# --description--

أحيانًا قد تحتاج إلى معرفة الحالة السابقة عند تحديث الحالة. مع ذلك، قد تكون تحديثات الحالة غير متزامنة - وهذا يعني أن React قد تجمع عدة استدعاءات لـ `setState()` في تحديث واحد. لذا، لا يمكنك الاعتماد على القيمة السابقة لـ `this.state` أو `this.props` عند حساب القيمة التالية. لذلك، لا يجب أن تستخدم شفرة مثل هذه:

```jsx
this.setState({
  counter: this.state.counter + this.props.increment
});
```

بدلاً من ذلك، يجب أن تمرر إلى `setState` دالة تتيح لك الوصول إلى الحالة والخصائص. استخدام دالة مع `setState` يضمن أنك تعمل بأحدث قيم الحالة والخصائص. هذا يعني أن الشفرة أعلاه يجب إعادة كتابتها كالتالي:

```jsx
this.setState((state, props) => ({
  counter: state.counter + props.increment
}));
```

يمكنك أيضًا استخدام صيغة بدون `props` إذا كنت تحتاج فقط إلى `state`:

```jsx
this.setState(state => ({
  counter: state.counter + 1
}));
```

لاحظ أنه يجب عليك تغليف كائن البيانات بين أقواس عادية، وإلا سيظن JavaScript أنه كتلة شفرة.

# --instructions--

لدى `MyComponent` خاصية `visibility` مُهيأة إلى `false`. تعيد طريقة العرض عرضًا واحدًا إذا كانت قيمة `visibility` صحيحة، وعرضًا مختلفًا إذا كانت خاطئة.

حاليًا، لا توجد طريقة لتحديث خاصية `visibility` في `state` للمكوِّن. يجب أن تتغير القيمة بين true و false عند كل تبديل. يوجد معالِج نقر على الزر يستدعي طريقة الصنف المسماة `toggleVisibility()`. مرِّر دالة إلى `setState` لتعريف هذه الطريقة بحيث يتغير `state` لـ `visibility` إلى القيمة المعاكسة عند استدعاء الطريقة. إذا كانت `visibility` هي `false`، تضبط الطريقة القيمة إلى `true`، والعكس صحيح.

أخيرًا، انقر على الزر لرؤية العرض الشرطي للمكوِّن بناءً على `state` الخاص به.

**تلميح:** لا تنسَ ربط الكلمة المفتاحية `this` بالطريقة في `constructor`!

# --hints--

يجب أن يعيد `MyComponent` عنصر `div` يحتوي على زر `button`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('div').find('button')
    .length,
  1
);
```

يجب أن تبدأ حالة `MyComponent` بخاصية `visibility` مضبوطة على `false`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).state('visibility'),
  false
);
```

يجب أن يؤدي النقر على عنصر الزر إلى تبديل خاصية `visibility` في الحالة بين `true` و `false` وعرض عنصر `h1` بشكل شرطي.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ visibility: false });
    return mockedComponent.state('visibility');
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const third = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(!firstValue && secondValue && !thirdValue);
})();
```

يجب تمرير دالة مجهولة إلى `setState`.

```js
const paramRegex = '[a-zA-Z$_]\\w*(,[a-zA-Z$_]\\w*)?';
assert(
  new RegExp(
    'this\\.setState\\((function\\(' +
      paramRegex +
      '\\){|([a-zA-Z$_]\\w*|\\(' +
      paramRegex +
      '\\))=>)'
  ).test(__helpers.removeWhiteSpace(code))
);
```

لا يجب استخدام `this` داخل `setState`

```js
assert(!/this\.setState\([^}]*this/.test(code));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    this.toggleVisibility = this.toggleVisibility.bind(this);
  }
  toggleVisibility() {
    this.setState(state => ({
      visibility: !state.visibility
    }));
  }
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```
