---
id: 5a24c314108439a4d403617a
title: تمرير الحالة كخصائص إلى المكوِّنات الفرعية
challengeType: 6
forumTopicId: 301403
dashedName: pass-state-as-props-to-child-components
---

# --description--

رأيت العديد من الأمثلة التي تمرر الخصائص إلى عناصر JSX الفرعية ومكوِّنات React الفرعية في التحديات السابقة. قد تتساءل من أين تأتي تلك الخصائص. نمط شائع هو وجود مكوِّن ذو حالة يحتوي على `state` المهم لتطبيقك، ثم يعرض مكوِّنات فرعية. تريد أن تصل هذه المكوِّنات إلى بعض أجزاء ذلك الـ `state`، التي تُمرر كخصائص.

على سبيل المثال، ربما لديك مكوِّن `App` يعرض مكوِّن `Navbar` بين مكوِّنات أخرى. في `App` لديك `state` يحتوي على الكثير من معلومات المستخدم، لكن `Navbar` يحتاج فقط إلى الوصول إلى اسم المستخدم ليعرضه. تمرر ذلك الجزء من الـ `state` إلى مكوِّن `Navbar` كخاصية.

يوضح هذا النمط بعض المبادئ المهمة في React. الأول هو *تدفق البيانات أحادي الاتجاه*. تتدفق الحالة في اتجاه واحد نزولًا عبر شجرة مكوِّنات التطبيق، من المكوِّن الأب ذو الحالة إلى المكوِّنات الفرعية. تستقبل المكوِّنات الفرعية فقط بيانات الحالة التي تحتاجها. الثاني هو أن التطبيقات المعقدة ذات الحالة يمكن تقسيمها إلى عدد قليل فقط، أو ربما مكوِّن واحد ذو حالة. بقية المكوِّنات تستقبل الحالة من الأب كخصائص، وتعرض واجهة مستخدم بناءً على تلك الحالة. يبدأ هذا في خلق فصل حيث تُدار الحالة في جزء من الشفرة وعرض الواجهة في جزء آخر. هذا المبدأ بفصل منطق الحالة عن منطق الواجهة هو من المبادئ الأساسية في React. عند استخدامه بشكل صحيح، يجعل تصميم التطبيقات المعقدة ذات الحالة أسهل بكثير في الإدارة.

# --instructions--

المكوِّن `MyApp` ذو حالة ويعرض مكوِّن `Navbar` كمكوِّن فرعي. مرر خاصية `name` في الـ `state` إلى المكوِّن الفرعي، ثم اعرض `name` في وسم `h1` الذي هو جزء من طريقة العرض في `Navbar`. يجب أن يظهر `name` بعد النص `Hello, my name is:`.

# --hints--

يجب أن يعرض المكوِّن `MyApp` مكوِّن `Navbar` بداخله.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return (
      mockedComponent.find('MyApp').length === 1 &&
      mockedComponent.find('Navbar').length === 1
    );
  })()
);
```

يجب أن يستقبل مكوِّن `Navbar` خاصية `name` من حالة `MyApp` كخصائص.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').props());
  };
  const navProps = await setState();
  assert(navProps.name === 'TestName');
```

يجب أن يعرض عنصر `h1` في `Navbar` خاصية `name`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const navH1Before = mockedComponent.find('Navbar').find('h1').text();
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').find('h1').text());
  };
  const navH1After = await setState();
  assert(new RegExp('TestName').test(navH1After) && navH1After !== navH1Before);
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         {/* Change code below this line */}
         <Navbar />
         {/* Change code above this line */}
       </div>
    );
  }
};

class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      {/* Change code below this line */}
      <h1>Hello, my name is: </h1>
      {/* Change code above this line */}
    </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         <Navbar name={this.state.name}/>
       </div>
    );
  }
};
class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      <h1>Hello, my name is: {this.props.name}</h1>
    </div>
    );
  }
};
```
