---
id: 5a24c314108439a4d403617e
title: إضافة مستمعي الأحداث
challengeType: 6
forumTopicId: 301377
dashedName: add-event-listeners
---

# --description--

تُعد طريقة `componentDidMount()` أفضل مكان لإرفاق أي مستمعي أحداث تحتاج لإضافتهم لوظائف محددة. يوفر React نظام أحداث اصطناعي يغلف نظام الأحداث الأصلي الموجود في المتصفحات. هذا يعني أن نظام الأحداث الاصطناعي يتصرف بنفس الطريقة تمامًا بغض النظر عن متصفح المستخدم - حتى لو كانت الأحداث الأصلية قد تتصرف بشكل مختلف بين المتصفحات المختلفة.

لقد كنت تستخدم بالفعل بعض معالجات الأحداث الاصطناعية مثل `onClick()`. نظام الأحداث الاصطناعي في React رائع للاستخدام لمعظم التفاعلات التي ستديرها على عناصر DOM. مع ذلك، إذا أردت إرفاق معالج حدث لكائنات document أو window، يجب أن تفعل ذلك مباشرة.

# --instructions--

أرفق مستمع حدث في طريقة `componentDidMount()` لأحداث `keydown` واجعل هذه الأحداث تُشغّل الاستدعاء الراجع `handleKeyPress()`. يمكنك استخدام `document.addEventListener()` التي تأخذ الحدث (بين علامات اقتباس) كالمعلمة الأولى والاستدعاء الراجع كالمعلمة الثانية.

ثم، في `componentWillUnmount()`، أزل نفس مستمع الحدث هذا. يمكنك تمرير نفس المعلمات إلى `document.removeEventListener()`. من الممارسات الجيدة استخدام هذه الطريقة في دورة حياة المكوِّن لتنظيف أي شيء قبل إلغاء تركيب المكوِّن وتدميره. إزالة مستمعي الأحداث مثال على مثل هذا التنظيف.

# --hints--

يجب أن يعرض `MyComponent` عنصر `div` يضم داخله وسم `h1`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').children().find('h1').length === 1;
  })()
);
```

يجب إرفاق مستمع `keydown` إلى المستند في `componentDidMount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const didMountString = mockedComponent
      .instance()
      .componentDidMount.toString();
    return new RegExp(
      'document.addEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(didMountString);
  })()
);
```

يجب إزالة مستمع `keydown` من المستند في `componentWillUnmount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const willUnmountString = mockedComponent
      .instance()
      .componentWillUnmount.toString();
    return new RegExp(
      'document.removeEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(willUnmountString);
  })()
);
```

بمجرد تركيب المكوِّن، يجب أن يؤدي الضغط على `enter` إلى تحديث حالته والوسم `h1` المعروض.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const beforeState = mockedComponent.state('message');
  const beforeText = mockedComponent.find('h1').text();
  const pressEnterKey = () => {
    mockedComponent.instance().handleKeyPress({ keyCode: 13 });
    return waitForIt(() => {
      mockedComponent.update();
      return {
        state: mockedComponent.state('message'),
        text: mockedComponent.find('h1').text()
      };
    });
  };
  const afterKeyPress = await pressEnterKey();
  assert(
    beforeState !== afterKeyPress.state && beforeText !== afterKeyPress.text
  );
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleEnter = this.handleEnter.bind(this);
    this.handleKeyPress = this.handleKeyPress.bind(this);
  }
  // Change code below this line
  componentDidMount() {

  }
  componentWillUnmount() {

  }
  // Change code above this line
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleKeyPress = this.handleKeyPress.bind(this);
    this.handleEnter = this.handleEnter.bind(this);  }
  componentDidMount() {
    // Change code below this line
    document.addEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  componentWillUnmount() {
    // Change code below this line
    document.removeEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```
