---
id: 5a24c314108439a4d403618a
title: استخدام Array.map() لعرض العناصر ديناميكيًا
challengeType: 6
forumTopicId: 301417
dashedName: use-array-map-to-dynamically-render-elements
---

# --description--

العرض الشرطي مفيد، لكن قد تحتاج أن تعرض مكوناتك عددًا غير معروف من العناصر. غالبًا في البرمجة التفاعلية، لا يستطيع المبرمج معرفة حالة التطبيق حتى وقت التشغيل، لأن الكثير يعتمد على تفاعل المستخدم مع البرنامج. يجب على المبرمجين كتابة شفرتهم للتعامل الصحيح مع تلك الحالة المجهولة مسبقًا. استخدام `Array.map()` في React يوضح هذا المفهوم.

على سبيل المثال، تنشئ تطبيق "قائمة مهام" بسيط. كمبرمج، لا يمكنك معرفة عدد العناصر التي قد يملكها المستخدم في قائمته. تحتاج إلى إعداد المكوِّن ليعرض ديناميكيًا العدد الصحيح من عناصر القائمة قبل أن يقرر المستخدم أن اليوم هو يوم الغسيل.

# --instructions--

محرر الشفرة يحتوي على معظم مكوِّن `MyToDoList` مُعدًّا. قد تبدو بعض هذه الشفرة مألوفة إذا أكملت تحدي النموذج المسيطر. ستلاحظ وجود `textarea` وزر `button`، مع بعض الطرق التي تتبع حالاتهما، لكن لا شيء يُعرض على الصفحة بعد.

داخل `constructor`، أنشئ كائن `this.state` وعرّف حالتين: يجب تهيئة `userInput` كسلسلة نصية فارغة، وتهيئة `toDoList` كمصفوفة فارغة. بعد ذلك، احذف القيمة `null` في دالة `render()` بجانب متغير `items`. بدلًا منها، استخدم التكرار على مصفوفة `toDoList` المخزنة في الحالة الداخلية للمكوِّن لعرض عنصر `li` ديناميكيًا لكل عنصر. جرّب إدخال السلسلة النصية `eat, code, sleep, repeat` في `textarea`، ثم انقر الزر لترى النتيجة.

**ملاحظة:** قد تعرف أن جميع العناصر الشقيقة التي تُنشأ بواسطة عملية التكرار مثل هذه يجب أن تُزوّد بخاصية `key` فريدة. لا تقلق، هذا موضوع التحدي التالي.

# --hints--

يجب أن يوجد مكوِّن MyToDoList ويُعرض على الصفحة.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return mockedComponent.find('MyToDoList').length === 1;
  })()
);
```

يجب أن يكون الطفل الأول لـ `MyToDoList` عنصر `textarea`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(0).type() ===
      'textarea'
    );
  })()
);
```

يجب أن يكون الطفل الثاني لـ `MyToDoList` عنصر `br`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(1).type() === 'br'
    );
  })()
);
```

يجب أن يكون الطفل الثالث لـ `MyToDoList` عنصر `button`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(2).type() ===
      'button'
    );
  })()
);
```

يجب تهيئة حالة `MyToDoList` بحيث تكون `toDoList` مصفوفة فارغة.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      Array.isArray(initialState.toDoList) === true &&
      initialState.toDoList.length === 0
    );
  })()
);
```

يجب تهيئة حالة `MyToDoList` بحيث تكون `userInput` سلسلة نصية فارغة.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      typeof initialState.userInput === 'string' &&
      initialState.userInput.length === 0
    );
  })()
);
```

عند نقر زر `Create List`، يجب أن يعيد مكوِّن `MyToDoList` ديناميكيًا قائمة غير مرتبة تحتوي على عنصر قائمة لكل عنصر من قائمة مفصولة بفواصل تم إدخالها في عنصر `textarea`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  const state_1 = () => {
    return mockedComponent.find('ul').find('li');
  };
  const setInput = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      'testA, testB, testC'
    );
  };
  const click = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_2 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const setInput_2 = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      't1, t2, t3, t4, t5, t6'
    );
  };
  const click_1 = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_3 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const awaited_state_1 = state_1();
  const awaited_setInput = setInput();
  const awaited_click = click();
  const awaited_state_2 = state_2();
  const awaited_setInput_2 = setInput_2();
  const awaited_click_1 = click_1();
  const awaited_state_3 = state_3();
  assert(
    awaited_state_1.length === 0 &&
      awaited_state_2.nodes.length === 3 &&
      awaited_state_3.nodes.length === 6 &&
      awaited_state_2.text === 'testAtestBtestC' &&
      awaited_state_3.text === 't1t2t3t4t5t6'
  );
})();
```

# --seed--

## --seed-contents--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = null; // Change this line
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      toDoList: [],
      userInput: ''
    };
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = this.state.toDoList.map((item, i) => {
      return <li key={i}>{item}</li>;
    });
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```
