---
id: 5a24c314108439a4d4036179
title: إنشاء نموذج مُتحكم به
challengeType: 6
forumTopicId: 301384
dashedName: create-a-controlled-form
---

# --description--

أظهر التحدٍّ الأخير أن React يمكنها التحكم في الحالة الداخلية لعناصر معينة مثل `input` و`textarea`، مما يجعلها مكوِّنات مُتحكم بها. ينطبق هذا أيضًا على عناصر النماذج الأخرى، بما في ذلك عنصر `form` العادي في HTML.

# --instructions--

تم إعداد مكوِّن `MyForm` بنموذج فارغ `form` مع معالِج للإرسال. سيتم استدعاء معالِج الإرسال عند تقديم النموذج.

أضفنا زرًا يُرسل النموذج. يمكنك أن ترى أن خاصية `type` مضبوطة على `submit` مما يدل على أنه الزر الذي يتحكم بالنموذج. أضف عنصر `input` داخل `form` واضبط خاصيتي `value` و`onChange()` كما في التحدي السابق. بعد ذلك، أكمل طريقة `handleSubmit` بحيث تضبط خاصية الحالة `submit` على القيمة الحالية للمدخل في الحالة المحلية `state`.

**ملاحظة:** يجب عليك أيضًا استدعاء `event.preventDefault()` في معالِج الإرسال لمنع السلوك الافتراضي لإرسال النموذج الذي سيُحدِث تحديثًا لصفحة الويب. لملاءمة المتدربين، تم تعطيل السلوك الافتراضي هنا لمنع التحديثات التي تعيد تعيين شفرة التحدي.

أخيرًا، أنشئ وسم `h1` بعد الـ `form` يعرض قيمة `submit` من `state` الخاص بالمكوِّن. بعد ذلك، يمكنك الكتابة في النموذج والنقر على الزر (أو الضغط على Enter)، وسترى المُدخَل معروضًا على الصفحة.

# --hints--

يجب أن يُرجع `MyForm` عنصر `div` يحتوي على `form` ووسم `h1`. يجب أن يتضمن النموذج عنصر `input` وزرًا `button`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyForm));
    return (
      mockedComponent.find('div').children().find('form').length === 1 &&
      mockedComponent.find('div').children().find('h1').length === 1 &&
      mockedComponent.find('form').children().find('input').length === 1 &&
      mockedComponent.find('form').children().find('button').length === 1
    );
  })()
);
```

يجب أن تبدأ حالة `MyForm` بخاصيتي `input` و`submit`، كلاهما مضبوطان كسلاسل نصية فارغة.

```js
assert(
  Enzyme.mount(React.createElement(MyForm)).state('input') === '' &&
    Enzyme.mount(React.createElement(MyForm)).state('submit') === ''
);
```

يجب أن يؤدي الكتابة في عنصر `input` إلى تحديث خاصية `input` في حالة المكوِّن.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return mockedComponent.state('input');
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return {
      state: mockedComponent.state('input'),
      inputVal: mockedComponent.find('input').props().value
    };
  };
  const before = _1();
  const after = _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
})();
```

يجب أن يؤدي إرسال النموذج إلى تشغيل `handleSubmit` التي تضبط خاصية `submit` في الحالة مساوية للقيمة الحالية للمدخل.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'SubmitInput' } });
    return mockedComponent.state('submit');
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.state('submit');
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'SubmitInput');
})();
```

يجب أن تستدعي `handleSubmit` الدالة `event.preventDefault`

```js
assert(
  __helpers.isCalledWithNoArgs(
    'event.preventDefault',
    MyForm.prototype.handleSubmit.toString()
  )
);
```

يجب أن يعرض عنصر العنوان `h1` قيمة حقل `submit` من حالة المكوِّن.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return mockedComponent.find('h1').text();
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.find('h1').text();
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'TestInput');
})();
```

# --seed--

## --seed-contents--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          {/* Change code below this line */}

          {/* Change code above this line */}
          <button type='submit'>Submit!</button>
        </form>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    event.preventDefault();
    this.setState(state => ({
      submit: state.input
    }));
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          <input value={this.state.input} onChange={this.handleChange} />
          <button type='submit'>Submit!</button>
        </form>
        <h1>{this.state.submit}</h1>
      </div>
    );
  }
}
```
