---
id: 5a24c314108439a4d4036168
title: كتابة مكوِّن React من الصفر
challengeType: 6
forumTopicId: 301424
dashedName: write-a-react-component-from-scratch
---

# --description--

الآن بعد أن تعلمت أساسيات JSX ومكوِّنات React، حان الوقت لكتابة مكوِّن بنفسك. مكوِّنات React هي اللبنات الأساسية لتطبيقات React، لذا من المهم أن تتقن كتابتها جيدًا. تذكر أن مكوِّن React النموذجي هو `class` في ES6 يمتد من `React.Component`. يحتوي على طريقة عرض (render) تُرجع HTML (من JSX) أو `null`. هذا هو الشكل الأساسي لمكوِّن React. بمجرد أن تفهم هذا جيدًا، ستكون مستعدًا لبدء بناء مشاريع React أكثر تعقيدًا.

# --instructions--

عرّف صنفًا `MyComponent` يمتد من `React.Component`. يجب أن تُرجع طريقة العرض (render) الخاصة به `div` يحتوي على وسم `h1` به النص: `My First React Component!`. استخدم هذا النص بالضبط، فالحالة وعلامات الترقيم مهمة. تأكد أيضًا من استدعاء دالة البناء (constructor) لمكوِّنك.

اعرض هذا المكوِّن في DOM باستخدام `ReactDOM.render()`. هناك `div` معرف بـ `id='challenge-node'` متاح للاستخدام.

# --hints--

يجب أن يكون هناك مكوِّن React اسمه `MyComponent`.

```js
assert(__helpers.removeWhiteSpace(code).includes('classMyComponentextendsReact.Component{'));
```

يجب أن يحتوي `MyComponent` على وسم `h1` به النص `My First React Component!`. الحالة وعلامات الترقيم مهمة.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('h1').text() === 'My First React Component!';
  })()
);
```

يجب أن يعرض `MyComponent` في DOM.

```js
assert(document.getElementById('challenge-node').childNodes.length === 1);
```

يجب أن يحتوي `MyComponent` على دالة بناء تستدعي `super` مع `props`.

```js
assert(
  MyComponent.toString().includes('constructor(props)') &&
    MyComponent.toString().includes('super(props)')
);
```

# --seed--

## --seed-contents--

```jsx
// Change code below this line
```

# --solutions--

```jsx
// Change code below this line
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>My First React Component!</h1>
      </div>
    );
  }
};

ReactDOM.render(<MyComponent />, document.getElementById('challenge-node'));
```
