---
id: 5a24c314108439a4d4036187
title: استخدام تعبير ثلاثي للعرض الشرطي
challengeType: 6
forumTopicId: 301414
dashedName: use-a-ternary-expression-for-conditional-rendering
---

# --description--

قبل الانتقال إلى تقنيات العرض الديناميكي، هناك طريقة أخيرة لاستخدام الشروط المدمجة في JavaScript لعرض ما تريد: <dfn>المعامل الثلاثي</dfn>. يُستخدم المعامل الثلاثي غالبًا كاختصار لعبارات `if/else` في JavaScript. ليست قوية مثل عبارات `if/else` التقليدية، لكنها شائعة جدًا بين مطوّري React. أحد الأسباب هو كيف تُجمّع JSX، إذ لا يمكن إدخال عبارات `if/else` مباشرة داخل شفرة JSX. ربما لاحظت هذا في تحديات سابقة — عندما كان مطلوبًا استخدام `if/else`، كان دائمًا *خارج* عبارة `return`. يمكن أن تكون التعبيرات الثلاثية بديلًا ممتازًا إذا أردت تنفيذ منطق شرطي داخل JSX. تذكر أن المعامل الثلاثي يتكون من ثلاثة أجزاء، لكن يمكنك دمج عدة تعبيرات ثلاثية معًا. إليك الصيغة النحوية الأساسية:

```jsx
condition ? expressionIfTrue : expressionIfFalse;
```

# --instructions--

يحتوي محرر الشفرة على ثلاث ثوابت معرفة داخل طريقة `render()` لمكوِّن `CheckUserAge`. تُسمى هذه الثوابت `buttonOne` و`buttonTwo` و`buttonThree`. كل منها مُعيَّن إلى تعبير JSX بسيط يمثل عنصر زر. أولًا، ابدأ بتهيئة حالة `CheckUserAge` بحيث تكون الخاصيتان `input` و`userAge` مُعيَّنتين إلى قيم سلسلة نصية فارغة.

بمجرد أن يعرض المكوِّن معلومات على الصفحة، يجب أن يكون لدى المستخدمين طريقة للتفاعل معه. داخل عبارة `return` للمكوِّن، أنشئ تعبيرًا ثلاثيًا ينفذ المنطق التالي: عند تحميل الصفحة لأول مرة، اعرض زر الإرسال `buttonOne` على الصفحة. ثم، عندما يدخل المستخدم عمره وينقر الزر، اعرض زرًا مختلفًا بناءً على العمر. إذا أدخل المستخدم رقمًا أقل من `18`، اعرض `buttonThree`. إذا أدخل رقمًا أكبر من أو يساوي `18`، اعرض `buttonTwo`.

# --hints--

يجب أن يعرض مكوِّن `CheckUserAge` عنصر `input` واحد وزر `button` واحد فقط.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('input')
    .length === 1 &&
    Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('button')
      .length === 1
);
```

يجب تهيئة حالة مكوِّن `CheckUserAge` بخاصيتين هما `userAge` و`input`، كلاهما مُعيَّن إلى قيمة سلسلة نصية فارغة.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).state().input === '' &&
    Enzyme.mount(React.createElement(CheckUserAge)).state().userAge === ''
);
```

عند عرض مكوِّن `CheckUserAge` لأول مرة في DOM، يجب أن يكون نص الزر `button` الداخلي هو Submit.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('button').text() ===
    'Submit'
);
```

عندما يُدخل رقم أقل من 18 في عنصر الـ `input` ويُنقر على الـ `button`، يجب أن يكون نص الـ `button` الداخلي هو `You Shall Not Pass`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter3AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '3' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter17AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '17' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge3 = enter3AndClickButton();
  const userAge17 = enter17AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge3 === 'You Shall Not Pass' &&
      userAge17 === 'You Shall Not Pass'
  );
})();
```

عندما يُدخل رقم أكبر من أو يساوي 18 في عنصر الـ `input` ويُنقر على الـ `button`، يجب أن يكون نص الـ `button` الداخلي هو `You May Enter`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter35AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '35' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const userAge35 = enter35AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge18 === 'You May Enter' &&
      userAge35 === 'You May Enter'
  );
})();
```

بمجرد إرسال رقم، وإذا تغيرت قيمة الـ `input` مرة أخرى، يجب أن يعود نص الـ `button` إلى القراءة `Submit`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const changeInputDontClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '5' } });
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter10AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '10' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const changeInput1 = changeInputDontClickButton();
  const userAge10 = enter10AndClickButton();
  const changeInput2 = changeInputDontClickButton();
  assert(
    userAge18 === 'You May Enter' &&
      changeInput1 === 'Submit' &&
      userAge10 === 'You Shall Not Pass' &&
      changeInput2 === 'Submit'
  );
})();
```

يجب ألا يحتوي شفرتك على أي عبارات `if/else`.

```js
assert(
  new RegExp(/(\s|;)if(\s|\()/).test(
    code
  ) === false
);
```

# --seed--

## --seed-contents--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userAge: '',
      input: ''
    };
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {this.state.userAge === ''
          ? buttonOne
          : this.state.userAge >= 18
          ? buttonTwo
          : buttonThree}
      </div>
    );
  }
}
```
