---
id: 5a24c314108439a4d4036182
title: إضافة أنماط مضمَّنة في React
challengeType: 6
forumTopicId: 301378
dashedName: add-inline-styles-in-react
---

# --description--

ربما لاحظت في التحدٍّ الأخير وجود عدة اختلافات في الصيغة النحوية عن أنماط HTML المضمَّنة بالإضافة إلى خاصية `style` المعينة لكائن JavaScript. أولًا، أسماء بعض خصائص CSS تستخدم صيغة camel case. على سبيل المثال، في التحدٍّ الأخير تم تعيين حجم الخط باستخدام `fontSize` بدلًا من `font-size`. الكلمات المرتبطة بعلامة الشرطة مثل `font-size` غير صالحة كصيغة نحوية لخصائص كائن JavaScript، لذلك يستخدم React صيغة camel case. كقاعدة، تُكتب جميع الخصائص المرتبطة بعلامة الشرطة باستخدام camel case في JSX.

جميع وحدات طول قيم الخصائص (مثل `height`، `width`، و`fontSize`) تُفترض بوحدة `px` ما لم يُذكر خلاف ذلك. إذا أردت استخدام `em`، على سبيل المثال، ضع القيمة والوحدة بين علامات اقتباس، مثل `{fontSize: "4em"}`. بخلاف قيم الطول التي تفترض `px`، يجب وضع جميع قيم الخصائص الأخرى بين علامات اقتباس.

# --instructions--

إذا كان لديك مجموعة كبيرة من الأنماط، فيمكنك إسناد `object` للأنماط إلى ثابت للحفاظ على تنظيم شفرتك. أعلن ثابت الأنماط كمتغير عام في أعلى الملف. هيّئ الثابت `styles` وأسنِد إليه `object` بثلاث خصائص نمط وقيمها. امنح `div` لونًا `purple` وحجم خط `40` وحدًا `2px solid purple`. ثم اضبط الخاصية `style` لتساوي الثابت `styles`.

# --hints--

يجب أن تكون المتغير `styles` عبارة عن `object` يحتوي على ثلاث خصائص.

```js
assert(Object.keys(styles).length === 3);
```

يجب أن يحتوي المتغير `styles` على خاصية `color` بقيمة `purple`.

```js
assert(styles.color === 'purple');
```

يجب أن تحتوي المتغير `styles` على خاصية `fontSize` بقيمة `40`.

```js
assert(styles.fontSize == 40);
```

يجب أن يحتوي المتغير `styles` على خاصية `border` بقيمة `2px solid purple`.

```js
assert(styles.border === '2px solid purple');
```

يجب أن يعرض المكوِّن عنصر `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return mockedComponent.type() === 'div';
  })()
);
```

يجب أن تكون أنماط عنصر `div` معرفة بواسطة كائن `styles`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return (
      mockedComponent.props().style.color === 'purple' &&
      mockedComponent.props().style.fontSize == 40 &&
      mockedComponent.props().style.border === '2px solid purple'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={{color: "yellow", fontSize: 24}}>Style Me!</div>
    );
    // Change code above this line
  }
};
```

# --solutions--

```jsx
const styles = {
  color: "purple",
  fontSize: 40,
  border: "2px solid purple"
};
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={styles}>Style Me!</div>
    );
    // Change code above this line
  }
};
```
