---
id: 5a24c314108439a4d4036163
title: إنشاء مكوِّن React
challengeType: 6
forumTopicId: 301386
dashedName: create-a-react-component
---

# --description--

الطريقة الأخرى لتعريف مكوِّن React هي باستخدام صيغة `class` في ES6. في المثال التالي، يرث `Kitten` من `React.Component`:

```jsx
class Kitten extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <h1>Hi</h1>
    );
  }
}
```

هذا ينشئ صنف ES6 باسم `Kitten` يرث من صنف `React.Component`. لذلك، أصبح صنف `Kitten` يمتلك وصولًا إلى العديد من ميزات React المفيدة، مثل الحالة المحلية (local state) ودورات حياة المكوِّن (lifecycle hooks). لا تقلق إذا لم تكن مألوفًا بهذه المصطلحات بعد، فسيتم شرحها بتفصيل أكبر في تحديات لاحقة. لاحظ أيضًا أن صنف `Kitten` يحتوي على `constructor` معرف بداخله يستدعي `super()`. يستخدم `super()` لاستدعاء دالة البناء للصنف الأب، وفي هذه الحالة `React.Component`. دالة البناء هي طريقة خاصة تُستخدم أثناء تهيئة الكائنات التي تُنشأ باستخدام كلمة `class`. من أفضل الممارسات استدعاء `constructor` للمكوِّن باستخدام `super`، وتمرير `props` لكليهما. هذا يضمن تهيئة المكوِّن بشكل صحيح. في الوقت الحالي، اعلم أن تضمين هذا الشفرة أمر قياسي. قريبًا سترى استخدامات أخرى لدالة البناء وكذلك لـ `props`.

# --instructions--

تم تعريف `MyComponent` في محرر الشفرة باستخدام صيغة الصنف. أكمل كتابة طريقة `render` بحيث تُعيد عنصر `div` يحتوي على `h1` بالنص `Hello React!`.

# --hints--

يجب أن يُعيد مكوِّن React عنصر `div`.

```js
assert(Enzyme.shallow(React.createElement(MyComponent)).type() === 'div');
```

يجب أن يعرض عنصر `div` المُعاد عنصر عنوان `h1` بداخله.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.shallow(React.createElement(MyComponent)).html()
  )
);
```

يجب أن يحتوي عنصر العنوان `h1` على السلسلة النصية `Hello React!`.

```js
assert(
  Enzyme.shallow(React.createElement(MyComponent)).html() ===
    '<div><h1>Hello React!</h1></div>'
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line



    // Change code above this line
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line
    return (
      <div>
        <h1>Hello React!</h1>
      </div>
    );
    // Change code above this line
  }
};
```
