---
id: 5a24c314108439a4d4036180
title: تحسين عمليات إعادة العرض باستخدام shouldComponentUpdate
challengeType: 6
forumTopicId: 301398
dashedName: optimize-re-renders-with-shouldcomponentupdate
---

# --description--

حتى الآن، إذا استقبل أي مكوِّن `state` جديد أو `props` جديدة، فإنه يُعيد العرض لنفسه ولكل مكوناته الفرعية. هذا عادةً مقبول. لكن React توفر طريقة دورة حياة يمكنك استدعاؤها عندما تستقبل المكونات الفرعية `state` أو `props` جديدة، وتُعلن تحديدًا إذا كان يجب تحديث المكونات أم لا. الطريقة هي `shouldComponentUpdate()`، وتأخذ `nextProps` و `nextState` كمعلمات.

هذه الطريقة مفيدة لتحسين الأداء. على سبيل المثال، السلوك الافتراضي هو أن مكونك يُعيد العرض عندما يستقبل `props` جديدة، حتى لو لم تتغير هذه `props`. يمكنك استخدام `shouldComponentUpdate()` لمنع ذلك بمقارنة `props`. يجب أن تُرجع الطريقة قيمة منطقية (`boolean`) تخبر React إذا كان يجب تحديث المكون أم لا. يمكنك مقارنة الخصائص الحالية (`this.props`) مع الخصائص القادمة (`nextProps`) لتحديد إذا كنت بحاجة للتحديث، وترجع `true` أو `false` بناءً على ذلك.

# --instructions--

تُضاف طريقة `shouldComponentUpdate()` في مكوِّن يُسمى `OnlyEvens`. حاليًا، تُرجع هذه الطريقة القيمة `true` بحيث يُعاد عرض `OnlyEvens` في كل مرة يستقبل فيها خصائص جديدة `props`. عدّل الطريقة بحيث يُحدّث `OnlyEvens` نفسه فقط إذا كانت قيمة `value` في خصائصه الجديدة زوجية. انقر على زر `Add` وراقب ترتيب الأحداث في وحدة تحكم المتصفح أثناء تفعيل دوال دورة الحياة.

# --hints--

يجب أن يعرض مكوِّن `Controller` المكوِّن `OnlyEvens` كمكوِّن فرعي.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Controller));
    return (
      mockedComponent.find('Controller').length === 1 &&
      mockedComponent.find('OnlyEvens').length === 1
    );
  })()
);
```

يجب تعريف طريقة `shouldComponentUpdate` على مكوِّن `OnlyEvens`.

```js
assert(
  (() => {
    const child = React.createElement(OnlyEvens)
      .type.prototype.shouldComponentUpdate.toString()
      .replace(/s/g, '');
    return child !== 'undefined';
  })()
);
```

يجب أن يُرجع مكوِّن `OnlyEvens` وسم `h1` يعرض قيمة `this.props.value`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 1000 });
    return mockedComponent.find('h1').html();
  };
  const second = () => {
    mockedComponent.setState({ value: 10 });
    return mockedComponent.find('h1').html();
  };
  const firstValue = first();
  const secondValue = second();
  assert(firstValue === '<h1>1000</h1>' && secondValue === '<h1>10</h1>');
})();
```

يجب أن يُعيد `OnlyEvens` العرض فقط عندما تكون قيمة `nextProps.value` زوجية.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 8 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ value: 7 });
    return mockedComponent.find('h1').text();
  };
  const third = () => {
    mockedComponent.setState({ value: 42 });
    return mockedComponent.find('h1').text();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(firstValue === '8' && secondValue === '8' && thirdValue === '42');
})();
```

# --seed--

## --seed-contents--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return true;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```

# --solutions--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return nextProps.value % 2 === 0;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```
