---
id: 5a24c314108439a4d4036177
title: كتابة عداد بسيط
challengeType: 6
forumTopicId: 301425
dashedName: write-a-simple-counter
---

# --description--

يمكنك تصميم مكوِّن ذو حالة أكثر تعقيدًا بدمج المفاهيم التي تم تغطيتها حتى الآن. تشمل هذه المفاهيم تهيئة `state`، وكتابة طرق تضبط `state`، وتعيين معالجات النقر لتحفيز هذه الطرق.

# --instructions--

يحفظ مكوِّن `Counter` قيمة `count` في `state`. هناك زران يستدعيان الطريقتين `increment()` و `decrement()`. اكتب هذه الطرق بحيث تزداد أو تنقص قيمة العداد بمقدار 1 عند نقر الزر المناسب. أيضًا، أنشئ طريقة `reset()` بحيث عند نقر زر إعادة التعيين، تُضبط القيمة على 0.

**ملاحظة:** تأكد من عدم تعديل `className`s للأزرار. وتذكر أيضًا إضافة الربط اللازم للطرق التي أنشأتها حديثًا في دالة البناء.

# --hints--

يجب أن يُرجع `Counter` عنصر `div` يحتوي على ثلاثة أزرار بنصوص مرتبة بهذا الترتيب: `Increment!`، `Decrement!`، `Reset`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Counter));
    return (
      mockedComponent.find('.inc').text() === 'Increment!' &&
      mockedComponent.find('.dec').text() === 'Decrement!' &&
      mockedComponent.find('.reset').text() === 'Reset'
    );
  })()
);
```

يجب أن يبدأ `Counter` بحالة تحتوي على خاصية `count` بقيمة `0`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
assert(mockedComponent.find('h1').text() === 'Current Count: 0');
```

يجب أن يزيد نقر زر الزيادة قيمة العداد بمقدار `1`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.inc').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: 1');
```

يجب أن ينقص نقر زر النقصان قيمة العداد بمقدار `1`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.dec').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: -1');
```

يجب أن يعيد نقر زر إعادة التعيين قيمة العداد إلى `0`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.setState({ count: 5 });
const currentCountElement = mockedComponent.find('h1');
assert(currentCountElement.text() === 'Current Count: 5');
mockedComponent.find('.reset').simulate('click');
assert(currentCountElement.text() === 'Current Count: 0');
```

# --seed--

## --seed-contents--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
  this.increment = this.increment.bind(this);
 this.decrement = this.decrement.bind(this);
 this.reset = this.reset.bind(this);
 }
  reset() {
    this.setState({
      count: 0
    });
  }
  increment() {
    this.setState(state => ({
      count: state.count + 1
    }));
  }
  decrement() {
    this.setState(state => ({
      count: state.count - 1
    }));
  }
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```
