---
id: 5a24c314108439a4d403616b
title: استخدام الخصائص الافتراضية
challengeType: 6
forumTopicId: 301418
dashedName: use-default-props
---

# --description--

يوفر React خيار تعيين الخصائص الافتراضية أيضًا. يمكنك تعيين الخصائص الافتراضية لمكوِّن كخاصية على المكوِّن نفسه، ويُعيّن React الخاصية الافتراضية عند الحاجة. هذا يسمح لك بتحديد قيمة الخاصية إذا لم تُقدّم قيمة صريحة. على سبيل المثال، إذا أعلنت `MyComponent.defaultProps = { location: 'San Francisco' }`، تكون قد عرّفت خاصية location بقيمة السلسلة النصية `San Francisco`، ما لم تحدد خلاف ذلك. يعيّن React الخصائص الافتراضية إذا كانت الخصائص غير معرفة، لكن إذا مرّرت `null` كقيمة لخاصية، ستظل `null`.

# --instructions--

يعرض محرر الشفرة مكوِّن `ShoppingCart`. عرّف الخصائص الافتراضية على هذا المكوِّن بحيث تحدد خاصية `items` بقيمة `0`.

# --hints--

يجب أن يعرض مكوِّن `ShoppingCart`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

يجب أن يحتوي مكوِّن `ShoppingCart` على خاصية افتراضية `{ items: 0 }`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    mockedComponent.setProps({ items: undefined });
    return mockedComponent.find('ShoppingCart').props().items === 0;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};
// Change code below this line
```

# --solutions--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};

// Change code below this line
ShoppingCart.defaultProps = {
  items: 0
}
```
