---
id: 5a24c314108439a4d4036184
title: العرض باستخدام شرط If-Else
challengeType: 6
forumTopicId: 301410
dashedName: render-with-an-if-else-condition
---

# --description--

تطبيق آخر لاستخدام JavaScript للتحكم في العرض هو ربط العناصر التي تُعرض بشرط معين. عندما يكون الشرط صحيحًا، يُعرض شكل معين. وعندما يكون خاطئًا، يُعرض شكل مختلف. يمكنك فعل ذلك باستخدام جملة `if/else` القياسية في طريقة `render()` لمكوِّن React.

# --instructions--

يحتوي MyComponent على متغير من نوع `boolean` في الحالة يتتبع ما إذا كنت تريد إظهار عنصر ما في واجهة المستخدم أم لا. يقوم الـ `button` بتبديل حالة هذه القيمة. حاليًا، يعرض نفس واجهة المستخدم في كل مرة. أعد كتابة طريقة `render()` مستخدمًا عبارة `if/else` بحيث إذا كانت قيمة `display` هي `true`، تُرجع الترميز الحالي. أما إذا لم تكن كذلك، فأرجع الترميز بدون عنصر `h1`.

**ملاحظة:** يجب كتابة جملة `if/else` لاجتياز الاختبارات. استخدام المعامل الثلاثي لن ينجح هنا.

# --hints--

يجب أن يوجد `MyComponent` ويُعرض.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length === 1;
  })()
);
```

عندما تكون قيمة `display` مضبوطة على `true`، يجب عرض عناصر `div` و`button` و`h1`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 2 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 1
  );
```

عندما تكون قيمة `display` مضبوطة على `false`، يجب عرض عناصر `div` و`button` فقط.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 0
  );
```

يجب أن تستخدم طريقة العرض عبارة `if/else` للتحقق من حالة `this.state.display`.

```js
assert(code.includes('if') && code.includes('else'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line

    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
 }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    if (this.state.display) {
      return (
         <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
           <h1>Displayed!</h1>
         </div>
      );
    } else {
      return (
        <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
         </div>
      );
    }
  }
};
```
