---
id: 5a24c314108439a4d4036188
title: العرض الشرطي باستخدام الخصائص
challengeType: 6
forumTopicId: 301405
dashedName: render-conditionally-from-props
---

# --description--

حتى الآن، رأيت كيف تستخدم `if/else`، و`&&`، ومعامل ثلاثي (`condition ? expressionIfTrue : expressionIfFalse`) لاتخاذ قرارات شرطية حول ما يجب عرضه ومتى. مع ذلك، هناك موضوع مهم واحد متبقٍ للنقاش يسمح لك بدمج أي من هذه المفاهيم أو كلها مع ميزة قوية أخرى في React: الخصائص. استخدام الخصائص للعرض الشرطي للشفرة شائع جدًا بين مطوّري React — أي أنهم يستخدمون قيمة خاصية معينة لاتخاذ قرارات تلقائية حول ما يجب عرضه.

في هذا التحدٍّ، ستُعد مكوّنًا فرعيًا لاتخاذ قرارات العرض بناءً على الخصائص. ستستخدم أيضًا المعامل الثلاثي، لكن يمكنك أن ترى كيف أن عدة مفاهيم أخرى تم تغطيتها في التحديات السابقة قد تكون مفيدة بنفس القدر في هذا السياق.

# --instructions--

محرر الشفرة يحتوي على مكوّنين معرّفين جزئيًا لك: مكوّن رئيسي اسمه `GameOfChance`، ومكوّن فرعي اسمه `Results`. يُستخدمان لإنشاء لعبة بسيطة يضغط فيها المستخدم زرًا ليرى إذا كان سيربح أو يخسر.

أولًا، تحتاج إلى تعبير بسيط يعيد قيمة مختلفة عشوائيًا في كل مرة يُنفذ فيها. يمكنك استخدام `Math.random()`. هذه الطريقة تعيد قيمة بين `0` (شاملة) و`1` (غير شاملة) في كل استدعاء. لذا للحصول على احتمالات 50/50، استخدم التعبير `Math.random() >= .5`. إحصائيًا، هذا التعبير سيعيد `true` بنسبة 50% من الوقت، و`false` النصف الآخر. في طريقة العرض، استبدل `null` بالتعبير السابق لإكمال إعلان المتغير.

الآن لديك تعبير يمكنك استخدامه لاتخاذ قرار عشوائي في الشفرة. بعد ذلك، تحتاج إلى تنفيذ ذلك. اعرض مكوّن `Results` كمكوّن فرعي من `GameOfChance`، ومرّر `expression` كخاصية اسمها `fiftyFifty`. في مكوّن `Results`، اكتب تعبيرًا ثلاثيًا لعرض عنصر `h1` بالنص `You Win!` أو `You Lose!` بناءً على خاصية `fiftyFifty` التي تُمرر من `GameOfChance`. أخيرًا، تأكد من أن طريقة `handleClick()` تعد كل دور بشكل صحيح ليعرف المستخدم عدد المرات التي لعب فيها. هذا أيضًا يخبر المستخدم أن المكوّن قد تم تحديثه فعليًا في حال فاز أو خسر مرتين متتاليتين.

# --hints--

يجب أن يوجد مكوّن `GameOfChance` ويُعرض على الصفحة.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('GameOfChance').length,
  1
);
```

يجب أن يُرجع `GameOfChance` عنصر زر واحد `button`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('button').length,
  1
);
```

يجب أن يُرجع `GameOfChance` نسخة واحدة من مكوّن `Results`، الذي يحتوي على خاصية اسمها `fiftyFifty`.

```js
assert(
  Enzyme.mount(React.createElement(GameOfChance)).find('Results').length ===
    1 &&
    Enzyme.mount(React.createElement(GameOfChance))
      .find('Results')
      .props()
      .hasOwnProperty('fiftyFifty') === true
);
```

يجب تهيئة حالة `GameOfChance` بمتغير `counter` بقيمة `1`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).state().counter,
  1
);
```

عند عرض مكوّن `GameOfChance` لأول مرة في DOM، يجب أن يُرجع عنصر `p` يحتوي على النص الداخلي `Turn: 1`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('p').text(),
  'Turn: 1'
);
```

في كل مرة يُنقر فيها الزر، يجب زيادة حالة العداد بمقدار 1، ويُعرض عنصر `p` واحد في DOM يحتوي على النص `Turn: N`، حيث `N` هي قيمة حالة العداد.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    count: comp.state('counter'),
    text: comp.find('p').text()
  });
  const _1 = () => {
    simulate();
    return result();
  };
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _1_val = _1();
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  assert(
    _1_val.count === 2 &&
      _1_val.text === 'Turn: 2' &&
      _2_val.count === 3 &&
      _2_val.text === 'Turn: 3' &&
      _3_val.count === 4 &&
      _3_val.text === 'Turn: 4' &&
      _4_val.count === 5 &&
      _4_val.text === 'Turn: 5' &&
      _5_val.count === 6 &&
      _5_val.text === 'Turn: 6'
  );
})();
```

عند تركيب مكوّن `GameOfChance` لأول مرة في DOM وفي كل مرة يُنقر فيها الزر بعد ذلك، يجب أن يُرجع عنصر `h1` واحد يعرض عشوائيًا إما `You Win!` أو `You Lose!`. ملاحظة: قد يفشل هذا عشوائيًا. إذا حدث ذلك، حاول مرة أخرى.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    h1: comp.find('h1').length,
    text: comp.find('h1').text()
  });
  const _1 = result();
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _6 = () => {
    simulate();
    return result();
  };
  const _7 = () => {
    simulate();
    return result();
  };
  const _8 = () => {
    simulate();
    return result();
  };
  const _9 = () => {
    simulate();
    return result();
  };
  const _10 = () => {
    simulate();
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  const _6_val = _6();
  const _7_val = _7();
  const _8_val = _8();
  const _9_val = _9();
  const _10_val = _10();
  const __text = new Set([
    _1.text,
    _2_val.text,
    _3_val.text,
    _4_val.text,
    _5_val.text,
    _6_val.text,
    _7_val.text,
    _8_val.text,
    _9_val.text,
    _10_val.text
  ]);
  const __h1 = new Set([
    _1.h1,
    _2_val.h1,
    _3_val.h1,
    _4_val.h1,
    _5_val.h1,
    _6_val.h1,
    _7_val.h1,
    _8_val.h1,
    _9_val.h1,
    _10_val.h1
  ]);
  assert(__text.size === 2 && __h1.size === 1);
})();
```

# --seed--

## --seed-contents--

```jsx
class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    {/* Change code below this line */}
    return <h1></h1>;
    {/* Change code above this line */}
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      // Complete the return statement:
      return {
        counter: prevState
      }
    });
  }
  render() {
    const expression = null; // Change this line
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        {/* Change code below this line */}

        {/* Change code above this line */}
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```

# --solutions--

```jsx
// We want this to be deterministic for testing purposes.
const randomSequence = [true, false, false, true, true, false, false, true, true, false];
let index = 0;
const fiftyFifty = () => randomSequence[index++ % randomSequence.length];

class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h1>{this.props.fiftyFifty ? 'You Win!' : 'You Lose!'}</h1>;
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      return {
        counter: prevState.counter + 1
      }
    });
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        <Results fiftyFifty={fiftyFifty()} />
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```
