---
id: 5a24c314108439a4d403616f
title: مراجعة استخدام الخصائص مع المكونات الوظيفية عديمة الحالة
challengeType: 6
forumTopicId: 301411
dashedName: review-using-props-with-stateless-functional-components
---

# --description--

باستثناء التحدي الأخير، كنت تمرر الخصائص إلى المكونات الوظيفية عديمة الحالة. هذه المكونات تعمل مثل الدوال النقية. تستقبل الخصائص كمُدخَل وتُعيد نفس العرض في كل مرة تُمرر فيها نفس الخصائص. قد تتساءل ما هي الحالة، والتحدي التالي سيغطيها بمزيد من التفصيل. قبل ذلك، إليك مراجعة لمصطلحات المكونات.

المكوِّن الوظيفي عديم الحالة هو أي دالة تكتبها تستقبل الخصائص وتُعيد JSX. أما المكوِّن عديم الحالة فهو صنف يرث من `React.Component` لكنه لا يستخدم الحالة الداخلية (سيُشرح في التحدي التالي). وأخيرًا، المكوِّن ذو الحالة هو صنف يحتفظ بحالته الداخلية. قد ترى المكونات ذات الحالة تُسمى ببساطة مكونات أو مكونات React.

النمط الشائع هو محاولة تقليل استخدام الحالة وإنشاء مكونات وظيفية عديمة الحالة حيثما أمكن. هذا يساعد على حصر إدارة الحالة في مساحة محددة من التطبيق. وبدوره، يحسن تطوير وصيانة التطبيق بجعل تتبع تأثير تغييرات الحالة على السلوك أسهل.

# --instructions--

يحتوي محرر الشفرة على مكوِّن `CampSite` يعرض مكوِّن `Camper` كطفل. عرّف المكوِّن `Camper` وعيّن له الخصائص الافتراضية `{ name: 'CamperBot' }`. داخل مكوِّن `Camper`، اعرض أي شفرة تريدها، لكن تأكد من وجود عنصر `p` واحد يحتوي فقط على قيمة `name` التي تُمرر كـ `prop`. أخيرًا، عرّف `propTypes` على مكوِّن `Camper` لتطلب توفير `name` كخاصية وتتحقق من أنه من نوع `string`.

# --hints--

يجب أن يعرض مكوِّن `CampSite`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('CampSite').length === 1;
  })()
);
```

يجب أن يعرض مكوِّن `Camper`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('Camper').length === 1;
  })()
);
```

يجب أن يتضمن مكوِّن `Camper` خصائص افتراضية تعيّن السلسلة النصية `CamperBot` للمفتاح `name`.

```js
assert(
  /Camper.defaultProps={name:(['"`])CamperBot\1,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

يجب أن يتضمن مكوِّن `Camper` أنواع الخصائص التي تطلب أن تكون خاصية `name` من نوع `string`.

```js
assert(
  /Camper.propTypes={name:PropTypes.string.isRequired,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

يجب أن يحتوي مكوِّن `Camper` على عنصر `p` يحتوي فقط على النص من خاصية `name`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return (
      mockedComponent.find('p').text() ===
      mockedComponent.find('Camper').props().name
    );
  })()
);
```

# --before-all--

```js
var PropTypes = {
   string: { isRequired: true }
};
```

# --seed--

## --seed-contents--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line
```

# --solutions--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line

const Camper = (props) => {
   return (
     <div>
       <p>{props.name}</p>
     </div>
   );
};

Camper.propTypes = {
  name: PropTypes.string.isRequired
};

Camper.defaultProps = {
  name: 'CamperBot'
};
```
