---
id: 5a24c314108439a4d403618b
title: أعطِ العناصر الشقيقة خاصية مفتاح فريدة
challengeType: 6
forumTopicId: 301394
dashedName: give-sibling-elements-a-unique-key-attribute
---

# --description--

أظهر التحدٍّ الأخير كيف تُستخدم طريقة `map` لعرض عدد من العناصر ديناميكيًا بناءً على مُدخل المستخدم. مع ذلك، كان هناك جزء مهم مفقود من ذلك المثال. عند إنشاء مصفوفة من العناصر، يحتاج كل عنصر إلى خاصية `key` تُعيّن إلى قيمة فريدة. يستخدم React هذه المفاتيح لتتبع العناصر التي أُضيفت أو تغيّرت أو أُزيلت. هذا يساعد على جعل عملية إعادة العرض أكثر كفاءة عند تعديل القائمة بأي شكل.

**ملاحظة:** يجب أن تكون المفاتيح فريدة فقط بين العناصر الشقيقة، ولا يلزم أن تكون فريدة على مستوى التطبيق بأكمله.

# --instructions--

يحتوي محرر الشفرة على مصفوفة بها بعض أُطُر العمل للواجهة الأمامية ومكوِّن دالي عديم الحالة اسمه `Frameworks()`. يحتاج `Frameworks()` إلى تحويل المصفوفة إلى قائمة غير مرتبة، تمامًا كما في التحدي السابق. أكمل كتابة استدعاء `map` لإرجاع عنصر `li` لكل إطار عمل في مصفوفة `frontEndFrameworks`. هذه المرة، تأكد من إعطاء كل عنصر `li` خاصية `key` بقيمة فريدة. يجب أن تحتوي عناصر `li` أيضًا على نص من `frontEndFrameworks`.

عادةً، تريد أن تجعل المفتاح شيئًا يميز العنصر المعروض بشكل فريد. كحل أخير يمكن استخدام فهرس المصفوفة، لكن من الأفضل عادةً محاولة استخدام معرف فريد.

# --hints--

يجب أن يوجد مكوِّن `Frameworks` ويُعرض على الصفحة.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('Frameworks').length === 1
);
```

يجب أن يعرض `Frameworks` عنصر `h1`.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('h1').length === 1);
```

يجب أن يعرض `Frameworks` عنصر `ul`.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('ul').length === 1);
```

يجب أن يعرض وسم `ul` ستة عناصر فرعية من نوع `li`.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('ul').children().length ===
    6 &&
    Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .childAt(0)
      .name() === 'li' &&
    Enzyme.mount(React.createElement(Frameworks)).find('li').length === 6
);
```

يجب أن يحتوي كل عنصر من عناصر القائمة على خاصية `key` فريدة.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(Frameworks)).find('ul');
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key(),
      ul.childAt(4).key(),
      ul.childAt(5).key()
    ]);
    return keys.size === 6;
  })()
);
```

يجب أن يحتوي كل عنصر من عناصر القائمة على نص من `frontEndFrameworks`.

```js
assert(
  (() => {
    const li = Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .children();
    return [...Array(5)].every((_, i) =>
      frontEndFrameworks.includes(li.at(i).text())
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = null; // Change this line
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```

# --solutions--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = frontEndFrameworks.map((fw, i) => <li key={i}>{fw}</li>);
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```
