---
id: 5a24c314108439a4d4036166
title: تركيب مكوِّنات React
challengeType: 6
forumTopicId: 301381
dashedName: compose-react-components
---

# --description--

مع استمرار التحديات في استخدام تراكيب أكثر تعقيدًا مع مكوِّنات React و JSX، هناك نقطة مهمة يجب ملاحظتها. عرض مكوِّنات الصنف بأسلوب ES6 داخل مكوِّنات أخرى لا يختلف عن عرض المكوِّنات البسيطة التي استخدمتها في التحديات السابقة. يمكنك عرض عناصر JSX، والمكوِّنات الوظيفية عديمة الحالة، ومكوِّنات الصنف بأسلوب ES6 داخل مكوِّنات أخرى.

# --instructions--

في محرر الشفرة، مكوِّن `TypesOfFood` يعرض بالفعل مكوِّنًا يسمى `Vegetables`. أيضًا، هناك مكوِّن `Fruits` من التحدي السابق.

قم بتسكين مكوِّنَين داخل `Fruits` — أولًا `NonCitrus`، ثم `Citrus`. كلا هذين المكوِّنَين متوفران لك خلف الكواليس. بعد ذلك، قم بتسكين مكوِّن الصنف `Fruits` داخل مكوِّن `TypesOfFood`، أسفل عنصر العنوان `h1` وفوق `Vegetables`. يجب أن تكون النتيجة سلسلة من المكوِّنات المسكنة، تستخدم نوعين مختلفين من المكوِّنات.

# --before-all--

```js
globalThis.NonCitrus = function NonCitrus() {
  return React.createElement('div', null);
};
globalThis.Citrus = function Citrus() {
  return React.createElement('div', null);
};
globalThis.Vegetables = function Vegetables() {
  return React.createElement('div', null);
};
```

# --hints--

يجب أن يُرجع مكوِّن `TypesOfFood` عنصر `div` واحدًا فقط.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

يجب أن يُرجع مكوِّن `TypesOfFood` مكوِّن `Fruits`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

يجب أن يُرجع مكوِّن `Fruits` مكوِّنَي `NonCitrus` و `Citrus`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return (
      mockedComponent.find('Fruits').children().find('NonCitrus').length ===
        1 &&
      mockedComponent.find('Fruits').children().find('Citrus').length === 1
    );
  })()
);
```

يجب أن يُرجع مكوِّن `TypesOfFood` مكوِّن `Vegetables` أسفل مكوِّن `Fruits`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
        <Vegetables />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }
        <NonCitrus />
        <Citrus />
        { /* Change code above this line */ }
      </div>
    )
  }
}

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
    render() {
      return (
        <div>
        <h1>Types of Food:</h1>
          { /* Change code below this line */ }
          <Fruits />
          { /* Change code above this line */ }
          <Vegetables />
        </div>
      );
    }
};
```
