---
id: 5a24c314108439a4d403616a
title: تمرير مصفوفة كخصائص
challengeType: 6
forumTopicId: 301401
dashedName: pass-an-array-as-props
---

# --description--

أظهر التحدٍّ الأخير كيفية تمرير المعلومات من مكوِّن أب إلى مكوِّن فرعي كـ `props` أو خصائص. هذا التحدٍّ يشرح كيف يمكن تمرير المصفوفات كـ `props`. لتمرير مصفوفة إلى عنصر JSX، يجب التعامل معها كجافاسكريبت وتغليفها بأقواس متعرجة.

```jsx
<ParentComponent>
  <ChildComponent colors={["green", "blue", "red"]} />
</ParentComponent>
```

بعدها يحصل المكوِّن الفرعي على وصول إلى خاصية المصفوفة `colors`. يمكن استخدام طرق المصفوفة مثل `join()` عند الوصول إلى الخاصية.

```jsx
const ChildComponent = (props) => <p>{props.colors.join(', ')}</p>
``` 

هذا سيجمع كل عناصر مصفوفة `colors` في سلسلة نصية مفصولة بفواصل وينتج: `<p>green, blue, red</p>`. لاحقًا، سنتعلم عن طرق شائعة أخرى لعرض مصفوفات البيانات في React.

# --instructions--

يوجد مكوِّنات `List` و `ToDo` في محرر الشفرة. عند عرض كل `List` من مكوِّن `ToDo`، مرر خاصية `tasks` مُعيَّنة إلى مصفوفة من مهام قائمة المهام، مثل `["walk dog", "workout"]`. ثم وصول إلى مصفوفة `tasks` هذه في مكوِّن `List`، وعرض قيمتها داخل عنصر `p`. استخدم `join(", ")` لعرض مصفوفة `props.tasks` في عنصر `p` كقائمة مفصولة بفواصل. يجب أن تحتوي قائمة اليوم على مهمتين على الأقل وقائمة الغد على ثلاث مهام على الأقل.

# --hints--

يجب أن يُرجع مكوِّن `ToDo` عنصر `div` خارجي واحد.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().type() === 'div';
  })()
);
```

يجب أن يكون الطفل الثالث لمكوِّن `ToDo` نسخة من مكوِّن `List`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(2).name() === 'List';
  })()
);
```

يجب أن يكون الطفل الخامس لمكوِّن `ToDo` نسخة من مكوِّن `List`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(4).name() === 'List';
  })()
);
```

يجب أن تحتوي كلتا النسختين من مكوِّن `List` على خاصية تسمى `tasks` ويجب أن تكون `tasks` من نوع مصفوفة.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      Array.isArray(mockedComponent.find('List').get(0).props.tasks) &&
      Array.isArray(mockedComponent.find('List').get(1).props.tasks)
    );
  })()
);
```

يجب أن يحتوي مكوِّن `List` الأول الذي يمثل مهام اليوم على عنصرين أو أكثر.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(0).props.tasks.length >= 2;
  })()
);
```

يجب أن يحتوي مكوِّن `List` الثاني الذي يمثل مهام الغد على ثلاثة عناصر أو أكثر.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(1).props.tasks.length >= 3;
  })()
);
```

يجب أن يعرض مكوِّن `List` القيمة من خاصية `tasks` داخل وسم `p`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      mockedComponent
        .find('p')
        .get(0)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(0)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',') &&
      mockedComponent
        .find('p')
        .get(1)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(1)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',')
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const List = (props) => {
  { /* Change code below this line */ }
  return <p>{}</p>
  { /* Change code above this line */ }
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        { /* Change code below this line */ }
        <List/>
        <h2>Tomorrow</h2>
        <List/>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const List= (props) => {
  return <p>{props.tasks.join(', ')}</p>
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        <List tasks={['study', 'exercise']} />
        <h2>Tomorrow</h2>
        <List tasks={['call Sam', 'grocery shopping', 'order tickets']} />
      </div>
    );
  }
};
```
