---
id: 5a24c314108439a4d403616c
title: تجاوز الخصائص الافتراضية
challengeType: 6
forumTopicId: 301399
dashedName: override-default-props
---

# --description--

إمكانية تعيين الخصائص الافتراضية ميزة مفيدة في React. طريقة تجاوز الخصائص الافتراضية هي تعيين قيم الخصائص صراحةً لمكوِّن معين.

# --instructions--

المكوِّن `ShoppingCart` يعرض الآن مكوِّنًا فرعيًا اسمه `Items`. هذا المكوِّن `Items` لديه خاصية افتراضية `quantity` مُعيّنة إلى العدد الصحيح `0`. تجاوز الخاصية الافتراضية بتمرير قيمة `10` لـ `quantity`.

**ملاحظة:** تذكّر أن الصيغة النحوية لإضافة خاصية إلى مكوِّن تشبه كيفية إضافة خصائص HTML. لكن بما أن قيمة `quantity` عدد صحيح، فلا توضع بين علامات اقتباس بل يجب أن تُغلف بأقواس متعرجة. على سبيل المثال، `{100}`. هذه الصيغة تخبر JSX بتفسير القيمة داخل الأقواس مباشرةً ككود جافاسكريبت.

# --hints--

يجب أن يعرض المكوِّن `ShoppingCart`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

يجب أن يعرض المكوِّن `Items`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

يجب أن يحتوي المكوِّن `Items` على خاصية `{ quantity: 10 }` يتم تمريرها من المكوِّن `ShoppingCart`.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
      return (
        mockedComponent.find('Items').props().quantity == 10 &&
        code.replace(/ /g, '')
          .includes('<Itemsquantity={10}/>')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items />
    { /* Change code above this line */ }
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items quantity = {10} />
    { /* Change code above this line */ }
  }
};
```
