---
id: 5a24c314108439a4d4036165
title: استخدام React لعرض مكوِّنات مسكنة
challengeType: 6
forumTopicId: 301420
dashedName: use-react-to-render-nested-components
---

# --description--

أظهر التحدي الأخير طريقة بسيطة لتركيب مكوِّنَين، لكن هناك طرقًا عديدة لتركيب المكوِّنات باستخدام React.

تكوين المكوِّنات هو إحدى الميزات القوية في React. عند العمل مع React، من المهم أن تبدأ بالتفكير في واجهة المستخدم الخاصة بك من حيث المكوِّنات مثل مثال التطبيق في التحدي السابق. تقوم بتقسيم واجهة المستخدم إلى وحداتها الأساسية، وتصبح تلك الوحدات مكوِّنات. هذا يساعد على فصل الكود المسؤول عن واجهة المستخدم عن الكود المسؤول عن منطق التطبيق. يمكن لهذا أن يبسط بشكل كبير تطوير وصيانة المشاريع المعقدة.

# --instructions--

هناك مكوِّنان وظيفيان معرفان في محرر الكود، اسمهما `TypesOfFruit` و `Fruits`. خذ مكوِّن `TypesOfFruit` وقم بتكوينه، أو *تسكينه*، داخل مكوِّن `Fruits`. ثم خذ مكوِّن `Fruits` وقم بتسكينه داخل مكوِّن `TypesOfFood`. يجب أن تكون النتيجة مكوِّنًا فرعيًا مسكنًا داخل مكوِّن أب، وهو بدوره مسكن داخل مكوِّن أب خاص به!

# --hints--

يجب أن يُرجع مكوِّن `TypesOfFood` عنصر `div` واحدًا فقط.

```js
assert(Enzyme.shallow(React.createElement(TypesOfFood)).type() === 'div');
```

يجب أن يُرجع مكوِّن `TypesOfFood` مكوِّن `Fruits`.

```js
assert(
  Enzyme.shallow(React.createElement(TypesOfFood)).props().children[1].type
    .name === 'Fruits'
);
```

يجب أن يُرجع مكوِّن `Fruits` مكوِّن `TypesOfFruit`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('h2').html() ===
    '<h2>Fruits:</h2>'
);
```

يجب أن يُرجع مكوِّن `TypesOfFruit` عنصري `h2` و `ul`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('ul').text() ===
    'ApplesBlueberriesStrawberriesBananas'
);
```

# --seed--

## --seed-contents--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }

      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }
        <TypesOfFruit />
      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }
        <Fruits />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
