---
id: 5a24c314108439a4d403617c
title: استخدام طريقة دورة الحياة componentWillMount
challengeType: 6
forumTopicId: 301423
dashedName: use-the-lifecycle-method-componentwillmount
---

# --description--

تحتوي مكوِّنات React على عدة طرق خاصة تتيح لك تنفيذ إجراءات في نقاط محددة من دورة حياة المكوِّن. تُسمى هذه الطرق طرق دورة الحياة أو خطافات دورة الحياة، وتسمح لك بالتقاط المكوِّنات في أوقات معينة. قد يكون ذلك قبل عرضها، قبل تحديثها، قبل استلام الخصائص، قبل إزالتها، وهكذا. إليك قائمة ببعض الطرق الرئيسية لدورة الحياة: `componentWillMount()` `componentDidMount()` `shouldComponentUpdate()` `componentDidUpdate()` `componentWillUnmount()` ستغطي الدروس التالية بعض الاستخدامات الأساسية لهذه الطرق.

**ملاحظة:** ستُلغى طريقة دورة الحياة `componentWillMount` في إصدار مستقبلي من 16.X وستُزال في الإصدار 17. تعرّف على المزيد في هذا <a href="https://www.freecodecamp.org/news/how-to-safely-use-reacts-life-cycles-with-fiber-s-async-rendering-fd4469ebbd8f/" target="_blank" rel="noopener noreferrer nofollow">المقال</a>

# --instructions--

تُستدعى طريقة `componentWillMount()` قبل طريقة `render()` عند تركيب المكوِّن في DOM. سجّل شيئًا في وحدة التحكم داخل `componentWillMount()` - قد ترغب في فتح وحدة تحكم المتصفح لرؤية النتيجة.

# --hints--

يجب أن يعرض `MyComponent` عنصر `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').length === 1;
  })()
);
```

يجب استدعاء `console.log` داخل `componentWillMount`.

```js
assert(
  (function () {
    const lifecycle = React.createElement(MyComponent)
      .type.prototype.componentWillMount.toString()
      .replace(/ /g, '');
    return lifecycle.includes('console.log(');
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line
    console.log('Component is mounting...');
    // Change code above this line
  }
  render() {
    return <div />
  }
};
```
