---
id: 5a24c314108439a4d4036161
title: تعرّف على وسوم JSX المغلقة ذاتيًا
challengeType: 6
forumTopicId: 301396
dashedName: learn-about-self-closing-jsx-tags
---

# --description--

حتى الآن، رأيت كيف يختلف JSX عن HTML بطريقة رئيسية باستخدام `className` بدلاً من `class` لتعريف أصناف HTML.

طريقة مهمة أخرى يختلف فيها JSX عن HTML هي فكرة الوسم المغلق ذاتيًا.

في HTML، تحتوي معظم الوسوم على وسم افتتاح ووسم إغلاق: `<div></div>`؛ وسم الإغلاق دائمًا يحتوي على شرطة مائلة قبل اسم الوسم الذي تغلقه. مع ذلك، هناك حالات خاصة في HTML تُسمى <dfn>العناصر الفارغة</dfn>، وهي عناصر لا تتطلب وسم افتتاح ووسم إغلاق قبل أن يبدأ عنصر آخر.

على سبيل المثال، يمكن كتابة وسم كسر السطر كـ `<br>` أو `<br />`، لكن لا يجب أبدًا كتابته كـ `<br></br>` لأنه لا يحتوي على أي محتوى.

في JSX، القواعد مختلفة قليلاً. يمكن كتابة أي عنصر JSX بوسم مغلق ذاتيًا، ويجب إغلاق كل عنصر. على سبيل المثال، يجب دائمًا كتابة وسم كسر السطر كـ `<br />` ليكون JSX صالحًا وقابلًا للتحويل. أما وسم `<div>`، فيمكن كتابته كـ `<div />` أو `<div></div>`. الفرق هو أنه في الصيغة الأولى لا يمكن تضمين أي شيء داخل `<div />`. سترى في التحديات التالية أن هذه الصيغة مفيدة عند عرض مكوِّنات React.

# --instructions--

صحح الأخطاء في محرر الشفرة ليصبح JSX صالحًا وقابلًا للتحويل بنجاح. تأكد من عدم تغيير أي محتوى - فقط أغلق الوسوم حيثما يلزم.

# --hints--

يجب أن تُرجع الثابتة `JSX` عنصر `div`.

```js
assert.strictEqual(JSX.type, 'div');
```

يجب أن يحتوي `div` على وسم `br`.

```js
assert(Enzyme.shallow(JSX).find('br').length === 1);
```

يجب أن يحتوي `div` على وسم `hr`.

```js
assert(Enzyme.shallow(JSX).find('hr').length === 1);
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h2>Welcome to React!</h2> <br >
    <p>Be sure to close all tags!</p>
    <hr >
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div>
  <h2>Welcome to React!</h2> <br />
  <p>Be sure to close all tags!</p>
  <hr />
</div>
);
```
