---
id: 5a24c314108439a4d403616e
title: الوصول إلى الخصائص باستخدام this.props
challengeType: 6
forumTopicId: 301375
dashedName: access-props-using-this-props
---

# --description--

غطت التحديات الأخيرة الطرق الأساسية لتمرير الخصائص إلى المكوِّنات الفرعية. لكن ماذا لو كان المكوِّن الفرعي الذي تمرر له خاصية هو مكوِّن صنف ES6، وليس مكوِّن دالة عديم الحالة؟ يستخدم مكوِّن صنف ES6 اصطلاحًا مختلفًا قليلًا للوصول إلى الخصائص.

في أي وقت تشير فيه إلى مكوِّن صنف داخل نفسه، تستخدم الكلمة المفتاحية `this`. للوصول إلى الخصائص داخل مكوِّن صنف، تسبق الشفرة التي تستخدمها للوصول إليها بـ `this`. على سبيل المثال، إذا كان لمكوِّن صنف ES6 خاصية تسمى `data`، تكتب `{this.props.data}` في JSX.

# --instructions--

اعرض نسخة من المكوِّن `Welcome` داخل المكوِّن الأب `App`. هنا، أعطِ `Welcome` خاصية `name` وعيّن لها قيمة سلسلة نصية. داخل المكوِّن الفرعي `Welcome`، وصول إلى الخاصية `name` داخل وسم `strong`.

# --hints--

يجب أن يُرجع المكوِّن `App` عنصر `div` واحدًا فقط.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.children().type() === 'div';
  })()
);
```

يجب أن يكون المكوِّن الفرعي لـ `App` هو المكوِّن `Welcome`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.children().childAt(0).name() === 'Welcome'
    );
  })()
);
```

يجب أن يحتوي المكوِّن `Welcome` على خاصية تسمى `name`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.find('Welcome').props().name;
  })()
);
```

يجب أن يعرض المكوِّن `Welcome` السلسلة النصية التي تمررها كخاصية `name` داخل وسم `strong`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.find('strong').text() ===
      mockedComponent.find('Welcome').props().name
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome />
            { /* Change code above this line */ }
        </div>
    );
  }
};

class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong></strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};
```

# --solutions--

```jsx
class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong>{this.props.name}</strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};

class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome name="Quincy"/>
            { /* Change code above this line */ }
        </div>
    );
  }
};
```
