---
id: 5a24c314108439a4d4036185
title: استخدم && لشرط أكثر إيجازًا
challengeType: 6
forumTopicId: 301413
dashedName: use--for-a-more-concise-conditional
---

# --description--

نجحت عبارات `if/else` في التحدٍّ السابق، لكن هناك طريقة أكثر إيجازًا لتحقيق نفس النتيجة. تخيل أنك تتابع عدة شروط في مكوِّن وتريد عرض عناصر مختلفة اعتمادًا على كل من هذه الشروط. إذا كتبت الكثير من عبارات `else if` لإرجاع واجهات مستخدم مختلفة قليلاً، قد تكرر شفرة مما يترك مجالًا للخطأ. بدلًا من ذلك، يمكنك استخدام العامل المنطقي `&&` لأداء المنطق الشرطي بطريقة أكثر إيجازًا. هذا ممكن لأنك تريد التحقق مما إذا كان الشرط `true`، وإذا كان كذلك، ترجع بعض الترميز. إليك مثالًا:

```jsx
{condition && <p>markup</p>}
```

إذا كان `condition` يساوي `true`، فسيتم إرجاع الترميز. إذا كان الشرط `false`، ستُرجع العملية فورًا `false` بعد تقييم `condition` ولن تُرجع شيئًا. يمكنك تضمين هذه العبارات مباشرة في JSX وربط عدة شروط معًا بكتابة `&&` بعد كل منها. هذا يسمح لك بالتعامل مع منطق شرطي أكثر تعقيدًا في طريقة `render()` دون تكرار الكثير من الشفرة.

# --instructions--

حل المثال السابق مرة أخرى، بحيث يتم عرض `h1` فقط إذا كان `display` يساوي `true`، لكن استخدم العامل المنطقي `&&` بدلًا من عبارة `if/else`.

# --hints--

يجب أن يوجد `MyComponent` ويعرض.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length;
  })()
);
```

عندما يكون `display` مضبوطًا على `true`، يجب عرض `div` و`button` و`h1`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 2 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 1
  );
```

عندما يكون `display` مضبوطًا على `false`، يجب عرض `div` و`button` فقط.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 1 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 0
  );
```

يجب أن تستخدم طريقة العرض العامل المنطقي `&&` للتحقق من شرط `this.state.display`.

```js
assert(code.includes('&&'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         {this.state.display && <h1>Displayed!</h1>}
       </div>
    );
  }
};
```
