---
id: 5a24c314108439a4d4036181
title: تقديم الأنماط المضمَّنة
challengeType: 6
forumTopicId: 301395
dashedName: introducing-inline-styles
---

# --description--

هناك مفاهيم معقدة أخرى تضيف قدرات قوية إلى شفرة React الخاصة بك. لكن قد تتساءل عن المشكلة الأبسط وهي كيفية تنسيق عناصر JSX التي تنشئها في React. من المحتمل أنك تعرف أنها لن تكون نفسها تمامًا كما في العمل مع HTML بسبب <a href="/learn/front-end-development-libraries/react/define-an-html-class-in-jsx" target="_blank" rel="noopener noreferrer nofollow">طريقة تطبيق الأصناف على عناصر JSX</a>.

إذا استوردت الأنماط من ملف أنماط، فلن يختلف الأمر كثيرًا. تطبق صنفًا على عنصر JSX باستخدام الخاصية `className`، وتطبق الأنماط على الصنف في ملف الأنماط. خيار آخر هو تطبيق الأنماط المضمَّنة، وهي شائعة جدًا في تطوير ReactJS.

تطبق الأنماط المضمَّنة على عناصر JSX بطريقة مشابهة لكيفية تطبيقها في HTML، لكن مع بعض الاختلافات في JSX. إليك مثالًا على نمط مضمَّن في HTML:

```jsx
<div style="color: yellow; font-size: 16px">Mellow Yellow</div>
```

تستخدم عناصر JSX الخاصية `style`، لكن بسبب طريقة تحويل JSX، لا يمكنك تعيين القيمة إلى `string`. بدلًا من ذلك، تعينها إلى كائن JavaScript `object`. إليك مثالًا:

```jsx
<div style={{color: "yellow", fontSize: 16}}>Mellow Yellow</div>
```

هل لاحظت كيف نستخدم تنسيق سنام الجمل في خاصية `fontSize`؟ ذلك لأن React لا يقبل المفاتيح بصيغة kebab-case في كائن الأنماط. React سيُطبق اسم الخاصية الصحيح في HTML نيابةً عنا.

# --instructions--

أضف خاصية `style` إلى عنصر `div` في محرر الشفرة لتعطي النص لونًا أحمر وحجم خط `72px`.

لاحظ أنه يمكنك اختيارًا تعيين حجم الخط كرقم بدون وحدة `px`، أو كتابته كـ `"72px"`.
# --hints--

يجب أن يعرض المكوِّن عنصر `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().type() === 'div';
  })()
);
```

يجب أن يكون لعنصر `div` لون `red`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().props().style.color === 'red';
  })()
);
```

يجب أن يكون لعنصر `div` حجم خط `72px`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return (
      mockedComponent.children().props().style.fontSize === 72 ||
      mockedComponent.children().props().style.fontSize === '72' ||
      mockedComponent.children().props().style.fontSize === '72px'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div>Big Red</div>
    );
  }
};
```

# --solutions--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div style={{color: "red", fontSize: 72}}>Big Red</div>
    );
  }
};
```
