---
id: 5a24c314108439a4d4036163
title: إنشاء مكوِّن React
challengeType: 6
forumTopicId: 301386
dashedName: create-a-react-component
---

# --description--

الطريقة الأخرى لتعريف مكوِّن React هي باستخدام صيغة `class` في ES6. في المثال التالي، يرث `Kitten` من `React.Component`:

```jsx
class Kitten extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <h1>Hi</h1>
    );
  }
}
```

هذا ينشئ صنف ES6 باسم `Kitten` يمدد صنف `React.Component`. الآن صنف `Kitten` لديه وصول إلى العديد من ميزات React المفيدة مثل الحالة المحلية وhooks دورة الحياة. لا تقلق إذا لم تكن معتادًا على هذه المصطلحات بعد، سيتم شرحها بالتفصيل في تحديات لاحقة. لاحظ أيضًا أن صنف `Kitten` يحتوي على دالة البناء `constructor` مُعرَّفة بداخله وتستدعي `super()`. تُستخدم `super()` لاستدعاء دالة البناء للصنف الأب، وفي هذه الحالة هو `React.Component`. دالة البناء هي طريقة خاصة تُستخدم أثناء تهيئة الكائنات التي تُنشأ بكلمة `class`. من أفضل الممارسات استدعاء دالة البناء للمكوِّن (`constructor`) مع `super`، وتمرير `props` لكل منهما. هذا يضمن تهيئة المكوِّن بشكل صحيح. في الوقت الحالي، اعلم أن تضمين هذا الكود هو أمر معتاد. قريبًا سترى استخدامات أخرى لدالة البناء وكذلك للـ `props`.

# --instructions--

تم تعريف `MyComponent` في محرر الكود باستخدام صيغة الصنف. أكمل كتابة طريقة `render` بحيث تُرجع عنصر `div` يحتوي على عنصر `h1` بنص `Hello React!`.

# --hints--

يجب أن يُعيد مكوِّن React عنصر `div`.

```js
assert(Enzyme.shallow(React.createElement(MyComponent)).type() === 'div');
```

يجب أن يعرض عنصر `div` المُعاد عنصر عنوان `h1` بداخله.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.shallow(React.createElement(MyComponent)).html()
  )
);
```

يجب أن يحتوي عنصر العنوان `h1` على السلسلة النصية `Hello React!`.

```js
assert(
  Enzyme.shallow(React.createElement(MyComponent)).html() ===
    '<div><h1>Hello React!</h1></div>'
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line



    // Change code above this line
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line
    return (
      <div>
        <h1>Hello React!</h1>
      </div>
    );
    // Change code above this line
  }
};
```
