---
id: 5a24c314108439a4d4036173
title: تعيين الحالة باستخدام this.setState
challengeType: 6
forumTopicId: 301412
dashedName: set-state-with-this-setstate
---

# --description--

التحديات السابقة تناولت `state` الخاص بالمكوِّن وكيفية تهيئة الحالة في `constructor`. هناك أيضًا طريقة لتغيير `state` الخاص بالمكوِّن. يوفر React طريقة لتحديث `state` الخاص بالمكوِّن تُسمى `setState`. تستدعي طريقة `setState` داخل صنف المكوِّن هكذا: `this.setState()`، مع تمرير كائن يحتوي على أزواج المفتاح-القيمة. المفاتيح هي خصائص الحالة والقيم هي بيانات الحالة المحدثة. على سبيل المثال، إذا كنا نخزن `username` في الحالة وأردنا تحديثه، فسيبدو الأمر هكذا:

```jsx
this.setState({
  username: 'Lewis'
});
```

يتوقع React ألا تعدل `state` مباشرة، بل استخدم دائمًا `this.setState()` عند حدوث تغييرات في الحالة. يجب أن تعلم أيضًا أن React قد يجمع عدة تحديثات للحالة لتحسين الأداء. هذا يعني أن تحديثات الحالة عبر طريقة `setState` قد تكون غير متزامنة. هناك صيغة بديلة لطريقة `setState` توفر حلًا لهذه المشكلة. نادرًا ما تحتاجها، لكن من الجيد أن تضعها في اعتبارك! يرجى مراجعة <a href="https://www.freecodecamp.org/news/what-is-state-in-react-explained-with-examples/" target="_blank" rel="noopener noreferrer nofollow">مقال React</a> لمزيد من التفاصيل.

# --instructions--

يوجد عنصر `button` في محرر الشفرة يحتوي على معالِج `onClick()`. يُفعّل هذا المعالِج عند نقر `button` في المتصفح، ويُشغّل الطريقة `handleClick` المعرفة في `MyComponent`. داخل طريقة `handleClick`، حدّث `state` الخاص بالمكوِّن باستخدام `this.setState()`. عيّن الخاصية `name` في `state` لتساوي السلسلة النصية `React Rocks!`.

انقر الزر وراقب تحديث العرض للحالة. لا تقلق إذا لم تفهم تمامًا كيف يعمل كود معالِج النقر في هذه المرحلة. سيتم تغطيته في التحديات القادمة.

# --hints--

يجب أن تبدأ حالة `MyComponent` بزوج المفتاح والقيمة `{ name: Initial State }`.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'Initial State'
);
```

يجب أن يعرض `MyComponent` عنصر عنوان `h1`.

```js
assert(Enzyme.mount(React.createElement(MyComponent)).find('h1').length === 1);
```

يجب أن يحتوي عنصر العنوان `h1` المعروض على نص مأخوذ من حالة المكوِّن.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(/<h1>TestName<\/h1>/.test(firstValue));
```

استدعاء طريقة `handleClick` على `MyComponent` يجب أن يعيّن خاصية name في الحالة لتساوي `React Rocks!`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'Before' });
    return waitForIt(() => mockedComponent.state('name'));
  };
  const second = () => {
    mockedComponent.instance().handleClick();
    return waitForIt(() => mockedComponent.state('name'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Before' && secondValue === 'React Rocks!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
     // Change code below this line
    this.setState({
      name: 'React Rocks!'
    });
    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```
