---
id: 5a24c314108439a4d4036189
title: تغيير CSS المضمَّن شرطياً بناءً على حالة المكوِّن
challengeType: 6
forumTopicId: 301380
dashedName: change-inline-css-conditionally-based-on-component-state
---

# --description--

حتى الآن، شاهدت عدة تطبيقات للعرض الشرطي واستخدام الأنماط المضمَّنة. إليك مثالًا آخر يجمع بين هذين الموضوعين. يمكنك أيضًا عرض CSS شرطياً بناءً على حالة مكوِّن React. للقيام بذلك، تتحقق من شرط، وإذا تحقق هذا الشرط، تعدل كائن الأنماط المُعيَّن لعناصر JSX في طريقة العرض.

هذا النموذج مهم لفهمه لأنه يمثل تحولًا كبيرًا عن النهج التقليدي لتطبيق الأنماط عبر تعديل عناصر DOM مباشرةً (وهو شائع جدًا مع jQuery، على سبيل المثال). في ذلك النهج، يجب متابعة متى تتغير العناصر والتعامل مع التعديل الفعلي مباشرةً. قد يصبح من الصعب متابعة التغييرات، مما قد يجعل واجهة المستخدم غير متوقعة. عند تعيين كائن نمط بناءً على شرط، تصف كيف يجب أن تبدو واجهة المستخدم كدالة لحالة التطبيق. هناك تدفق واضح للمعلومات يتحرك في اتجاه واحد فقط. هذه هي الطريقة المفضلة عند كتابة التطبيقات باستخدام React.

# --instructions--

يحتوي محرر الشفرة على مكوِّن مُدخَل بسيط مُتحكم به مع حدود منسَّقة. تريد تنسيق هذه الحدود باللون الأحمر إذا كتب المستخدم أكثر من 15 حرفًا في مربع المُدخَل. أضف شرطًا للتحقق من ذلك، وإذا تحقق الشرط، عيّن نمط حدود المُدخَل إلى `3px solid red`. يمكنك تجربته بإدخال نص في المُدخَل.

# --hints--

يجب أن يعرض مكوِّن `GateKeeper` عنصر `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.find('div').length === 1;
  })()
);
```

يجب تهيئة مكوِّن `GateKeeper` بمفتاح حالة `input` مُعيَّن إلى سلسلة فارغة.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.state().input === '';
  })()
);
```

يجب أن يعرض مكوِّن `GateKeeper` وسم `h3` ووسم `input`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('h3').length === 1 &&
      mockedComponent.find('input').length === 1
    );
  })()
);
```

يجب أن يكون لوسم `input` في البداية نمط حدود `1px solid black` لخاصية `border`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('input').props().style.border === '1px solid black'
    );
  })()
);
```

يجب تنسيق وسم `input` بحدود `3px solid red` إذا كانت قيمة المُدخَل في الحالة أطول من 15 حرفًا.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  let initialStyle = mockedComponent.find('input').props().style.border;
  const state_1 = () => {
    mockedComponent.setState({ input: 'this is 15 char' });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const state_2 = () => {
    mockedComponent.setState({
      input: 'A very long string longer than 15 characters.'
    });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const style_1 = await state_1();
  const style_2 = await state_2();
  assert(
    initialStyle === '1px solid black' &&
      style_1 === '1px solid black' &&
      style_2 === '3px solid red'
  );
```

# --seed--

## --seed-contents--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    // Change code below this line

    // Change code above this line
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    if (this.state.input.length > 15) {
      inputStyle.border = '3px solid red';
    };
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```
