---
id: 5a24c314108439a4d403618d
title: عرض React على الخادم باستخدام renderToString
challengeType: 6
forumTopicId: 301407
dashedName: render-react-on-the-server-with-rendertostring
---

# --description--

حتى الآن، كنت تعرض مكوّنات React على العميل. عادةً، هذا هو ما ستفعله دائمًا. مع ذلك، هناك بعض الحالات التي يكون من المنطقي فيها عرض مكوّن React على الخادم. بما أن React مكتبة عرض بلغة JavaScript ويمكنك تشغيل JavaScript على الخادم باستخدام Node، فهذا ممكن. في الواقع، توفر React طريقة `renderToString()` يمكنك استخدامها لهذا الغرض.

هناك سببان رئيسيان لاستخدام العرض على الخادم في تطبيق حقيقي. أولًا، بدون هذا، ستتكون تطبيقات React الخاصة بك من ملف HTML فارغ نسبيًا وحزمة كبيرة من JavaScript عند تحميلها أول مرة في المتصفح. هذا قد لا يكون مثاليًا لمحركات البحث التي تحاول فهرسة محتوى صفحاتك ليتمكن الناس من العثور عليك. إذا عرضت الترميز الأولي لـ HTML على الخادم وأرسلته إلى العميل، فإن تحميل الصفحة الأولي يحتوي على كل ترميز الصفحة الذي يمكن لمحركات البحث زحفه. ثانيًا، هذا يخلق تجربة تحميل صفحة أولية أسرع لأن HTML المعروض أصغر من شفرة JavaScript الخاصة بالتطبيق بأكمله. ستظل React قادرة على التعرف على تطبيقك وإدارته بعد التحميل الأولي.

# --instructions--

طريقة `renderToString()` متوفرة على `ReactDOMServer`، وهو كائن عام متاح هنا. تأخذ الطريقة معلمة واحدة وهي عنصر React. استخدم هذا لعرض `App` كسلسلة نصية.

# --hints--

يجب أن يعرض مكوّن `App` كسلسلة نصية باستخدام `ReactDOMServer.renderToString`.

```js
  assert(
    code.replace(/ /g, '')
      .includes('ReactDOMServer.renderToString(<App/>)') &&
      Enzyme.mount(React.createElement(App)).children().name() === 'div'
  );
```

# --before-all--

```js
var ReactDOMServer = { renderToString(x) { return null; } };
```

# --seed--

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <div/>
  }
};

// Change code below this line
```

# --solutions--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <div/>
  }
};

// Change code below this line
ReactDOMServer.renderToString(<App/>);
```
