---
id: 5a24c314108439a4d403616d
title: استخدام PropTypes لتعريف الخصائص التي تتوقعها
challengeType: 6
forumTopicId: 301419
dashedName: use-proptypes-to-define-the-props-you-expect
---

# --description--

يوفر React ميزات مفيدة لفحص النوع للتحقق من أن المكونات تستقبل الخصائص من النوع الصحيح. على سبيل المثال، يقوم تطبيقك باستدعاء API لاسترجاع بيانات تتوقع أن تكون في مصفوفة، ثم تُمرر إلى مكوِّن كخاصية. يمكنك تعيين `propTypes` على مكوّنك لتطلب أن تكون البيانات من نوع `array`. هذا سيُصدر تحذيرًا مفيدًا عندما تكون البيانات من أي نوع آخر.

يُعتبر من أفضل الممارسات تعيين `propTypes` عندما تعرف نوع الخاصية مسبقًا. يمكنك تعريف خاصية `propTypes` لمكوِّن بنفس الطريقة التي عرّفت بها `defaultProps`. هذا سيتحقق من وجود الخصائص ذات المفتاح المعين وبنوع معين. إليك مثالًا لطلب نوع `function` لخاصية تسمى `handleClick`:

```js
MyComponent.propTypes = { handleClick: PropTypes.func.isRequired }
```

في المثال أعلاه، الجزء `PropTypes.func` يتحقق من أن `handleClick` هي دالة. إضافة `isRequired` تخبر React أن `handleClick` خاصية مطلوبة لذلك المكوِّن. سترى تحذيرًا إذا لم تُزوّد هذه الخاصية. لاحظ أيضًا أن `func` تمثل `function`. من بين الأنواع الأولية السبعة في JavaScript، `function` و`boolean` (المكتوبة كـ `bool`) هما الوحيدان اللذان يستخدمان تهجئة غير معتادة. بالإضافة إلى الأنواع الأولية، هناك أنواع أخرى متاحة. على سبيل المثال، يمكنك التحقق من أن الخاصية هي عنصر React. يرجى الرجوع إلى <a href="https://reactjs.org/docs/typechecking-with-proptypes.html#proptypes" target="_blank" rel="noopener noreferrer nofollow">التوثيق</a> لجميع الخيارات.

**ملاحظة:** اعتبارًا من React v15.5.0، يتم استيراد `PropTypes` بشكل مستقل عن React، هكذا: `import PropTypes from 'prop-types';`

# --instructions--

عرّف `propTypes` لمكوِّن `Items` لتطلب خاصية `quantity` وتتحقق من أن نوعها هو `number`.

# --hints--

يجب أن يعرض مكوِّن `ShoppingCart`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

يجب أن يعرض مكوِّن `Items`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

يجب أن يتضمن مكوِّن `Items` فحص `propTypes` لطلب قيمة لـ `quantity` والتأكد من أن قيمتها عدد.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(noWhiteSpace.includes('quantity:PropTypes.number.isRequired') && noWhiteSpace.includes('Items.propTypes='));
```

# --before-all--

```js
var PropTypes = {
  number: { isRequired: true }
};
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line

// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line
Items.propTypes = {
  quantity: PropTypes.number.isRequired
};
// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```
