---
id: 5f3326b143638ee1a09ff1e3
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Para preparar seu código para criar conteúdo, adicione um elemento `body` abaixo do elemento `head`.

# --hints--

Você deve ter uma tag de abertura `<body>`.

```js
assert(code.match(/<body>/i));
```

Você deve ter uma tag de fechamento `</body>`.

```js
assert(code.match(/<\/body>/i));
```

Você não deve alterar seu elemento `head`. Certifique-se de que não deletou sua tag de fechamento.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

O elemento `body` deve estar depois do elemento `head`.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

