---
id: 5f344fad8bf01691e71a30eb
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Até agora, você tinha controle limitado sobre a apresentação e aparência do seu conteúdo.Para mudar isso, adicione um elemento `style` com o elemento `head`.

# --hints--

Seu código deve ter uma tag de abertura `<style>`.

```js
assert(code.match(/<style\s*>/i));
```

Seu código deve ter uma tag de fechamento `</style>`.

```js
assert(code.match(/<\/style\s*>/));
```

O elemento `style` deve estar dentro do elemento `head`.

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
