---
id: 66db529d37ad966480ebb633
title: Tengeneza programu ya kuchagua chakula cha mchana
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

Katika maabara hii, utatengeneza programu inayosaidia kusimamia chaguzi za chakula cha mchana. Utafanya kazi na safu ya data ya chakula cha mchana, kuongeza na kuondoa vitu kutoka kwenye safu ya data, na kuchagua kwa bahati moja ya chaguo la chakula cha mchana.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kutangaza tofauti `lunches` na kuwekea thamani safu ya data tupu itakayotumika kuhifadhi vitu vya chakula cha mchana.

2. Unapaswa kutengeneza kitendakazi `addLunchToEnd` kinachopokea safu ya data kama hoja ya kwanza na mfuatano wa herufi kama hoja ya pili. Kitendakazi kinapaswa:
    - Ongeza mfuatano wa herufi mwishoni mwa safu ya data.
    - Andika mfuatano wa herufi `[Lunch Item] added to the end of the lunch menu.` kwenye konsoli, ambapo `[Lunch Item]` ni mfuatano wa herufi uliotumwa kwa kitendakazi.
    - Rudisha safu ya data iliyosasishwa.

3. Unapaswa kutengeneza kitendakazi `addLunchToStart` kinachopokea safu ya data kama hoja ya kwanza na mfuatano wa herufi kama hoja ya pili. Kitendakazi kinapaswa:
    - Ongeza mfuatano wa herufi mwanzoni mwa safu ya data.
    - Andika mfuatano wa herufi `[Lunch Item] added to the start of the lunch menu.` kwenye konsoli, ambapo `[Lunch Item]` ni mfuatano wa herufi uliotumwa kwa kitendakazi.
    - Rudisha safu ya data iliyosasishwa.

4. Unapaswa kutengeneza kitendakazi `removeLastLunch` kinachopokea safu ya data kama hoja yake. Kitendakazi kinapaswa:
   - Ondoa kipengele cha mwisho kutoka kwenye safu ya data.
   - Ikiwa kuondoa kumefanikiwa, andika mfuatano wa herufi `[Lunch Item] removed from the end of the lunch menu.` kwenye konsoli, ambapo `[Lunch Item]` ni kipengele kilichoondolewa kutoka kwenye safu ya data.
   - Ikiwa safu ya data ni tupu, andika mfuatano wa herufi `"No lunches to remove."` kwenye konsoli.
   - Rudisha safu ya data iliyosasishwa.

5. Unapaswa kutengeneza kitendakazi `removeFirstLunch` kinachopokea safu ya data kama hoja yake. Kitendakazi kinapaswa:
   - Ondoa kipengele cha kwanza kutoka kwenye safu ya data.
   - Ikiwa kuondoa kumefanikiwa, andika mfuatano wa herufi `[Lunch Item] removed from the start of the lunch menu.` kwenye konsoli, ambapo `[Lunch Item]` ni kipengele kilichoondolewa kutoka kwenye safu ya data.
   - Ikiwa safu ya data ni tupu, andika mfuatano wa herufi `"No lunches to remove."` kwenye konsoli.
   - Rudisha safu ya data iliyosasishwa.

6. Unapaswa kutengeneza kitendakazi `getRandomLunch` kinachopokea safu ya data kama hoja yake. Kitendakazi kinapaswa:
   - Chagua kipengele cha bahati kutoka kwenye safu ya data.
   - Ikiwa imetokea mafanikio, andika mfuatano wa herufi `Randomly selected lunch: [Lunch Item]` kwenye konsoli, ambapo `[Lunch Item]` ni kipengele cha bahati kutoka kwenye safu ya data.
   - Ikiwa safu ya data ni tupu, andika mfuatano wa herufi `"No lunches available."` kwenye konsoli.

7. Unapaswa kutengeneza kitendakazi `showLunchMenu` kinachopokea safu ya data kama hoja yake na:
   - Ikiwa kuna vipengele kwenye safu ya data, andika mfuatano wa herufi `Menu items: [Lunch Item], [Lunch Item]...` kwenye konsoli, ambapo kila `[Lunch item]` ni mojawapo ya vipengele kwenye safu ya data, kwa mpangilio.
   - Ikiwa safu ya data ni tupu, andika mfuatano wa herufi `"The menu is empty."` kwenye konsoli.

# --hints--

Unapaswa kutangaza tofauti `lunches` na kuwekea thamani safu ya data tupu kuhifadhi vitu vya chakula cha mchana.

```js
assert.isArray(lunches);
```

Unapaswa kufafanua kitendakazi `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

Kitendakazi `addLunchToEnd` kinapaswa kuwa na vigezo viwili.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` inapaswa kuandika mfuatano wa herufi `"Tacos added to the end of the lunch menu."` kwenye konsoli.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` inapaswa kurudisha `["Pizza", "Tacos", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

Unapaswa kufafanua kitendakazi `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

Kitendakazi `addLunchToStart` kinapaswa kuwa na vigezo viwili.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` inapaswa kuandika mfuatano wa herufi `"Sushi added to the start of the lunch menu."` kwenye konsoli.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` inapaswa kurudisha `["Pizza", "Burger", "Sushi"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

Unapaswa kufafanua kitendakazi `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

Kitendakazi `removeLastLunch` kinapaswa kuwa na kigezo kimoja.

```js
assert.lengthOf(removeLastLunch, 1);
```

Wakati safu ya data ya ingizo ni tupu, kitendakazi `removeLastLunch` kinapaswa kuandika mfuatano wa herufi `"No lunches to remove."` kwenye konsoli.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` inapaswa kuandika mfuatano wa herufi `"Toast removed from the end of the lunch menu."` kwenye konsoli.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` inapaswa kurudisha `["Sushi", "Pizza"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

Unapaswa kufafanua kitendakazi `removeFirstLunch`.

```js
assert.isFunction(removeFirstLunch);
```

Kitendakazi `removeFirstLunch` kinapaswa kuwa na kigezo kimoja.

```js
assert.lengthOf(removeFirstLunch, 1);
```

Wakati safu ya data ya ingizo ni tupu, kitendakazi `removeFirstLunch` kinapaswa kuandika mfuatano wa herufi `"No lunches to remove."` kwenye konsoli.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` inapaswa kuandika mfuatano wa herufi `"Salad removed from the start of the lunch menu."` kwenye konsoli.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` inapaswa kurudisha `["Pizza", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`, `addLunchToStart`, `removeLastLunch`, na `removeFirstLunch` zinapaswa kurudisha safu ile ile iliyotumwa kama hoja baada ya kuiboresha.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

Unapaswa kufafanua kitendakazi `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

Kitendakazi `getRandomLunch` kinapaswa kuwa na kigezo kimoja.

```js
assert.lengthOf(getRandomLunch, 1);
```

Wakati safu ya data ya ingizo ni tupu, kitendakazi `getRandomLunch` kinapaswa kuandika mfuatano wa herufi `"No lunches available."` kwenye konsoli.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

Wakati safu ya data si tupu, kitendakazi `getRandomLunch` kinapaswa kuandika mfuatano wa herufi katika muundo wa `Randomly selected lunch: [Lunch Item]` kwenye konsoli.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

Kitendakazi `getRandomLunch` hakipaswi kubadilisha safu ya data iliyotumwa kama hoja.

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

Unapaswa kufafanua kitendakazi `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

Kitendakazi `showLunchMenu` kinapaswa kuwa na kigezo kimoja.

```js
assert.lengthOf(showLunchMenu, 1);
```

Wakati safu ya data ya ingizo ni tupu, kitendakazi `showLunchMenu` kinapaswa kuandika mfuatano wa herufi `"The menu is empty."` kwenye konsoli.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` inapaswa kuandika `"Menu items: Greens, Corns, Beans"` kwenye konsoli.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` inapaswa kuandika `"Menu items: Pizza, Burger, Fries, Salad"` kwenye konsoli.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
