---
id: 6823e637c1c0ed56f781b4fc
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Infine, JavaScript rende la tua pagina web interattiva — ti permette di dire alla pagina cosa fare quando qualcuno clicca un pulsante, invia un modulo o fa molte altre cose.

Per l’ultimo passaggio del workshop, aggiungi un altro elemento `h3` e `p` alla pagina che descrive JavaScript.

Per prima cosa, aggiungi un elemento `h3` con il testo:

```md
Introduction to JavaScript
```

Poi, sotto quell’elemento `h3`, aggiungi un elemento `p` con il seguente testo:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

Dovresti aggiungere un terzo elemento `h3` alla pagina.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Il testo del tuo elemento `h3` dovrebbe essere `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

Dovresti aggiungere un quarto elemento `p` alla pagina.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Il tuo elemento `p` dovrebbe avere il testo `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
