---
id: 6823d6244511f252c8300eed
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Quando devi aggiungere un paragrafo a una pagina web, puoi usare l'elemento `p` in questo modo:

```html
<p>This is a paragraph element.</p>
```

Trasforma `Learn the skills to become a full-stack developer` in un elemento paragrafo.

# --hints--

Il tuo elemento `p` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

Il tuo elemento `p` dovrebbe avere un tag di chiusura. I tag di chiusura hanno questa sintassi: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

Il testo del tuo elemento `p` dovrebbe essere `Learn the skills to become a full-stack developer`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

Il tuo elemento `p` dovrebbe essere sotto il tuo elemento `h2`.

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
