---
id: 6823d9ac8bdc3853df65a1ff
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Ci sono sei elementi di intestazione in HTML: `h1` fino a `h6`. Vengono usati per mostrare l'importanza delle sezioni nella tua pagina web, con `h1` che è il più importante e `h6` il meno importante.

Sotto il tuo elemento `p`, aggiungi un'intestazione `h3` con il testo:

```md
Introduction to HTML
```

# --hints--

Il tuo elemento `h3` dovrebbe avere un tag di apertura `<h3>`.

```js
assert.exists(document.querySelector("h3"));
```

Il tuo elemento `h3` dovrebbe avere un tag di chiusura `</h3>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

Il testo del tuo elemento `h3` dovrebbe essere `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
