---
id: 6823e036ea4b71553558c01b
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Nota che ogni intestazione appare un po' diversa nell'anteprima. Lo stesso vale per il paragrafo.

Poi, aggiungi un altro elemento `p` che mostri il seguente testo:

```md
HTML represents the content and structure of a webpage
```

# --hints--

Dovresti avere un secondo elemento `p`.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Il tuo elemento `p` dovrebbe contenere il testo `HTML represents the content and structure of a webpage`. Puoi copiare il testo dalle istruzioni.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
