---
id: 6823ac607bfdbc46331b2559
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML sta per HyperText Markup Language. È il codice che definisce la struttura e il contenuto di una pagina web. Questo è il tuo editor di codice, dove scriverai HTML.

Trova la riga 1 nell'editor e digita questo testo:

```md
Welcome to freeCodeCamp
```

Quando hai finito, clicca sul pulsante "controlla il tuo codice" per vedere se è corretto.

# --hints--

Dovresti avere il testo `Welcome to freeCodeCamp` nel tuo editor. Controlla bene l'ortografia.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
