---
id: 682ba2318000b62f179bdf04
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

L'HTML è composto da elementi. Il primo che userai è l'elemento `h1`:

```html
<h1>Welcome to freeCodeCamp</h1>
```

Inizia con un tag di apertura (`<h1>`), termina con un tag di chiusura (`</h1>`) e ha il testo che verrà mostrato tra i tag.

Trasforma il tuo testo `Welcome to freeCodeCamp` in un elemento `h1` aggiungendo un tag di apertura davanti e un tag di chiusura dopo.

# --hints--

Il tuo elemento `h1` dovrebbe avere un tag di apertura `<h1>`.

```js
assert.exists(document.querySelector("h1"));
```

Il tuo elemento `h1` dovrebbe avere un tag di chiusura `</h1>`.

```js
assert.match(code, /<\/h1\s*\>/);
```

Il tuo elemento `h1` dovrebbe essere così: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
