---
id: 682cd206883fc7b25eb539c5
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Nota che l'HTML che scrivi nell'editor appare nell'anteprima. In questo workshop, scriverai l'HTML per una pagina parziale di un curriculum.

Sotto il tuo elemento `h1`, digita quanto segue sulla riga vuota:

```md
Full-Stack Curriculum
```

# --hints--

Dovresti avere il testo `Full-Stack Curriculum` nel tuo editor. Controlla bene l'ortografia.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
