---
id: 683921c4769fd23dbadec2fe
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Mentre l'HTML definisce la struttura e il contenuto di una pagina web, il CSS serve per aggiungere lo stile — come colori, caratteri tipografici, spaziature e layout.

Sotto il `h3` che hai appena aggiunto, aggiungi un altro elemento paragrafo con il testo:

```md
CSS is used to style a webpage
```

# --hints--

Dovresti aggiungere un secondo elemento `h3` alla pagina.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Il testo del tuo elemento `h3` dovrebbe essere `Introduction to CSS`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

Dovresti aggiungere un terzo elemento `p` alla pagina.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Il tuo elemento `p` dovrebbe contenere il testo `CSS is used to style a webpage`. Puoi copiare il testo dalle istruzioni.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
