---
id: 6823d6244511f252c8300eed
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Якщо на сторінку потрібно додати абзац, можна скористатись елементом `p`:

```html
<p>This is a paragraph element.</p>
```

Turn `Learn the skills to become a full stack developer` into a paragraph element.

# --hints--

Елемент `p` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<назваЕлемента>`.

```js
assert.exists(document.querySelector("p"));
```

Елемент `p` повинен мати кінцевий тег. Кінцеві теги мають такий синтаксис: `</назваЕлемента>`.

```js
assert.match(code, /<\/p\s*\>/);
```

Текстом елемента `p` має бути `Learn the skills to become a full stack developer`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full\s*stack\s*developer\s*\<\/p\s*\>/i);
```

Під елементом `h2` має бути елемент `p`.

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full stack developer
--fcc-editable-region--
```
