---
id: 6823ac607bfdbc46331b2559
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML stands for HyperText Markup Language. It's the code that defines the structure and content of a webpage. This is your code editor, where you'll write HTML.

Знайдіть перший рядок в редакторі та напишіть цей текст:

```md
Welcome to freeCodeCamp
```

Коли закінчите, натисніть на кнопку «Перевірити код», щоб переконатись, чи все правильно.

# --hints--

В редакторі має бути текст `Welcome to freeCodeCamp`. Двічі перевірте написання.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
