---
id: 6823e637c1c0ed56f781b4fc
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Finally, JavaScript makes your webpage interactive — it lets you tell the page what to do when someone clicks a button, submits a form, or many other things.

For the last step of the workshop, add another `h3` and `p` element to the page describing JavaScript.

Спочатку додайте елемент `h3` з таким текстом:

```md
Introduction to JavaScript
```

Потім — під елементом `h3` — додайте елемент `p` з таким текстом:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

Додайте третій елемент `h3` до сторінки.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Текстом елемента `h3` має бути `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

Додайте четвертий елемент `p` до сторінки.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Елемент `p` повинен мати текст `JavaScript adds interactivity to a webpage`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn all of the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
