---
id: 6823d9ac8bdc3853df65a1ff
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

У HTML існує шість заголовків: від `h1` до `h6`. Вони представляють важливість розділів вебсторінки, де `h1` — найважливіший, а `h6` — найменш важливий.

Під елементом `p` додайте заголовок `h3` з таким текстом:

```md
Introduction to HTML
```

# --hints--

Елемент `h3` повинен мати початковий тег `<h3>`.

```js
assert.exists(document.querySelector("h3"));
```

Елемент `h3` повинен мати кінцевий тег `</h3>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

Текстом елемента `h3` має бути `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
