---
id: 682ba2318000b62f179bdf04
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML складається з елементів. Першим ви використаєте елемент `h1`:

```html
<h1>Welcome to freeCodeCamp</h1>
```

Він починається з початкового тегу (`<h1>`) і закінчується кінцевим тегом (`</h1>`), а між ними містить текст, який буде відтворюватись.

Перетворіть текст `Welcome to freeCodeCamp` на елемент `h1`, додавши початковий тег перед ним і кінцевий тег після нього.

# --hints--

Елемент `h1` повинен мати початковий тег `<h1>`.

```js
assert.exists(document.querySelector("h1"));
```

Елемент `h1` повинен мати кінцевий тег `</h1>`.

```js
assert.match(code, /<\/h1\s*\>/);
```

Елемент `h1` має виглядати ось так: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
