---
id: 682cd206883fc7b25eb539c5
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Зверніть увагу, що HTML, який ви пишете в редакторі, з’являється у попередньому перегляді. У цьому практичному занятті ви напишете HTML для частини навчальної програми.

У порожньому рядку під елементом `h1` напишіть такий текст:

```md
Full Stack Curriculum
```

# --hints--

В редакторі має бути текст `Full Stack Curriculum`. Двічі перевірте написання.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full\s*stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
