---
id: 6551f6f9ec6db270eb83f3db
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

您可以在终端上看到数字 `11` 被打印出来，因为 `'Hello World'` 包含11个字符。

另一个有用的内置函数是 `type()`, 它返回一个变量的数据类型。 修改您的代码`print()`来调用并打印`text`的数据类型.

# --hints--

你应该调用 `type()` 功能。

```js
assert.match(code, /type\s*\(.*\)/)
```

你应该将 `text` 传递到 `type()` 函数，在括号之间包含它。

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

你应该通过在括号之间添加一个 `type(text)` 到 `print()` 函数中。 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

你的代码中不应包含 `print(len(text))` 。

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
