---
id: 6551fe3b1df7c9740f13f270
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

Python变量命名的关键是：

- 有些词是系统保留的关键词(例如 `for`, `while`, `True`)。 他们在 Python 语言里有特殊的含义，所以你不能将它们用作变量名称。
- 变量名称不能以数字开头，它们只能包含字母、数字和下划线。
- 变量名称是区分大小写的，即 `my_var` 不同于 `my_Var` 和 `MY_VAR`
- 最后，用 `snake_case`写变量名称是一个惯例， 其中每个空格被下划线字符所取代，单词用小写字母写出。

删除两个 `print()` ,并定义另一个变量名为 `alphabet`。 将字符串 `'abcdefghijklmnopqrstuvwxyz'` 分配到此变量。

# --hints--

你的代码中不应包含 `print(type(text))` 。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

你的代码中不应包含 `print(type(shift))` 。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

你应该声明一个名为 `alphabet` 的变量。 注意将变量名放在行首。

```js
assert.match(code, /^alphabet\s*=/m)
```

你应该将字符串 `'abcdefghijklmnopqrstuvwxyz'` 分配给你的 `alphabet` 变量。 请记住使用单引号或双引号来括住字符串。

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

你的代码中包含无效的句法和/或无效的缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
