---
id: 6551fb531ce7f773071a0b84
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

如你所见，打印 `type(text)` 的输出是 `<class 'str'>`，这意味着你的变量是一个字串，表示为 `str`。

现在另起一行并创建一个名为 `shift`的变量，并赋值为 `3` 。

# --hints--

你应该声明一个名为 `shift` 的变量。 注意将变量名放在行首。

```js
assert.match(code, /^shift\s*=/m)
```

你应该将值 `3` 分配给你的 `shift` 变量。

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
