---
id: 6551f22653e6ce6e4f1d1370
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

修改您的代码`print(shift)`来调用并打印`shift`的数据类型.

# --hints--

你应该保留现有的 `print(type(text))` 调用。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

你应该将 `shift` 传递给 `type()` 函数。

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

你应该将 `type(shift)` 传递给 `print()` 函数。 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

你的代码中不应该有 `print(shift)`。

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
