---
id: 6551f1e8737dbe6e0cb08ea4
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

现在打印您的新变量。

# --hints--

你不应删除现有的 `print(type(text))` 调用。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

你应该调用传入 `shift` 变量的 `print()` 函数。 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
