---
id: 6551ecf597038c36637f3759
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

变量是Python和任何编程语言的重要部分。 变量是一个引用某对象或指向某对象的名称。 您可以在赋值运算符`=`的左侧写入变量的名称以设一个变量，并在赋值运算符右侧指定分配给该变量的值。

```py
variable_name = value
```

请创建一个名为`number`的变量，并分配`5`为该变量的值。

# --hints--

你应设一个名为`number`的变量。 注意将变量名放在行首。

```js
assert.match(code, /^number\s*=/m)
```

您应将`5`分配为变量`number`。

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
