---
id: 6551f4259d66026ee28cb205
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

您也可以从字符串结尾处开始访问字符。 最后一个字符的索引为 `-1`, 倒数第二个是 `-2` ，以此类推。

现在修改您现有的 `print()` 来调用并打印你字符串里的最后一个字符。

# --hints--

你仍应调用 `print()` 函数。

```js
assert.match(code, /print\s*\(.*\)/)
```

你应该将 `text[-1]` 包含在括号之间，传递给 `print()` 函数。 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
