---
id: 6551f5756c20146fc94f8675
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

您可以使用内置的 `len()` 函数来统计字符串中的字符数（长度）。

修改你现有的 `print()` 通话，通过 `len(text)` 而不是 `text[-1]`

# --hints--


你应该调用 `len()` 函数。

```js
assert.match(code, /len\s*\(.*\)/)
```

你应该通过将变量 `text` 包含在括号之间来将其传递给 `len()` 函数。

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

你应该将 `len(text)` 传递给 `print()` 函数，方法是将其包含在括号之间。 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

你的代码中不应该有 `print(text[-1])`

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
