---
id: 6551ef504d91936d2d4e54f8
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

*argument*是传递给函数的对象或表达式 - 在调用时添加在左括号和右括号之间：

```py
greet = 'Hello!'
print(greet)
```

上面例子中的代码将打印字符串 `'Hello!'`，它是作为参数传递给 `print()` 的变量 `greet` 的值。

通过将 `text` 变量作为参数传递给 `print()` 函数，将 `text` 变量打印到屏幕上。

# --hints--

你应该通过在左括号和右括号内添加此变量的名称，将 `text` 传递给 `print()` 函数。 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
