---
id: 6552008c0d9d9075cbec9772
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

你将使用 `.find()` 方法来查找消息中每个字母在字母表中的位置。 方法类似于函数，但它属于对象。

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

上面，`.find()` 方法*调用*在`sentence`（要搜索的字符串）上，并且 `'r'`（要定位的字符）作为参数传递。 `sentence.find('r')` 调用将返回 `4`，它是 `'r'` 在 `sentence`中第一次出现的索引。

在代码末尾，对 `alphabet` 调用 `.find()` 并将 `'z'` 作为参数传递给该方法。

# --hints--

你应该调用 `.find()` 方法。

```js
assert.match(code, /find\s*\(.*\)/)
```

你应该对 `alphabet` 变量调用 `.find()` 方法。

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


你应该对 `alphabet` 变量调用 `.find()` 方法，并将 `'z'` 传递给该方法。 注意将函数调用放在行首。

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
