---
id: 65520e6f2b9678799977f24d
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

删除最后一个 `print()` 调用。 然后，将 `text[0].lower()` 作为参数传递给 `.find()` 调用，而不是 `text[0]`查看输出。

# --hints--

你的代码中不应包含 `print(text.lower())`

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

你仍然应该打印你的 `index` 变量

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

你应该更新 `alphabet.find(text[0])` 调用以使用 `text[0].lower()` 作为参数。 注意将函数调用放在行首。

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
