---
id: 65524b3aa6a1938a069a91ab
title: 步骤 34
challengeType: 20
dashedName: step-34
---

# --description--

接下来，打印 `new_char` 并查看输出。

# --hints--

你应该打印你的 `new_char` 变量。

```js
assert.match(code, /print\s*\(\s*new_char\s*\)/)
```

你应该在你的循环身体末尾打印你的 `new_char` 变量。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*new_char\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
    new_char = alphabet[new_index]
--fcc-editable-region--
```
