---
id: 655220953ba90d80514d7ee2
title: 步骤 27
challengeType: 20
dashedName: step-27
---

# --description--

目前，`print()` 函数采用单个参数 `char`，但它可以采用多个参数，以逗号分隔。

向 `print(char)` 添加第二个参数，以便打印字符及其在字母表中的索引。

# --hints--

你应该将 `index` 添加为现有 `print(char)` 调用的第二个参数。 别忘了用逗号分隔参数。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
