---
id: 65520f8282faf57a0db4f7fe
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

从输出中可以看到，`'h'` 位于 `alphabet` 字符串中的索引 `7` 处。 现在你需要找到索引 `7` 处的字母以及 `shift` 的值。 为此，你可以使用加法运算符 `+`，就像使用它进行数学加法一样。

修改 `shifted` 变量，使其将 `alphabet` 的值存储在索引 `index + shift` 处。

# --hints--

你应该将 `alphabet[index + shift]` 分配给你的 `shifted` 变量

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
