---
id: 65524b790ba8558a2f1c9fe5
title: 步骤35
challengeType: 20
dashedName: step-35
---

# --description--

清理输出。 删除 `print(char, index)`，并将最后一个 `print()` 调用变为 `print('char:', char, 'new char:', new_char)`。

# --hints--

你的代码中不应包含 `print(char, index)`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/print\s*\(\s*char\s*,\s*index\s*\)/.test(commentless_code))
```

你应该将 `print(new_char)` 更改为 `print('char:', char, 'new char:', new_char)`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/print\s*\(\s*new_char\s*\)/.test(commentless_code));
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')new\schar:\2\s*,\s*new_char\s*\)\s*$/));
```

你的代码包含无效语法和/或无效缩进

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
    new_char = alphabet[new_index]
    print(new_char)
--fcc-editable-region--
```
