---
id: 65522255d5b9cd80f335c6f2
title: 步骤 29
challengeType: 20
dashedName: step-29
---

# --description--

在循环体的末尾，声明一个名为 `new_index` 的变量，并将 `index + shift` 的值分配给该变量。

# --hints--

你应该在 `for` 循环中声明一个名为 `new_index` 的变量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

你应该在 `for` 循环体末尾将 `index + shift` 分配给新变量。

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
