---
id: 65524d2a1a253b8bb5197ae2
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

`find` 再次针对大写字母和空格字符返回 `-1`。 稍后你将负责处理该空间。

现在，不要迭代 `text`，而是将 `for` 循环更改为迭代 `text.lower()`。

# --hints--

你应该修改 `for` 循环以迭代 `text.lower()` 而不是 `text`。

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
