---
id: 65525e359ca28d938baa82c5
title: 步骤 39
challengeType: 20
dashedName: step-39
---

# --description--

你可以使用加法赋值运算符获得与 `a = a + b` 相同的效果：

```py
a += b
```

加法赋值运算符使你能够将值添加到变量，然后将结果赋给该变量。

使用`+=` 运算符添加一个值，同时将其分配给`encrypted_text`。

# --hints--

你应该使用加法赋值运算符将 `alphabet[new_index]` 添加到循环体内 `encrypted_text` 的当前值。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
