---
id: 655261b2e1f2c197093f3993
title: 步骤 41
challengeType: 20
dashedName: step-41
---

# --description--

目前，空格被加密为 `'c'`。 要保持纯消息中的原始间距，你需要条件 `if` 语句。 它由 `if` 关键字、条件和冒号 `:` 组成。

```py
if x != 0:
    print(x)
```

在上面的示例中，`if` 语句的条件是 `x != 0`。 `if` 语句主体内的代码 `print(x)` 仅当条件计算结果为 `True` 时运行（在本例中，这意味着 `x` 不同于零）。

At the top of your for loop, replace `print(char == ' ')` with an `if` statement.  The condition of this `if` statement should evaluate to `True` if `char` is an empty space and `False` otherwise. 在 `if` 主体内，打印字符串 `'space!'` 记住缩进此行。

# --hints--

你的代码中不应包含 `print(char == ' ')` 。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

你应该将 `print(char == ' ')` 替换为 `if` 语句，该语句在 `char == ' '` 时触发。 不要使用括号将 `if` 条件括起来，并记住包含最后的冒号。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

你应该在新的 `if` 语句中打印字符串 `'space!'`。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
