---
id: 655208d59b131e7816f18c96
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

`print()` 函数仅在控制台中提供输出，但函数和方法可以具有可在代码中使用的返回值。

现在将 `alphabet.find(text[0])` 分配给名为 `index` 的变量。 这样，`index`将存储`alphabet.find(text[0])`返回的值。

# --hints--

你应该声明一个名为 `index` 的变量。 注意将变量名放在行首。

```js
assert.match(code, /^index\s*=/m)
```

你应该将 `alphabet.find(text[0])` 返回的值分配给 `index` 变量。

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
