---
id: 65521fc818947e800bffe48a
title: 步骤 26
challengeType: 20
dashedName: step-26
---

# --description--

在 `for` 循环内，在打印当前字符之前，声明一个名为 `index` 的变量，并将 `alphabet.find(char)` 返回的值分配给这个变量。

# --hints--

你应该在 `for` 循环的开头声明一个名为 `index` 的新变量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

你应该将 `alphabet.find(char)` 分配给新的 `index` 变量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
