---
id: 655251308f31958d06cdf267
title: 步骤 36
challengeType: 20
dashedName: step-36
---

# --description--

目前，加密字符在每次迭代中都会更新。 最好将加密的字符串存储在新变量中。 在 `for` 循环之前，声明一个名为 `encrypted_text` 的变量，并为该变量分配一个空字符串 (`''`)。

# --hints--

你应该在 `for` 循环之前声明一个名为 `encrypted_text` 的变量。

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

你应该为 `encrypted_text` 变量分配一个空字符串.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
