---
id: 655209a4a27dd37873c4cac3
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` 返回字符串中匹配字符的索引 如果未找到该字符，则返回`-1`。 正如你所看到的，`text` 中的第一个字符（大写的 `'H'`）未找到，因为 `alphabet` 仅包含小写字母。

你可以使用 `.lower()` 方法将字符串转换为其等效的小写形式。 添加另一个 `print()` 调用来打印 `text.lower()` 并查看输出。

# --hints--

你的代码中仍应包含 `print(index)` 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

你应该打印 `text.lower()`。 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
