---
id: 6553995f412dd8122ed38e4a
title: 第 43 步
challengeType: 20
dashedName: step-43
---

# --description--

条件语句还可以有 `else` 子句。 可以将此子句添加到 `if` 语句的末尾，以便在 `if` 语句的条件为假时执行替代代码：

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

正如你在输出中看到的，当循环迭代到达空格时，会在加密字符串中添加一个空格。 然后执行 `if` 块外部的代码，并将 `c` 添加到加密字符串中。

要修复此问题，请在 `encrypted_text += char` 之后添加 `else` 子句，并缩进除 `print()` 调用之外的所有后续代码行。

# --hints--

你应该创建一个 `else` 子句。 请记住包含最后的冒号。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

除了 `print()` 调用之外，你应该在 `else` 子句之后缩进代码行。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
