---
id: 6553980e0527fa115c705646
title: 步骤 42
challengeType: 20
dashedName: step-42
---

# --description--

现在，不再打印 `'space!'`，而是使用加法赋值运算符将空格（当前存储在 `char` 中）添加到 `encrypted_text`的当前值.

# --hints--

你的代码中不应包含 `print('space!')` 。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

你应该使用 `+=` 运算符将 `char` 添加到新 `if` 语句中 `encrypted_text` 的当前值。

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
