---
id: 6553f6b4ec51112d44d737c8
title: 步骤 48
challengeType: 20
dashedName: step-48
---

# --description--

在 `print` 调用之前，添加另一个并将 `'plain text:', text` 作为参数传递给 `print()`。 使用相同的缩进方式。

# --hints--

你应该在 for 循环之后打印 `'plain text:', text` 。

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*^print/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
