---
id: 6553a7d8d05cbb1ae335a665
title: 步骤 45
challengeType: 20
dashedName: step-45
---

# --description--

当循环到达字母 `Z` 时，`index + shift` 之和超过了字符串 `alphabet` 的最后一个索引。 因此，`alphabet[new_index]` 尝试使用无效索引，这会导致抛出 `IndexError` 异常。

你可以注意到终端中的输出停在 `Z` 之前的空格处，即引发错误之前的最后一个 `print`。

在这种情况下，模运算符 (`%`) 可用于返回两个数字相除的余数。 例如：`5 % 2` 等于 `1`，因为 5 除以 2 的商为 2，余数为 1。

将 `index + shift` 用括号括起来，并用 `26`（即 `alphabet` 长度）对表达式取模。

# --hints--

你的 `else` 语句中应该有 `new_index = (index + shift) % 26` 。

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
