---
id: 6553f3fc92741c2bf8ded140
title: 步骤54
challengeType: 20
dashedName: step-54
---

# --description--

在函数体内，将 `text` 和 `shift` 变量分别重命名为 `message` 和 `offset`。

# --hints--

你应该将所有出现的 `text` 重命名为 `message`。

```js
assert.match(code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*:/);
assert.match(code, /print\s*\(\s*("|')(plain text|(plain )?message):\1\s*,\s*message\s*\)/);
```

你应该将 `shift` 重命名为 `offset`。

```js
assert.match(code, /new_index\s*=\s*\(\s*index\s*\+\s*offset\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/);
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
--fcc-editable-region--
```
