---
id: 6553ed69ece88d29594748aa
title: 第 52 步
challengeType: 20
dashedName: step-52
---

# --description--

要执行，需要通过在函数名称后附加一对括号来调用（或调用）函数，如下所示:

```py
function_name()
```

在代码末尾，调用 `caesar` 函数。 注意缩进.

# --hints--

你应该调用你的 `caesar` 功能。 请确保在行开始时写入函数调用。

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
