---
id: 6553efd6ada3f42aa2d75448
title: 步骤 53
challengeType: 20
dashedName: step-53
---

# --description--

现在你应该再次看到输出。 虽然这种方法行之有效，但它并没有显著提高代码的可再使用性。 重复调用你的函数将导致相同的结果。 但是，可以使用*parameters*来声明函数，以引入多功能性和自定义功能：

```py
def function_name(param_1, param_2):
    <code>
```

参数是你可以在函数内使用的变量。 一个函数可以用不同数目的参数声明. 在上述示例中， `param_1` 和 `param_2` 是参数。

修改函数声明，使其采用两个名为 `message` 和 `offset` 的参数。

之后, 你会看到终端中出现一个错误。 你会看到如何在下一步中解决它。

# --hints--

你的 `caesar` 函数应该以 `message` 和 `offset` 作为参数。 请记住用逗号分隔参数。

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
