---
id: 6553a755879b131a445e664c
title: 步骤 44
challengeType: 20
dashedName: step-44
---

# --description--

尝试将字符串 `'Hello Zaira'` 分配给你的 `text` 变量，然后看看终端中会发生什么。

你将看到`string index out of range`异常。 别担心，你将很快确定如何修复它！

# --hints--

你应该有一个 `text` 变量。

```js
assert.match(code, /^text\s*=/m)
```

你应该将字符串 `'Hello Zaira'` 分配给你的 `text` 变量。

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
