---
id: 6553f4f66099802c6ae94613
title: 步骤 55
challengeType: 20
dashedName: step-55
---

# --description--

目前，你的代码会引发 `TypeError`，因为 `caesar` 函数是使用两个参数（`message` 和 `offset`）定义的，因此预计会使用两个*参数*来调用。

不带参数调用 `caesar()` 就停止执行代码。

给 `message` 和 `offset` 提供内容，通过将 `text` 和 `shift` 作为参数传递给 `caesar`函数调用。

# --hints--

你应该将 `text` 和 `shift` 作为参数传递给函数调用，方法是将它们包含在括号内。 别忘了用逗号分隔参数。

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
