---
id: 6553f6086add4b2cbb99fd78
title: 步骤 56
challengeType: 20
dashedName: step-56
---

# --description--

在代码底部，在现有的 `caesar(text, shift)` 调用之后，再次调用你的函数。

这次，传递 `text` 变量和整数 `13` 作为参数。

# --hints--

你应该再次调用你的函数，这次传递 `text` 和 `13` 作为参数。

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

你应该保持现有函数调用。

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
