---
id: 6553f8c570f9982e013a8886
title: 步骤 57
challengeType: 20
dashedName: step-57
---

# --description--

目前，根据指定的偏移量，每个字母总是用相同的字母加密。 如果每个字母的偏移量都不同呢？ 这样就更难解密了。 该算法称为维吉尼亚密码，其中每个字母的偏移量由另一个文本（称为*key*）确定。

通过删除两个函数调用，开始将凯撒密码转换为维吉尼亚密码

# --hints--

你应该删除两个 `caesar()` 函数调用。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
