---
id: 65549f90cf78131c96ebcf28
title: 步骤 61
challengeType: 20
dashedName: step-61
---

# --description--

当编码时，可读性是关键的。 注释是有效的备注，解释了代码背后的逻辑。 当一段时间后返回项目时，它们会非常有用，并且还能帮助你的同事理解代码。

在 Python 中，你可以使用 `#` 写注释。 `#` 之后出现的任何事情都不会被执行。

在 `if` 语句之前，添加一条注释：`Append space to the message`。

# --hints--

你应该添加一条注释，说明`Append space to the message`。 不要忘记开始时的 `#`。

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--

        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
