---
id: 6554e0adc7bb193cbfdb36d5
title: 步骤 73
challengeType: 20
dashedName: step-73
---

# --description--

检查该函数是否可以将字符串解密回纯文本。

声明另一个名为 `decryption` 的变量，并将其分配给 `vigenere(encryption, custom_key, -1)`。

# --hints--

你应该调用 `vigenere` 并传递 `encryption`、`custom_key` 和 `-1` 作为参数。

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/) })
```

你应该声明一个`decryption`变量

```js
({ test: () => assert(__userGlobals.has("decryption")) })
```

你的 `decryption` 变量的值应为 `vigenere(encryption, custom_key, -1)`。

```js
({ test: () => assert.match(code, /^decryption\s*=\s*vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--    
encryption = vigenere(text, custom_key, 1)
print(encryption)

--fcc-editable-region--
```
