---
id: 6554d0332949b133a0b35eaa
title: 步骤 68
challengeType: 20
dashedName: step-68
---

# --description--

调用你的函数，传递 `text` 和 `custom_key` 作为参数 将函数调用的返回值存储在名为 `encryption` 的变量中。

# --hints--

你应该调用 `vigenere`，传递 `text` 和 `custom_key` 作为参数。

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

你应该有一个`encryption`变量。

```js
({ test: () => assert(__userGlobals.has("encryption")) })
```

你应该将函数调用分配给`encryption`变量。

```js
({ test: () => assert.match(code, /encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():

        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]

    return encrypted_text
--fcc-editable-region--

--fcc-editable-region--
```
