---
id: 6554a88d5af937226f4a9121
title: 步骤 65
challengeType: 20
dashedName: step-65
---

# --description--

`.index()` 方法与 `.find()` 方法相同，但如果无法找到子字符串，则会抛出 `ValueError` 异常。

`ValueError` 是一种内置异常，当将类型正确但值不合适的参数传递给函数时会引发该异常。

递增 `key_index` 后，声明一个名为 `offset` 的变量 找到`key_char`在`alphabet`中的索引并将其分配给`offset` 使用 `.index()` 来查找索引。

# --hints--

你应该声明一个名为 `offset` 的变量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

你的 `offset` 变量应存储 `alphabet.index(key_char)` 的值。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
