---
id: 6554ad2463b8892748f8efdd
title: 步骤 67
challengeType: 20
dashedName: step-67
---

# --description--

目前，你的函数会打印一些字符串，但代码的其他部分无法使用这些值来执行任何操作。

为此，你需要使用 `return` 语句：

```python
def foo():
    return 'spam'
```

你需要编写 `return` ，后跟一个空格和函数应返回的值 当找到 `return` 语句时，该值将返回，并且函数的执行将停止， 转到函数调用后的下一行代码。 在上面的示例中，`foo` 函数返回字符串 `'spam'`。

从函数中删除两个 `print()` 调用并返回 `encrypted_text`。

# --hints--

你应该从函数中删除两个 `print()` 调用。

```js
({ 
  test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/print\s*\(\s*("|')plain\stext:\1\s*,\s*message\s*\)/.test(commentless_code))
    assert.isFalse(/print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/.test(commentless_code))
  }
})
```

你的函数应该返回 `encrypted_text`

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*/);
    const regex = /return\s+encrypted_text/
    assert(function_body.match(regex));
    assert.notMatch(block_body, regex);
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
