---
id: 65549561463f0016876e852c
title: 步骤 60
challengeType: 20
dashedName: step-60
---

# --description--

由于你的密钥比你必须加密的文本短，因此你需要重复它才能生成整个加密文本。 在函数体的开头，声明一个 `key_index` 变量并将其设置为零。

# --hints--

你应该在函数体的开头声明一个名为 `key_index` 的变量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

你应该将 `0` 分配给你的 `key_index` 变量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
