---
id: 6554930320d70414e7b6acc6
title: 步骤 59
challengeType: 20
dashedName: step-59
---

# --description--

删除你的 `shift` 变量。 然后，声明一个名为 `custom_key` 的新变量，并将字符串 `'python'` 的值分配给该变量。

# --hints--

你应该删除 `shift` 变量及其值。

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

你应该声明一个名为 `custom_key` 的变量。

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

你应该将字符串 `'python'` 分配给你的 `custom_key` 变量。

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
