---
id: 6554a334a40edb1fb4eff827
title: 步骤62
challengeType: 20
dashedName: step-62
---

# --description--

接下来，在 `else` 块内，声明一个名为 `key_char` 的变量，并将索引 `key_index` 处的 `key` 值 mod(`%`) `key` 的长度分配给它。

# --hints--

你应该在 `else` 子句主体的开头声明一个名为 `key_char` 的变量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");
assert(block_body.match(/^\s*key_char\s*=/));
```

你应该将 `key[key_index % len(key)]` 分配给 `key_char` 变量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");    
assert(block_body.match(/key_char\s*=\s*key\s*\[\s*key_index\s*%\s*len\s*\(\s*key\s*\)\s*\]/));
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--        
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
--fcc-editable-region--        
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
