---
id: 655491bd5b98b813fa5bedca
title: 步骤 58
challengeType: 20
dashedName: step-58
---

# --description--

现在修改你的函数声明：将函数名称更改为 `vigenere`，将第二个参数更改为 `key`。

# --hints--

你应该将函数名称修改为`vigenere`。

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

你的 `vigenere` 函数应将 `message` 和 `key` 作为参数。

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
