---
id: 6555201d9b7fc917399f9f0b
title: 步骤 80
challengeType: 20
dashedName: step-80
---

# --description--

如果调用该方法的字符串中的所有字符都是字母，则 `.isalpha()` 方法将返回 `True`。 例如，下面的代码返回 `True`：

```py
'freeCodeCamp'.isalpha()
# True
```

删除条件 `char == ' '` 并通过调用 `char` 上的 `.isalpha()` 方法来替换它。

# --hints--

你的代码中应该有 `if char.isalpha():`

```js
({ test: () => assert.match(code, /if\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--    
        # Append space to the message
        if char == ' ':
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
