---
id: 6555db4a7b788e15795674e5
title: 步骤 89
challengeType: 20
dashedName: step-89
---

# --description--

两个或更多字符串可以通过使用 `+` 操作器进行装配。 例如：`'Hello' + ' there!'` 结果为 `'Hello there!'`

调用 `print()` 函数并使用 `+` 运算符将 `text` 变量连接到字符串 `'Encrypted text: '`. 注意间隔。

# --hints--

你应该调用 `print()` 并传递 `'Encrypted text: ' + text` 作为参数。 

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^print\s*\(\s*("|')Encrypted\stext:\s\1\s*\+\s*text\s*\)/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--    

--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
