---
id: 655520c20cb1e6177b0641d6
title: 第 81 步
challengeType: 20
dashedName: step-81
---

# --description--

`not` 运算符用于对表达式求反。 当放置在真值（计算结果为 `True` 的值）之前时，它会返回 `False` ，反之亦然。

在 `if` 条件的开头添加 `not` 运算符来检查字符是否不是字母。

# --hints--

你应该在 `if` 语句的条件中使用 `not` 运算符。

```js
({ test: () => assert.match(code, /if\s+not\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char.isalpha():
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
