---
id: 6555df0c81300b175308557d
title: 步骤 90
challengeType: 20
dashedName: step-90
---

# --description--

在刚刚添加的 `print()` 调用下方，添加另一个 `print()` 调用，通过连接字符串 `Key: python``'Key: '` 和 `custom_key` 变量的值。

# --hints--

在刚刚添加的 print() 调用下方，添加另一个 `print()` 调用，通过连接字符串 `Key: python``'Key: '` 和 `custom_key` 变量的值。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^print\s*\(\s*("|')Key:\s\1\s*\+\s*custom_key\s*\)/m);
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)

--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
