---
id: 6555e43e783ed31a0532b1b2
title: 步骤 93
challengeType: 20
dashedName: step-93
---

# --description--

新行字符 `\n` 是一个用于表示新行的特殊序列。 你可以在字符串中作为正常字符串的序列写一个反斜线 `\` 之后, 然后是一个 `n`, 在程序运行时, 它将被输出中的一个新行替换.

在你的第一个 `print` 调用开始放置换行文并查看输出。

# --hints--

你应该将第一个打印调用修改为 `print(f'\nEncrypted text: {text}')`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*f("|')\\nEncrypted\stext:\s\{\s*text\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print(f'Encrypted text: {text}')
print(f'Key: {custom_key}')
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
