---
id: 6555d6e12c95701172b55709
title: 步骤 77
challengeType: 20
dashedName: step-77
---

# --description--

函数可以用默认参数调用。 缺省参数指的是参数未传递时函数应占用的值。 例如，下面的函数接受了三个参数，但你可以通过两个参数来调用它。 第三者默认情况下将假设指定的值：

```py
def foo(a, b, c=value):
    <code>
```

修改 `vigenere` 函数，使其 `direction` 参数具有默认值 `1`。

# --hints--

`vigenere` 函数的 `direction` 参数应具有默认值 `1`。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

--fcc-editable-region--
def vigenere(message, key, direction):
--fcc-editable-region--
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
