---
id: 655619327c7b364166f8dd6f
title: 步骤46
challengeType: 20
dashedName: step-46
---

# --description--

如果你希望将其他字符合并到`alphabet`字符串中，例如数字或特殊字符，你会发现有必要手动修改模运算的正确操作数。

将 `26` 替换为 `len(alphabet)` 以避免此问题。

# --hints--

你应该修改 `new_index` 值，将 `26` 替换为 `len(alphabet)`。

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
  assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/));
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % 26
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
