---
id: 655a2a7210094920069b117c
title: 步骤 40
challengeType: 20
dashedName: step-40
---

# --description--

比较操作符允许你根据两个对象的值进行比较。 你可以使用比较操作符，在你想要比较的对象之间放置它。 它们返回一个 *布尔值* — 即 `True` 或 `False` — 取决于表达式的真实性。

Python 有以下比较操作符：

<table>
  <thead>
    <tr>
      <th>运算符</th>
      <th>含义</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>等于</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>不等于</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>大于</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>小于</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>大于或等于</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>小于或等于</td>
    </tr>
  </tbody>
</table>

在你的循环实体开始时，打印对 `char` 和空格进行比较的结果(`' '`)。 为此使用相等运算符 `==`。

# --hints--

你应该在 `for` 循环中使用相等运算符将 `char` 与空格进行比较。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

你应该打印 `char` 与 `for` 循环内的空格进行比较的结果。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

你应该在循环开头打印 `char` 与空格的比较结果。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
