---
id: 65688b5a1655a7a6caede847
title: 步骤 31
challengeType: 20
dashedName: step-31
---

# --description--

当你尝试像上一步中那样更改字符串的各个字符时，你会收到 `TypeError`，当你的代码中使用了不适当类型的对象时，就会出现这种情况。

从错误消息中可以看到，字符串不支持项目分配，因为它们是不可变的。 然而，一个变量可以重新分配另一个字符串：

```py
message = 'Hello World'
message = 'Hello there!'
```

删除 `text[0]` 行并重新为 `text` 分配字符串 `'Albatross'`。

# --hints--

你的代码中不应包含 `text[0] = 'C'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

你应该重新分配 `text` 字符串 `'Albatross'`。 不要修改 `text = 'Hello World'`

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
