---
id: 65687d2f8c7ee27b0446cef3
title: 步骤 24
challengeType: 20
dashedName: step-24
---

# --description--

每次迭代执行的代码 — 放置在 `:` 之后 — 构成循环体 此代码必须缩进. 在Python，建议每个缩进级别使用4个空格。 此缩进级别是一个代码块。

```py
for i in text:
    print(i)
```

Python 依靠缩进来指示代码块。 一行末尾的彩色是一个信号，即将出现一个新的缩进代码块。

因此，当最后一个冒号之后没有找到缩进块时，代码执行将停止并抛出`IndentationError`。 此代码不会显示输出，而是引发 `IndentationError`：

```py
for i in text:
print(i)
```

通过添加对 `print(i)` 的调用为你的 `for` 循环提供主体。 记住缩进循环体。

# --hints--

你应该将 `print(i)` 添加到 `for` 循环体中。 注意缩进. 

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
