---
id: 656886ec4e4d9b96ded3c565
title: 步骤 25
challengeType: 20
dashedName: step-25
---

# --description--

迭代变量可以有任何有效的名称，但给它一个有意义的名称会很方便。

将你的 `i` 变量重命名为 `char`。

# --hints--

你应该将 `i` 替换为 `char`。 

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
