---
id: 6568c86dc9193000d11ca5e0
title: 步骤 38
challengeType: 20
dashedName: step-38
---

# --description--

不要将 `alphabet[new_index]` 分配给 `encrypted_text`，而是将 `encrypted_text` 的当前值加上 `alphabet[new_index]` 分配给这个变量。

# --hints--

你应该将 `encrypted_text + alphabet[new_index]` 分配给你的 `encrypted_text` 变量。

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*=\s*encrypted_text\s*\+\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
