---
id: 6568b7a40c7ba0ccbc4b4425
title: 步骤 32
challengeType: 20
dashedName: step-32
---

# --description--

正如你可以看到的那样，重新分配的字符串的每个字符都会在循环中打印。

通过删除 `text` 重新分配返回到原始字符串。

# --hints--

你的代码中不应包含 `text = 'Albatross'`

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

你仍然应该有一个值为 `'Hello World'` 的 `text` 变量

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
