---
id: 65688a50e6c998a21d8e41d3
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

字符串是不可变的，这意味着一旦创建，它们就不能被更改。 例如，你可能认为以下代码将 `my_string` 的值更改为字符串 `'train'`，但这是无效的：

```py
my_string = 'brain'
my_string[0] = 't'
```

通过使用括号表示法访问 `text` 中的第一个字母来确认这一点，并尝试将其更改为你选择的字符。 你将看到输出消失并出现错误。

# --hints--

你的代码中仍应包含 `text = 'Hello World'`。

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

你应该使用 `text[0]` 访问 `text` 中的第一个字母。

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

你应该使用 `=` 运算符将你选择的字符分配给 `text[0]`。 不要忘记以单引号或双引号显示字符。

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
