---
id: 65951b637ddec63611c482c0
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

你可以使用内置函数 `print()` 在终端上打印代码的输出。

函数是可重用的代码块，你可以在需要时调用或调用它们来运行其代码。 要调用函数，只需在其名称旁边写一对括号即可。 你将很快了解更多关于函数的信息。

现在转到一个新行，并在 `print()` 函数中添加一个空调用。 你不应该看到任何输出。

# --hints--

你的代码中应该有 `print()` 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
