---
id: 659526f3e5424b3eac4b6aa7
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

接下来，将 `index` 变量打印到控制台。

# --hints--

你应该打印你的 `index` 变量。 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
