---
id: 65952b70924edb449faf1950
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

声明一个名为 `shifted` 的新变量。 使用括号表示法访问索引 `index` 处的 `alphabet` 值并将其分配给新变量。

# --hints--

你应该声明一个名为 `shifted` 的变量。 注意将变量名放在行首。

```js
assert.match(code, /^shifted\s*=/m)
```

你应该将 `alphabet[index]` 分配给你的 `shifted` 变量。

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
