---
id: 659530ca823d6c47f61daa8a
title: 步骤 20
challengeType: 20
dashedName: 步骤-20
---

# --description--

打印你的 `shifted` 变量。

# --hints--

你应该打印你的 `shifted` 变量。 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

你的代码包含无效语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
