---
id: 6624e0ec0bf47017eed84b1f
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

你将要构建的第一种密码称为*Caesar*密码。具体来说，你将对消息中的每个字母，查找它在字母表中的位置，取字母表中向后 3 个位置的字母，并用新字母替换原始字母。

要实现此目的，你将使用上一步中讨论的 `.find()` 方法 修改现有的 `.find()` 调用，将其传递为 `text[0]` 作为参数，而不是 `'z'`。

# --hints--

你应该修改现有的 `alphabet.find('z')` 调用，将 `text[0]` 传递给该方法。

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
