---
id: 666064e915eba7aa1de03f6b
title: 步骤51
challengeType: 20
dashedName: step-51
---

# --description--

现在，通过删除尝试在 `caesar` 函数外部打印 `alphabet` 变量的行来修复错误。

# --hints--

你应该删除 `print(alphabet)` 行。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--
print(alphabet)
--fcc-editable-region--
```
