---
id: 66605a64a9822f01ec5ae8ce
title: 步骤 50
challengeType: 20
dashedName: step-50
---

# --description--

在 Python 中，变量的*范围*决定了可以访问该变量的位置：

- 在函数外部定义的变量具有*全局*作用域，并且可以从代码的任何部分访问它们。

- 函数内部定义的变量是该函数的*本地*，并且不能在函数外部访问。

要查看其实际效果，请尝试在代码末尾打印 `alphabet` 变量。 这将引发 `NameError` 异常

你应该看到一条错误消息，指示 `alphabet` 未定义。 这是因为 `alphabet` 是在 `caesar` 函数内部定义的，并且无法在其外部访问。

# --hints--

你应该尝试在 caesar 函数外部打印 `alphabet` 变量。

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
