---
id: 66cc095dfe1682753d2ab030
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

أحد المفاهيم التي تعلمتها في الدروس السابقة هو كيفية إرجاع القيم من دالة.

إليك تذكير بكيفية إرجاع قيمة من دالة:

```js
function myFunction() {
  return "Hello World";
}
```

داخل دالتك `addTwoAndSeven`، أرجع مجموع `2` و `7`.

# --hints--

يجب أن تُرجع دالتك `addTwoAndSeven` مجموع `2` و `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

يجب أن تُرجع دالتك `addTwoAndSeven` المجموع باستخدام العامل `+`.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
