---
id: 66cc12fa504b0479dac479a0
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

ميزة الدوال هي أنه يمكن استدعاؤها بمعلمات مختلفة، مما يسمح لك بإعادة استخدام نفس الشفرة مع قيم متنوعة.

إليك بعض أمثلة استدعاءات الدوال مع معلمات سلسلة نصية مختلفة:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

أضف `console.log` يستدعي الدالة `calculateSum` بالمعلمات `2` و `5`.

# --hints--

يجب أن يكون لديك `console.log` يستدعي الدالة `calculateSum` بالمعلمات `2` و `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateSum(2,5))", "console.log(calculateSum(5,2))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
