---
id: 66cc1a3a39aef47d6473cb2f
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

الخطوة التالية هي إضافة وظيفة القسمة إلى الآلة الحاسبة الخاصة بك.

أعلن عن دالة باسم `calculateQuotient` تأخذ معلمتين، `num1` و `num2`.

داخل الدالة، أرجع حاصل القسمة للمعلمتين.

ثم أضف `console.log` يستدعي الدالة `calculateQuotient` بالمعاملين `7` و `11`.

# --hints--

يجب إعلان دالة باسم `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

يجب أن تحتوي دالة `calculateQuotient` على معلمتين باسم `num1` و `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateQuotient } = explorer.allFunctions;
assert.lengthOf(calculateQuotient?.parameters, 2);
assert.equal(calculateQuotient?.parameters[0].toString(), "num1");
assert.equal(calculateQuotient?.parameters[1].toString(), "num2");
```

يجب أن تُرجع دالة `calculateQuotient` حاصل القسمة للمعلمتين.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

يجب أن يكون لديك `console.log` يستدعي الدالة `calculateQuotient` بالمعاملين `7` و `11`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateQuotient(7,11))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
