---
id: 66cc0a9e06e00b75d6782be9
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

في الدروس السابقة، تعلمت كيف تستدعي (تنفذ) دالة. استدعاء الدالة يعني تنفيذ الشفرة داخل الدالة.

إليك تذكير بكيفية استدعاء دالة:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

أضف عبارة `console.log`، وداخلها استدعِ الدالة `addTwoAndSeven`.

يجب أن ترى الآن مجموع `2` و `7` مسجلاً في وحدة التحكم.

# --hints--

يجب أن تحتوي شفرتك على `console.log`.

```js
assert.match(code, /console\.log(.*)/);
```

يجب أن تستدعي الدالة `addTwoAndSeven` داخل `console.log`.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
