---
id: 66cc16abc30b367ba1d1f9ef
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

حان الوقت الآن لاختبار دالة `calculateDifference` الخاصة بك.

ابدأ بإضافة `console.log` يستدعي دالة `calculateDifference` بالمعلمتين `22` و `5`.

ثم أضف `console.log` ثاني يستدعي دالة `calculateDifference` بالمعلمتين `12` و `1`.

أخيرًا، أضف `console.log` ثالث يستدعي دالة `calculateDifference` بالمعلمتين `17` و `9`.

# --hints--

يجب أن يكون لديك `console.log` يستدعي دالة `calculateDifference` بالمعلمتين `22` و `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(22,5))";
assert.include(strippedCode, expectedStatement);
```

يجب أن يكون لديك `console.log` يستدعي دالة `calculateDifference` بالمعلمتين `12` و `1`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(12,1))";
assert.include(strippedCode, expectedStatement);
```

يجب أن يكون لديك `console.log` يستدعي دالة `calculateDifference` بالمعلمتين `17` و `9`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(17,9))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
