---
id: 66cc281d687975858049fd8d
title: الخطوة 19
challengeType: 1
dashedName: step-19
---

# --description--

في الخطوة النهائية من ورشة العمل، ستحتاج إلى اختبار دالة `calculateSquareRoot` الخاصة بك.

أضف أول `console.log` يستدعي الدالة `calculateSquareRoot` مع المعلمة `25`.

ثم أضف ثاني `console.log` يستدعي الدالة `calculateSquareRoot` مع المعلمة `100`.

وبهذه التغييرات، يصبح تطبيق الآلة الحاسبة الخاص بك مكتملًا!

# --hints--

يجب أن يكون لديك `console.log` يستدعي الدالة `calculateSquareRoot` مع المعلمة `25`.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquareRoot\s*\(\s*25\s*\)\s*\)/);
```

يجب أن يكون لديك `console.log` يستدعي الدالة `calculateSquareRoot` مع المعلمة `100`.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquareRoot\s*\(\s*100\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

function calculateSquareRoot(num) {
  return Math.sqrt(num);
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

function calculateSquareRoot(num) {
  return Math.sqrt(num);
}

console.log(calculateSquareRoot(25));
console.log(calculateSquareRoot(100));
```
