---
id: 66cc0ba5881acb7692cfc4de
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

حان الوقت الآن لإضافة دالة أخرى.

أعلن دالة باسم `addThreeAndFour` تُرجع مجموع `3` و `4`.

ثم استدعِ دالة `addThreeAndFour` داخل `console.log` لرؤية النتيجة.

# --hints--

يجب أن تعلن دالة باسم `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

يجب أن تُرجع دالة `addThreeAndFour` مجموع `3` و `4`.

```js
assert.equal(addThreeAndFour(), 7);
```

يجب أن تُرجع دالة `addThreeAndFour` المجموع باستخدام العامل `+`.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

يجب أن تستدعي دالة `addThreeAndFour` داخل `console.log`.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
