---
id: 66cc0de2920bec775f610424
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

على الرغم من أن الدوال تعمل كما هو متوقع، هناك الكثير من التكرار في كودك.

بما أنك تبني آلة حاسبة، لا تريد إنشاء دالة لكل تركيبة ممكنة من الأرقام التي تريد جمعها.

من الأفضل إنشاء دالة واحدة قابلة لإعادة الاستخدام يمكنها جمع أي رقمين معًا.

احذف كل الكود الذي كتبته حتى الآن.

في الخطوة التالية، ستراجع كيفية العمل مع المعلمات والمعلمات في الدوال.

# --hints--

لا يجب أن يكون لديك دالة باسم `addTwoAndSeven`.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
assert.notExists(addTwoAndSeven);
```

لا يجب أن يكون لديك استدعاء دالة لـ `addTwoAndSeven`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addTwoAndSeven())";
assert.notInclude(strippedCode, unexpectedStatement);
```

لا يجب أن يكون لديك دالة باسم `addThreeAndFour`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
assert.notExists(addThreeAndFour);
```

لا يجب أن يكون لديك استدعاء دالة لـ `addThreeAndFour`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addThreeAndFour())";
assert.notInclude(strippedCode, unexpectedStatement);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
