---
id: 66cc14ca14e65e7a9ebecb80
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

الآن بعد أن أصبح بإمكان الآلة الحاسبة جمع رقمين معًا، حان الوقت لإنشاء دالة تطرح رقمين.

أعلن دالة باسم `calculateDifference` تأخذ معلمتين، `num1` و `num2`.

داخل الدالة، أرجع الفرق بين المعلمتين.

# --hints--

يجب إعلان دالة باسم `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

يجب أن تحتوي دالة `calculateDifference` على معلمتين باسم `num1` و `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateDifference } = explorer.allFunctions;
assert.lengthOf(calculateDifference?.parameters, 2);
assert.equal(calculateDifference?.parameters[0].toString(), "num1");
assert.equal(calculateDifference?.parameters[1].toString(), "num2");
```

يجب أن تُرجع دالة `calculateDifference` الفرق بين المعلمتين.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
