---
id: 66cc0f1ae40802781b2ea972
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

في الدروس السابقة، تعلمت كيف تتعامل مع معلمات الدالة.

معلمة الدالة هي متغير يُعرّف في إعلان الدالة ويعمل كعنصر نائب.

إليك مثال على دالة تحتوي على معلمة:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

أعلن دالة باسم `calculateSum` تأخذ معلمتين، `num1` و `num2`.

# --hints--

يجب أن تعلن دالة باسم `calculateSum`.

```js
assert.isFunction(calculateSum);
```

يجب أن تأخذ دالة `calculateSum` معلمتين، `num1` و `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateSum } = explorer.allFunctions;
assert.lengthOf(calculateSum?.parameters, 2);
assert.equal(calculateSum?.parameters[0].toString(), "num1");
assert.equal(calculateSum?.parameters[1].toString(), "num2");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
