---
id: 66cc06012e46aa6bc9b8c001
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستراجع كيفية العمل مع الدوال من خلال بناء تطبيق آلة حاسبة.

في الدروس السابقة، تعلمت كيفية إعلان الدوال بهذه الطريقة:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

ابدأ بإنشاء دالة تسمى `addTwoAndSeven`.

يمكنك اختيار استخدام صيغة الدالة العادية أو صيغة الدالة السهمية.

# --hints--

يجب إعلان دالة تسمى `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
