---
id: 66cc172af7159a7c67804544
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

حان الوقت الآن لإضافة وظيفة الضرب إلى الآلة الحاسبة الخاصة بك.

أعلن عن دالة باسم `calculateProduct` تأخذ معلمتين، `num1` و `num2`.

داخل الدالة، أرجع حاصل ضرب المعلمتين.

ثم أضف `console.log` يستدعي الدالة `calculateProduct` بالمعاملين `13` و `5`.

# --hints--

يجب أن يكون لديك دالة باسم `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

يجب أن تحتوي دالة `calculateProduct` على معلمتين باسم `num1` و `num2`.

```js
const regex = __helpers.functionRegex('calculateProduct', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تُرجع دالة `calculateProduct` حاصل ضرب المعلمتين.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

يجب أن يكون لديك `console.log` يستدعي الدالة `calculateProduct` بالمعاملين `13` و `5`.

```js
assert.match(code, /console\.log\s*\(\s*calculateProduct\s*\(\s*13\s*,\s*5\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
