---
id: 66cc0de2920bec775f610424
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

رغم أن الدوال تعمل كما هو متوقع، هناك الكثير من التكرار في شفرتك.

بما أنك تبني آلة حاسبة، لا تريد إنشاء دالة لكل تركيبة ممكنة من الأرقام التي تريد جمعها.

من الأفضل إنشاء دالة واحدة قابلة لإعادة الاستخدام يمكنها جمع أي رقمين معًا.

احذف كل الشفرة التي كتبتها حتى الآن.

في الخطوة التالية، ستراجع كيفية العمل مع المعلمات والمعلمات في الدوال.

# --hints--

لا يجب أن يكون لديك دالة باسم `addTwoAndSeven`.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

لا يجب أن يكون لديك استدعاء دالة لـ `addTwoAndSeven`.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

لا يجب أن يكون لديك دالة باسم `addThreeAndFour`.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

لا يجب أن يكون لديك استدعاء دالة لـ `addThreeAndFour`.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
