---
id: 66cc0a9e06e00b75d6782be9
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

في الدروس السابقة، تعلمت كيف تستدعي (تنفذ) دالة. استدعاء الدالة يعني تنفيذ الكود داخل الدالة.

إليك تذكير بكيفية استدعاء دالة:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

أضف عبارة `console.log`، وداخلها استدعِ الدالة `addTwoAndSeven`.

يجب أن ترى الآن مجموع `2` و `7` مسجلاً في وحدة التحكم.

# --hints--

يجب استدعاء الدالة `addTwoAndSeven` داخل `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addTwoAndSeven())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
