---
id: 66cc1a3a39aef47d6473cb2f
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

الخطوة التالية هي إضافة وظيفة القسمة إلى الآلة الحاسبة الخاصة بك.

أعلن عن دالة باسم `calculateQuotient` تأخذ معلمتين، `num1` و `num2`.

داخل الدالة، أرجع حاصل القسمة للمعلمتين.

ثم أضف `console.log` يستدعي الدالة `calculateQuotient` بالمعاملين `7` و `11`.

# --hints--

يجب أن يكون لديك دالة باسم `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

يجب أن تحتوي دالة `calculateQuotient` على معلمتين باسم `num1` و `num2`.

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تُرجع دالة `calculateQuotient` حاصل القسمة للمعلمتين.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

يجب أن يكون لديك `console.log` يستدعي الدالة `calculateQuotient` بالمعاملين `7` و `11`.

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
