---
id: 66cc0ba5881acb7692cfc4de
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

حان الوقت الآن لإضافة دالة أخرى.

أعلن دالة باسم `addThreeAndFour` تُرجع مجموع `3` و `4`.

ثم استدعِ دالة `addThreeAndFour` داخل `console.log` لرؤية النتيجة.

# --hints--

يجب أن تعلن دالة باسم `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

يجب أن تُرجع دالة `addThreeAndFour` مجموع `3` و `4`.

```js
assert.strictEqual(addThreeAndFour(), 7);
```

يجب أن تُرجع دالة `addThreeAndFour` المجموع باستخدام العامل `+`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

يجب أن تستدعي دالة `addThreeAndFour` داخل `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
