---
id: 66cc172af7159a7c67804544
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

حان الوقت الآن لإضافة وظيفة الضرب إلى الآلة الحاسبة الخاصة بك.

أعلن عن دالة باسم `calculateProduct` تأخذ معلمتين، `num1` و `num2`.

داخل الدالة، أرجع حاصل ضرب المعلمتين.

ثم أضف `console.log` يستدعي الدالة `calculateProduct` بالمعاملين `13` و `5`.

# --hints--

يجب إعلان دالة باسم `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

يجب أن تحتوي دالة `calculateProduct` على معلمتين باسم `num1` و `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateProduct } = explorer.allFunctions;
assert.lengthOf(calculateProduct?.parameters, 2);
assert.equal(calculateProduct?.parameters[0].toString(), "num1");
assert.equal(calculateProduct?.parameters[1].toString(), "num2");
```

يجب أن تُرجع دالة `calculateProduct` حاصل ضرب المعلمتين.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

يجب أن يكون لديك `console.log` يستدعي الدالة `calculateProduct` بالمعاملين `13` و `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateProduct(13,5))", "console.log(calculateProduct(5,13))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
