---
id: 66cc2739f687e484d50bb6f1
title: الخطوة 18
challengeType: 1
dashedName: step-18
---

# --description--

آخر وظيفة ستضيفها إلى الآلة الحاسبة هي وظيفة الجذر التربيعي.

المربع هو عدد يُضرب في نفسه. إذن، الجذر التربيعي هو العدد الذي عندما يُضرب في نفسه يعطي العدد الأصلي.

على سبيل المثال، الجذر التربيعي للعدد `9` هو `3` لأن `3 * 3 = 9`.

لحساب الجذر التربيعي لعدد في JavaScript، يمكنك استخدام الطريقة <dfn>Math.sqrt()</dfn>.

إليك مثالاً:

```js
// result: 3
Math.sqrt(9);
```

أعلن دالة باسم `calculateSquareRoot` تأخذ معلمة `num` وتُرجع الجذر التربيعي لـ `num`.

# --hints--

يجب أن يكون لديك دالة باسم `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

يجب أن تأخذ دالة `calculateSquareRoot` معلمة `num`.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تُرجع دالة `calculateSquareRoot` الجذر التربيعي لـ `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
