---
id: 66cc21d23238dc8240a8a182
title: الخطوة 16
challengeType: 1
dashedName: step-16
---

# --description--

سيكون من الجيد أن يحسب الآلة الحاسبة مربع الرقم. مربع الرقم هو الرقم مضروب في نفسه.

لحساب مربع رقم في JavaScript، يمكنك استخدام الطريقة <dfn>Math.pow()</dfn>. تأخذ الطريقة `Math.pow()` معلمتين: الرقم الأساسي والأس.

يمكنك أيضًا استخدام عامل الأس (`**`) لحساب مربع الرقم.

إليك مثالًا:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

أعلن دالة باسم `calculateSquare` تأخذ معلمة `num` وتُرجع مربع `num`.

# --hints--

يجب أن يكون لديك دالة باسم `calculateSquare`.

```js
assert.isFunction(calculateSquare);
```

يجب أن تأخذ دالة `calculateSquare` معلمة `num`.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تُرجع دالة `calculateSquare` مربع `num`.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
