---
id: 66cdf76685e4cb5a8726e27b
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

لاختبار دالتك `calculateSum` بشكل أعمق، يجب استدعاؤها بمعلمات مختلفة.

أضف `console.log` آخر يستدعي الدالة `calculateSum` بالمعلمات `10` و `10`.

أسفل ذلك `console.log`، أضف `console.log` آخر يستدعي الدالة `calculateSum` بالمعلمات `5` و `5`.

# --hints--

يجب أن يكون لديك `console.log` يستدعي الدالة `calculateSum` بالمعلمات `10` و `10`.

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

يجب أن يكون لديك `console.log` يستدعي الدالة `calculateSum` بالمعلمات `5` و `5`.

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
--fcc-editable-region--

--fcc-editable-region--
```
