---
id: 65f0694fb2296f3caadf8347
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Die Länge eines Vektors $\mathbf{a}$, oder Norm, wird typischerweise als $\\| \mathbf{a} \\|$ angegeben. Sie kann als Quadratwurzel der Summe seiner quadrierten Komponenten berechnet werden:
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
 \\]

Berechnen Sie die Vektornorm und geben Sie das Ergebnis aus Ihrer `norm`-Methode zurück.

# --hints--

Ihre `norm`-Methode sollte die angegebene Formel verwenden, um eine Zahl zu berechnen und zurückzugeben, die die Norm der aktuellen Vektorinstanz darstellt. Nähern Sie den Wert nicht an.

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
