---
id: 65f058e06f34fd1f0ee6e55d
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Erstellen Sie als Nächstes innerhalb der `Vector`-Klasse eine leere `norm`-Methode und geben Sie ihr einen `self`-Parameter.

# --hints--

Ihre `Vector`-Klasse sollte eine `norm`-Methode enthalten.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

Ihre `norm`-Methode sollte einen `self`-Parameter haben.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
