---
id: 65f03d9f92eac9183a4d3281
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Ein Vektor kann durch zwei Koordinaten, `x` und `y`, in der euklidischen Ebene definiert werden. Der Abstand zwischen dem Ursprung der Achsen und dem Punkt `(x, y)` ist seine Länge oder Norm. Und die Richtung des Vektors zeigt auf `(x, y)`.

<img alt="a 2-dimensional vector of coordinates (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Erstellen Sie innerhalb der `Vector`-Klasse eine `__init__`-Methode und geben Sie ihr drei Parameter: `self`, `x` und `y`.

# --hints--

Sie sollten eine `__init__`-Methode innerhalb der `Vector`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

Ihre `__init__`-Methode sollte drei Parameter annehmen: `self`, `x` und `y`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
