---
id: 65f01969115f933073b6be03
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Ein Vektor ist ein Objekt, das eine Länge (oder Größe) und eine Richtung hat und nicht durch eine einzelne Zahl ausgedrückt werden kann. In der Physik werden Vektoren häufig verwendet, um Kräfte, Geschwindigkeiten, Beschleunigungen und andere Größen darzustellen.

In diesem Projekt werden Sie einen Vektorraum aufbauen, eine Menge, in der eine Reihe von Operationen zwischen den Elementen dieser Menge definiert ist. Alle Details dazu lernen Sie sehr bald.

Für den Anfang starten Sie das Projekt, indem Sie eine leere Klasse namens `Vector` deklarieren.

# --hints--

Sie sollten eine leere Klasse namens `Vector` deklarieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
