---
id: 65f06a8e5a57673d700c79c3
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Außerhalb der `Vector`-Klasse erstellen Sie eine Instanz von `Vector`, indem Sie die ganzen Zahlen `2` und `3` als die Werte für `x` und `y` übergeben, und weisen Sie diese einer Variablen namens `v1` zu.

# --hints--

Sie sollten eine Variable `v1` deklarieren und ihr `Vector(2, 3)` zuweisen.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
