---
id: 65f055b9190fc41ca35549b8
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Python bietet verschiedene Methoden an, deren Namen sowohl mit einem führenden als auch einem abschließenden doppelten Unterstrich versehen sind. Einige davon kennen Sie vielleicht schon, wie `__init__` und `__str__`. Diese Methoden, die dem `__<name>__`-Namensmuster folgen, werden als spezielle Methoden, magische Methoden oder *dunder*-Methoden (was für double underscore steht) bezeichnet.

Das Definieren spezieller Methoden in einer Klasse beeinflusst das Verhalten dieser Klasse. Sie werden im Hintergrund in bestimmten Situationen aufgerufen (z. B. `__init__` während der Instanziierung, `__str__` wenn das Objekt ausgegeben oder an `str()` übergeben wird). In diesem Projekt lernen Sie einige der am häufigsten verwendeten speziellen Methoden kennen.

Weisen Sie vorerst `x` dem `x`-Attribut des `Vector`-Objekts zu.

# --hints--

Sie sollten `x` `self.x` zuweisen.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
