---
id: 65f3fa60a93b84110b3f1708
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Sie möchten eine aussagekräftige String-Darstellung zurückgeben. Zum Beispiel möchten Sie im Fall von `v1` einen String zurückgeben, der das Tuple mit den Werten der Vektorkomponenten enthält: `(2, 3)`.

Geben Sie aus der `__str__`-Methode einen String zurück, der den Vektor als Tuple darstellt und die Vektorkomponenten in der Reihenfolge enthält. Gehen Sie dann außerhalb der `Vector`-Klasse und geben Sie `v1` aus, um das Ergebnis zu überprüfen.

# --hints--

Ihre `__str__`-Methode sollte einen String zurückgeben, der den Vektor als Tuple darstellt und die Vektorkomponenten in der Reihenfolge enthält.

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

Sie sollten `v1` ausgeben.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
