---
id: 65f3fa097f26b510db6c710b
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Super, `norm` funktioniert wie erwartet. Wenn Sie jetzt versuchen, `v1` auszugeben, erhalten Sie die Standard-String-Darstellung eines Objekts (etwas wie `<__main__.Vector object at 0x11eb778>`).

Deklarieren Sie innerhalb der `Vector`-Klasse eine leere `__str__`-Methode, um eine lesbare String-Darstellung zu implementieren. Denken Sie daran, ihr einen `self`-Parameter zu geben.

Achten Sie darauf, `v1` nicht auszugeben, bevor `__str__` einen String zurückgibt, sonst erhalten Sie einen `TypeError`, da `__str__` immer einen String zurückgeben muss.

# --hints--

Sie sollten eine `__str__`-Methode innerhalb der `Vector`-Klasse definieren.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

Ihre `__str__`-Methode sollte einen `self`-Parameter haben.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
