---
id: 65f3fe07cc763212efe91285
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Ein Vektor kann eine Anzahl `n` von Dimensionen (Komponenten) haben. Hier ist eine Darstellung eines dreidimensionalen Vektors:

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Bisher haben Sie einen zweidimensionalen Vektor erstellt. Sie möchten Vektoren mit einer unterschiedlichen Anzahl von Dimensionen darstellen können, ohne den notwendigen Code für jeden speziellen Fall neu schreiben zu müssen. Dafür verwenden Sie Vererbung.

Beginnen Sie damit, die `Vector`-Klasse in `R2Vector` umzubenennen, um anzugeben, dass diese Klasse sich mit zweidimensionalen Vektoren beschäftigt. Denken Sie daran, auch die Instanziierung von `v1` zu ändern.

# --hints--

Sie sollten die `Vector`-Klasse in `R2Vector` umbenennen. Denken Sie auch daran, die Instanziierung von `v1` zu ändern.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
