---
id: 65f4535bbdb28d436ff3ddc9
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Wie Sie an der Ausgabe sehen können, enthält `__dict__` die Werte Ihrer Instanzattribute. Anstatt die Quadrate von `self.x` und `self.y` explizit zu addieren, werden Sie die Werte, die in `__dict__` gespeichert sind, durchlaufen, um den Wert der Norm zu berechnen.

Ersetzen Sie im Rumpf der `norm`-Methode den Inhalt der runden Klammern durch einen Generatorausdruck, der jeden Wert `val` in `self.__dict__.values()` mit der Potenz `2` erhöht. Übergeben Sie diesen Generatorausdruck außerdem als Argument an die `sum`-Funktion, damit alle Quadrate addiert werden, bevor die Quadratwurzel berechnet wird.

# --hints--

Sie sollten `sum(val**2 for val in self.__dict__.values())**0.5` aus der `norm()`-Methode zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
