---
id: 65f4520e363e2642f8112e33
title: Schritt 19
challengeType: 20
dashedName: step-19
---

# --description--

Wie Sie sehen, stimmt etwas nicht ganz. Die `norm`- und `__str__`-Methoden, die von `R2Vector` geerbt wurden, können sich nicht an einen dreidimensionalen Vektor anpassen. Ihre Implementierung muss flexibler sein.

Jedes Objekt in Python hat ein spezielles Attribut namens `__dict__`, das ein Wörterbuch ist und die Objektattribute speichert.

Entfernen Sie die vorhandenen `print`-Aufrufe. Geben Sie dann das `__dict__`-Attribut Ihrer `v1`- und `v2`-Vektoren aus, um zu sehen, wie sie aussehen.

# --hints--

Sie sollten keine `print(v2)` und `print(v2.norm())` in Ihrem Code haben.

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

Sie sollten das `__dict__`-Attribut von `v1` und `v2` ausgeben.

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
