---
id: 65f40401e6ef53173c04e27d
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Deklarieren Sie innerhalb Ihrer neuen Klasse eine `__init__`-Methode. Geben Sie ihr `self`, `x`, `y` und `z` als Parameter.

# --hints--

Sie sollten eine `__init__`-Methode innerhalb Ihrer `R3Vector`-Klasse definieren.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").has_function("__init__")
    `))
})
```

Ihre `__init__`-Methode sollte vier Parameter haben: `self`, `x`, `y` und `z`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, x, y, z")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
