---
id: 65f40051d6b09a139f253e8e
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Vererbung ermöglicht es Ihnen, eine Klasse von einer bestehenden Klasse abzuleiten. Die neue Klasse, Kind genannt, erbt alle Methoden und Eigenschaften der bestehenden Klasse, die Elternklasse genannt wird.

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

Im obigen Beispiel erbt die Kind-Klasse `Oak` von `Tree` und übernimmt die `sprout`-Methode von der Elternklasse `Tree`.

Erstellen Sie eine neue Klasse mit dem Namen `R3Vector` und folgen Sie dem obigen Beispiel, um sie von der `R2Vector`-Klasse erben zu lassen.

# --hints--

Sie sollten eine neue Klasse `R3Vector` definieren.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

Ihre neue Klasse `R3Vector` sollte von `R2Vector` erben.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
