---
id: 65f412c208c3791fee305acf
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

In Python können Sie die Verwendung von nur mit Schlüsselwortargumenten erzwingen, indem Sie einen `*` als zusätzlichen Parameter zur Signatur der Funktion oder Methode hinzufügen. Ändern Sie beide `__init__`-Methoden, indem Sie einen `*` als zweiten Parameter (nach `self`) hinzufügen. Jeder Parameter, der danach folgt, erfordert die Verwendung eines Schlüsselwortarguments beim Aufruf der Funktion/Methode.

Das bedeutet, dass Sie auch den `super().__init__(x, y)`-Aufruf anpassen müssen. Tun Sie dies, indem Sie `x` den Wert `x` und `y` den Wert `y` zuweisen.

Schließlich ändern Sie die Instanziierung von `v1` und `v2`, indem Sie Schlüsselwortargumente verwenden.

# --hints--

Sie sollten die Verwendung von Schlüsselwortargumenten erzwingen, indem Sie einen `*` als zweiten Parameter in beiden `__init__`-Methoden hinzufügen.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

Sie sollten Ihren `super.__init__(x, y)`-Aufruf so ändern, dass Schlüsselwortargumente verwendet werden.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

Sie sollten die Zuweisung von `v1` so ändern, dass sie `R2Vector(x=2, y=3)` ist.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

Sie sollten die Zuweisung von `v2` so ändern, dass sie `R3Vector(x=2, y=2, z=3)` ist.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
