---
id: 65f407ea37ad6e181b90462e
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Sie könnten jeden `i` Parameter wie zuvor `self.i` zuweisen. Obwohl Sie in diesem Fall nur wenige Zeilen wiederholen müssten, ist eine Möglichkeit, diese Wiederholung zu vermeiden, die Verwendung der `super()` Funktion. `super()` ermöglicht es Ihnen, implizit auf die Elternklasse zu verweisen: `super().__init__(x, y)` ruft die `__init__` Methode der Elternklasse auf.

Fügen Sie einen `super().__init__(x, y)` Aufruf zu Ihrer `__init__` Methode hinzu.

# --hints--

Sie sollten einen `super().__init__(x, y)` Aufruf in Ihrer `__init__` Methode haben.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
