---
id: 65f876d17832001e8e1abb05
title: Schritt 28
challengeType: 20
dashedName: step-28
---

# --description--

Deklarieren Sie eine Variable `arg_list`. Weisen Sie ihr den Wert einer List Comprehension zu, die `key` und `val` für jedes Schlüssel-Wert-Paar in `vars(self).items()` durchläuft und bei jedem Durchlauf den f-String `f'{key}={val}'` berechnet.

# --hints--

Sie sollten eine Variable `arg_list` deklarieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_variable("arg_list")`)) })
```

Sie sollten eine List Comprehension zu `vars(self).items()` zuweisen, die `arg_list` durchläuft.

```js
({ test: () => runPython(`
import ast
var = _Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list")
assert isinstance(var.tree.value, ast.ListComp), "Expected arg_list to be a list comprehension"
assert (actual := var.find_comp_iters()[0].is_equivalent("vars(self).items()")), f"Expected vars(self).items(), got {actual}"
`) })
```

Die List Comprehension, die `arg_list` zugewiesen wird, sollte `key` und `val` verwenden, um `vars(self).items()` zu durchlaufen.

```js
({ test: () => assert(runPython(`
_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_targets()[0].is_equivalent("key, val")
`)) })
```

Die List Comprehension, die `arg_list` zugewiesen wird, sollte für jedes Schlüssel-Wert-Paar in `f'{key}={val}'` den String `vars(self).items()` berechnen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_expr().is_equivalent("f'{key}={val}'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
