---
id: 65f8709620b2ce1a62608f5a
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Wenn Sie dynamisch auf einige Attribute zugreifen müssen, ausgehend von einer String-Eingabe, ist die eingebaute `getattr()`-Funktion genau das, was Sie brauchen. Sie nimmt ein Objekt als erstes Argument und einen String, der den Attributnamen enthält, als zweites Argument.

Beginnen Sie damit, die `__str__`-Methode zu korrigieren, indem Sie den von `__str__()` zurückgegebenen String durch einen Generatorausdruck ersetzen, der über die Objektattribute iteriert und für jedes Attribut `getattr()` aufruft `i`.

# --hints--

Sie sollten einen Generatorausdruck zurückgeben, der über `vars(self)` iteriert.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

Sie sollten einen Generatorausdruck zurückgeben, der `i` als Iterationsvariable verwendet.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

Sie sollten einen Generatorausdruck zurückgeben, der für jedes Element `getattr(self, i)` in `i` `vars(self)` aufruft.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
