---
id: 65f8749b13774b1d2e4a7fba
title: Schritt 27
challengeType: 20
dashedName: step-27
---

# --description--

Während die `__str__`-Methode eine für Menschen lesbare String-Darstellung eines Objekts zurückgibt, soll die `__repr__`-Methode den String zurückgeben, der benötigt wird, um das Objekt zu instanziieren.

Fügen Sie innerhalb der `__repr__`-Klasse eine `self`-Methode mit einem `R2Vector`-Parameter hinzu.

# --hints--

Ihre `R2Vector`-Klasse sollte eine `__repr__`-Methode haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__repr__")`)) })
```

Ihre `__repr__`-Methode sollte einen `self`-Parameter haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
