---
id: 65f9b9710cca621244d3bde1
title: Schritt 40
challengeType: 20
dashedName: step-40
---

# --description--

In Python ist `NotImplemented` ein spezieller Wert, der verwendet wird, um anzuzeigen, dass eine Operation für einen bestimmten Fall nicht implementiert ist.

`NotImplemented` löst nicht sofort eine Ausnahme aus. Stattdessen signalisiert es dem anderen Operanden, wie die Operation ausgeführt werden soll. Wenn die Anfrage nicht erfüllt werden kann, wird standardmäßig ein `TypeError` zurückgegeben.

Da Sie zwei Vektoren nur dann addieren möchten, wenn sie zur gleichen Klasse gehören, geben Sie `NotImplemented` aus der `if`-Anweisung zurück, die Sie im vorherigen Schritt erstellt haben.

# --hints--

Sie sollten `NotImplemented` aus dem Körper von `if` zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
