---
id: 65f9cb678070ca1668898c70
title: Schritt 44
challengeType: 20
dashedName: step-44
---

# --description--

Auf die gleiche Weise, wie `__add__` unter der Haube aufgerufen wird, wenn zwei Objekte addiert werden, wird die `__sub__`-Methode implizit im Fall einer Subtraktion aufgerufen.

Definieren Sie nun eine leere `__sub__`-Methode und geben Sie zwei Parameter an: `self` und `other`. Erstellen Sie innerhalb Ihrer neuen Methode eine if-Anweisung, um zu prüfen, ob `self` und `other` nicht zur gleichen Klasse gehören, und geben Sie `NotImplemented` zurück, wie Sie es zuvor getan haben.

# --hints--

Sie sollten eine `__sub__`-Methode innerhalb der `R2Vector`-Klasse definieren.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

Ihre `__sub__`-Methode sollte zwei Parameter haben, `self` und `other`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

Sie sollten eine `if`-Anweisung erstellen, die überprüft, ob `self` und `other` nicht zur gleichen Klasse gehören.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

Sie sollten `NotImplemented` aus dem Körper von `if` zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
