---
id: 65f9bc10a9fb1612e066e9e8
title: Schritt 42
challengeType: 20
dashedName: step-42
---

# --description--

Beim Addieren von zwei Vektoren wird jede Komponente eines Vektors zur gleichen Komponente des anderen Vektors addiert. Zum Beispiel erzeugt das Addieren von `(1, 2)` und `(2, 4)` einen dritten Vektor `(3, 6)`, wobei `3` die Summe von `1` und `2` ist und `6` die Summe von `2` und `4`.

Das `kwargs`-Dictionary wird die Schlüssel-Wert-Paare enthalten, die benötigt werden, um einen neuen Vektor derselben Klasse wie die beiden addierten Vektoren zu instanziieren.

Verwandeln Sie das leere Dictionary in eine Dictionary-Comprehension, die `vars(self)` durchläuft und für jeden Schlüssel (`i`) ein Schlüssel-Wert-Paar erstellt, wobei der Schlüssel derselbe Schlüssel der aktuellen Iteration ist und sein Wert die Summe der Werte der `i`-Attribute der beiden Operanden ist.

# --hints--

Sie sollten das leere Dictionary, das `kwargs` zugewiesen ist, in eine Dictionary-Comprehension umwandeln, die `vars(self)` durchläuft.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

Die Dictionary-Comprehension, die `kwargs` zugewiesen ist, sollte die Variable `i` verwenden, um `vars(self)` zu durchlaufen.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

Die Dictionary-Comprehension, die `kwargs` zugewiesen ist, sollte `i` als Schlüssel verwenden.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

Die Dictionary-Comprehension, die `kwargs` zugewiesen ist, sollte die Summe des `i`-Attributs von `self` und des `i`-Attributs von `other` dem Schlüssel `i` für jeden `i` in `vars(self)` zuweisen. Verwenden Sie dafür die `getattr()`-Funktion.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
