---
id: 65f93d4647ce2670dc6f095c
title: Schritt 31
challengeType: 20
dashedName: step-31
---

# --description--

Die `__repr__`-Methode wird im Hintergrund von der `repr`-Funktion aufgerufen. Überprüfen Sie, ob `__repr__` richtig funktioniert, indem Sie `f'\nrepr = {repr(v1)}'` und `f'\nrepr = {repr(v2)}'` als zweiten Argumente zu Ihren ersten und zweiten `print`-Aufrufen hinzufügen.

# --hints--

Sie sollten `f'\nrepr = {repr(v1)}'` als zweiten Argument zu Ihrem `print(f'v1 = {v1}')`-Aufruf hinzufügen.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v1 = {v1}', f'\\\\nrepr = {repr(v1)}')")`)) })
```

Sie sollten `f'\nrepr = {repr(v2)}'` als zweiten Argument zu Ihrem `print(f'v2 = {v2}')`-Aufruf hinzufügen.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v2 = {v2}', f'\\\\nrepr = {repr(v2)}')")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
--fcc-editable-region--
```
