---
id: 65f9dd6e5a08af19c196c2df
title: Schritt 45
challengeType: 20
dashedName: step-45
---

# --description--

Der Vektor, der sich aus der Subtraktion eines Vektors von einem anderen ergibt, wird berechnet, indem die Differenz jeder ihrer Komponenten ermittelt wird. Zum Beispiel erzeugt die Subtraktion von `(2, 4)` von `(7, 3)` einen dritten Vektor `(5, -1)`, wobei `5` die Differenz zwischen `7` und `2` ist und `-1` die Differenz zwischen `3` und `4`.

Wie Sie es zuvor in der `__add__`-Methode gemacht haben, deklarieren Sie eine `kwargs`-Variable nach der `if`-Anweisung. Weisen Sie ihr eine Dictionary-Comprehension zu, die die Objektattribute durchläuft und für jeden Schlüssel `i` ein Schlüssel-Wert-Paar erstellt, wobei der Schlüssel derselbe Schlüssel wie in der aktuellen Iteration ist und dessen Wert die Differenz zwischen den Werten der `i`-Attribute der beiden Operanden ist.

# --hints--

Sie sollten eine Variable `kwargs` deklarieren und ihr eine Dictionary-Comprehension zuweisen, die die Objektattribute durchläuft.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

Die der Variable `kwargs` zugewiesene Dictionary-Comprehension sollte die Variable `i` verwenden, um die Objektattribute zu durchlaufen.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

Die der Variable `kwargs` zugewiesene Dictionary-Comprehension sollte `i` als Schlüssel verwenden.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

Die der Variable `kwargs` zugewiesene Dictionary-Comprehension sollte für jeden `i` die Differenz zwischen dem `self`-Attribut von `i` und dem `other`-Attribut von `i` dem Schlüssel `i` zuweisen.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
