---
id: 65f93b67169a9c703264458a
title: Schritt 30
challengeType: 20
dashedName: step-30
---

# --description--

Da die Methode den String zurückgeben soll, um das Objekt für `R2Vector` sowie `R3Vector` bei Vererbung zu instanziieren, können Sie den String nicht mit dem Klassennamen fest codieren.

Sie können mit `__class__.__name__` auf den Namen einer Klasse zugreifen. Fügen Sie eine `return`-Anweisung zur `__repr__`-Methode hinzu und geben Sie den String zurück, der benötigt wird, um das Objekt zu instanziieren.

# --hints--

Sie sollten `__class__.__name__` verwenden, um den f-String zu erstellen, der benötigt wird, um ein Vektorobjekt zu instanziieren, und ihn aus der `__repr__`-Methode zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_return("f'{self.__class__.__name__}({args})'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
