---
id: 65f9b17a638f0d0dcce8c354
title: Schritt 38
challengeType: 20
dashedName: step-38
---

# --description--

Um einen Vektorraum zu erstellen, müssen Sie definieren, wie sich Vektoren in verschiedenen Fällen verhalten sollen. Vektoren können addiert werden, wodurch ein neuer Vektor entsteht.

Die spezielle Methode `__add__` kann implementiert werden, um zu überschreiben, was standardmäßig passiert, wenn zwei Objekte mit dem `+`-Operator addiert werden.

Momentan würde das Hinzufügen von zwei Instanzen von `R2Vector` oder `R3Vector` eine Ausnahme auslösen. Erstellen Sie eine leere `__add__`-Methode innerhalb der `R2Vector`-Klasse und geben Sie ihr zwei Parameter: `self` und `other`.

# --hints--

Sie sollten eine `__add__`-Methode innerhalb der `R2Vector`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

Ihre `__add__`-Methode sollte zwei Parameter haben, `self` und `other`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
