---
id: 65fa8d5cf0ef3b141010f5d8
title: Schritt 48
challengeType: 20
dashedName: step-48
---

# --description--

Es ist Zeit zu überprüfen, ob die Additions- und Subtraktionsoperationen wie erwartet funktionieren. Deklarieren Sie eine weitere Variable `v3` und weisen Sie ihr die Summe von `v1` plus `v2` zu.

# --hints--

Sie sollten eine Variable `v3` deklarieren und ihr die Summe von `v1` und `v2` zuweisen.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_stmt("v3 = v1 + v2") or _Node(_code).has_stmt("v3 = v2 + v1")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) - getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
--fcc-editable-region--
```
