---
id: 66019977710caa516276c0a8
title: Schritt 71
challengeType: 20
dashedName: step-71
---

# --description--

Das *Kreuzprodukt* oder *Vektorprodukt* ist zwischen dreidimensionalen Vektoren definiert und ergibt einen dritten Vektor, der senkrecht zu beiden steht.

Die `R3Vector`-Klasse erbt von `R2Vector`, das bedeutet, sie kann auf alle Methoden und Eigenschaften zugreifen, die in `R2Vector` definiert sind. Eine Kindklasse kann zusätzliche Funktionen implementieren. Sie haben bereits eine Möglichkeit gesehen, die Implementierung einer Methode zu ändern. Jetzt werden Sie der Kindklasse `R3Vector` stattdessen eine neue Methode geben.

Definieren Sie innerhalb der `R3Vector`-Klasse eine `cross`-Methode und geben Sie ihr zwei Parameter: `self` und `other`.

# --hints--

Sie sollten eine `cross`-Methode innerhalb der `R3Vector`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").has_function("cross")`)) })
```

Ihre `cross`-Methode sollte zwei Parameter haben, `self` und `other`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("cross").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) - getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __mul__(self, other):
        if type(other) in (int, float):
            kwargs = {i: getattr(self, i) * other for i in vars(self)}
            return self.__class__(**kwargs)        
        elif type(self) == type(other):
            args = [getattr(self, i) * getattr(other, i) for i in vars(self)]
            return sum(args)            
        return NotImplemented

    def __eq__(self, other):
        if type(self) != type(other):
            return NotImplemented
        return all(getattr(self, i) == getattr(other, i) for i in vars(self))
        
    def __ne__(self, other):
        return not self == other

    def __lt__(self, other):
        if type(self) != type(other):
            return NotImplemented
        return self.norm() < other.norm()

    def __gt__(self, other):
        if type(self) != type(other):
            return NotImplemented
        return self.norm() > other.norm()

    def __le__(self, other):
        return not self > other

    def __ge__(self, other):
        return not self < other
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')
v5 = v1 * v2
print(f'v1 * v2 = {v5}')
```
