---
id: 66682150151af29efec9727d
title: Schritt 33
challengeType: 20
dashedName: step-33
---

# --description--

Die `__getattribute__`-Methode wird automatisch aufgerufen, sobald Sie versuchen, auf ein Instanzattribut zuzugreifen. Wenn das Attribut auf Instanzebene nicht gefunden wird, sucht die Methode auf Klassenebene und schließlich in übergeordneten Klassen danach.

Definieren Sie innerhalb Ihrer Klasse eine `__getattribute__`-Methode mit zwei Parametern, `self` und `attr`, und lassen Sie sie den String `'calling __getattribute__'` zurückgeben. Sie werden die Standardimplementierung gleich überschreiben, um zu sehen, wie der Zugriff auf Attribute funktioniert.

# --hints--

Sie sollten eine Methode namens `__getattribute__` mit zwei Parametern, `self` und `attr`, innerhalb der `R2Vector`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

Ihre `__getattribute__`-Methode sollte den String `'calling __getattribute__'` zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
