---
id: 66680ddfd0f8c76782923cb0
title: Schritt 22
challengeType: 20
dashedName: step-22
---

# --description--

Die `norm()`-Methode gibt die korrekten Werte zurück, aber es gibt noch etwas, das Sie verbessern können: die Lesbarkeit.

Die eingebaute `vars()`-Funktion nimmt ein Objekt als Argument und gibt das `__dict__`-Attribut dieses Objekts zurück.

Anstatt direkt auf das `__dict__`-Attribut von `self` zuzugreifen, ändern Sie die `norm`-Methode so, dass sie die `vars()`-Funktion verwendet.

# --hints--

Sie sollten Ihre `norm`-Methode so ändern, dass sie `vars(self)` anstelle von `self.__dict__` verwendet.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in vars(self).values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
