---
id: 68e97fe79367ad7b5dd6c9cd
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Fügen Sie in diesem Schritt ein `h1`-Element mit dem Text `XYZ Bookstore` hinzu.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
assert.exists(document.querySelector('h1'));
```

Der Text Ihres `h1`-Elements sollte `XYZ Bookstore` sein. Überprüfen Sie Rechtschreibung und Groß-/Kleinschreibung.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

Ihr `h1`-Element sollte sich innerhalb des `body`-Elements befinden.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
