---
id: 68ebdbacdd3fa474132cc975
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Fügen Sie unter dem `h1`-Element ein `p`-Element mit diesem Text hinzu: `Browse our collection of amazing books!`.

# --hints--

Sie sollten ein `p`-Element haben.

```js
assert.exists(document.querySelector('p'));
```

Der Text Ihres `p`-Elements sollte lauten: `Browse our collection of amazing books!`.

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

Ihr `p`-Element sollte unter Ihrem `h1`-Element sein.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
