---
id: 68eca11e99e3c5c894ca9d69
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Das `button`-Element wird verwendet, um anklickbare Buttons auf einer Webseite zu erstellen. Buttons sind interaktive Elemente, auf die Nutzer klicken können, um Aktionen auszuführen.

Fügen Sie ein `button`-Element innerhalb des Elements hinzu, das ein `class` mit dem Wert `card` besitzt, geben Sie dem Button ein `class`-Attribut mit dem Wert `btn` und den Text `Buy Now`.

# --hints--

Sie sollten ein `button`-Element verschachtelt innerhalb des Elements haben, das die Klasse `card` besitzt.

```js
assert.exists(document.querySelector('.card button'));
```

Ihr `button`-Element sollte ein `class`-Attribut haben.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

Ihr `button`-Element sollte ein `class` mit dem Wert `btn` besitzen.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

Der Text Ihres `button`-Elements sollte `Buy Now` sein.

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
