---
id: 68eca5c0065c82d256d29ca3
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Fügen Sie innerhalb der zweiten Karte ein `button`-Element mit dem `class`-Attribut hinzu, das auf `btn` gesetzt ist, und dem Text `Buy Now`.

Beide `button`-Elemente teilen sich nun dasselbe `class`, was bedeutet, dass sie gemeinsam einheitlich gestylt werden können.

# --hints--

Sie sollten ein `button`-Element verschachtelt innerhalb des zweiten Elements haben, das die Klasse `card` besitzt.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

Das `button`-Element Ihrer zweiten Karte sollte ein `class`-Attribut haben.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

Das `button`-Element Ihrer zweiten Karte sollte ein `class` mit dem Wert `btn` besitzen.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

Der Text des `button`-Elements Ihrer zweiten Karte sollte `Buy Now` sein.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
