---
id: 68eca6e5a6759ed4ea0034dc
title: Schritt 19
challengeType: 0
dashedName: step-19
---

# --description--

Fügen Sie innerhalb des Elements mit der Klasse `btn-container` zwei `button`-Elemente hinzu:

Erster Button:

- Id: `view-cart-btn`
- Klasse: `btn`
- Text: `View Cart`

Zweiter Button:

- Id: `checkout-btn`
- Klasse: `btn`
- Text: `Checkout`


Herzlichen Glückwunsch! Sie haben erfolgreich die Struktur einer Buchhandlungsseite mit divs, Klassen und Ids erstellt, um Ihren Inhalt zu organisieren.

# --hints--

Sie sollten zwei `button`-Elemente verschachtelt innerhalb des `.btn-container`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

Ihr erster Button innerhalb des Elements mit der Klasse `btn-container` sollte ein `id`-Attribut mit dem Wert `view-cart-btn` besitzen.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

Ihr erster Button innerhalb des Elements mit der Klasse `btn-container` sollte ein `class`-Attribut mit dem Wert `btn` besitzen.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

Der Text Ihres ersten Buttons sollte `View Cart` sein.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

Ihr zweiter Button innerhalb des Elements mit der Klasse `btn-container` sollte ein `id`-Attribut mit dem Wert `checkout-btn` besitzen.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

Ihr zweiter Button innerhalb des Elements mit `class="btn-container"` sollte ein `class`-Attribut mit dem Wert `btn` besitzen.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

Der Text Ihres zweiten Buttons sollte `Checkout` sein.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
  <p>Review your selections and continue to checkout.</p>
  
  <div class="btn-container">
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
