---
id: 68ec99e478211dc578699944
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Fügen Sie unter dem `h2`-Element im ersten Element mit der Klasse `card` ein `p`-Element mit folgendem Text hinzu:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

Sie sollten ein `p`-Element verschachtelt innerhalb des Elements mit der Klasse `card` haben.

```js
assert.exists(document.querySelector('.card p'));
```

Der Text Ihres `p`-Elements sollte `This is an epic story of Sally and her dog Rex as they navigate through other worlds.` sein.

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Ihr `p`-Element sollte unter Ihrem `h2`-Element stehen.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
