---
id: 68ec9332a9b5b2b32487bd00
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Das `id`-Attribut fügt einem HTML-Element eine eindeutige Kennung hinzu. Jedes `id` sollte innerhalb einer Seite einzigartig sein und nur einmal verwendet werden.

`id`-Werte dürfen keine Leerzeichen enthalten und sollten nur Buchstaben, Ziffern, Unterstriche und Bindestriche enthalten.

Hier ist ein Beispiel:

```html
<p id="para">example paragraph</p>
```

Fügen Sie Ihrem Element mit der Klasse `id` ein `card`-Attribut hinzu und setzen Sie dessen Wert auf `sally-adventure-book`.

# --hints--

Ihr Element mit der Klasse `card` sollte ein `id`-Attribut haben.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

Ihr Element mit der Klasse `card` sollte ein `id` mit dem Wert `sally-adventure-book` haben.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
