---
id: 68eca4c4b0f952cf09fabe09
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Fügen Sie innerhalb des zweiten Elements mit der Klasse `card` ein `h2`-Element mit dem Text `Dave's Cooking Adventure` hinzu.

# --hints--

Sie sollten ein `h2`-Element verschachtelt innerhalb des zweiten Elements haben, das die Klasse `card` besitzt.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

Das `h2`-Element innerhalb des zweiten Elements mit der Klasse `card` sollte den Text `Dave's Cooking Adventure` enthalten.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
