---
id: 68eca5412b5dd9d06b3d6404
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Fügen Sie unter dem `h2`-Element in der zweiten Karte ein `p`-Element mit diesem Text hinzu:

```md
This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.
```

# --hints--

Sie sollten ein `p`-Element verschachtelt innerhalb des zweiten Elements haben, das eine Klasse von `card` besitzt.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('p'));
```

Der Text des `p`-Elements Ihrer zweiten Karte sollte `This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.` sein.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('p')?.innerText.trim(), "This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.");
```

Das `p`-Element Ihrer zweiten Karte sollte unter Ihrem `h2`-Element stehen.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
