---
id: 68eca6364c5616d393389a30
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Denken Sie daran, ein HTML-Element sieht so aus:

```html
<element attribute="value">
    inner text
</element>
```

Fügen Sie unter dem Element mit der Klasse `card-container` ein neues `p`-Element mit diesem Text hinzu:

```md
Review your selections and continue to checkout.
```

Erstellen Sie unter dem `p`-Element ein `div`-Element mit dem `class`-Attribut, das auf `btn-container` gesetzt ist. Dieser Container gruppiert Ihre Navigationsschaltflächen.

# --hints--

Das neu hinzugefügte `p`-Element sollte unter dem Element mit der Klasse `card-container` stehen.

```js
assert.exists(document.querySelector('.card-container + p'));
```

Der Text Ihres neuen `p`-Elements sollte `Review your selections and continue to checkout.` sein.

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

Sie sollten ein `div`-Element mit einem `class` von `btn-container` haben.

```js
assert.exists(document.querySelector('div.btn-container'));
```

Ihr `div`-Element mit einem `class` von `btn-container` sollte nach dem neu hinzugefügten `p`-Element kommen.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
