---
id: 68ec98b38d83a3c28dd30efe
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Fügen Sie innerhalb des ersten Elements mit der Klasse `card` ein `h2`-Element mit dem Text `Sally's SciFi Adventure` hinzu.

# --hints--

Sie sollten ein `h2`-Element verschachtelt innerhalb des Elements mit der Klasse `card` haben.

```js
assert.exists(document.querySelector('.card h2'));
```

Der Text Ihres `h2`-Elements sollte `Sally's SciFi Adventure` sein.

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
