---
id: 68eca2a795b333ca5fee30a8
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Erstellen Sie nun eine zweite Buchkarte. Fügen Sie ein weiteres `div`-Element mit dem `class`-Attribut hinzu, das auf `card` gesetzt ist. Beachten Sie, wie Sie denselben Klassennamen für mehrere Elemente wiederverwenden können, um ein einheitliches Styling anzuwenden.

# --hints--

Sie sollten zwei `div`-Elemente mit einem `class`-Attribut haben, dessen Wert `card` ist.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

Das zweite Element mit einer Klasse von `card` sollte sich innerhalb des Elements mit einer Klasse von `card-container` befinden.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

Die beiden Elemente mit einer Klasse von `card` sollten Geschwister sein.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
