---
id: 68ec6cd5b7e8f5a8f7319e32
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Das `div`-Element wird als Container verwendet, um andere HTML-Elemente zu gruppieren. Sie werden hauptsächlich das `div`-Element verwenden, wenn Sie HTML-Elemente gruppieren möchten, die einen gemeinsamen Satz von CSS-Stilen teilen.

Fügen Sie unter dem `p`-Element ein `div`-Element hinzu. Dieses `div` wird ein Container für Ihre Buchkarten sein.

**Hinweis**: Dieser Workshop wendet kein CSS an. Klassen und gruppierte Elemente sind für das CSS-Styling nützlich, aber in diesem Workshop dienen sie nur dazu, Inhalte zu strukturieren und zu gruppieren. Wie Styling funktioniert, lernen Sie in einem späteren Modul.

# --hints--

Sie sollten ein `div`-Element haben.

```js
assert.exists(document.querySelector('div'));
```

Ihr `div`-Element sollte unter Ihrem `p`-Element sein.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
