---
id: 68ec6d9a315221aa31e54816
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Das `class`-Attribut wird verwendet, um ein oder mehrere Elemente für das Styling zu identifizieren. Im Gegensatz zum `id`-Attribut müssen Klassennamen nicht einzigartig sein: Mehrere Elemente können dieselbe Klasse teilen.

Hier ist ein Beispiel:

```html
<p class="example">example paragraph</p>
```

Fügen Sie Ihrem `div`-Element ein `class`-Attribut hinzu und setzen Sie dessen Wert auf `card-container`.

# --hints--

Ihr `div`-Element sollte ein `class`-Attribut haben.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Ihr `div`-Element sollte ein `class`-Attribut mit dem Wert `card-container` haben.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
