---
id: 69134b6cb43aa5f254950a10
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop erstellen Sie eine Buchladen-Seite, indem Sie Buchkarten anlegen, die Informationen über verschiedene Bücher anzeigen. Sie üben, Inhalte mit `div`-Elementen, Klassen und IDs zu organisieren.

Beginnen Sie Ihre Buchladen-Seite, indem Sie den HTML-Boilerplate-Code erstellen.

Fügen Sie die `<!DOCTYPE html>`-Deklaration sowie die `html`- und `head`-Elemente hinzu.

Fügen Sie dem `html`-Element ein `lang`-Attribut hinzu und setzen Sie es auf `"en"`.

# --hints--

Ihr Code sollte eine `<!DOCTYPE html>`-Deklaration enthalten.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Ihr Code sollte sowohl öffnende als auch schließende `<html>`-Tags enthalten.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

Ihr `html`-Element sollte ein `lang`-Attribut mit dem Wert `"en"` haben.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

Ihr Code sollte sowohl öffnende als auch schließende `<head>`-Tags enthalten.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Ihr `head`-Element sollte innerhalb eines `html`-Elements verschachtelt sein.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
