---
id: 69134d14150d51f8066a3246
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Verbessern Sie nun die Struktur Ihres HTML-Dokuments, um sicherzustellen, dass Ihre Seite korrekt kodiert ist.

Fügen Sie innerhalb des `head`-Elements das `<meta charset="UTF-8">`-Element hinzu.

Fügen Sie zuletzt ein `body`-Element unterhalb des `head`-Abschnitts hinzu. Hier wird der gesamte sichtbare Inhalt Ihrer Seite platziert.

# --hints--

Ihr Code sollte ein `<meta charset="UTF-8">`-Element enthalten.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Ihr Code sollte öffnende und schließende `<body>`-Tags enthalten.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Ihr `body`-Element sollte unterhalb des `head`-Elements stehen.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
