---
id: 691f85085787a8c4ca589fbb
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie das `title`-Element innerhalb des `head`-Elements ein.

Setzen Sie den Seitentitel auf `XYZ Bookstore Page`.

# --hints--

Ihr Code sollte sowohl öffnende als auch schließende `<title>`-Tags enthalten.

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

Ihr Titel sollte `XYZ Bookstore Page` sein.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

Ihr `title`-Element sollte innerhalb des `head`-Elements verschachtelt sein.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
