---
id: a105e963526e7de52b219be9
title: Sortierte Vereinigung
challengeType: 1
forumTopicId: 16077
dashedName: sorted-union
---

# --description--

Schreiben Sie eine Funktion, die zwei oder mehr Arrays entgegennimmt und ein neues Array mit eindeutigen Werten in der Reihenfolge der ursprünglich bereitgestellten Arrays zurückgibt.

Mit anderen Worten: Alle Werte aus allen Arrays sollten in ihrer ursprünglichen Reihenfolge enthalten sein, jedoch ohne Duplikate im endgültigen Array.

Die eindeutigen Zahlen sollten nach ihrer ursprünglichen Reihenfolge sortiert sein, aber das endgültige Array darf nicht numerisch sortiert sein.

Überprüfen Sie die Zusicherungstests für Beispiele.

# --hints--

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` sollte `[1, 3, 2, 5, 4]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` sollte `[1, 2, 3, 5]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` sollte `[1, 2, 3, 5, 4, 6, 7, 8]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` sollte `[1, 3, 2, 5, 4, 6]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` sollte `[1, 3, 2, 5, 4]` zurückgeben.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js
function uniteUnique(arr) {
  return arr;
}

uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]);
```

# --solutions--

```js
function uniteUnique(arr) {
  return [].slice.call(arguments).reduce(function(a, b) {
    return [].concat(
      a, 
      b.filter(function(e, currentIndex) {
        return b.indexOf(e) === currentIndex && a.indexOf(e) === -1;
      }));
  }, []);
}
```
