---
id: a2f1d72d9b908d0bd72bb9f6
title: Eine Person erstellen
challengeType: 1
forumTopicId: 16020
dashedName: make-a-person
---

# --description--

Füllen Sie den Objektkonstruktor mit den folgenden Methoden aus:

```js
getFirstName()
getLastName()
getFullName()
setFirstName(first)
setLastName(last)
setFullName(first, last)
```

Führen Sie die Tests aus, um die erwartete Ausgabe für jede Methode zu sehen. Diese Methoden müssen die einzigen verfügbaren Mittel sein, um mit dem Objekt zu interagieren. Jeder Test wird eine neue `Person`-Instanz als `new Person('Bob', 'Ross')` deklarieren.

# --hints--

Sie sollten die Funktionssignatur nicht ändern.

```js
assert.match(code, /const\s+Person\s*=\s*function\s*\(\s*first\s*,\s*last\s*\)\s*{/);
```

Sie sollten den `first`-Parameter nicht neu zuweisen.

```js
assert.notMatch(code, /\bfirst\s*=\s*/);
```

Sie sollten den `last`-Parameter nicht neu zuweisen.

```js
assert.notMatch(code, /\blast\s*=\s*/);
```

Es sollten keine Eigenschaften hinzugefügt werden. `Object.keys(Person).length` sollte immer 6 zurückgeben.

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
_person.setLastName('Curry');
_person.setFullName('John', 'Smith');
assert.lengthOf(Object.keys(_person), 6);
```

Sie sollten in der Lage sein, Ihr `Person`-Objekt zu instanziieren.

```js
const _person = new Person('Bob', 'Ross');
assert.instanceOf(_person, Person);
```

Ihr `Person`-Objekt sollte keine `firstName`-Eigenschaft haben.

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'firstName');
```

Ihr `Person`-Objekt sollte keine `lastName`-Eigenschaft haben.

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'lastName');
```

Die `.getFirstName()`-Methode sollte den String `Bob` zurückgeben.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFirstName(), 'Bob');
```

Die `.getLastName()` sollte den String `Ross` zurückgeben.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getLastName(), 'Ross');
```

Die `.getFullName()`-Methode sollte den String `Bob Ross` zurückgeben.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFullName(), 'Bob Ross');
```

Die `.getFullName()`-Methode sollte den String `Haskell Ross` zurückgeben, nachdem `.setFirstName('Haskell')` aufgerufen wurde.

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
assert.strictEqual(_person.getFullName(), 'Haskell Ross');
```

Die `.getFullName()`-Methode sollte den String `Bob Curry` zurückgeben, nachdem `.setLastName('Curry')` aufgerufen wurde.

```js
const _person = new Person('Bob', 'Ross');
_person.setLastName('Curry');
assert.strictEqual(_person.getFullName(), 'Bob Curry');
```

Die `.getFullName()`-Methode sollte den String `Haskell Curry` zurückgeben, nachdem `.setFullName('Haskell', 'Curry')` aufgerufen wurde.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFullName(), 'Haskell Curry');
```

Die `.getFirstName()`-Methode sollte den String `Haskell` zurückgeben, nachdem `.setFullName('Haskell', 'Curry')` aufgerufen wurde.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFirstName(), 'Haskell');
```

Die `.getLastName()`-Methode sollte den String `Curry` zurückgeben, nachdem `.setFullName('Haskell', 'Curry')` aufgerufen wurde.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getLastName(), 'Curry');
```

Die `.getFullName()`-Methode sollte den String `Emily Martinez de la Rosa` zurückgeben, nachdem `.setFullName('Emily Martinez', 'de la Rosa')` aufgerufen wurde.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFullName(), 'Emily Martinez de la Rosa');
```

Die `.getFirstName()`-Eigenschaft sollte den String `Emily Martinez` zurückgeben, nachdem `.setFullName('Emily Martinez', 'de la Rosa')` aufgerufen wurde.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFirstName(), 'Emily Martinez');
```

Die `.getLastName()`-Eigenschaft sollte den String `de la Rosa` zurückgeben, nachdem `.setFullName('Emily Martinez', 'de la Rosa')` aufgerufen wurde.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getLastName(), 'de la Rosa');
```

# --seed--

## --seed-contents--

```js
const Person = function(first, last) {
  this.getFullName = function() {
    return "";
  };
  return "";
};
```

# --solutions--

```js
const Person = function(first, last) {
  let firstName = first;
  let lastName = last;

  this.getFirstName = function(){
    return firstName;
  };

  this.getLastName = function(){
    return lastName;
  };

  this.getFullName = function(){
    return firstName + " " + lastName;
  };

  this.setFirstName = function(str){
    firstName = str;
  };


  this.setLastName = function(str){
    lastName = str;
  };

  this.setFullName = function(first, last){
    firstName = first;
    lastName = last;
  };
};
```
