---
id: aa7697ea2477d1316795783b
title: Pig Latin
challengeType: 1
forumTopicId: 16039
dashedName: pig-latin
---

# --description--

Pig Latin ist eine Methode, englische Wörter zu verändern. Die Regeln lauten wie folgt:

\- Beginnt ein Wort mit einem Konsonanten, nehmen Sie den ersten Konsonanten oder die erste Konsonantengruppe, verschieben Sie sie ans Ende des Wortes und fügen Sie `ay` hinzu.

\- Beginnt ein Wort mit einem Vokal, fügen Sie einfach `way` am Ende hinzu.

# --instructions--

Übersetzen Sie den gegebenen String in Pig Latin. Die Eingabestrings sind garantiert englische Wörter in Kleinbuchstaben.

# --hints--

`translatePigLatin("california")` sollte den String `aliforniacay` zurückgeben.

```js
assert.deepEqual(translatePigLatin('california'), 'aliforniacay');
```

`translatePigLatin("paragraphs")` sollte den String `aragraphspay` zurückgeben.

```js
assert.deepEqual(translatePigLatin('paragraphs'), 'aragraphspay');
```

`translatePigLatin("glove")` sollte den String `oveglay` zurückgeben.

```js
assert.deepEqual(translatePigLatin('glove'), 'oveglay');
```

`translatePigLatin("algorithm")` sollte den String `algorithmway` zurückgeben.

```js
assert.deepEqual(translatePigLatin('algorithm'), 'algorithmway');
```

`translatePigLatin("eight")` sollte den String `eightway` zurückgeben.

```js
assert.deepEqual(translatePigLatin('eight'), 'eightway');
```

Sollte Wörter behandeln, bei denen der erste Vokal in der Mitte des Wortes steht. `translatePigLatin("schwartz")` sollte den String `artzschway` zurückgeben.

```js
assert.deepEqual(translatePigLatin('schwartz'), 'artzschway');
```

Sollte Wörter ohne Vokale behandeln. `translatePigLatin("rhythm")` sollte den String `rhythmay` zurückgeben.

```js
assert.deepEqual(translatePigLatin('rhythm'), 'rhythmay');
```

# --seed--

## --seed-contents--

```js
function translatePigLatin(str) {
  return str;
}

translatePigLatin("consonant");
```

# --solutions--

```js
function translatePigLatin(str) {
  if (isVowel(str.charAt(0))) return str + "way";
  var front = [];
  str = str.split('');
  while (str.length && !isVowel(str[0])) {
    front.push(str.shift());
  }
  return [].concat(str, front).join('') + 'ay';
}

function isVowel(c) {
  return ['a', 'e', 'i', 'o', 'u'].indexOf(c.toLowerCase()) !== -1;
}
```
