---
id: 67ed03ac474c48692f41749e
title: Eine Hash-Tabelle erstellen
challengeType: 27
dashedName: build-a-hash-table
saveSubmissionToDB: true
---

# --description--

In diesem Labor bauen Sie eine Hash-Tabelle von Grund auf neu. Eine Hash-Tabelle ist eine Datenstruktur, die Schlüssel-Wert-Paare speichert. Eine Hash-Tabelle funktioniert, indem sie den Schlüssel als Eingabe nimmt und diesen Schlüssel dann gemäß einer bestimmten Hash-Funktion hasht.

Für dieses Labor wird die Hash-Funktion einfach sein: Sie summiert die Unicode-Werte jedes Zeichens im Schlüssel. Der Hash-Wert wird dann als tatsächlicher Schlüssel verwendet, um den zugehörigen Wert zu speichern. Derselbe Hash-Wert wird auch verwendet, um den Wert abzurufen und zu löschen, der mit dem Schlüssel verknüpft ist.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Klasse namens `HashTable` definieren, mit einem `collection`-Attribut, das beim Erstellen einer neuen Instanz von `HashTable` als leeres Wörterbuch initialisiert wird. Das `collection`-Wörterbuch soll Schlüssel-Wert-Paare basierend auf dem gehashten Wert des Schlüssels speichern.

2. Die `HashTable`-Klasse sollte vier Instanzmethoden haben: `hash`, `add`, `remove` und `lookup`.

3. Die `hash`-Methode sollte:

    - Einen String als Parameter nehmen.
    - Einen gehashten Wert zurückgeben, der als Summe der Unicode-(ASCII-)Werte jedes Zeichens im String berechnet wird. Sie können dafür die `ord`-Funktion verwenden.

4. Die `add`-Methode sollte:
    - Zwei Argumente annehmen, die ein Schlüssel-Wert-Paar darstellen, und den Hash des Schlüssels berechnen.
    - Den berechneten Hash-Wert als Schlüssel verwenden, um ein Wörterbuch mit dem Schlüssel-Wert-Paar im `collection`-Wörterbuch zu speichern.
    - Wenn mehrere Schlüssel denselben Hash-Wert erzeugen, sollten ihre Schlüssel-Wert-Paare im bestehenden verschachtelten Wörterbuch unter demselben Hash-Wert gespeichert werden.

5. Die `remove`-Methode sollte:

    - Einen Schlüssel als Argument nehmen und dessen Hash berechnen.
    - Überprüfen, ob der Schlüssel in der Sammlung existiert.
    - Das entsprechende Schlüssel-Wert-Paar aus der Hash-Tabelle entfernen.
    - Wenn der Schlüssel nicht in der Sammlung existiert, sollte kein Fehler ausgelöst und nichts gelöscht werden.

6. Die `lookup`-Methode sollte:

    - Einen Schlüssel als Argument nehmen.
    - Den Hash des Schlüssels berechnen und den entsprechenden Wert zurückgeben, der in der Hash-Tabelle gespeichert ist.
    - Wenn der Schlüssel nicht in der Sammlung existiert, sollte `None` zurückgegeben werden.

# --hints--

Sie sollten eine `HashTable`-Klasse definieren.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_class("HashTable")
    `)) 
})
```

Wenn eine neue Instanz der `HashTable`-Klasse erstellt wird, sollte ihr `collection`-Attribut als leeres Wörterbuch initialisiert werden.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.collection == {}
    `) 
})
```

Die `HashTable`-Klasse sollte eine `hash`-Methode haben.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('hash')`))}})
```

Die `hash`-Methode sollte einen String als Parameter nehmen.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.hash)
    len(sig.parameters) == 1
  `))
})
```

Die `hash`-Methode sollte einen String als Argument nehmen und die Summe der Unicode-Werte jedes Zeichens im String zurückgeben.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("fcc")
        assert hash_result == 300

        # prevent hardcoding
        assert ht.hash("golf") == 424
        assert ht.hash("read") == 412
`) })
```

Die `HashTable`-Klasse sollte eine `add`-Methode haben.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('add')`))}})
```

Die `add`-Methode sollte einen Schlüssel und einen Wert als Parameter nehmen.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.add)
    len(sig.parameters) == 2
  `))
})
```

Die `HashTable`-Klasse sollte eine `remove`-Methode haben.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('remove')`))}})
```

Die `remove`-Methode sollte einen Schlüssel als Parameter nehmen.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.remove)
    len(sig.parameters) == 1
  `))
})
```

Wenn Sie versuchen, einen Schlüssel zu entfernen, der nicht in der Sammlung existiert, sollte kein Fehler ausgelöst und nichts gelöscht werden.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("rose", "flower")
        index = ht.hash("rose")
        original = ht.collection.copy()

        ht.remove("tulip")
        ht.remove("sore")

        assert ht.collection == original
        assert "rose" in ht.collection[index]
        assert "tulip" not in ht.collection.get(index, {})
    `) 
})
```

Wenn mehrere Schlüssel zum selben Index hashen, sollte die `remove`-Methode nur das spezifische Schlüssel-Wert-Paar löschen und nicht das gesamte Wörterbuch an diesem Index.

```js
({
  test: () => runPython(`
    ht = HashTable()
    ht.add("rose", "flower")
    ht.add("sore", "pain")  # "rose" and "sore" both hash to the same index

    index = ht.hash("rose")
    ht.remove("rose")

    assert index in ht.collection
    assert "rose" not in ht.collection[index]
    assert "sore" in ht.collection[index]
    assert ht.collection[index]["sore"] == "pain"
  `)
})
```

Die `HashTable`-Klasse sollte eine `lookup`-Methode haben.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('lookup')`))}})
```

Die `lookup`-Methode sollte einen Schlüssel als Parameter nehmen.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.lookup)
    len(sig.parameters) == 1
  `))
})
```

`HashTable().hash('golf')` sollte `424` zurückgeben.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("golf")
        assert hash_result == 424

        # prevent hardcoding
        assert ht.hash("dear") == 412
        assert ht.hash("cat") == 312
    `) 
})
```

`HashTable().add('golf', 'sport')` sollte das Schlüssel-Wert-Paar in der Sammlung am Schlüssel `424` hinzufügen.

```js

({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")
    expected_value = {424: {'golf': 'sport'}}

    assert ht.collection == expected_value
    `) 
})
```

`HashTable().add('dear', 'friend')` und `HashTable().add('read', 'book')` sollten beide Schlüssel-Wert-Paare in der Sammlung am Index `412` als verschachteltes Wörterbuch hinzufügen.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("dear", "friend")
    ht.add("read", "book")

    expected_value = {
        "dear": "friend",
        "read": "book"
    }

    assert ht.collection.get(412) == expected_value
    `) 
})
```

Wenn ein Schlüssel in der Hash-Tabelle existiert, sollte die `remove()`-Methode den gegebenen Schlüssel und den zugehörigen Wert aus der Sammlung entfernen.

```js
({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")

    expected_value_before_removal = {
        "golf": "sport"
    }

    index = ht.hash("golf")

    assert ht.collection.get(index) == expected_value_before_removal

    ht.remove("golf")

    assert "golf" not in ht.collection.get(index, {})
    `) 
})
```

Wenn das `'golf', 'sport'`-Schlüssel-Wert-Paar in der Hash-Tabelle existiert, sollte `HashTable().lookup('golf')` `sport` zurückgeben.

```js
({ 
    test: () => runPython(`
    ht = HashTable()

    ht.add("golf", "sport")

    expected_value = "sport"
    assert ht.lookup("golf") == expected_value
    `) 
})

```

Wenn das `'golf', 'sport'`-Schlüssel-Wert-Paar nicht in der Sammlung existiert, sollte `HashTable().lookup('golf')` `None` zurückgeben.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.lookup("golf") is None
`) })

```

Wenn der `'fcc'`-Schlüssel in der Sammlung existiert, sollte `HashTable().lookup('cfc')` `None` zurückgeben.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("fcc", "coding")
        assert ht.lookup("cfc") is None
`) })
```

Wenn Sie `('rose', 'flower')` zur Hash-Tabelle hinzufügen, sollte ihr `collection`-Attribut so aussehen: `{ 441: { 'rose': 'flower' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("rose", "flower")

    expected_value = {
        "rose": "flower"
    }

    assert ht.collection.get(441) == expected_value

    # prevent hardcoding

    ht2 = HashTable()

    ht2.add("kebab", "food")
    expected_value = {
        "kebab": "food"
    }

    assert ht2.collection.get(501) == expected_value
    `)
})
```

Wenn Sie einen Schlüssel hinzufügen, der zum selben Wert hasht wie ein vorhandener Schlüssel, wie `fcc` und `cfc`, sollte das `collection` so aussehen: `{ 300: { 'fcc': 'coding', 'cfc':  'chemical' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("fcc", "coding")
    ht.add("cfc", "chemical")

    expected_value = {
        "fcc": "coding",
        "cfc": "chemical"
    }

    assert ht.collection.get(300) == expected_value

    # prevent hardcoding

    ht2 = HashTable()
    ht2.add("cat", "animal")
    ht2.add("act", "verb")

    expected_value = {
        "cat": "animal",
        "act": "verb"
    }

    assert ht2.collection.get(312) == expected_value
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class HashTable:
    def __init__(self):
        self.collection = {}

    def hash(self, string):
        hashed = 0
        for char in string:
            hashed += ord(char)  
        return hashed

    def add(self, key, val):
        the_hash = self.hash(key)
        if the_hash not in self.collection:
            self.collection[the_hash] = {}
        self.collection[the_hash][key] = val

    def remove(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            del self.collection[the_hash][key]
            if not self.collection[the_hash]: 
                del self.collection[the_hash]

    def lookup(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            return self.collection[the_hash][key]
        return None
```
