---
id: 660699119472f332798860ad
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

现在，用 `print` 调用替换 `pass`，并将字符串 `'Spam!'` 传递给它。

# --hints--

你应该删除 `pass` 并在 `spam` 方法中打印字符串 `'Spam!'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
