---
id: 6606c2d203a8124c83b2234b
title: 步骤 31
challengeType: 20
dashedName: step-31
---

# --description--

你需要检查给定的数字是否不等于当前行指定列中的数字。

为此，用一个生成器表达式替换 `pass`，该表达式在从 `0` 到 `8`（包含）的范围内进行迭代，并且对于每一行 `row`，评估板上指定 `row` 和列 `col` 处的数字是否与 `num` 不同。

# --hints--

你应该删除 `pass` 并创建一个生成器表达式 `(self.board[row][col] != num for row in range(9))` 。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").find_body().is_equivalent("(self.board[row][col] != num for row in range(9))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    def valid_in_col(self, col, num):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
