---
id: 66069e5759b800364707988e
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

回到 `__init__` 方法，它需要一个额外的参数来表示要解决的难题。

向 `__init__` 方法添加第二个名为 `board` 的参数，并通过将 `puzzle` 列表传递给 `gameboard` 来修复其实例化，就像将参数传递给函数调用一样。

# --hints--

`__init__` 方法应该有两个参数按如下顺序传入： `self` 和 `board` 。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self, board")`)) })
```

`gameboard` 变量的值应该为 `Board(puzzle)` 。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board(puzzle)")`)) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board()
--fcc-editable-region--
```
