---
id: 6606b224a69a293f98f8db8f
title: Step 17
challengeType: 20
dashedName: step-25
---

# --description--

当值未被找到时，`.index()` 方法会引起 `ValueError` 异常。 为了防止程序停止执行，你将要在 `try` 块中嵌套这行代码。 `try` 语句用于封装可能引起异常的代码。 另一方面，`except` 子句提供了在发生异常时执行的替代代码：

```py
try:
    <code>
except:
    <code>
```

将 `col` 赋值代码放到 `try` 代码块中。 然后，创建一条 `except` 子句，并在里面填充 `pass` 。

# --hints--

你应该把 `col` 赋值语句放到 `try` 代码块中，并创建 `except` 子句。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
