---
id: 6606b961ebcf04460f8af76e
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

在 `gameboard` 上调用 `find_empty_cell` 方法并打印结果来测试该方法是否正常工作。

请注意，尽管 `find_empty_cell` 定义了一个参数，但你不能通过向函数调用传递参数来赋予它值，因为 `self` 会自动作为您调用该方法的对象传入。

# --hints--

你应该打印 `gameboard.find_empty_cell()`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.find_empty_cell())")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
