---
id: 660699aabc59c532f2d556e5
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

要调用实例方法，你需要使用点符号：

```py
instance_name.method_name()
```

其中 `instance_name` 是实例或对象，`method_name` 是你想要调用的方法。

调用 `gameboard` 对象中的 `spam` 方法。

# --hints--

你应该使用 `gameboard.spam()` 调用 `gameboard` 对象中的 `spam` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
