---
id: 66069b0b36053733a2f012fe
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

现在，删除 `spam` 调用。

# --hints--

代码中不应该有 `gameboard.spam()` 。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
gameboard.spam()
--fcc-editable-region--
```
