---
id: 6606a3ccb1eea93a23c066bf
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

现在你将研究一种在数独板中查找空白单元格的方法。

在 `Board` 类中，创建一个名为 `find_empty_cell` 的空方法，并给其传递一个 `self` 参数。

# --hints--

你应该在 `Board` 类中定义一个名为 `find_empty_cell` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("find_empty_cell")`)) })
```

你的 `find_empty_cell` 方法应该有一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
