---
id: 6606a219f9efbf38ad496f67
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

你也可以使用点符号访问实例属性。

在 `Board` 类外，在 `gameboard` 对象初始化后，使用 `gameboard.board` 来访问`gameboard` 对象中的 `board` 属性，并将结果输出在屏幕。

# --hints--

你应该输出 `gameboard.board` 。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.board)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
