---
id: 6606bf4561f8794a0d345919
title: 步骤 27
challengeType: 20
dashedName: step-27
---

# --description--

现在是测试 `valid_in_row` 方法的时候了。 在 `gameboard` 上调用 `valid_in_row` 。 将 `0` 和 `8` 作为参数并打印结果。

再次注意该方法是如何用三个参数定义的，但它只用两个参数调用，因为 `self` 被自动传递为调用该方法的对象。

# --hints--

你应该在 `gameboard` 上打印调用 `valid_in_row(0, 8)` 的结果。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.valid_in_row(0, 8))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
