---
id: 66069b992c1c5e3451f3deb0
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

实例化创建了一个空对象。 `__init__` 方法允许你实例化一个自定义的对象。 当一个类实现 `__init__` 方法时，实例化时会自动调用 `__init__`。

在你的 `Board` 类中，删除 `spam` 方法并将其替换为包含 `self` 参数的 `__init__` 方法。

# --hints--

在 `Board` 类中不应该存在 `spam` 方法。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

你应该在 `Board` 类中定义一个 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

`__init__` 方法应该接收一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
