---
id: 6606927d010be4300a4e5330
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

实例化创建了一个空对象。 但是类可以有方法，它们就像每个实例的本地函数一样。 在类中，方法声明如下：

```py
class ClassName:
    def method_name():
        pass
```

在 `Board` 类中，用空方法 `spam` 替换 `pass`。

# --hints--

你应该在 `Board` 类中定义一个名为 `spam` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass

gameboard = Board()
--fcc-editable-region--
```
