---
id: 6606a4641ec48b3a9fe8c2fc
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

`enumerate` 内置函数以可迭代对象作为参数，并返回可迭代的枚举对象。 它提供了计数(默认始于零)和可迭代的值。

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

上面示例中的循环将输出元组 `0, a` 、`1, b` and `2, c` 。

在 `find_empty_cell` 方法中，用 `for` 循环替换 `pass`，该循环使用 `enumerate()` 函数遍历数独板中的每一行。 使用 `row` 作为当前行的索引，使用 `contents` 作为当前行的元素。


# --hints--

你应该将 `pass` 替换为遍历 `enumerate(self.board)` 的 `for` 循环。 使用 `row` 和 `contents` 作为循环变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
