---
id: 66069f86f58f85371d47123e
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

属性是与对象关联的变量，用于将数据存储为常规变量。

在 `__init__` 方法中，使用 `self.board` 将 `board` 参数（在创建 `Board` 类的实例时传递）分配给实例属性 `board`。

`self.board` 指的是类实例的 `board` 属性。 它是一个属于从 `Board` 类创建的对象的变量。

# --hints--

你应该删除 `pass`，并将 `board` 参数分配给 `__init__` 方法内的 `self.board`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
