---
id: 6606933d6813a8308c962dd1
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

为了成为实例方法，方法需要一个特殊的参数，按照惯例名为 `self`。 此参数是对该类实例的引用，并且必须始终是第一个参数。

添加 `self` 方法到 `spam` 方法。

# --hints--

`spam` 方法应该有一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
