---
id: 6606bc4e5535c0484990ccd5
title: 步骤 24
challengeType: 20
dashedName: step-24
---

# --description--

接下来，你将设计一种方法，检查是否可以将给定的数字插入到数独板的指定行中。

在 `Board` 类中，创建一个名为 `valid_in_row` 的方法，并给它三个参数：`self`、`row` 和 `num` 。 其中 `self` 表示类的实例，`row` 和 `num` 分别是行索引和要检查的数量。


# --hints--

你应该在 `Board` 类中创建一个名为 `valid_in_row` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

`valid_in_row` 方法应该有三个参数：`self`、`row` 和 `num`，按此顺序。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
