---
id: 66068fb0bfddba2b7977eb60
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在这个项目中，你将通过构建数独谜题解答器来学习类和对象。

在 Python 中，类是创建对象的模板。 从类创建的对象是该类的实例。 你可以使用此语法创建一个类：

```py
class ClassName:
    pass
```

`class` 关键字用于定义类，`ClassName` 是类的名字，通常以 *PascalCase* 格式命名。

首先创建一个 `Board` 类。

# --hints--

你应该创建一个 `Board` 类。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
