---
id: 6606beade9200b49aaeecd94
title: 步骤 26
challengeType: 20
dashedName: step-26
---

# --description--

如果 `num` 不在该行中，则表达式计算结果为 `True`，表示该数字可以有效插入。

如果 `num` 在该行中，则表达式计算结果为 `False`，并且插入违反规则。

在 `valid_in_row` 方法体内的表达式前面添加 `return` 关键字，以便检查数字的有效性。

# --hints--

你应该返回 `num not in self.board[row]` 。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").find_body().is_equivalent("return num not in self.board[row]")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    def valid_in_row(self, row, num):
        num not in self.board[row]
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
