---
id: 66069167b3307b2f4067b22b
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

使用函数符号来创建类的新实例，即在类名后附加一对括号。

在类的定义外，创建类 `Board` 的实例并将其赋值给 `gameboard` 变量。

# --hints--

你应该在 `Board` 类外声明一个 `gameboard` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

`gameboard` 变量的值应该为 `Board()` 。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
