---
id: 6606c14182435d4bab0de2ee
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

接下来，你将创建一种方法，通过检查数字是否尚未出现在该列中来检查是否可以将数字插入数独板的指定列中。

在 `Board` 类中，创建一个名为 `valid_in_col` 的方法，并给它三个参数：`self`、`col` 和 `num` 。 `col` 和 `num` 分别是要检查的列索引和数字。

# --hints--

你应该在 `Board` 类中创建一个名为 `valid_in_col` 的新方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

你的 `valid_in_col` 方法应该有三个参数： `self`、 `col` 和 `num`，按此顺序排列。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
