---
id: 69708d59833ceded7d00541f
title: 연결 리스트 연산 구현
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

이 실습에서는 노드를 추가하고 제거하는 기본 기능을 바탕으로 연결 리스트 자료 구조에 대한 추가 연산을 구현할 것입니다.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리**

1. 연결 리스트와 요소를 인수로 받는 `contains` 함수를 구현하세요. 지정한 요소가 연결 리스트에 존재하면 `true`을, 그렇지 않으면 `false`를 반환해야 합니다.
1. 연결 리스트와 인덱스를 인수로 받는 `getAt` 함수를 구현하세요. 연결 리스트에서 주어진 인덱스에 있는 요소를 반환해야 합니다. 인덱스가 범위를 벗어나면 `undefined`를 반환해야 합니다.
1. 연결 리스트, 인덱스, 요소를 인수로 받는 `insertAt` 함수를 구현하세요. 지정한 위치에 요소를 삽입해야 합니다. 인덱스가 범위를 벗어나면 리스트를 수정하지 않아야 합니다.
1. 연결 리스트와 인덱스를 인수로 받는 `removeAt` 함수를 구현하세요. 주어진 인덱스에 있는 노드를 제거해야 합니다. 인덱스가 범위를 벗어나면 리스트를 수정하지 않아야 합니다.
1. 연결 리스트를 인수로 받는 `clear` 함수를 구현하세요. 연결 리스트의 모든 요소를 제거하여 빈 상태로 초기화해야 합니다.

**참고:** 이후 테스트는 이전 메서드에 의존하는 경우가 있습니다. 예를 들어 `getAt`이 올바르게 구현되지 않으면 `insertAt`와 `removeAt` 같은 함수의 테스트가 거의 맞더라도 실패할 수 있습니다.

# --hints--

`initList`은 `head`가 `null`이고 `length`가 `0`인 객체를 반환해야 합니다.

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

`add`은 리스트 길이를 1만큼 증가시켜야 합니다.

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

`remove`은 리스트 길이를 1만큼 감소시켜야 합니다.

```js
const list = initList();  
add(list, 'apple');  
add(list, 'banana');  
remove(list, 'apple');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(list.head.element, 'banana');  
assert.isNull(list.head.next);
```

`contains` 함수를 구현해야 합니다.

```js
assert.isFunction(contains);
```

`contains`는 요소가 리스트에 존재하면 `true`를 반환해야 합니다.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains`는 요소가 리스트에 없으면 `false`를 반환해야 합니다.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains`는 빈 리스트에 대해 `false`를 반환해야 합니다.

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

`getAt` 함수를 구현해야 합니다.

```js
assert.isFunction(getAt);
```

`getAt`는 주어진 인덱스의 요소를 반환해야 합니다.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt`는 잘못된 인덱스에 대해 `undefined`를 반환해야 합니다.

```js
const list = initList();  
add(list, 'apple');  
assert.isUndefined(getAt(list, -1));  
assert.isUndefined(getAt(list, 1));  
assert.isUndefined(getAt(list, 5));
```

`getAt`는 빈 리스트에 대해 `undefined`를 반환해야 합니다.

```js
const list = initList();  
assert.isUndefined(getAt(list, 0));
```

`insertAt` 함수를 구현해야 합니다.

```js
assert.isFunction(insertAt);
```

`insertAt`는 리스트 시작 부분에 요소를 삽입해야 합니다.

```js
const list = initList();  
add(list, 'second');  
add(list, 'third');  
insertAt(list, 0, 'first');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt`는 리스트 중간에 요소를 삽입해야 합니다.

```js
const list = initList();  
add(list, 'first');  
add(list, 'third');  
insertAt(list, 1, 'second');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt`는 리스트 끝에 요소를 삽입해야 합니다.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt`는 빈 리스트에 인덱스 `0`에서 삽입해야 합니다.

```js  
const list = initList();  
insertAt(list, 0, 'first');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.isNull(list.head.next); 
```

`insertAt`는 잘못된 인덱스에 대해 리스트를 수정하지 않아야 합니다.

```js
const list = initList();  
add(list, 'apple');  
insertAt(list, -1, 'invalid');  
insertAt(list, 10, 'invalid');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
assert.isUndefined(getAt(list, 1));
```

`removeAt` 함수를 구현해야 합니다.

```js
assert.isFunction(removeAt);
```

`removeAt`는 지정한 인덱스의 요소를 제거해야 합니다.

```js
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'third');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt`는 첫 번째 요소를 올바르게 제거해야 합니다.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

`removeAt`는 마지막 요소를 올바르게 제거해야 합니다.

```js  
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 2);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt`는 유일한 요소를 제거할 때 리스트를 초기화해야 합니다.

```js
const list = initList();  
add(list, 'only');  
removeAt(list, 0);  
assert.strictEqual(list.length, 0);  
assert.isNull(list.head);  
assert.isTrue(isEmpty(list));
```

`removeAt`는 잘못된 인덱스에 대해 리스트를 수정하지 않아야 합니다.

```js
const list = initList();  
add(list, 'apple');  
removeAt(list, -1);  
removeAt(list, 10);  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
```

`clear` 함수를 구현해야 합니다.

```js
assert.isFunction(clear);
```

`clear`는 리스트의 모든 요소를 제거해야 합니다.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`clear`는 리스트를 재사용할 수 있게 해야 합니다.

```js  
const list = initList();  
add(list, 'first');  
clear(list);  
clear(list);  
add(list, 'second');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'second');  
assert.isNull(list.head.next); 
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  
}

function getAt(list, index) {
 
}

function insertAt(list, index, element) {
  
}

function removeAt(list, index) {
 
}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
