---
id: 6735a7370e0ae93a4577c689
title: Zeitungslayout entwerfen
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

In diesem Labor entwerfen Sie ein Zeitungslayout mit CSS Grid, einschließlich Konzepten wie Grid-Zeilen und Grid-Spalten.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Ihre Seite sollte ein `main`-Element mit der Klasse `newspaper-layout` enthalten.
1. Das `.newspaper-layout` sollte ein `header` mit der Klasse `title` enthalten, das ein `h1` zur Anzeige des Zeitungsnamens enthält.
1. Das `.newspaper-layout` sollte ein `article` mit der Klasse `feature-article` für den Hauptnachrichtenartikel enthalten.
1. Das `.feature-article` sollte ein `h2`-Element für den Artikeltitel enthalten, gefolgt von einem `p`-Element für den Artikelinhalt.
1. Sie sollten drei weitere `article`-Elemente für kleinere Artikel hinzufügen, mit den Klassen `small-article1`, `small-article2` und `small-article3`.
1. Jeder der kleineren Artikel sollte ein `h3`-Element für den Artikeltitel enthalten, gefolgt von einem `p`-Element für den Artikelinhalt.
1. Das `.newspaper-layout` sollte ein `article`-Element mit der Klasse `secondary-article` für einen zusätzlichen Nachrichtenabschnitt enthalten.
1. Das `.secondary-article` sollte ein `h2`-Element für den Artikeltitel enthalten, gefolgt von einem `p`-Element für den Artikelinhalt.
1. Das `.newspaper-layout` sollte ein `figure` mit der Klasse `cover-image` enthalten, um ein Bild darzustellen, das den Inhalt der Zeitung repräsentiert.
1. Die Elemente mit den Klassen `title`, `feature-article`, `secondary-article`, `cover-image`, `small-article1`, `small-article2` und `small-article3` sollten eine `grid-area`-Eigenschaft mit demselben Klassennamen haben.
1. Ihr `.newspaper-layout` sollte CSS Grid verwenden mit einer `grid-template-areas`-Eigenschaft, um die Anordnung der Abschnitte zu definieren:
   - Das `.title` sollte sich über die obere Zeile erstrecken.
   - Das `.feature-article` sollte sich in der zweiten Zeile über zwei Spalten erstrecken, mit dem `.cover-image` in der dritten Spalte.
   - Das `.secondary-article` sollte sich in der dritten Zeile über zwei Spalten erstrecken, mit dem `.cover-image` in der dritten Spalte.
   - Die drei kleinen Artikel sollten in der vierten Zeile erscheinen.
1. Das `.newspaper-layout` sollte eine `grid-template-columns`-Eigenschaft haben, die den Raum in drei gleiche Spalten unterteilt.
1. Sie sollten die `.newspaper-layout`-`grid-template-rows`-Eigenschaft auf `auto` für die erste Zeile setzen und den verbleibenden Raum für die anderen Zeilen in gleiche Teile aufteilen.
1. Sie sollten einen Abstand zwischen den Grid-Elementen hinzufügen.
1. Stellen Sie sicher, dass das Bild innerhalb von `.cover-image` in den vorgesehenen Raum passt, indem Sie seine `max-width`-Eigenschaft auf `100%` setzen.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verlinken und Ihr CSS anwenden.

# --hints--

Sie sollten ein `main`-Element haben.

```js
assert.exists(document.querySelector('main'));
```

Ihr `main`-Element sollte eine `class` von `newspaper-layout` haben.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

Sie sollten ein `header`-Element mit der `class` von `title` innerhalb Ihres `.newspaper-layout`-Elements haben.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

Ihr `header`-Element sollte ein `h1`-Element enthalten.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

Ihr `.title`-Element sollte Text enthalten.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

Sie sollten für jeden Artikel ein `article`-Element mit den Klassen `feature-article`, `small-article1`, `small-article2`, `small-article3` und `secondary-article` haben.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

Das `.feature-article` sollte ein `h2`-Element gefolgt von einem `p`-Element enthalten.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`, `.small-article2` und `.small-article3` sollten jeweils ein `h3`-Element gefolgt von einem `p`-Element enthalten.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

Das `.secondary-article` sollte ein `h2`-Element gefolgt von einem `p`-Element enthalten.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

Sie sollten ein `figure`-Element mit der Klasse `cover-image` innerhalb Ihres `.newspaper-layout` haben.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

Ihr `.cover-image` sollte ein `img`-Element enthalten.

```js
assert.exists(document.querySelector(".cover-image img"));
```

Die `display`-Eigenschaft des `.newspaper-layout`-Elements sollte auf `grid` gesetzt sein.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

Ihr `.newspaper-layout`-Element sollte eine `grid-template-areas`-Eigenschaft mit `title` haben, die alle drei Spalten der ersten Zeile überspannt.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

Die zweite Zeile der Grid-Vorlage sollte ein `feature-article` enthalten, das zwei Spalten überspannt, und ein `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

Die dritte Zeile der Grid-Vorlage sollte ein `secondary-article` enthalten, das zwei Spalten überspannt, und ein `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

Die vierte Zeile der Grid-Vorlage sollte `small-article1`, `small-article2` und `small-article3` enthalten.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

Ihr `.newspaper-layout`-Selektor sollte die `grid-template-columns`-Eigenschaft verwenden, um den Raum in drei gleiche Teile zu unterteilen. Denken Sie daran, dass Sie entweder die `repeat`-Funktion verwenden oder jede Bruch-Einheit manuell definieren können.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Ihr `.newspaper-layout`-Selektor sollte eine `grid-template-rows`-Eigenschaft haben, die auf `auto 1fr 1fr 1fr` gesetzt ist.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

Sie sollten einen Abstand zwischen den Grid-Elementen hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` sollte eine `grid-area` von `title` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` sollte eine `grid-area` von `feature-article` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` sollte eine `grid-area` von `cover-image` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` sollte eine `grid-area` von `secondary-article` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` sollte eine `grid-area` von `small-article1` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` sollte eine `grid-area` von `small-article2` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` sollte eine `grid-area` von `small-article3` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

Das `img` innerhalb von `.cover-image` sollte eine `max-width` von `100%` haben.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
