---
id: 67ed03ac474c48692f41749e
title: Jenga jedwali la hash
challengeType: 27
dashedName: build-a-hash-table
saveSubmissionToDB: true
---

# --description--

Katika maabara hii, u-tajenga jedwali la hash kutoka mwanzo. Jedwali la hash ni muundo wa data unaohifadhi jozi za ufunguo-thamani. Jedwali la hash hufanya kazi kwa kuchukua ufunguo kama ingizo kisha kuufanyia hashing kulingana na kitendakazi maalum cha hashing.

Kwa madhumuni ya maabara hii, kitendakazi cha hashing kitakuwa rahisi: kitajumlisha thamani za Unicode za kila herufi katika ufunguo. Thamani ya hash itatumika kama ufunguo halisi kuhifadhi thamani inayohusiana. Thamani hiyo hiyo ya hash itatumika pia kupata na kufuta thamani inayohusiana na ufunguo.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kufafanua darasa la vitu liitwalo `HashTable` lenye sifa ya kipengele `collection` iliyowekwa kuwa kamusi tupu wakati mfano mpya wa `HashTable` unapotengenezwa. Kamusi ya `collection` inapaswa kuhifadhi jozi za ufunguo-thamani kulingana na thamani ya hash ya ufunguo.

2. Darasa la `HashTable` linapaswa kuwa na njia nne za mfano: `hash`, `add`, `remove`, na `lookup`.

3. Njia ya `hash` inapaswa:

    - Kuchukua mfuatano wa herufi kama kigezo.
    - Kurudisha thamani ya hash iliyohesabiwa kama jumla ya thamani za Unicode (ASCII) za kila herufi katika mfuatano wa herufi. Unaweza kutumia kitendakazi cha `ord` kwa hesabu hii.

4. Njia ya `add` inapaswa:
    - Kuchukua hoja mbili zinazo wakilisha jozi ya ufunguo-thamani, na kuhesabu hash ya ufunguo.
    - Kutumia thamani ya hash iliyohesabiwa kama ufunguo kuhifadhi kamusi inayojumuisha jozi ya ufunguo-thamani ndani ya kamusi ya `collection`.
    - Ikiwa funguo nyingi zitatoa thamani sawa ya hash, jozi zao za ufunguo-thamani zinapaswa kuhifadhiwa katika kamusi zilizopangwa ndani zilizopo chini ya thamani hiyo ya hash.

5. Njia ya `remove` inapaswa:

    - Kuchukua ufunguo kama hoja na kuhesabu hash yake.
    - Kuhakiki kama ufunguo upo katika mkusanyiko.
    - Kufuta jozi ya ufunguo-thamani inayolingana kutoka kwenye jedwali la hash.
    - Ikiwa ufunguo haupo katika mkusanyiko, haipaswi kutoa kosa wala kufuta kitu chochote.

6. Njia ya `lookup` inapaswa:

    - Kuchukua ufunguo kama hoja.
    - Kuhesabu hash ya ufunguo, na kurudisha thamani inayohifadhiwa ndani ya jedwali la hash.
    - Ikiwa ufunguo haupo katika mkusanyiko, inapaswa kurudisha `None`.

# --hints--

Unapaswa kufafanua darasa la `HashTable`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_class("HashTable")
    `)) 
})
```

Wakati mfano mpya wa darasa la `HashTable` unapotengenezwa, sifa yake ya kipengele `collection` inapaswa kuwekwa kuwa kamusi tupu.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.collection == {}
    `) 
})
```

Darasa la `HashTable` linapaswa kuwa na njia ya `hash`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('hash')`))}})
```

Njia ya `hash` inapaswa kuchukua mfuatano wa herufi kama kigezo.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.hash)
    len(sig.parameters) == 1
  `))
})
```

Njia ya `hash` inapaswa kuchukua mfuatano wa herufi kama hoja na kurudisha jumla ya thamani za Unicode za kila herufi katika mfuatano huo.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("fcc")
        assert hash_result == 300

        # prevent hardcoding
        assert ht.hash("golf") == 424
        assert ht.hash("read") == 412
`) })
```

Darasa la `HashTable` linapaswa kuwa na njia ya `add`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('add')`))}})
```

Njia ya `add` inapaswa kuchukua ufunguo na thamani kama vigezo.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.add)
    len(sig.parameters) == 2
  `))
})
```

Darasa la `HashTable` linapaswa kuwa na njia ya `remove`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('remove')`))}})
```

Njia ya `remove` inapaswa kuchukua ufunguo kama kigezo.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.remove)
    len(sig.parameters) == 1
  `))
})
```

Unapojaribu kufuta ufunguo usiopo katika mkusanyiko, haipaswi kutoa kosa wala kufuta kitu chochote.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("rose", "flower")
        index = ht.hash("rose")
        original = ht.collection.copy()

        ht.remove("tulip")
        ht.remove("sore")

        assert ht.collection == original
        assert "rose" in ht.collection[index]
        assert "tulip" not in ht.collection.get(index, {})
    `) 
})
```

Ikiwa funguo nyingi zina thamani sawa ya hash, njia ya `remove` inapaswa kufuta jozi ya ufunguo-thamani husika tu na si kamusi nzima katika kielezo hicho.

```js
({
  test: () => runPython(`
    ht = HashTable()
    ht.add("rose", "flower")
    ht.add("sore", "pain")  # "rose" and "sore" both hash to the same index

    index = ht.hash("rose")
    ht.remove("rose")

    assert index in ht.collection
    assert "rose" not in ht.collection[index]
    assert "sore" in ht.collection[index]
    assert ht.collection[index]["sore"] == "pain"
  `)
})
```

Darasa la `HashTable` linapaswa kuwa na njia ya `lookup`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('lookup')`))}})
```

Njia ya `lookup` inapaswa kuchukua ufunguo kama kigezo.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.lookup)
    len(sig.parameters) == 1
  `))
})
```

`HashTable().hash('golf')` inapaswa kurudisha `424`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("golf")
        assert hash_result == 424

        # prevent hardcoding
        assert ht.hash("dear") == 412
        assert ht.hash("cat") == 312
    `) 
})
```

`HashTable().add('golf', 'sport')` inapaswa kuongeza jozi ya ufunguo-thamani katika mkusanyiko kwa ufunguo `424`.

```js

({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")
    expected_value = {424: {'golf': 'sport'}}

    assert ht.collection == expected_value
    `) 
})
```

`HashTable().add('dear', 'friend')` na `HashTable().add('read', 'book')` zinapaswa kuongeza jozi zote za ufunguo-thamani katika mkusanyiko kwa kielezo `412` kama kamusi zilizopangwa ndani.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("dear", "friend")
    ht.add("read", "book")

    expected_value = {
        "dear": "friend",
        "read": "book"
    }

    assert ht.collection.get(412) == expected_value
    `) 
})
```

Wakati ufunguo upo katika jedwali la hash, njia ya `remove()` inapaswa kufuta ufunguo uliotolewa na thamani yake inayolingana kutoka kwenye mkusanyiko.

```js
({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")

    expected_value_before_removal = {
        "golf": "sport"
    }

    index = ht.hash("golf")

    assert ht.collection.get(index) == expected_value_before_removal

    ht.remove("golf")

    assert "golf" not in ht.collection.get(index, {})
    `) 
})
```

Wakati jozi ya ufunguo-thamani ya `'golf', 'sport'` ipo katika jedwali la hash, `HashTable().lookup('golf')` inapaswa kurudisha `sport`.

```js
({ 
    test: () => runPython(`
    ht = HashTable()

    ht.add("golf", "sport")

    expected_value = "sport"
    assert ht.lookup("golf") == expected_value
    `) 
})

```

Wakati jozi ya ufunguo-thamani ya `'golf', 'sport'` haipo katika mkusanyiko, `HashTable().lookup('golf')` inapaswa kurudisha `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.lookup("golf") is None
`) })

```

Wakati ufunguo wa `'fcc'` upo katika mkusanyiko, `HashTable().lookup('cfc')` inapaswa kurudisha `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("fcc", "coding")
        assert ht.lookup("cfc") is None
`) })
```

Unapoongeza `('rose', 'flower')` kwenye jedwali la hash, sifa yake ya kipengele `collection` inapaswa kuonekana hivi: `{ 441: { 'rose': 'flower' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("rose", "flower")

    expected_value = {
        "rose": "flower"
    }

    assert ht.collection.get(441) == expected_value

    # prevent hardcoding

    ht2 = HashTable()

    ht2.add("kebab", "food")
    expected_value = {
        "kebab": "food"
    }

    assert ht2.collection.get(501) == expected_value
    `)
})
```

Unapoongeza ufunguo unaotoa thamani sawa ya hash na ufunguo uliopo, kama `fcc` na `cfc`, `collection` inapaswa kuonekana hivi: `{ 300: { 'fcc': 'coding', 'cfc':  'chemical' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("fcc", "coding")
    ht.add("cfc", "chemical")

    expected_value = {
        "fcc": "coding",
        "cfc": "chemical"
    }

    assert ht.collection.get(300) == expected_value

    # prevent hardcoding

    ht2 = HashTable()
    ht2.add("cat", "animal")
    ht2.add("act", "verb")

    expected_value = {
        "cat": "animal",
        "act": "verb"
    }

    assert ht2.collection.get(312) == expected_value
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class HashTable:
    def __init__(self):
        self.collection = {}

    def hash(self, string):
        hashed = 0
        for char in string:
            hashed += ord(char)  
        return hashed

    def add(self, key, val):
        the_hash = self.hash(key)
        if the_hash not in self.collection:
            self.collection[the_hash] = {}
        self.collection[the_hash][key] = val

    def remove(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            del self.collection[the_hash][key]
            if not self.collection[the_hash]: 
                del self.collection[the_hash]

    def lookup(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            return self.collection[the_hash][key]
        return None
```
