---
id: ac6993d51946422351508a41
title: Tekeleza algoriti ya kukata mfuatano wa herufi
challengeType: 26
dashedName: truncate-a-string
---

# --description--

Katika maabara hii, utazoeza kukata mfuatano wa herufi hadi urefu fulani.

**Lengo:** Timiza hadithi za mtumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za mtumizi:**

1. Unapaswa kuwa na kitendakazi `truncateString` kinachokubali hoja mbili, ya kwanza ni mfuatano wa herufi, ya pili ni nambari.
1. Ikiwa urefu wa mfuatano wa herufi ni zaidi ya nambari iliyotolewa, mfuatano huo unapaswa kukatwa ili kupunguza urefu wake hadi kufanana na nambari hiyo, na `...` inapaswa kuongezwa mwishoni mwa mfuatano uliokatwa.
1. Ikiwa urefu wa mfuatano wa herufi ni sawa au chini ya nambari iliyotolewa, mfuatano huo unapaswa kurudishwa bila kubadilika.

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` inapaswa kurudisha mfuatano `A-tisket...`.

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` inapaswa kurudisha mfuatano `Peter Piper...`.

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` inapaswa kurudisha mfuatano `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` inapaswa kurudisha mfuatano `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` inapaswa kurudisha mfuatano `A...`.

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` inapaswa kurudisha mfuatano `Ab...`.

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
