---
id: 683ec95f2f8781355556ff7a
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

خوارزمية البحث الثنائي التي تبنيها لن تُرجع فقط `True` أو `False` لتبيان ما إذا وُجدت القيمة أم لا. بل سترجع أيضًا المسار الذي تتبعه للعثور على تلك القيمة. ولهذا، سيكون لديك متغير `path_to_target` لتتبع المسار.

داخل دالتك `binary_search`، استبدل الكلمة المفتاحية `pass` بمتغير `path_to_target` وابدأ بتهيئته إلى قائمة فارغة.

# --hints--

يجب أن تنشئ متغيرًا باسم `path_to_target` وتضبطه على قائمة فارغة داخل دالة `binary_search`. لا تنسَ إزالة الكلمة المفتاحية `pass`.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
