---
id: 683ec95f2f8781355556ff83
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

يمكنك أن ترى أن الاستدعاء الأول يُرجع `[3]`. ذلك لأن `3` هو القيمة الوسطى في `[1, 2, 3, 4, 5]`.

من ناحية أخرى، الاستدعاء الثاني يُرجع قائمة فارغة. يحدث هذا لأن الحلقة بعد التحقق من القيمة الوسطى تفتقر حاليًا إلى المنطق الذي يضيق نطاق البحث بتحديث المتغيرين `low` أو `high`. هي فقط تقوم بفحص نقطة الوسط مرة واحدة ثم تنتهي دون متابعة البحث.

للسماح لخوارزمية البحث الثنائي بمواصلة تضييق نطاق البحث، أضف كتلة `elif` تتحقق مما إذا كانت `value` أكبر من `value_at_middle`. أضف كلمة `pass` داخل `elif` في الوقت الحالي.

# --hints--

يجب أن تضيف كتلة `elif` عندما تكون `value` أكبر من `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        
--fcc-editable-region--
        break
    
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
