---
id: 683ec95f2f8781355556ff80
title: خطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

الخطوة التالية هي إنشاء شرط يتحقق مما إذا كانت القيمة المستهدفة `value` في المنتصف.

داخل حلقة `while`، أنشئ جملة `if` تتحقق مما إذا كانت القيمة المستهدفة `value` تساوي `value_at_middle`. إذا كانت كذلك، أرجع المتغير `path_to_target`.

# --hints--

يجب أن تنشئ جملة `if` تتحقق مما إذا كانت القيمة `value` تساوي `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

يجب أن تُرجع المتغير `path_to_target` داخل جملة `if`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
--fcc-editable-region--
        
--fcc-editable-region--
```
