---
id: 683ec95f2f8781355556ff7b
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

في البحث الثنائي، يوجد أدنى فهرس ممكن وأعلى فهرس ممكن.

يمثل أدنى فهرس ممكن الحد الأيسر في الجزء الحالي من القيم التي يُبحث فيها، في هذه الحالة `search_list`. أما أعلى فهرس ممكن، فيمثل الحد الأيمن في القيم التي يُبحث فيها.

لأخذ هذين الحدين في الاعتبار، يجب أن تراعي القائمة الكاملة التي يُبحث فيها. لذا، داخل دالة `binary_search`، عرّف متغيرًا باسم `low` بقيمة `0`، ومتغيرًا باسم `high` بقيمة تمثل آخر فهرس في القائمة التي يُبحث فيها.

# --hints--

يجب أن يكون لديك متغير باسم `low` في دالة `binary_search` الخاصة بك.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

يجب تعيين متغير `low` الخاص بك إلى القيمة `0`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

يجب أن يكون لديك متغير باسم `high` في دالة `binary_search` الخاصة بك.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

يجب تعيين متغير `high` الخاص بك إلى `len(search_list) - 1`.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
--fcc-editable-region--
    
--fcc-editable-region--
```
