---
id: 683ec95f2f8781355556ff7f
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

أضف الآن `value_at_middle` إلى قائمة `path_to_target` لتتبع الخطوات التي اتُخذت أثناء البحث، بغض النظر عن كونها الهدف أم لا.

# --hints--

يجب أن تستخدم طريقة `append()` لإضافة `value_at_middle` إلى قائمة `path_to_target`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid] \\npath_to_target.append(value_at_middle)")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
--fcc-editable-region--
        
--fcc-editable-region--
```
