---
id: 683ec95f2f8781355556ff82
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

تذكّر أن الدالة تأخذ قائمة تسمى `search_list` والقيمة `value` التي تبحث عنها.

الآن، استدعِ الدالة مع `binary_search([1, 2, 3, 4, 5], 3)` و `binary_search([1, 2, 3, 4, 5, 9], 4)` واطبع الاستدعاءات فورًا.

# --hints--

يجب أن تستدعي الدالة مع `binary_search([1, 2, 3, 4, 5], 3)` وتطبعها فورًا.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

يجب أن تستدعي الدالة مع `binary_search([1, 2, 3, 4, 5, 9], 4)` وتطبعها فورًا.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
        break
    
    return []

--fcc-editable-region--

--fcc-editable-region--
```
