---
id: 683ec95f2f8781355556ff81
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

تحتاج إلى اختبار الأمور لتفهم تدفق الخوارزمية في هذه المرحلة الأولية.

للقيام بذلك، يجب أن تخرج أولاً من الحلقة. ذلك لأن التنفيذ الحالي يسمح بتكرار واحد فقط، فإذا لم يتحقق الشرط، ستحدث حلقة لا نهائية.

بعد جملة `if` مباشرة، استخدم الكلمة المفتاحية `break` للخروج من حلقة `while`. ثم، بعد حلقة `while`، أعد قائمة فارغة للدلالة على أن القيمة لم تُعثر عليها.

# --hints--

يجب أن تستخدم الكلمة المفتاحية `break` للخروج من حلقة `while` بعد جملة `if`.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

يجب أن تعيد قائمة فارغة بعد حلقة `while` الخاصة بك.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        
--fcc-editable-region--
```
