---
id: 683ec95f2f8781355556ff7c
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

الآن بعد أن عرّفت حدود البحث، حان الوقت لإنشاء الحلقة التي ستنفذ خوارزمية البحث الثنائي.

يعمل البحث الثنائي عن طريق تضييق مساحة البحث بشكل متكرر. تستمر هذه العملية طالما هناك نطاق صالح من العناصر للفحص.

يمكنك التعبير عن ذلك باستخدام حلقة `while` تستمر طالما أن مؤشر `low` أقل من أو يساوي مؤشر `high`.

إذا أصبح `low` أكبر من `high` في أي وقت، فهذا يعني أن مساحة البحث أصبحت فارغة، وأن `value` غير موجود في القائمة.

داخل دالة `binary_search`، أنشئ حلقة `while` بشرط يتحقق إذا كان `low` أقل من أو يساوي high. داخل الحلقة، أضف الكلمة المفتاحية `pass` كعنصر نائب مؤقتًا.

# --hints--

يجب أن تنشئ حلقة `while` تعمل طالما أن `low` أقل من أو يساوي `high`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
