---
id: 683ec95f2f8781355556ff7d
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

لتبدأ مع الحلقة، عليك إيجاد فهرس المنتصف لمساحة البحث الحالية.

يمكنك الحصول على فهرس المنتصف بحساب متوسط المتغيرين `low` و`high`. يجب استخدام القسمة الأرضية لحساب المتوسط بعد جمع القيمتين، ليكون الناتج دائمًا عددًا صحيحًا مقربًا للأسفل، لأن الفهارس يجب أن تكون أعدادًا صحيحة.

احذف الكلمة المفتاحية `pass` من حلقة `while` وأنشئ متغيرًا `mid` يعادل متوسط `low` و`high`.

# --hints--

يجب حذف الكلمة المفتاحية `pass` من حلقة while وإنشاء متغير `mid` يعادل `(low + high) // 2`.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
--fcc-editable-region--
        pass
--fcc-editable-region--
```
