---
id: 683ec39565dbc7e3375263b1
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في درس سابق، تعلمت عن البحث الثنائي وكيف يختلف عن البحث الخطي.

في هذه الورشة، ستطبق ذلك من خلال بناء خوارزمية بحث ثنائي.

لبدء العمل، أنشئ دالة `binary_search` بمعلمتين هما `search_list` و `value`.

`search_list` هي قائمة العناصر التي ستبحث الدالة ضمنها، و`value` هو الهدف الذي ستحاول الدالة إيجاده داخل `search_list`.

أضف الكلمة المفتاحية `pass` داخل الدالة مؤقتًا حتى ينجح الاختبار.

# --hints--

يجب أن تنشئ دالة `binary_search`.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

يجب أن تحتوي دالة `binary_search` على معلمتين هما `search_list` و `value`.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
