---
id: 683ec95f2f8781355556ff7e
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

بعد أن حسبت فهرس `mid`، عليك فعليًا استرجاع القيمة من `search_list` عند ذلك الفهرس. هذه القيمة هي التي ستقارنها مع معلمة `value`، وهي الهدف الذي تبحث عنه.

عرّف متغيرًا باسم `value_at_middle` وعيّن له العنصر من `search_list` الموجود عند فهرس `mid`.

# --hints--

يجب أن تعرّف متغيرًا باسم `value_at_middle` وتعيّن له القيمة `search_list[mid]`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
--fcc-editable-region--
        
--fcc-editable-region--
```
