---
id: 684169b00a073c9b9f0a703c
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

هل لاحظت أن الاستدعاء الثاني للدالة يحتوي الآن على القيم `3, 5, 4` في القائمة؟ هذا يدل على أن البحث الثنائي يعمل كما هو مقصود.

هكذا حدث الأمر: الخوارزمية فحصت أولًا `3` كمنتصف القائمة الأولية. بما أن `4` أكبر من `3`، تحوّل البحث إلى النصف الأيمن. ثم فحصت `5` كمنتصف جديد. وبما أن `4` أقل من `5`، تحرك البحث إلى اليسار، وفي النهاية حددت `4` كمنتصف النطاق النهائي.

لاختبار الدالة مرة أخرى، استدعها مع `[1, 3, 5, 9, 14, 22], 10` واطبع الاستدعاء فورًا. هذه حالة لن يتم فيها العثور على القيمة.

# --hints--

يجب أن تستدعي الدالة بـ `binary_search([1, 3, 5, 9, 14, 22], 10)` وتطبعها فورًا.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1
    
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
