---
id: 686386e76197458b7dfd52ef
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Sasa utaunda kitufe cha CTA. Wito wa Kuchukua Hatua (CTA) ni kitufe au kiungo kilichoundwa kuhamasisha mtumizi kuchukua hatua fulani, katika kesi hii, kuanza kusikiliza.

Nje ya `div` ya ndani kabisa, unda kipengele cha nanga chenye `href` ya `#` na maandishi `Start listening`.

Hiyo inahitimisha HTML ya kadi ya bei ya kwanza.

# --hints--

Unapaswa kuunda kipengele cha nanga.

```js
assert.exists(document.querySelector("a"))
```

Kipengele chako cha nanga kinapaswa kuwa na `href` iliyowekwa kuwa `#`.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

Kipengele chako cha nanga kinapaswa kuwa na maandishi `Start listening`.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
