---
id: 686386e76197458b7dfd52eb
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Unda kipengele cha `div` chini ya `h1`. Ndani ya `div` hiyo, unda kipengele kingine cha `div`, na panga ndani kipengele kingine cha `div` ndani yake.

Na ndani ya `div` ya ndani kabisa, unda kipengele cha `h2` chenye maandishi ya `Listener`. Hii itawakilisha mpango wa kwanza wa bei.

# --hints--

Unapaswa kuunda kipengele cha `div`.

```js
assert.exists(document.querySelector("div"));
```

Unapaswa kuunda kipengele kingine cha `div` ndani ya `div` ya nje.

```js
assert.exists(document.querySelector("div > div"));
```

Unapaswa kuunda kipengele kingine cha `div` ndani ya `div` ya ndani.

```js
assert.exists(document.querySelector("div > div > div"));
```

Unapaswa kuunda `h2` yenye maandishi ya `Listener` ndani ya kipengele cha `div` cha ndani kabisa.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
