---
id: 686386e76197458b7dfd52ee
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Ndani ya `ul`, tengeneza vitu viwili vya orodha. Ndani ya kipengee cha orodha cha kwanza, tengeneza kipengee cha `span` na ongeza kielezi cha HTML `&#10003;` kama maandishi yake. Mara moja baada ya `span` lakini ndani ya kipengee cha orodha, ongeza maandishi `Ad-supported streaming`.

Ndani ya kipengee cha orodha cha pili, tengeneza kipengee cha `span` na ongeza kielezi `&#10003;` kama maandishi yake. Mara moja baada ya `span`, lakini ndani ya kipengee cha orodha, ongeza maandishi `Curated playlists`.

# --hints--

Unapaswa kuwa na vitu viwili vya orodha ndani ya kipengee chako cha `ul`.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

Kipengee chako cha kwanza cha `li` kinapaswa kuwa na kipengee cha `span` chenye kielezi cha tabia ya herufi cha HTML `&#10003;` ndani yake.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

Kipengee chako cha kwanza cha `li` kinapaswa kuwa na maandishi `Ad-supported streaming` mara moja baada ya kipengee cha `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

Kipengee chako cha pili cha `li` kinapaswa kuwa na kipengee cha `span` chenye kielezi cha tabia ya herufi cha HTML `&#10003;` ndani yake.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

Kipengee chako cha pili cha `li` kinapaswa kuwa na maandishi `Curated playlists` mara moja baada ya kipengee cha `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
