---
id: 686383abac6c625a84458ea9
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafundishwa zaidi kuhusu Tailwind CSS kwa kutumia madarasa yake ya msaada wa gridi kujenga sehemu ya bei.

Muundo wa msingi wa HTML umeandaliwa kwa ajili yako. Anza kwa kuunda kipengele cha `h1` chenye maandishi `Choose your listening plan`.

# --hints--

Unapaswa kuunda kipengele cha `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `Choose your listening plan`.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
