---
id: 697fc395f3d85f14b05cd9ca
title: Restaurar uma narrativa coerente a partir de um array de fragmentos de história
challengeType: 26
dashedName: lab-story-fragment-restoration
---

# --description--

Neste laboratório, você vai restaurar uma narrativa coerente a partir de um array corrompido de fragmentos de história.

Você vai praticar o uso de loops implementando algoritmos fundamentais de array do zero.

Você vai trabalhar com arrays de objetos fragmentos de história. Cada objeto fragmento tem as seguintes propriedades:

| Propriedade | Descrição | Valor de exemplo |
|-|-|-|
| `id` | Um inteiro positivo que indica a posição do fragmento na história | `3` |
| `text` | O conteúdo real da história | `"and I use Arch btw.\""` |

Neste laboratório, você recebe um array pré-preenchido chamado `shuffledFragments`.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**User stories**

1. Você não deve alterar o array pré-preenchido `shuffledFragments`.
2. Você deve criar uma função chamada `compactFragments` que recebe um array de fragmentos e retorna um novo array com todos os elementos undefined removidos. Se a função remover algum elemento undefined, ela deve registrar uma mensagem no console. A mensagem deve começar com o prefixo `[COMPACTED]`.
3. Você deve declarar uma variável chamada `compactedShuffledFragments` e atribuir a ela o resultado da chamada de `compactFragments` com o array `shuffledFragments`.
4. Você deve criar uma função chamada `sortFragments` que recebe um array de fragmentos sem elementos `undefined` e retorna um novo array ordenado pela propriedade `id` em ordem crescente, mantendo os fragmentos que compartilham o mesmo `id` na ordem original. Você não deve usar o método `sort` nativo do JavaScript.
5. Você deve declarar uma variável chamada `sortedFragments` e atribuir a ela o resultado da chamada de `sortFragments` com o array `compactedShuffledFragments`.
6. Você deve criar uma função chamada `dedupeFragments` que recebe um array ordenado de fragmentos e retorna um novo array com os duplicados removidos, mantendo apenas a primeira ocorrência. Você deve definir duplicados como dois ou mais fragmentos que compartilham o mesmo `id`. Para cada `id` desduplicado, a função deve registrar uma mensagem no console. A mensagem deve começar com o prefixo `[DEDUPED]`.
7. Você deve declarar uma variável chamada `dedupedFragments` e atribuir a ela o resultado da chamada de `dedupeFragments` com o array `sortedFragments`.
8. Você deve criar uma função chamada `fillMissingFragments` que recebe um array ordenado de fragmentos e retorna um novo array com os fragmentos ausentes preenchidos com objetos de espaço reservado. Você deve definir fragmentos ausentes como lacunas na sequência entre o menor e o maior `id`. Os objetos de espaço reservado devem ter o formato `{ id: missingId, text: "[...]" }`. Para cada espaço reservado adicionado, a função deve registrar uma mensagem no console. A mensagem deve começar com o prefixo `[FILLED]`.
9. Você deve declarar uma variável chamada `filledFragments` e atribuir a ela o resultado da chamada de `fillMissingFragments` com o array `dedupedFragments`.
10. Você deve criar uma função chamada `assembleStory` que recebe um array ordenado de fragmentos e retorna uma única string contendo todos os textos dos fragmentos, separados por quebras de linha.
11. Você deve usar `assembleStory` com seus `filledFragments` para exibir a história completa no console.
12. Suas funções `compactFragments`, `sortFragments`, `dedupeFragments`, `fillMissingFragments` e `assembleStory` não devem modificar o array com o qual são chamadas.

**Exemplo**

Aqui está um exemplo de um array contendo fragmentos de história:

```js
const exampleArray = [
  { id: 3, text: "and I use Arch btw.\"" },
  ,
  { id: 1, text: "Naomi said:" },
  { id: 3, text: "and I use Arch btw.\"" },
];
```

Depois de restaurar a história a partir de `exampleArray`, ela ficaria assim:

```md
Naomi said:
[...]
and I use Arch btw."
```

# --hints--

Você não deve alterar o array pré-preenchido `shuffledFragments`.

```js
const shuffledFragments__unchanged = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

assert.deepEqual(shuffledFragments, shuffledFragments__unchanged);
```

Você deve ter uma função `compactFragments`.

```js
assert.isFunction(compactFragments);
```

`compactFragments([{ id: 1, text: "Hello" }, undefined])` deve registrar uma mensagem que começa com `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Quando a função `compactFragments` é chamada com um array contendo elementos `undefined`, ela deve registrar uma mensagem no console que começa com `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" }, undefined, { id: 11, text: "bar" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" },,])` deve registrar uma mensagem que começa com `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" },,]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Quando a função `compactFragments` é chamada com um array contendo espaços vazios, ela deve registrar uma mensagem no console que começa com `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" },,{ id: 11, text: "bar" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Quando a função `compactFragments` é chamada com o array `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`, ela não deve registrar nenhuma mensagem no console que comece com `[COMPACTED]`.

```js
const shuffledFragments__withoutUndefined = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const spy = __helpers.spyOn(console, 'log');
compactFragments(shuffledFragments__withoutUndefined);
assert.isFalse(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }])` deve retornar `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`.

```js
assert.deepEqual(
  compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }]),
  [{ id: 1, text: "Hello" }, { id: 2, text: "World" }]
);
```

`compactFragments` não deve modificar o array passado como argumento.

```js
const input = [{ id: 1, text: "a" }, undefined, { id: 2, text: "b" }];
const copy = structuredClone(input);
compactFragments(input);
assert.deepEqual(input, copy);
```

`compactedShuffledFragments` deve ser um array sem elementos `undefined`.

```js
assert.isArray(compactedShuffledFragments);
assert.isFalse(compactedShuffledFragments.some(f => f === undefined));
```

Você deve ter uma função `sortFragments`.

```js
assert.isFunction(sortFragments);
```

`sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }])` deve retornar `[{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]`.

```js
assert.deepEqual(
  sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }]),
  [{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]
);
```

`sortFragments` não deve usar o método embutido `.sort()`.

```js
function isBuiltInSortUsed() {
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => { sortUsed = true; };
  try {
    sortFragments([{ id: 2, text: "b" }, { id: 1, text: "a" }]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

`sortFragments` não deve modificar o array passado como argumento.

```js
const input = [{ id: 3, text: "c" }, { id: 1, text: "a" }];
const copy = structuredClone(input);
sortFragments(input);
assert.deepEqual(input, copy);
```

`sortedFragments` deve ter o `id` de cada elemento menor ou igual ao `id` do próximo elemento.

```js
for (let i = 0; i < sortedFragments.length - 1; i++) {
  assert.isAtMost(sortedFragments[i].id, sortedFragments[i + 1].id);
}
```

Você deve ter uma função `dedupeFragments`.

```js
assert.isFunction(dedupeFragments);
```

`dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }])` deve retornar `[{ id: 1, text: "first" }, { id: 2, text: "second" }]`.

```js
assert.deepEqual(
  dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]),
  [{ id: 1, text: "first" }, { id: 2, text: "second" }]
);
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }])` deve registrar uma mensagem `[DEDUPED]` para cada `id` duplicado.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 2);
spy.restore();
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }])` deve registrar uma mensagem `[DEDUPED]` quando o último `id` for duplicado.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 1);
spy.restore();
```

Quando `dedupeFragments` é chamada com um array contendo duplicados, ela deve registrar uma mensagem começando com `[DEDUPED]` para cada `id` duplicado.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[DEDUPED\]/)));
spy.restore();
```

`dedupeFragments` não deve modificar o array passado como argumento.

```js
const input = [{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }];
const copy = structuredClone(input);
dedupeFragments(input);
assert.deepEqual(input, copy);
```

`dedupedFragments` não deve ter valores duplicados de `id`.

```js
const ids = dedupedFragments.map(f => f.id);
assert.strictEqual(ids.length, new Set(ids).size);
```

Você deve ter uma função `fillMissingFragments`.

```js
assert.isFunction(fillMissingFragments);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` deve retornar `[{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]`.

```js
assert.deepEqual(
  fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]),
  [{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]
);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` deve registrar uma mensagem que começa com `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

Quando `fillMissingFragments` adiciona um marcador de posição, ela deve registrar uma mensagem começando com `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 5, text: "e" }, { id: 8, text: "h" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

`fillMissingFragments` não deve modificar o array passado como argumento.

```js
const input = [{ id: 1, text: "a" }, { id: 3, text: "c" }];
const copy = structuredClone(input);
fillMissingFragments(input);
assert.deepEqual(input, copy);
```

`filledFragments` não deve ter gaps na sequência de `id`.

```js
for (let i = 0; i < filledFragments.length - 1; i++) {
  assert.strictEqual(filledFragments[i + 1].id, filledFragments[i].id + 1);
}
```

Você deve ter uma função `assembleStory`.

```js
assert.isFunction(assembleStory);
```

`assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }])` deve retornar a string `"Hello\nWorld"`.

```js
assert.strictEqual(
  assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }]),
  "Hello\nWorld"
);
```

`assembleStory` não deve modificar o array passado como argumento.

```js
const input = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const copy = structuredClone(input);
assembleStory(input);
assert.deepEqual(input, copy);
```

Você deve chamar `console.log` com o resultado de `assembleStory(filledFragments)`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(.*assembleStory\s*\(\s*filledFragments\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];
```

# --solutions--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

const compactFragments = (fragments) => {
  let compacted = [];
  let numberOfUndefined = 0;

  for (const fragment of fragments) {
    if (fragment === undefined) {
      numberOfUndefined++;
    } else {
      compacted.push(fragment);
    }
  }

  if (numberOfUndefined > 0) {
    console.log(`[COMPACTED] Removed ${numberOfUndefined} undefined fragment(s).`);
  }

  return compacted;
};

const compactedShuffledFragments = compactFragments(shuffledFragments);

const sortFragments = (fragments) => {
  let sorted = [...fragments];

  for (let i = 0; i < sorted.length; i++) {
    for (let j = i; j > 0; j--) {
      if (sorted[j - 1].id > sorted[j].id) {
        swap(sorted, j - 1, j);
      }
    }
  }

  function swap(array, index1, index2) {
    [sorted[index1], sorted[index2]] = [sorted[index2], sorted[index1]];
  }

  return sorted;
};

const sortedFragments = sortFragments(compactedShuffledFragments);

const dedupeFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let deduped = [shallowCopy.shift()];
  let duplicates = [];
  let currentIdOccurrence = 1;

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = deduped.at(-1);
    if (fragmentToInsert.id === previousFragment.id) {
      currentIdOccurrence++;
    } else {
      deduped.push(fragmentToInsert);

      if (currentIdOccurrence > 1) {
        duplicates.push({ id: previousFragment.id, occurrence: currentIdOccurrence });
      }

      currentIdOccurrence = 1;
    }
  }

  if (currentIdOccurrence > 1) {
    duplicates.push({ id: deduped.at(-1).id, occurrence: currentIdOccurrence });
  }

  for (const duplicate of duplicates) {
    console.log(`[DEDUPED] Fragment #${duplicate.id} appeared ${duplicate.occurrence} times, kept first occurrence.`);
  }

  return deduped;
};

const dedupedFragments = dedupeFragments(sortedFragments);

const fillMissingFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let filled = [shallowCopy.shift()];

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = filled.at(-1);
    const numberOfMissingIDs = fragmentToInsert.id - 1 - previousFragment.id;
    for (let i = 1; i <= numberOfMissingIDs; i++) {
      const missingId = previousFragment.id + i;
      filled.push({ id: missingId, text: "[...]" });
      console.log(`[FILLED] Fragment #${missingId} missing, placeholder inserted.`)
    }
    filled.push(fragmentToInsert);
  }

  return filled;
};

const filledFragments = fillMissingFragments(dedupedFragments);

const assembleStory = (sortedFragments) => {
  const texts = [];
  for (const fragment of sortedFragments) {
    texts.push(fragment.text);
  }
  return texts.join("\n");
};

console.log("\nHere is the assembled story:\n");
console.log(assembleStory(filledFragments));
```
