---
id: 66e45c8140f9fda5c105ae26
title: Crie uma página de anúncio de evento
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um elemento `header` dentro do corpo.
1. Seu elemento `header` deve conter uma imagem para o seu evento e um `h1` com o título do seu evento, nessa ordem. Você pode usar esta imagem se desejar: `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. Você deve ter um elemento `main` dentro do body.
1. Seu elemento `main` deve conter pelo menos dois elementos `section` dentro dele, exibindo os recursos do evento.
1. Seus elementos `section` devem conter um `h2` dentro deles.
1. Seu corpo deve ter um preenchimento superior e inferior de `50px`.
1. Seu corpo deve ter uma margem superior e inferior de `0`, e uma margem esquerda e direita que se centralize.
1. Seu corpo deve ter uma largura definida em relação à largura da janela de visualização.
1. Seu corpo deve usar a função `calc` para definir sua propriedade `min-height` como 100% da altura da viewport menos todo o padding aplicado ao topo e à parte inferior do body.
1. Você deve ter pelo menos um `hr` dentro do seu flyer.
1. Você deve definir a `width` de todos os elementos `hr` e `section` para um valor percentual relativo ao seu elemento pai.

**Nota:** Certifique-se de vincular sua folha de estilos no seu HTML e aplicar seu CSS.

# --hints--

Você deve ter um elemento `header` dentro do elemento `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Você deve ter um elemento `img` como o primeiro filho no `header`.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

Você deve ter um elemento `h1` como o segundo filho no `header`.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

Você deve ter apenas um elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Seu `h1` deve conter o título do seu evento.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

Você deve ter um elemento `main` dentro do elemento `body`.

```js
assert.exists(document.querySelector('body > main'));
```

Você deve ter pelo menos dois elementos `section` dentro do seu elemento `main`.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

Seus elementos `section` devem conter cada um um `h2` dentro deles.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

Seus elementos `h2` não devem estar vazios.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

Seu elemento `body` deve ter um padding superior e inferior de `50px`.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

Seu elemento `body` deve ter uma margem superior e inferior de `0`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

Seu elemento `body` deve ter uma margem esquerda e direita de `auto`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

Seu elemento `body` deve definir sua largura usando `vw`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

Seu `body` deve usar `calc` para definir seu `min-height` como `100vh - 100px`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

Você deve ter pelo menos um elemento `hr`.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

A largura dos seus elementos `hr` deve ser definida usando um valor percentual.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

A largura dos seus elementos `section` deve ser definida usando um valor percentual.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
