---
id: 6a3e6fd08125dfa615cb431b
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Da Ihre `use_eggs` Funktion jetzt fertig ist, können Sie sie testen. Wenn Sie eine Funktion aufrufen, die **Parameter** benötigt, übergeben Sie die konkreten Werte innerhalb der runden Klammern. Diese Werte werden als **Argumente** bezeichnet.

Hier ist ein Beispiel, wie man eine Funktion mit Argumenten aufruft:

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

In diesem Beispiel sind die ersten beiden Zeilen die Funktionsdefinition, wobei `a` und `b` die Parameter sind. Die letzte Zeile, `calculate_sum(3, 1)`, ist der Funktionsaufruf. Die Werte `3` und `1` sind die Argumente. Sie werden den Parametern positionsabhängig zugeordnet, sodass `3` `a` und `1` `b` zugewiesen wird.

Rufen Sie am Ende Ihres Codes Ihre `use_eggs` Funktion auf und übergeben Sie ihr zwei Argumente: die globale Variable `available_eggs` und die Zahl `1` (die Anzahl der Eier, die Sie verwenden möchten).

# --hints--

Sie sollten die `use_eggs` Funktion aufrufen und `available_eggs` sowie `1` als Argumente übergeben.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
