---
id: 6a3e5475c35dbc138d77c6f4
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Da Sie nun Ihre Gesamtsumme berechnet haben, sollten Sie den Bestand anzeigen. Erinnern Sie sich aus vorherigen Lektionen, dass Sie die `print()`-Funktion mit **f-strings** verwenden können, um Variablen in Ihrer Ausgabe einzubinden.

Fügen Sie innerhalb Ihrer `check_kitchen_stock` Funktion unterhalb der Variable `total_items` genau diese vier print-Anweisungen hinzu:

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

Sie sollten die Anweisung `print(f'The kitchen has {total_items} total items:')` in Ihrer Funktion haben.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

Die Anweisung `print(f'The kitchen has {total_items} total items:')` sollte der erste `print()`-Aufruf in Ihrer Funktion sein.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

Sie sollten die Anweisung `print(f'- {available_eggs} eggs')` in Ihrer Funktion haben.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

Sie sollten die Anweisung `print(f'- {available_flour} flour')` in Ihrer Funktion haben.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

Sie sollten die Anweisung `print(f'- {available_sugar} sugar')` in Ihrer Funktion haben.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
