---
id: 6a3e4ae26f2b7e4b5b026cf8
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Es ist Zeit, eine Funktion zu erstellen, um zu überprüfen, was Sie in Ihrer Küche haben. In Python definieren Sie Funktionen mit dem Schlüsselwort `def`, gefolgt vom Funktionsnamen, runden Klammern `()` und einem Doppelpunkt `:`.

Hier ist ein Beispiel, wie man eine einfache Funktion definiert:

```py
def my_function():
    pass
```

Definieren Sie eine leere Funktion namens `check_kitchen_stock()`. Verwenden Sie vorerst das Schlüsselwort `pass` im Funktionskörper, damit Python keinen Syntaxfehler wirft.

Denken Sie daran, dass Python Einrückungen verwendet, um Codeblöcke zu definieren. Stellen Sie also sicher, dass Sie den Funktionskörper einrücken.

# --hints--

Sie sollten eine Funktion namens `check_kitchen_stock` definieren.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
