---
id: 6a3e5ccb99fd6053af9032a1
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Jetzt, da Sie wissen, was sich in Ihrer Küche befindet, können Sie diese Zutaten zum Kochen verwenden. Erstellen Sie dazu eine neue Funktion, die Parameter akzeptiert. Denken Sie an Parameter als Platzhalter für die Daten, die Sie beim Aufrufen der Funktion übergeben.

Hier ist eine Funktion, die zwei Parameter, `a` und `b`, akzeptiert und deren Summe im Terminal ausgibt.

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

Definieren Sie unterhalb Ihrer Funktion `check_kitchen_stock` eine leere Funktion namens `use_eggs`, die zwei Parameter akzeptiert: `available_eggs` und `eggs_to_use` in dieser Reihenfolge.

Beachten Sie, dass wenn Sie einem Parameter denselben Namen wie einer globalen Variable geben (wie `available_eggs`), die Funktion ihren eigenen lokalen Parameter `available_eggs` anstelle der globalen Variable verwendet. Dies wird als Name Shadowing bezeichnet, da die lokale Variable die globale "überschattet".

Verwenden Sie vorerst das Schlüsselwort `pass` im Funktionskörper.

# --hints--

Sie sollten eine Funktion namens `use_eggs` definieren.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

Ihre Funktion `use_eggs` sollte `available_eggs` und `eggs_to_use` als Parameter haben.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
