---
id: 6a3e58891e1ff13d07f4ab48
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Python liest den `def`-Block und merkt sich die Funktion, aber es führt den Code darin nicht sofort aus.

Um den Code auszuführen und Ihre Ausgabe zu sehen, müssen Sie die Funktion **aufrufen** (oder ausführen). Tun Sie dies, indem Sie den Namen der Funktion gefolgt von runden Klammern eingeben.

Zum Beispiel:

```py
my_function()
```

Rufen Sie am Ende Ihres Codes, außerhalb der Funktion, Ihre `check_kitchen_stock`-Funktion auf.

# --hints--

Sie sollten die `check_kitchen_stock`-Funktion am Ende Ihres Codes aufrufen. Stellen Sie sicher, dass sie nicht eingerückt ist, damit sie global ausgeführt wird.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
