---
id: 6a3e73eabddd614f91cc235a
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Jetzt werden Sie das Rätsel von vorhin lösen! Als Ihre `use_eggs`-Funktion die verbleibende Anzahl der Eier zurückgab, haben Sie diesen neuen Wert nirgendwo gespeichert. Er verschwand einfach im Nichts!

Um Ihren Vorrat tatsächlich zu aktualisieren, müssen Sie diesen zurückgegebenen Wert erfassen, indem Sie ihn einer Variablen zuweisen.

Verwenden Sie innerhalb Ihrer `make_fried_egg`-Funktion eine `if/else`-Anweisung, um die Variable `has_enough_eggs` zu überprüfen.

- Wenn `has_enough_eggs` `True` ist, weisen Sie Ihre Variable `available_eggs` neu zu mit dem Ergebnis von `use_eggs(available_eggs, 1)` und geben Sie die Nachricht `Made a fried egg. Yummy!` aus.

- Andernfalls geben Sie innerhalb eines `else`-Blocks die Nachricht `Could not make a fried egg. Not enough eggs!` aus.

# --hints--

Sie sollten eine `if`-Anweisung haben und `has_enough_eggs` als Bedingung verwenden.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

Innerhalb Ihres `if`-Blocks sollten Sie `available_eggs` neu zuweisen mit dem Ergebnis von `use_eggs(available_eggs, 1)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

Sie sollten die Nachricht `Made a fried egg. Yummy!` ausgeben, wenn `has_enough_eggs` `True` ist.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

Sie sollten einen `else`-Block haben, der die Nachricht `Could not make a fried egg. Not enough eggs!` ausgibt, wenn `has_enough_eggs` `False` ist.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
