---
id: 6a3e2b9b7eeca084dd4fb857
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Willkommen in der Küche! So wie ein Koch sich auf Rezepte verlässt, um eine Mahlzeit zuzubereiten, werden Sie sich auf Funktionen verlassen, um Software zu erstellen. In diesem Workshop lernen Sie, wie Sie wiederverwendbare Python-Funktionen schreiben und den Gültigkeitsbereich meistern, indem Sie genau verstehen, welche Zutaten (Variablen) zu einem bestimmten Zeitpunkt in Ihrem Code verfügbar sind.

Um zu beginnen, sollten Sie einige grundlegende Daten für Ihr Kücheninventar definieren. Erstellen Sie drei Variablen: `available_eggs`, `available_flour` und `available_sugar` und weisen Sie ihnen die Werte `1`, `2` und `3` zu.

# --hints--

Sie sollten eine Variable namens `available_eggs` definieren und ihr den Wert `1` zuweisen.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

Sie sollten eine Variable namens `available_flour` definieren und ihr den Wert `2` zuweisen.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

Sie sollten eine Variable namens `available_sugar` definieren und ihr den Wert `3` zuweisen.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
