---
id: 6a3e71a0a3a44260e675b157
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Um das Rätsel zu lösen, warum sich Ihre Eieranzahl nicht geändert hat, und tatsächlich das Frühstück zu kochen, erstellen Sie eine neue Funktion. Denken Sie an diese Funktion als Ihr Rezept, das andere Funktionen aufruft, um die Aufgabe zu erledigen.

Definieren Sie unter Ihrer `print(available_eggs)`-Anweisung eine leere Funktion namens `make_fried_egg`, die `available_eggs` als einzigen Parameter übernimmt.

Verwenden Sie vorerst das Schlüsselwort `pass` im Funktionskörper.

# --hints--

Sie sollten eine Funktion namens `make_fried_egg` definieren.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

Ihre Funktion `make_fried_egg` sollte `available_eggs` als Parameter haben.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
