---
id: 6a3e524fc3394ef52ed5b797
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Funktionen haben ihren eigenen lokalen Gültigkeitsbereich. Das bedeutet, dass Variablen, die innerhalb einer Funktion erstellt werden, wie Zutaten sind, die in der Speisekammer dieser Küche aufbewahrt werden: Sie können nur innerhalb dieser spezifischen Funktion aufgerufen werden.

Funktionen sind jedoch nicht völlig isoliert. Sie können globale Variablen lesen, die außerhalb ihres Gültigkeitsbereichs definiert sind. Sie können diese Variablen verwenden, um Berechnungen durchzuführen, Entscheidungen zu treffen oder neue lokale Variablen zu erstellen.

Ersetzen Sie innerhalb Ihrer `check_kitchen_stock`-Funktion das Schlüsselwort `pass` durch eine neue Variable namens `total_items` und weisen Sie ihr die Summe von `available_eggs`, `available_flour` und `available_sugar` zu.

# --hints--

Sie sollten innerhalb Ihrer `check_kitchen_stock`-Funktion eine Variable namens `total_items` definieren.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

Ihre Variable `total_items` sollte die Summe von `available_eggs`, `available_flour` und `available_sugar` sein.

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
