---
id: 6a3ea8a55317ba7a08e21d5b
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Jetzt können Sie endlich die frühere Frage beantworten: Warum blieb die globale Anzahl der Eier bei `1`? Wenn eine Funktion einen Wert zurückgibt, überschreibt sie nicht automatisch Ihre bestehenden Variablen. Sie müssen diesen neuen Wert ausdrücklich speichern!

Sie verwenden die Variable `available_eggs` innerhalb der Funktion `make_fried_egg`, aber es ist nicht dieselbe Variable, die am Anfang Ihres Codes definiert ist. Die innerhalb der Funktion ist eine lokale Variable, die nur existiert, solange die Funktion läuft. Wenn die Funktion endet, wird die lokale Variable verworfen, und die globale Variable `available_eggs` bleibt unverändert.

Um Ihren tatsächlichen Küchenbestand zu aktualisieren, überschreiben Sie die globale Variable `available_eggs` mit der neuen Anzahl an Eiern, die von Ihrer Funktion `make_fried_egg` zurückgegeben wird.

Rufen Sie ganz unten in Ihrem Code die Funktion `make_fried_egg` auf, übergeben Sie `available_eggs` als Argument und weisen Sie das Ergebnis der globalen Variable `available_eggs` wieder zu.

# --hints--

Sie sollten die Variable `available_eggs` neu zuweisen, und zwar mit dem Ergebnis von `make_fried_egg(available_eggs)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
