---
id: 6a3e6ad6b9554e31dcc30cd6
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Wenn Sie nicht genug Eier haben, sollten Sie nicht versuchen zu kochen! Sie brauchen eine Möglichkeit, die Funktion daran zu hindern, weiterzulaufen.

Dafür können Sie das Schlüsselwort `return` verwenden. Wenn eine Funktion auf eine `return`-Anweisung trifft, hört sie sofort auf auszuführen und gibt den zurückgegebenen Wert an die Stelle zurück, von der sie aufgerufen wurde.

Zum Beispiel:

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

Verwenden Sie innerhalb Ihres `if`-Blocks unterhalb der Ausgabeanweisung das Schlüsselwort `return`, um den Wert von `available_eggs` zurückzugeben.

# --hints--

Sie sollten `available_eggs` innerhalb Ihrer `if`-Anweisung zurückgeben.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

Sie sollten `available_eggs` nach der `print`-Anweisung innerhalb Ihres `if`-Blocks zurückgeben.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
