---
id: 6a3e6c0eb5d050ce07781709
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Wenn die `if`-Bedingung zu `False` ausgewertet wird (was bedeutet, dass Sie *genug* Eier haben), überspringt die Funktion den `if`-Block und führt den restlichen Code weiter aus.

Jetzt müssen Sie Ihre Daten verarbeiten und einen neuen berechneten Wert zurückgeben.

Außerhalb Ihres `if`-Blocks (auf derselben Einrückungsebene wie das `if`-Schlüsselwort) verwenden Sie einen `print()`-Aufruf und einen f-String, um `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` auszugeben.

Geben Sie dann in der nächsten Zeile die verbleibende Anzahl an Eiern zurück, indem Sie `eggs_to_use` von `available_eggs` subtrahieren, mit `return`.

# --hints--

Sie sollten eine print-Anweisung haben, die in Ihrer `use_eggs`-Funktion `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` ausgibt.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

Sie sollten die verbleibende Anzahl an Eiern zurückgeben, indem Sie `eggs_to_use` von `available_eggs` subtrahieren.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

Sie sollten `available_eggs - eggs_to_use` nach der `print`-Anweisung zurückgeben.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
