---
id: 6a3e70b94f67d9ee3ecd8486
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Sie haben gerade `use_eggs(available_eggs, 1)` aufgerufen, um ein Ei zu verwenden. Das bedeutet, dass Sie jetzt `0` Eier in Ihrer Küche haben sollten, richtig?

Das werden Sie herausfinden! Verwenden Sie am Ende Ihres Codes die `print()`-Funktion, um den Wert der Variable `available_eggs` auszugeben.

Wenn Sie den Code ausführen, werden Sie etwas Überraschendes bemerken: Die Ausgabe ist immer noch `1`! Warum hat sich die Anzahl der Eier nicht geändert? Dieses Rätsel werden Sie in den nächsten Schritten erkunden.

# --hints--

Sie sollten die Variable `available_eggs` am Ende Ihres Codes ausgeben.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
