---
id: 6a3e67fc5ef4045201c28219
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Innerhalb Ihrer `use_eggs` Funktion müssen Sie prüfen, ob Sie genug Eier zum Kochen haben. Das können Sie mit einer `if`-Anweisung machen.

Eine `if`-Anweisung bewertet eine Bedingung. Wenn die Bedingung wahr ist, wird der Code innerhalb des `if`-Blocks ausgeführt.

Ersetzen Sie das Schlüsselwort `pass` durch eine `if`-Anweisung, die überprüft, ob `eggs_to_use` größer (`>`) als `available_eggs` ist. Innerhalb des `if`-Blocks geben Sie den String `The kitchen does not have enough eggs.` aus.

# --hints--

Sie sollten eine `if`-Anweisung haben, die überprüft, ob `eggs_to_use > available_eggs` gilt.

```js
({
  test: () => runPython(`
  cond = _Node(_code).find_function("use_eggs").find_ifs()[0].find_conditions()[0]
  assert cond.is_equivalent("eggs_to_use > available_eggs") or cond.is_equivalent("available_eggs < eggs_to_use")`)
})
```

Sie sollten innerhalb Ihrer `if`-Anweisung `The kitchen does not have enough eggs.` ausgeben.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_call("print('The kitchen does not have enough eggs.')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
