---
id: 6a3e7284e89ce3367e709331
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Bevor Sie ein Spiegelei machen, müssen Sie zuerst prüfen, ob Sie die benötigten Zutaten haben. Das können Sie tun, indem Sie die verfügbare Anzahl an Eiern mit der für das Rezept benötigten Menge vergleichen.

Ersetzen Sie innerhalb Ihrer `make_fried_egg`-Funktion das Schlüsselwort `pass` durch eine Variable namens `has_enough_eggs`. Weisen Sie ihr das Ergebnis der Überprüfung zu, ob `available_eggs` größer oder gleich `1` ist.

# --hints--

Sie sollten `has_enough_eggs` das Ergebnis der Überprüfung zuweisen, ob `available_eggs` größer oder gleich `1` ist.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_variables("has_enough_eggs")[0]
assert node.is_equivalent("has_enough_eggs = available_eggs >= 1") or node.is_equivalent("has_enough_eggs = 1 <= available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
