---
id: 7a97b79c096ecb000d03bd36
title: Build an Nth Fibonacci Number Calculator
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Objetivo:** Cumple las historias de usuario a continuación y haz que todas las pruebas pasen para completar el laboratorio.

**User Stories:**

1. Debes crear una función llamada `fibonacci`.
2. Debes definir un arreglo llamado `sequence` dentro de la función `fibonacci`, y debe inicializarse con los valores `[0, 1]`.
3. La función `fibonacci` debe aceptar un parámetro, un entero positivo `n`.
4. Llamar a `fibonacci(n)` debe usar un enfoque de programación dinámica para calcular y devolver el `n`-ésimo número de la secuencia de Fibonacci, donde cada número es la suma de los dos números anteriores.
5. Cada número calculado en la posición `n` en la secuencia de Fibonacci debe almacenarse en el arreglo `sequence` en el índice `n - 1`.

# --hints--

Debes tener una función llamada `fibonacci`.

```js
assert.isFunction(fibonacci);
```

Tu función `fibonacci` debe tomar un solo parámetro.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

Debes tener un arreglo llamado `sequence` dentro de la función `fibonacci` inicializado a `[0, 1]`.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` should return `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` should return `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` should return `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` should return `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` should return `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` should return `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` should return `610`.

```js
assert.equal(fibonacci(15), 610);
```

No debes usar recursión en tu código.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
