---
id: 68e97fe79367ad7b5dd6c9cd
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

このステップでは、テキストが`h1`の`XYZ Bookstore`要素を追加してください。

# --hints--

`h1`要素があるはずです。

```js
assert.exists(document.querySelector('h1'));
```

`h1`要素のテキストは`XYZ Bookstore`であるべきです。スペルや大文字・小文字の間違いがないか再確認してください。

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

`h1`要素は`body`要素の中にあるべきです。

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
