---
id: 68ebdbacdd3fa474132cc975
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`h1` 要素の下に、次のテキストを含む `p` 要素を追加してください：`Browse our collection of amazing books!`。

# --hints--

`p` 要素があるはずです。

```js
assert.exists(document.querySelector('p'));
```

`p` 要素のテキストは `Browse our collection of amazing books!` であるべきです。

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

`p` 要素は `h1` 要素の下にあるべきです。

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
