---
id: 68ec6cd5b7e8f5a8f7319e32
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`div` 要素は、他のHTML要素をグループ化するためのコンテナとして使われます。主に、CSSスタイルのセットを共有するHTML要素をグループ化したいときに`div` 要素を使います。

`p` 要素の下に、`div` 要素を追加してください。この`div` はあなたの本のカードのコンテナになります。

**注意**：このワークショップではCSSは適用されません。クラスやグループ化された要素はCSSスタイリングに役立ちますが、このワークショップではコンテンツの構造化とグループ化のためだけに使われます。スタイリングの仕組みは後のモジュールで学びます。

# --hints--

`div` 要素があるはずです。

```js
assert.exists(document.querySelector('div'));
```

`div` 要素は`p` 要素の下にあるべきです。

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
