---
id: 68ec99e478211dc578699944
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

`h2` クラスを持つ最初の要素の `card` 要素の下に、次のテキストを持つ `p` 要素を追加してください。

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

`p` クラスを持つ要素の中に `card` 要素がネストされているはずです。

```js
assert.exists(document.querySelector('.card p'));
```

`p` 要素のテキストは `This is an epic story of Sally and her dog Rex as they navigate through other worlds.` であるべきです。

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

`p` 要素は `h2` 要素の下にあるべきです。

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
