---
id: 68eca5c0065c82d256d29ca3
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

2つ目のカードの中に、`button` 要素を追加し、`class` 属性を `btn` に設定し、テキストを `Buy Now` にしてください。

両方の `button` 要素は同じ `class` を共有しているため、一緒に一貫してスタイルを付けることができます。

# --hints--

2つ目の要素の中に、クラスが `button` の `card` 要素がネストされているはずです。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

2つ目のカードの `button` 要素には `class` 属性が必要です。

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

2つ目のカードの `button` 要素には、値が `class` の `btn` が必要です。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

2つ目のカードの `button` 要素のテキストは `Buy Now` であるべきです。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
