---
id: 68ec98b38d83a3c28dd30efe
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`card` クラスを持つ最初の要素の中に、テキストが `h2` の `Sally's SciFi Adventure` 要素を追加してください。

# --hints--

`h2` クラスを持つ要素の中に `card` 要素がネストされているはずです。

```js
assert.exists(document.querySelector('.card h2'));
```

`h2` 要素のテキストは `Sally's SciFi Adventure` である必要があります。

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
