---
id: 68eca2a795b333ca5fee30a8
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

2つ目の本のカードを作成してください。`div` 要素をもう1つ追加し、`class` 属性を `card` に設定してください。同じクラス名を複数の要素で使い回して、一貫したスタイルを適用できることに注目しましょう。

# --hints--

`div` 要素が2つあり、`class` 属性の値が `card` になっているはずです。

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

`card` クラスを持つ2つ目の要素は、`card-container` クラスを持つ要素の中に入っている必要があります。

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

`card` クラスを持つ2つの要素は兄弟要素である必要があります。

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
