---
id: 68ec9332a9b5b2b32487bd00
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`id` 属性はHTML要素に一意の識別子を追加します。各`id`はページ内で一意である必要があり、一度だけ使用してください。

`id` の値にはスペースを含めることはできず、文字、数字、アンダースコア、およびダッシュのみを含めるべきです。

例を示します：

```html
<p id="para">example paragraph</p>
```

クラスが`id`の要素に`card`属性を追加し、その値を`sally-adventure-book`に設定してください。

# --hints--

クラスが`card`の要素には`id`属性が必要です。

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

クラスが`card`の要素には、値が`id`の`sally-adventure-book`が必要です。

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
