---
id: 68eca11e99e3c5c894ca9d69
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`button` 要素は、ウェブページ上でクリック可能なボタンを作成するために使われます。ボタンはユーザーがクリックして操作を実行できるインタラクティブな要素です。

`button` が `class` の要素の中に `card` 要素を追加し、ボタンに `class` 属性を `btn` に設定し、テキストを `Buy Now` にしてください。

# --hints--

`button` のクラスを持つ要素の中に `card` 要素がネストされている必要があります。

```js
assert.exists(document.querySelector('.card button'));
```

`button` 要素には `class` 属性が必要です。

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

`button` 要素には `class` があり、その値は `btn` である必要があります。

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

`button` 要素のテキストは `Buy Now` であるべきです。

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
