---
id: 68eca6364c5616d393389a30
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

HTML要素は次のような形をしています：

```html
<element attribute="value">
    inner text
</element>
```

クラスが`card-container`の要素の下に、次のテキストを持つ新しい`p`要素を追加してください：

```md
Review your selections and continue to checkout.
```

`p`要素の下に、`div`属性が`class`に設定された`btn-container`要素を作成してください。このコンテナはナビゲーションボタン要素をグループ化します。

# --hints--

新しく追加する`p`要素は、クラスが`card-container`の要素の下に配置してください。

```js
assert.exists(document.querySelector('.card-container + p'));
```

新しい`p`要素のテキストは`Review your selections and continue to checkout.`にしてください。

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

`div`要素には`class`が`btn-container`である必要があります。

```js
assert.exists(document.querySelector('div.btn-container'));
```

`div`要素は、`class`が`btn-container`である新しく追加した`p`要素の後に配置してください。

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
