---
id: 68eca4c4b0f952cf09fabe09
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

`card` クラスを持つ2番目の要素の中に、テキストが `h2` の `Dave's Cooking Adventure` 要素を追加してください。

# --hints--

`h2` クラスを持つ2番目の要素の中に `card` 要素がネストされているはずです。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

`h2` クラスを持つ2番目の要素の中の `card` 要素は、テキストが `Dave's Cooking Adventure` である必要があります。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
