---
id: 68eca6e5a6759ed4ea0034dc
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

`btn-container` クラスを持つ要素の中に、2つの `button` 要素を追加してください。

最初のボタン：

- Id: `view-cart-btn`
- クラス: `btn`
- テキスト: `View Cart`

2つ目のボタン：

- Id: `checkout-btn`
- クラス: `btn`
- テキスト: `Checkout`


おめでとうございます！ div、クラス、idを使ってコンテンツを整理し、本屋のページの構造を無事に作成できました。

# --hints--

`button` 要素の中に、2つの `.btn-container` 要素がネストされているはずです。

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

`btn-container` クラスを持つ要素の中の最初のボタンには、値が `id` の `view-cart-btn` 属性が必要です。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

`btn-container` クラスを持つ要素の中の最初のボタンには、値が `class` の `btn` 属性が必要です。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

最初のボタンのテキストは `View Cart` であるべきです。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

`btn-container` クラスを持つ要素の中の2つ目のボタンには、値が `id` の `checkout-btn` 属性が必要です。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

`class="btn-container"` を持つ要素の中の2つ目のボタンには、値が `class` の `btn` 属性が必要です。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

2つ目のボタンのテキストは `Checkout` であるべきです。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
  <p>Review your selections and continue to checkout.</p>
  
  <div class="btn-container">
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
