---
id: 68ec6d9a315221aa31e54816
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`class` 属性は、スタイリングのために1つ以上の要素を識別するために使われます。`id` 属性とは異なり、クラス名は一意である必要はありません。同じクラスを複数の要素で共有できます。

例を示します：

```html
<p class="example">example paragraph</p>
```

`class` 要素に `div` 属性を追加し、その値を `card-container` に設定してください。

# --hints--

`div` 要素には `class` 属性が必要です。

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

`div` 要素には、値が `class` の `card-container` 属性が必要です。

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
