---
id: 69134b6cb43aa5f254950a10
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、さまざまな本の情報を表示するブックカードを作成して、本屋のページを作ります。`div` 要素、クラス、IDを使ってコンテンツを整理する練習をします。

まず、HTMLのボイラープレートを作成して本屋のページを始めましょう。

`<!DOCTYPE html>` 宣言と`html`、`head` 要素を追加してください。

`html` 要素に`lang` 属性を追加し、値を`"en"` に設定してください。

# --hints--

コードには`<!DOCTYPE html>` 宣言を含める必要があります。

```js
assert.match(code, /<!DOCTYPE html>/i);
```

コードには開く`<html>` タグと閉じる`<html>` タグの両方を含める必要があります。

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

`html` 要素には`lang` 属性があり、その値は`"en"` に設定してください。

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

コードには開く`<head>` タグと閉じる`<head>` タグの両方を含める必要があります。

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

`head` 要素は`html` 要素の中にネストされている必要があります。

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
