---
id: 69134d14150d51f8066a3246
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

HTMLドキュメントの構造を改善して、ページが正しくエンコードされるようにしましょう。

`head` 要素の中に `<meta charset="UTF-8">` 要素を追加してください。

最後に、`head` セクションの下に `body` 要素を追加してください。ここにページのすべての表示コンテンツが入ります。

# --hints--

コードには `<meta charset="UTF-8">` 要素が含まれている必要があります。

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

コードには開く `<body>` タグと閉じる `<body>` タグが含まれている必要があります。

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

`body` 要素は `head` 要素の下にある必要があります。

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
