---
id: 691f85085787a8c4ca589fbb
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`title` 要素の中に `head` 要素を追加してください。

ページタイトルを `XYZ Bookstore Page` に設定してください。

# --hints--

コードには `<title>` の開始タグと終了タグの両方を含めてください。

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

タイトルは `XYZ Bookstore Page` にしてください。

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

`title` 要素は `head` 要素の中にネストしてください。

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
