---
id: 64dce7ee2494e82bf4205a3f
title: Passo 24
challengeType: 20
dashedName: step-24
---

# --description--

O outro caso em que você precisa mover o disco necessariamente de `source` para `target` é quando a lista source **não** está vazia **e** o último disco em `source` é menor que o último disco em `target`.

Adicione uma declaração `elif` para verificar esta condição.Então, defina a variável `forward` como `True` se a condição for atendida.

# --hints--

Você deve ter uma declaração `elif` para verificar se `rods[source]` não está vazio e `rods[source][-1] < rods[target][-1]`.

```js
const allowedRes = [
    "rods\\s*\\[\\s*source\\s*\\]",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>\\s*0",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>=\\s*1",
];
const re = new RegExp(`elif\\s+((${allowedRes.join(")|(")}))\\s+and\\s+(rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*<\\s*rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]|rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*>\\s*rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\])\\s*:`);
assert.match(code, re);
```

Você deve definir `forward` como `True` dentro do seu novo comando `elif`.Você não deve alterar as atribuições anteriores de `forward`.

```js
const allowedRes = [
    "rods\\s*\\[\\s*source\\s*\\]",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>\\s*0",
    "len\\(rods\\s*\\[\\s*source\\s*\\]\\s*\\)\\s*>=\\s*1",
];
const re = new RegExp(`elif\\s+((${allowedRes.join(")|(")}))\\s+and\\s+(rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*<\\s*rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]|rods\\s*\\[\\s*target\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\]\\s*>\\s*rods\\s*\\[\\s*source\\s*\\]\\s*\\[\\s*-\\s*1\\s*\\])\\s*:\\s+forward\\s*=\\s*True`);
assert.match(code, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
--fcc-editable-region--            
            if not rods[target]:
                forward = True
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
