---
id: 64dc985e6720a23edac63b51
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Os movimentos permitidos do disco entre as hastes exibem um padrão repetitivo que ocorre a cada três movimentos.Por exemplo, movimentos entre a haste `A` e a haste `C` são permitidos no primeiro, no quarto e no sétimo movimento, onde o resto da divisão entre o número do movimento e 3 é igual a 1.

Dentro do `for` criado anteriormente, substitua a chamada existente de `print()` por uma declaração `if` que é acionada quando `(i + 1) % 3 == 1`.Dentro desta declaração `if`, imprima `f'Move {i + 1} allowed between {source} and {target}'` usando uma f-string.Por favor, note que `i + 1` é o número do movimento já que `i` é zero durante a primeira iteração.

# --hints--

Você deve ter uma declaração `if` que deve ser executada quando `(i + 1) % 3 == 1`. 

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

Você deve imprimir a string `f'Move {i + 1} allowed between {source} and {target}'` usando uma f-string.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
