---
id: 64dc8b1faf9c60304ca8a9b6
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste projeto, você irá resolver o quebra-cabeça matemático conhecido como a Torre de Hanoi. O quebra-cabeça consiste em três hastes e um número de discos de diferentes diâmetros.

O objetivo deste quebra-cabeça é mover os discos da primeira haste para a terceira haste, seguindo regras específicas que restringem colocar um disco maior sobre um menor.

Comece criando um dicionário vazio chamado `rods` para representar as hastes. 

# --hints--

Você deve ter uma variável chamada `rods`.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

Sua variável `rods` deve ser um dicionário vazio.

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
