---
id: 64dc922df2919e3c38cead61
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

No início do seu código, declare uma variável chamada `NUMBER_OF_DISKS` para armazenar o número de discos e atribua a ela o valor `3`.Então, substitua o primeiro argumento passado para a função `range()` pela sua nova variável.

# --hints--

Você deve declarar uma variável chamada `NUMBER_OF_DISKS`.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

A variável `NUMBER_OF_DISKS` deve ter o valor `3`.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

Você deve alterar o valor do primeiro argumento da função `range()` para a variável `NUMBER_OF_DISKS`.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
