---
id: 64dc97005bc7943e2978df0a
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Agora você pode remover sua chamada `print()`. Então, dentro da função `move()`, remova a palavra-chave `pass` e imprima o conteúdo do seu dicionário `rods`.

# --hints--

Você não deve ter `print(number_of_moves)` no seu código.

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

Você deve remover a palavra-chave `pass` e imprimir o dicionário `rods`.

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
