---
id: 64dc8ea01436383a88256d7f
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

O quebra-cabeça começa com os discos empilhados no primeiro pino, em tamanho decrescente, com o disco menor no topo e o disco maior na base.Você precisa preencher sua primeira lista com números que representam os vários tamanhos de disco.

Em vez de adicionar os itens manualmente à primeira lista, gere uma sequência de números decrescentes de `3` a `1` usando a função `range()` e atribua-a a `rods['A']`.Aqui, `3` representa o maior disco na parte inferior da pilha e `1` representa o menor disco na parte superior da pilha.

A sintaxe é `range(x, y, h)`, onde `x` é o inteiro inicial (inclusivo), `y` é o último inteiro (não inclusivo) e `h` é a diferença entre um número e o próximo na sequência.

# --hints--

Você deve usar a função `range()` para atribuir uma sequência de números a `rods['A']`.A sintaxe para chamar a função `range()` é `range(x, y, h)`.

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

Você deve usar a função `range()` para atribuir a sequência de números de `3` a `1` para `rods['A']`.

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
