---
id: 64dc916138dbe03bdfe6cec6
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

O objetivo da Torre de Hanoi é mover todos os discos para o último pino. Para fazer isso, você deve seguir três regras simples:

1. Você pode mover apenas os discos mais superiores
2. Você pode mover apenas um disco por vez
3. Você não pode colocar discos maiores em cima de discos menores

Abaixo do seu código existente, declare uma função vazia chamada `move`. Mais tarde, você usará essa função para mover os discos entre as hastes. Por enquanto, para evitar erros, use a palavra-chave `pass` dentro do corpo da função.

# --hints--

Você deve declarar uma função vazia chamada `move`. Lembre-se de usar a palavra-chave `pass` dentro do corpo da função com a indentação correta.

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
