---
id: 64dc92a9718fb53ca3d1bad7
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

O quebra-cabeça Torre de Hanoi pode ser resolvido em 2<sup>n</sup> - 1 movimentos, onde n é o número de discos.Declare uma variável chamada `number_of_moves` e atribua o número total de movimentos a essa variável.

O operador de potência em Python é `**`.

# --hints--

Você deve declarar uma variável chamada `number_of_moves`.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

O valor de `number_of_moves` deve ser a expressão para calcular o número de movimentos.

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
