---
id: 64dc90949ff85e3b37be40d0
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

A função `range()` retorna uma sequência imutável de números. Como você pode ver, o tipo de dado de `rods['A']` é `range`, mas você quer que seja uma lista.

Passe sua chamada `range()` para a função `list()` para fazer isso.

# --hints--

Você deve passar sua chamada `range()` para a função `list()`.

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
