---
id: 64dc9004f658183afa85f4f1
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora verifique o tipo de dado da sua chave `'A'` passando-a para a função `type()` e imprima o resultado no terminal.

# --hints--

Você deve passar sua tecla `'A'` para a função `type()`.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

Você deve imprimir o tipo de `rods['A']` usando a função `print()`.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
