---
id: 64ddce6f6e8066173208235b
title: Passo 27
challengeType: 20
dashedName: step-27
---

# --description--

Quando `forward` é `False`, o disco precisa ser movido na direção oposta. Escreva uma cláusula `else` para isso. Imprima o movimento e altere o conteúdo das listas de acordo.

# --hints--

Você deve ter uma cláusula `else` para mover os discos na direção oposta.

```js
({ test: () => assert.match(code, /rods\s*\[\s*target\s*\]\s*\.append\s*\(\s*rods\s*\[\s*source\s*\]\s*\.pop\s*\(\s*\)\s*\)\s+else\s*:/) })
```

Você deve ter `print(f'Moving disk {rods[target][-1]} from {target} to {source}')` dentro da sua cláusula `else`.

```js
({ test: () => assert.match(code, /else\s*:\s+print\s*\(\s*f('|")Moving\sdisk\s\{\s*rods\s*\[\s*target\s*\]\s*\[\s*-\s*1\s*\]\s*\}\sfrom\s\{\s*target\s*\}\sto\s\{\s*source\s*\}\1\s*\)/) })
```

Você deve remover o último elemento de `rods[target]`.

```js
({ test: () => assert.match(code, /rods\s*\[\s*target\s*\]\s*\.pop\s*\(\s*\)/) })
```

Você deve ter `rods[source].append(rods[target].pop())` no seu código.

```js
({ test: () => assert.match(code, /rods\s*\[\s*source\s*\]\s*\.append\s*\(\s*rods\s*\[\s*target\s*\]\s*\.pop\s*\(\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if not rods[target]:
                forward = True
            elif rods[source] and rods[source][-1] < rods[target][-1]:
                forward = True
--fcc-editable-region--                
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
