---
id: 64ddcc2ee182e0164c3539f8
title: Passo 26
challengeType: 20
dashedName: step-26
---

# --description--

Após imprimir o movimento, você precisa remover o último elemento da haste de origem e adicioná-lo à haste de destino.Use o método `.pop()` e o método `.append()` para isso.

# --hints--

Você deve remover o último elemento de `rods[source]`.

```js
({ test: () => assert.match(code, /rods\s*\[\s*source\s*\]\s*\.pop\s*\(\s*\)/) })
```

Você deve ter `rods[target].append(rods[source].pop())` no seu código.

```js
({ test: () => assert.match(code, /rods\s*\[\s*target\s*\]\s*\.append\s*\(\s*rods\s*\[\s*source\s*\]\s*\.pop\s*\(\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if not rods[target]:
                forward = True
            elif rods[source] and rods[source][-1] < rods[target][-1]:
                forward = True
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
