---
id: 64de1469fccce61940f498c4
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

O dicionário `rods` representará os três `rods` com seus discos.Dê as strings `'A'`, `'B'` e `'C'` como chaves e defina cada uma delas como uma lista vazia.

# --hints--

Seu dicionário `rods` deve ter uma chave `'A'`.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` deve ser uma lista vazia.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

Seu dicionário `rods` deve ter uma chave `'B'`.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` deve ser uma lista vazia.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

Seu dicionário `rods` deve ter uma chave `'C'`.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` deve ser uma lista vazia.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
