---
id: 64df496c6a8ddf1b38db1ed6
title: Passo 50
challengeType: 20
dashedName: step-50
---

# --description--

Em uma etapa anterior, você escreveu o código para mover o maior disco do subproblema para a haste de destino.

Agora, tudo o que você precisa fazer é adicionar outra chamada recursiva para mover os `n - 1` discos que você já deslocou. Copie a primeira chamada recursiva e cole-a no final do bloco `if`.

Observe que os argumentos da função não estão na ordem correta. Tente descobrir a ordem correta.

# --hints--

Você deve chamar `move(n - 1,  auxiliary, source, target)` no final da sua função `move`.

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
