---
id: 64df3f1011888113fbd3d81b
title: Passo 44
challengeType: 20
dashedName: step-44
---

# --description--

Para resolver o quebra-cabeça com recursão, a primeira coisa a fazer é dividir o problema original em subproblemas menores.

A configuração final com `n` discos empilhados até o terceiro pino em ordem decrescente pode ser obtida movendo:

- `n - 1` discos da origem para o pino auxiliar
- o maior disco da origem para o destino
- e então os `n - 1` discos da haste auxiliar para o alvo.

Então, a primeira coisa que a função `move` deve fazer é chamar a si mesma com `n - 1` como o primeiro argumento.Mas se você tentar fazer isso sem definir um caso base, você receberá um `RecursionError`.Isso acontece porque a função continua se chamando indefinidamente.

Antes do seu comentário e da sua chamada `print()`, adicione a chamada da função recursiva com `n - 1` como o primeiro argumento e certifique-se de que o corpo da função seja executado somente quando `n` for maior que zero.Por enquanto, deixe os outros argumentos na mesma ordem.

# --hints--

O corpo da função `move` deve começar com uma instrução `if` que é acionada quando `n` é maior que zero.

```js
({ test: () => assert.match(code, /def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:\s+if\s+n\s*>\s*0\s*:/) })
```

Você deve chamar `move(n - 1, source, auxiliary, target)` dentro da instrução `if`.Lembre-se de indentar sua chamada `print()`.

```js
const transformedCode = "\n" + code.replace(/\r/g, '');
const move = __helpers.python.getDef(transformedCode, "move");
const { function_parameters, function_body } = move;

assert.match(function_parameters, /n\s*,\s*source\s*,\s*auxiliary\s*,\s*target/);
const functionBodyIndent = function_body.match(/ +/)[0];
const re = new RegExp(`${functionBodyIndent}if\\s+n\\s*>\\s*0\\s*:\\s+^${functionBodyIndent}( +)move\\(\\s*n\\s*-\\s*1\\s*,\\s*source\\s*,\\s*auxiliary\\s*,\\s*target\\s*\\).+?^${functionBodyIndent}\\1print\\s*\\(\\s*rods\\s*,\\s*("|')\\\\n\\2\\s*\\)`, "ms");
assert.match(function_body, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
