---
id: 64df47b32b92301a815d5ef8
title: Passo 49
challengeType: 20
dashedName: step-49
---

# --description--

A princípio, a chamada recursiva que você acabou de adicionar lida com o subproblema de mover `n - 1` discos para a segunda haste.

Por essa razão, o argumento `target` corresponde à sua segunda haste, enquanto o argumento `auxiliary` é a terceira haste.Tenha em mente que esses continuarão trocando conforme a recursão avança.

Corrija a ordem dos argumentos trocando `target` e `auxiliary` na sua chamada recursiva.

# --hints--

Você deve modificar a ordem dos argumentos na sua chamada `move(n - 1, source, auxiliary, target)`.

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
