---
id: 650c6082e5586f9e3acfcd3b
title: Passo 52
challengeType: 20
dashedName: step-52
---

# --description--

Agora, delete o dicionário `rods` e transforme suas chaves nas variáveis `A`, `B` e `C`, respectivamente, mantendo seus valores. Refatore seu código para refletir essas mudanças. Se você vir a saída no terminal, você fez esta etapa corretamente. 

# --hints--

Você deve transformar `rods['A']` em uma variável chamada `A` mantendo seu valor.

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

Você deve transformar `rods['B']` em uma variável chamada `B` mantendo seu valor.

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

Você deve transformar `rods['C']` em uma variável chamada `C` mantendo seu valor.

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

Você deve modificar sua chamada `print()` para imprimir `A`, `B` e `C`, em vez do objeto `rods`. Mantenha o caractere de nova linha na chamada `print()`.

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

Você deve ter `target.append(source.pop())` no seu código.

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```

