---
id: 686b9720ee1d032bd77a480a
title: Tafuta makosa kwenye msahihishaji wa ISBN
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

ISBN (Nambari ya Kimataifa ya Kitabu) ni kitambulisho cha kipekee kilichopewa vitabu vya kibiashara. Ni mfuatano wa herufi wenye urefu wa herufi 10 au 13. Herufi ya mwisho ni nambari ya ukaguzi, ambayo huhesabiwa kutoka kwa nambari zilizo mbele yake, zinazojulikana kama nambari kuu.

Camperbot amejaribu kujenga msahihishaji wake wa ISBN. Hata hivyo, wamefanya makosa kadhaa njiani.

Katika maabara hii, utarekebisha msimbo uliopo na kuufanya ufanye kazi ipasavyo.

Tabia inayotarajiwa:
Wakati mtumizi anapoendesha programu, itaonyesha ujumbe wa kuomba `Enter ISBN and length: `. Mtumizi anaweza kuingiza msimbo wa ISBN anayetaka kuuthibitisha kwa muundo wa `ISBN,length`. Ingiza ISBN bila alama za mkato, ikifuatiwa na koma na urefu wake (`10` au `13`).

Mifano ya viingizo:
`1530051126,10` kwa misimbo ya ISBN-10.
`9781530051120,13` kwa misimbo ya ISBN-13.

Jinsi ya kupata tarakimu ya ukaguzi:

- Nambari kuu ni nambari 9 za kwanza za msimbo wa ISBN-10, au nambari 12 za kwanza za msimbo wa ISBN-13. Zinapaswa zote kuwa nambari.
- Herufi ya mwisho ni nambari ya ukaguzi iliyotolewa. Nambari ya ukaguzi inayotarajiwa huhesabiwa kutoka kwa nambari kuu, na msimbo wa ISBN ni halali tu ikiwa hizi mbili zinafanana.
- Kwa mfano, katika msimbo wa ISBN-10 `1530051126`, nambari kuu ni `153005112` na nambari ya ukaguzi iliyotolewa ni `6`. Katika msimbo wa ISBN-13 `9781530051120`, nambari kuu ni `978153005112` na nambari ya ukaguzi iliyotolewa ni `0`.
- Huhitaji kujua mantiki ya kina ya hesabu katika maabara hii. Vitendakazi `calculate_check_digit_10` na `calculate_check_digit_13` vitahakikisha hesabu na kurudisha nambari ya ukaguzi inayotarajiwa kama mfuatano wa herufi.
- Nambari ya ukaguzi kwa misimbo ya ISBN-10 inaweza kuwa nambari kutoka `0` hadi `9` au herufi kubwa `X`.
- Nambari ya ukaguzi kwa misimbo ya ISBN-13 inaweza kuwa nambari kutoka `0` hadi `9`.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kurekebisha `IndentationError` katika msimbo uliopo.
1. Hata kama mtumizi hataingizi thamani zilizoachwa kwa koma, programu inapaswa kushughulikia `IndexError` bila kuanguka.
1. Wakati mtumizi hataingizi thamani zilizoachwa kwa koma, wanapaswa kuona ujumbe `Enter comma-separated values.` kwenye konsoli, na programu inapaswa maliza.
1. Hata kama mtumizi ataingiza thamani isiyo ya nambari kwa urefu, programu inapaswa kushughulikia `ValueError` bila kuanguka.
1. Wakati mtumizi ataingiza thamani isiyo ya nambari kwa urefu, wanapaswa kuona ujumbe `Length must be a number.` kwenye konsoli, na programu inapaswa maliza.
1. Unapaswa kurekebisha kosa la moja katika kitendakazi `validate_isbn`.
1. Unapaswa kurekebisha `TypeError` katika msimbo uliopo unaotokea wakati mtumizi anaingiza msimbo halali wa ISBN.
1. Unapaswa kurekebisha `IndexError` katika msimbo uliopo wakati mtumizi anaingiza msimbo halali wa ISBN.
1. Hata kama mtumizi ataingiza msimbo usio sahihi wa ISBN wenye herufi zisizo za nambari katika nambari kuu, programu inapaswa kushughulikia `ValueError` inayotokea bila kuanguka.
1. Wakati mtumizi ataingiza msimbo usio sahihi wa ISBN wenye herufi zisizo za nambari katika nambari kuu, wanapaswa kuona ujumbe `Invalid character was found.` kwenye konsoli.
1. Wakati mtumizi ataingiza `1530051126,10`, wanapaswa kuona ujumbe `Valid ISBN Code.`
1. Wakati mtumizi ataingiza `9781530051120,13`, wanapaswa kuona ujumbe `Valid ISBN Code.`

**Muhimu**: utahitaji kutoa maoni kwenye wito wa `main()` katika nafasi ya jumla ili vipimo vifanye kazi ipasavyo.

Ukipata mradi huu, mtumizi anapaswa kuona ujumbe zifuatazo kulingana na thamani wanazoingiza.

| Msimbo wa ISBN | Urefu | Ujumbe | Mfano wa ingizo |
|-----------|--------|---------| ------------- |
| Halali | Halali | `Valid ISBN Code.` | `1530051126,10` |
| Nambari isiyo halali | Halali | `Invalid ISBN Code.` | `1530051125,10` |
| Haifai urefu uliotajwa au imeachwa tupu | Halali | `ISBN-10 code should be 10 digits long.` au `ISBN-13 code should be 13 digits long.`, kulingana na urefu waliouingiza. | `9781530051120,10` au `1530051126,13` |
| Ina herufi zisizo za nambari katika nambari kuu | Halali | `Invalid character was found.` | `15-0051126,10` |
| Yoyote | Nambari isiyo halali | `Length should be 10 or 13.` | `1530051126,9` |
| Yoyote | Ina herufi zisizo za nambari au imeachwa tupu | `Length must be a number.` | `1530051125,A` |
| Haijaachwa kwa koma | Haijaachwa kwa koma | `Enter comma-separated values.` | `1530051125` |

Unaweza kutumia thamani zifuatazo kujaribu msimbo wako kwa mikono ikiwa unataka.

Mifano ya viingizo kwa misimbo halali ya ISBN-10:
`1530051126,10`
`9971502100,10`
`080442957X,10`

Mifano ya viingizo kwa misimbo halali ya ISBN-13:
`9781530051120,13`
`9781947172104,13`

# --hints--

Unapaswa kutoa maoni kwenye wito wa kitendakazi cha `main` ili kuruhusu vipimo vingine kufanya kazi ipasavyo.

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

Unapaswa kuwa na kitendakazi cha `validate_isbn`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

Unapaswa kuwa na kitendakazi cha `calculate_check_digit_10`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

Unapaswa kuwa na kitendakazi cha `calculate_check_digit_13`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

Unapaswa kuwa na kitendakazi cha `main`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

Wakati mtumizi ataingiza thamani ambazo hazijaachwa kwa koma, unapaswa kuona ujumbe `Enter comma-separated values.` kwenye konsoli.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Wakati mtumizi anaingiza thamani isiyo ya nambari kwa urefu, unapaswa kuona ujumbe `Length must be a number.` kwenye konsoli.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Wakati mtumizi ataingiza msimbo usio sahihi wa ISBN wenye herufi zisizo za nambari katika nambari kuu, unapaswa kuona ujumbe `Invalid character was found.` kwenye konsoli.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Wakati mtumizi anaingiza `1530051126,10`, unapaswa kuona ujumbe `Valid ISBN Code.` kwenye konsoli.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Wakati mtumizi anaingiza `9781530051120,13`, unapaswa kuona ujumbe `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Wakati mtumizi ataingiza `1530051125,10`, unapaswa kuona ujumbe `Invalid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

Wakati mtumizi anaingiza `9781530051120,10`, unapaswa kuona ujumbe `ISBN-10 code should be 10 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

Wakati mtumizi anaingiza `1530051126,13`, unapaswa kuona ujumbe `ISBN-13 code should be 13 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

Wakati mtumizi anaingiza `15-0051126,10`, unapaswa kuona ujumbe `Invalid character was found.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Wakati mtumizi anaingiza `1530051126,9`, unapaswa kuona ujumbe `Length should be 10 or 13.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

Wakati mtumizi anaingiza `1530051125,A`, unapaswa kuona ujumbe `Length must be a number.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Wakati mtumizi anaingiza `1530051125`, unapaswa kuona ujumbe `Enter comma-separated values.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Wakati mtumizi anaingiza `9971502100,10`, unapaswa kuona ujumbe `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Wakati mtumizi anaingiza `080442957X,10`, unapaswa kuona ujumbe `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Wakati mtumizi anaingiza `9781947172104,13`, unapaswa kuona ujumbe `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
