---
id: a26cbbe9ad8655a977e1ceb5
title: 建置一個最長字搜尋應用程式
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

在本實驗中，你將建置一個函式（程式），該函式（程式）會傳回所提供句子中最長字的長度。

例如，在句子 `"The quick brown fox jumped over the lazy dog"` 中，最長的字是 `"jumped"`，長度為 6。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `findLongestWordLength` 的函式（程式），該函式以字串作為引數。
2. 該函式（程式）應傳回字串中最長單字的長度。

# --hints--

你應該創建一個名為 `findLongestWordLength` 的函式（程式）。

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength` 應該只有一個參數。

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` 應該傳回一個數字。

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` 應該傳回 `6`。

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")` 應該傳回 `5`。

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")` 應該傳回 `6`。

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("Googling do a barrel roll")` 應該傳回 `8`。

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` 應該傳回 `8`。

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` 應該傳回 `19`。

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
