---
id: 669e2f60e83c011754f711f9
title: Jenga ukurasa wa wakala wa usafiri
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na tamko la `DOCTYPE`.
1. Unapaswa kuwa na kipengele cha `html` chenye `lang` kimesetwa kuwa `en`.
1. Unapaswa kuwa na kipengele cha `head` chenye kipengele cha `meta` kisicho na maudhui chenye `charset` kimesetwa kuwa `utf-8` na `title` chenye maandishi ya `Travel Agency Page`.
1. Unapaswa kuwa na lebo ya `meta` katika kipengele chako cha `head` chenye maelezo mafupi ya tovuti yako kwa ajili ya SEO.
1. Unapaswa kuwa na kipengele cha `h1` kuwasilisha maeneo yako ya kusafiria.
1. Unapaswa kuwa na aya chini ya kipengele cha `h1` inayotambulisha fursa za usafiri.
1. Unapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Packages`.
1. Unapaswa kuwa na kipengele cha `p` kinachotambulisha kwa kifupi vifurushi mbalimbali.
1. Unapaswa kuwa na kipengele cha orodha isiyopangwa chenye vitu viwili vya orodha. Vitu viwili vya orodha vinapaswa kuwa na maandishi ya `Group Travels` na `Private Tours`, mtawalia. Maandishi ya kila kipengele cha orodha yanapaswa kufunikwa na kipengele cha nanga.
1. Unapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Top Itineraries`.
1. Unapaswa kuwa na angalau vipengele vitatu vya `figure`, kila kikiwa na kipengele cha nanga na kipengele cha `figcaption`.
1. Vipengele vitatu vya nanga vinapaswa kuwa na kipengele cha `img` chenye sifa ya kipengele inayofaa ya `alt` na sifa ya kipengele ya `src` iliyowekwa kuwa picha halali kama maudhui yao. Unaweza kutumia `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg`, na `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` kama unavyotaka.
1. Lebo zako zote tano za nanga zinapaswa kuwa na sifa ya kipengele ya `href` yenye thamani ya `https://www.freecodecamp.org/learn` na sifa ya kipengele ya `target` yenye thamani ya `_blank`.

# --hints--

Ukurasa wako wa wakala wa usafiri unapaswa kuwa na tamko la `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Unapaswa kuwa na kipengele cha `html` chenye `lang` kimesetwa kuwa `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

Unapaswa kuwa na kipengele cha `head` ndani ya kipengele cha `html`.

```js
assert.match(
  code,
  /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Unapaswa kuwa na vipengele viwili vya `meta` ndani ya kipengele chako cha `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*(<\s*meta[\s\S]*>[\s\S]*){2,}<\/\s*head\s*>/i);
```

Kipengele kimoja cha `meta` kinapaswa kuwa na sifa ya kipengele ya `name` yenye thamani ya `description` na sifa ya kipengele ya `content` isiyo tupu.

 ```js
 assert.match(code,
  /<\s*meta\s+[^>]*name\s*=\s*["']description["'][^>]*content\s*=\s*["'][^"']+?["'][^>]*>/i,
);
```

One `meta` element should have its `charset` attribute set to `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

You should have `title` element within your `head` element.

```js
assert.match(
  code,
  /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i
);
```

Your `title` element should have your travel agency name.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

You should have a `body` element within your `html` element.

```js
assert.match(
  code,
  /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i
);
```

You should have an `h1` element to present your travel destinations.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

You should only have one `h1` element.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

You should have a `p` element below the `h1` element.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Your first `p` element should introduce the travel opportunities.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Your first `h2` element should have the text `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

You should have a `p` element below your first `h2` element.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Your second `p` element should introduce briefly the various packages.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

You should have an unordered list element below your second `p` element.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

You should have two items in your unordered list.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Both your list items should contain an anchor element.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

The anchor element of your first list item should wrap the text `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

The anchor element of your second list item should wrap the text `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

You should have an `h2` element after your unordered list.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Your second `h2` element should have the text `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

You should have at least three `figure` elements.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Each `figure` element should contain an anchor element as its first nested element.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Each `figure` element should contain a `figcaption` element as its second nested element.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Each `figcaption` should contain some text.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Each of the `a` elements nested in your `figure` elements should contain an image.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Each `img` element should have a valid `src` attribute.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Each `img` element should have an `alt` attribute with an appropriate value.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Each `a` element should have an `href` attribute with the value of `https://www.freecodecamp.org/learn`. Don't forget the links in the list items.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Each `a` element should have a `target` attribute with the value of `_blank`. Don't forget the links in the list items.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
