---
id: 669e2f60e83c011754f711f9
title: Jenga ukurasa wa wakala wa usafiri
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na tamko la `DOCTYPE`.
1. Unapaswa kuwa na kipengele cha `html` chenye `lang` imewekwa kuwa `en`.
1. Unapaswa kuwa na kipengele cha `head` chenye kipengele cha `meta` kisicho na maudhui chenye `charset` imewekwa kuwa `utf-8` na `title` yenye maandishi `Travel Agency Page`.
1. Unapaswa kuwa na lebo ya `meta` ndani ya kipengele chako cha `head` inayojumuisha maelezo mafupi ya tovuti yako kwa ajili ya SEO.
1. Unapaswa kuwa na kipengele cha `h1` kuwasilisha maeneo yako ya kusafiria. Hiki kinapaswa kuwa kipengele pekee cha `h1` kwenye ukurasa wako.
1. Unapaswa kuwa na aya chini ya kipengele cha `h1` inayotambulisha fursa za kusafiri.
1. Unapaswa kuwa na kipengele cha `h2` chenye maandishi `Packages`.
1. Unapaswa kuwa na kipengele cha `p` kinachotambulisha kwa ufupi vifurushi mbalimbali.
1. Unapaswa kuwa na kipengele cha orodha isiyopangwa chenye vitu viwili vya orodha. Vitu viwili vya orodha vinapaswa kuwa na maandishi `Group Travels` na `Private Tours`, mtawalia. Maandishi ya kila kipengele cha orodha yanapaswa kuzungukwa na kipengele cha nanga.
1. Unapaswa kuwa na kipengele cha `h2` chenye maandishi `Top Itineraries`.
1. Unapaswa kuwa na angalau vipengele vitatu vya `figure`, kila kikiwa na kipengele cha nanga na kipengele cha `figcaption` kinachoelezea picha ya mchoro.
1. Vipengele vitatu vya nanga vinapaswa kuwa na kipengele cha `img` chenye sifa ya `alt` inayofaa na sifa ya `src` iliyowekwa kwa picha halali kama maudhui yao. Unaweza kutumia `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg`, na `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` kama unavyotaka.
1. Vitu vyote vitano vya nanga vinapaswa kuwa na sifa ya `href` yenye thamani `https://www.freecodecamp.org/learn` na sifa ya `target` yenye thamani `_blank`.

# --hints--

Ukurasa wako wa wakala wa usafiri unapaswa kuwa na tamko la `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Unapaswa kuwa na kipengele cha `html` chenye `lang` kimesetwa kuwa `en`.

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/html\s*>/i);
```

Unapaswa kuwa na kipengele cha `head` ndani ya kipengele cha `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

Kipengele chako cha `head` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/head\s*>/i);
```

Unapaswa kuwa na kipengele cha `title` ndani ya kipengele chako cha `head`.

```js
assert.isNotNull(document.querySelector('head title'));
```

Kipengele chako cha `title` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/title\s*>/i);
```

Unapaswa kuwa na vipengele viwili vya `meta` ndani ya kipengele chako cha `head`.

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

Kipengele kimoja cha `meta` kinapaswa kuwa na sifa ya kipengele ya `name` yenye thamani ya `description` na sifa ya kipengele ya `content` isiyo tupu.

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

Kipengele kimoja cha `meta` kinapaswa kuwa na sifa ya `charset` imewekwa kuwa `utf-8`.

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Kipengele chako cha `title` kinapaswa kuwa na jina la wakala wako wa usafiri.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Unapaswa kuwa na kipengele cha `body` ndani ya kipengele chako cha `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

Kipengele chako cha `body` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/body\s*>/i);
```

Unapaswa kuwa na kipengele cha `h1` kuwasilisha maeneo yako ya kusafiria.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Unapaswa kuwa na kipengele kimoja tu cha `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Unapaswa kuwa na kipengele cha `p` chini ya kipengele cha `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Kipengele chako cha kwanza cha `p` kinapaswa kuanzisha fursa za usafiri.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Kipengele chako cha kwanza cha `h2` kinapaswa kuwa na maandishi `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

Unapaswa kuwa na kipengele cha `p` chini ya kipengele chako cha kwanza cha `h2`.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Kipengele chako cha pili cha `p` kinapaswa kuanzisha kwa kifupi vifurushi mbalimbali.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Unapaswa kuwa na kipengele cha orodha isiyopangwa chini ya kipengele chako cha pili cha `p`.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

Unapaswa kuwa na vitu viwili katika orodha yako isiyopangwa.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Vitu vyote viwili vya orodha yako vinapaswa kuwa na kipengele cha tagi ya nanga.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

Kipengele cha tagi ya nanga cha kipengele chako cha kwanza cha orodha kinapaswa kuzunguka maandishi `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

Kipengele cha tagi ya nanga cha kipengele chako cha pili cha orodha kinapaswa kuzunguka maandishi `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

Unapaswa kuwa na kipengele cha `h2` baada ya orodha isiyopangwa.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Kipengele chako cha pili cha `h2` kinapaswa kuwa na maandishi `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

Unapaswa kuwa na angalau vipengele vitatu vya `figure`.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Kipengele kila `figure` kinapaswa kuwa na kipengele cha nanga kama kipengele chake cha kwanza kilichopangwa ndani.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Kila kipengele cha `figure` kinapaswa kuwa na kipengele cha `figcaption` kama kipengele chake cha pili kilichopangwa ndani.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Kila `figcaption` inapaswa kuwa na maandishi fulani.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Kila moja ya vipengele vya `a` vilivyopangwa ndani ya vipengele vyako vya `figure` inapaswa kuwa na picha.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Kipengele kila cha `img` kinapaswa kuwa na sifa ya `src` halali.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Kipengele kila cha `img` kinapaswa kuwa na sifa ya `alt` yenye thamani inayofaa.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Kipengele kila cha `a` kinapaswa kuwa na sifa ya `href` yenye thamani ya `https://www.freecodecamp.org/learn`. Usisahau viungo katika vitu vya orodha.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Kipengele kila cha `a` kinapaswa kuwa na sifa ya `target` yenye thamani ya `_blank`. Usisahau viungo katika vitu vya orodha.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
