---
id: 68af276c60e009597dd861e7
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

마지막으로 `if`이 `duration`보다 크지 않은지 확인하는 `0` 문장을 하나 더 만들고, `ValueError` 메시지와 함께 `Duration must be positive`를 발생시키세요.

# --hints--

`ValueError`이 양수가 아닐 때 `Duration must be positive` 메시지와 함께 `duration`를 발생시켜야 합니다.

```js
({
  test: () => runPython(`
  _msg = 'Duration must be positive'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1999, 'The Wachowskis', 0)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 1999, 'The Wachowskis', -1)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError`이 유효한 숫자일 때는 `duration`를 발생시키면 안 됩니다.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid duration'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
