---
id: 68af275860e009597dd861e5
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

첫 번째 영화는 1895년에 만들어졌습니다. `if`이 `year`보다 작은지 확인하는 또 다른 `1895` 문장을 만들고, `ValueError` 메시지와 함께 `Year must be 1895 or later`를 발생시키세요.

# --hints--

`ValueError`이 `Year must be 1895 or later`보다 작을 때 `year` 메시지와 함께 `1895`를 발생시켜야 합니다.

```js
({
  test: () => runPython(`
  _msg = 'Year must be 1895 or later'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1894, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 0, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError`이 유효한 숫자일 때는 `year`를 발생시키면 안 됩니다.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1895, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid year'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
