---
id: 68af277560e009597dd861e8
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

이제 `print(movie1)` 줄을 삭제하세요. 그런 다음 `movie1` 할당을 `try` 블록 안에 넣고 `except` 절을 만들어 `ValueError as e`를 잡아 `Validation Error:`와 공백, 그리고 오류 메시지를 출력하는 f-문자열을 사용하세요.

`movie1`를 인스턴스화할 때 전달하는 인수를 변경하여 다양한 오류 메시지를 확인해도 좋습니다.

# --hints--

코드에 `print(movie1)`가 없어야 합니다.

```js
({
  test: () => runPython(`assert not _Node(_code).has_call("print(movie1)")`)
})
```

`try`가 포함된 `movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)` 블록이 있어야 합니다.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)")`)
})
```

`except`를 잡는 `ValueError as e` 절이 있어야 합니다.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].has_except("ValueError", "e")`)
})
```

`Validation Error:` 블록 내에서 `except`와 공백, 그리고 오류 메시지를 출력하는 f-문자열을 사용해야 합니다.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_except("ValueError", "e").has_call("print(f'Validation Error: {e}')")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
--fcc-editable-region--
```
