---
id: 68af273960e009597dd861e2
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

디버깅과 출력을 위해 읽기 쉬운 표시를 원합니다. `__str__` 메서드를 만들어 `title (year) - duration min, director` 형식의 문자열을 반환하세요 (여기서 `title`, `year`, `duration`, `director`는 해당 속성 값으로 교체되어야 합니다). 이렇게 하면 영화를 출력할 때 깔끔한 요약이 나옵니다.

# --hints--

`Movie` 클래스에는 `__str__` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

`__str__` 메서드는 `title (year) - duration min, director` 형식의 문자열을 반환해야 하며, `title`, `year`, `duration`, `director`는 해당 속성 값으로 교체되어야 합니다.

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    
--fcc-editable-region--
```
