---
id: 68af278a60e009597dd861ea
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`MediaCatalogue` 클래스에 미디어 항목을 카탈로그에 추가하는 메서드가 필요합니다. 항목 `add`를 받아 `media_item` 목록에 추가하는 `items` 메서드를 만드세요.

# --hints--

`MediaCatalogue`에는 `add` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").has_function("add")`)
})
```

`add` 메서드에는 두 개의 매개변수 `self`과 `media_item`이 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("add").has_args("self, media_item")`)
})
```

`add` 메서드는 전달받은 항목을 `items` 속성에 추가해야 합니다.

```js
({
    test: () => runPython(`
    _catalogue = MediaCatalogue()
    _catalogue.add("first")
    assert _catalogue.items == ["first"]
    _catalogue.add("second")
    assert _catalogue.items == ["first", "second"]
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
