---
id: 68af274560e009597dd861e3
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

이제 원하는 인수를 전달하여 `Movie`의 인스턴스를 만들고 `movie1`이라는 변수에 할당하세요. 그런 다음 `movie1`를 콘솔에 출력하세요.

# --hints--

`movie1`라는 변수가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

`Movie`의 인스턴스를 `movie1` 변수에 할당해야 합니다.

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

`movie1`를 콘솔에 출력해야 합니다.

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
```
