---
id: 68af277e60e009597dd861e9
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

이제 `MediaCatalogue`라는 새 클래스를 만들고, `__init__` 메서드를 추가해 `items` 속성을 빈 리스트로 설정하세요. 이 속성은 카탈로그의 항목 목록을 나타냅니다.

# --hints--

`MediaCatalogue`라는 이름의 클래스가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("MediaCatalogue")`)
})
```

`MediaCatalogue` 클래스에는 `__init__` 메서드가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_class("MediaCatalogue").has_function("__init__")
    `)
})
```

`__init__` 메서드는 `items` 속성을 빈 리스트로 설정해야 합니다.

```js
({
    test: () => runPython(`
    _catalogue = MediaCatalogue()
    assert _catalogue.items == []
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
