---
id: 68af276360e009597dd861e6
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

`if`이 비어 있거나 공백만 포함하는지 확인하는 또 다른 `director` 문장을 만들고, 메시지 `ValueError`와 함께 `Director cannot be empty`를 발생시키세요.

# --hints--

`ValueError`이 비어 있거나 공백만 포함할 때 메시지 `Director cannot be empty`와 함께 `director`를 발생시켜야 합니다.

```js
({
  test: () => runPython(`
  _msg = 'Director cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1999, ' ', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 1999, '  ', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError`이 유효한 문자열일 때는 `director`를 발생시키면 안 됩니다.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid director'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
