---
id: 68af272f60e009597dd861e1
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 미디어 카탈로그를 만들 것입니다. `Movie`라는 이름의 클래스를 정의하는 것부터 시작하세요.

클래스에 `__init__` 메서드를 만들고 다음 매개변수를 포함하세요: `self`, `title`, `year`, `director`, `duration`. `__init__` 메서드 안에서 `title`, `year`, `director`, `duration` 매개변수를 같은 이름의 인스턴스 속성에 할당하세요.

# --hints--

`Movie`라는 이름의 클래스를 만들어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

`Movie` 클래스에 `__init__` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

`__init__` 메서드에 다음 매개변수가 있어야 합니다: `self`, `title`, `year`, `director`, `duration`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

`title` 메서드 안에서 `self.title`을 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

`year` 메서드 안에서 `self.year`를 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

`director` 메서드 안에서 `self.director`을 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

`duration` 메서드 안에서 `self.duration`을 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
