---
id: 68af274e60e009597dd861e4
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이전 단계에서 기본 영화 객체를 만들었습니다. 이제 영화가 잘못된 속성으로 인스턴스화되지 않도록 유효성 검사를 추가할 것입니다.

`__init__` 메서드 시작 부분에 `if`가 비어 있거나 공백만 포함하는지 확인하는 `title` 문을 만드세요 (조건을 만들 때 `strip` 메서드를 사용할 수 있습니다).

그럴 경우, 메시지 `ValueError`와 함께 `Title cannot be empty`를 발생시키세요.

# --hints--

`ValueError`가 비어 있거나 공백만 포함할 때 메시지 `Title cannot be empty`와 함께 `title`를 발생시켜야 합니다.

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError`가 유효한 문자열일 때는 `title`를 발생시키면 안 됩니다.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
