---
id: 68b6d9692a6526838265ad7f
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

`TVSeries` 객체는 부모 `Movie`에 비해 추가 속성을 가진 상태로 인스턴스화됩니다. 따라서 다음 매개변수를 가진 `__init__` 메서드를 만드세요: `self`, `title`, `year`, `director`, `duration`, `seasons`, 그리고 `total_episodes`.

`__init__` 메서드 안에서 `seasons`과 `total_episodes`를 같은 이름의 속성에 할당하세요.

# --hints--

`TVSeries` 클래스는 `__init__` 메서드를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").has_function("__init__")`)
})
```

`__init__` 메서드는 다음 매개변수를 가져야 합니다: `self`, `title`, `year`, `director`, `duration`, `seasons`, 그리고 `total_episodes`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_args("self, title, year, director, duration, seasons, total_episodes")`)
})
```

`seasons`를 `self.seasons`에 `__init__` 메서드 안에서 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("self.seasons = seasons")`)
})
```

`total_episodes`를 `self.total_episodes`에 `__init__` 메서드 안에서 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("self.total_episodes = total_episodes")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    pass
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
