---
id: 68b6cc518177b0a11c3fe3b1
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

`try` 블록 앞에 `MediaCatalogue`의 인스턴스를 생성하고 `catalogue`라는 변수에 할당하세요.

그런 다음 `add` 메서드를 사용해 `movie1`를 카탈로그에 추가하세요. 마지막으로 `catalogue`를 콘솔에 출력하세요.

# --hints--

`MediaCatalogue` 블록 전에 `catalogue`의 인스턴스를 `try`라는 변수에 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_stmt("catalogue = MediaCatalogue()")`)
})
```

`catalogue.add(movie1)` 블록 안에 `try`가 있어야 합니다.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("catalogue.add(movie1)")`)
})
```

`catalogue` 블록 끝에서 `try`를 출력해야 합니다.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().is_ordered("catalogue.add(movie1)", "print(catalogue)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result
--fcc-editable-region--


try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)

--fcc-editable-region--
except ValueError as e:
    print(f'Validation Error: {e}')
```
