---
id: 68b6d5244e6ce639e52fd1e7
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

만약 다른 유형의 미디어 항목, 예를 들어 TV 시리즈를 나타내고 싶다면 어떻게 할까요? 또 다른 클래스를 만들고 `Movie`에서 상속받아 모든 메서드와 속성에 접근할 수 있게 할 수 있습니다.

빈 클래스를 `TVSeries`라는 이름으로 만들고 `Movie`에서 상속받게 하세요. `Child` 클래스가 `Parent` 클래스를 상속받는 방법은 다음과 같습니다.

```py
class Parent:
    ...

class Child(Parent):
    ...
```

# --hints--

`TVSeries`라는 이름의 클래스가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("TVSeries")`)
})
```

`TVSeries` 클래스는 `Movie`에서 상속받아야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").inherits_from("Movie")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
