---
id: 68b704862dd6559e890c2a86
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

`if` 문장 다음에 `result`이라는 변수를 만드세요. 이 변수는 메서드에서 반환할 문자열입니다.

그래서 `Media Catalogue (N items):` 형식의 f-문자열을 할당하세요. `N`을 `items` 리스트의 항목 수로 바꾸고 문자열이 두 개의 줄바꿈 문자로 끝나게 만드세요.

# --hints--

`result` 메서드 안에 `__str__`라는 변수가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_variable("result")`)
})
```

`f'Media Catalogue ({len(self.items)} items):\n\n'` 변수에 `result`을 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_stmt("result = f'Media Catalogue ({len(self.items)} items):\\\\n\\\\n'")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)
--fcc-editable-region--
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'
        
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
