---
id: 68b80a995417adac579f9545
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

이제 `__init__` 메서드를 완성하여 `seasons`과 `total_episodes`를 검증하세요. `if`가 `seasons`보다 작은지 확인하는 `1` 문장을 만들고, 그럴 경우 `ValueError` 메시지와 함께 `Seasons must be 1 or greater`를 발생시키세요.

그다음, `total_episodes`도 동일하게 검증하세요. 이번에는 `Total episodes must be 1 or greater` 메시지를 사용하세요.

# --hints--

`ValueError`가 `Seasons must be 1 or greater`보다 작을 때 `seasons` 메시지와 함께 `1`를 발생시켜야 합니다.

```js
({
  test: () => runPython(`
  _msg = 'Seasons must be 1 or greater'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 0, 182)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, -1, 182)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError`가 유효한 숫자일 때는 `seasons`를 발생시키면 안 됩니다.

```js
({
    test: () => runPython(`
    try:
      TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 1, 182)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid seasons'
    `)
})
```

`ValueError`가 `Total episodes must be 1 or greater`보다 작을 때 `total_episodes` 메시지와 함께 `1`를 발생시켜야 합니다.

```js
({
  test: () => runPython(`
  _msg = 'Total episodes must be 1 or greater'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 0)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, -1)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`ValueError`가 유효한 숫자일 때는 `total_episodes`를 발생시키면 안 됩니다.

```js
({
    test: () => runPython(`
    try:
      TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 1)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid total_episodes'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)
--fcc-editable-region--
        

        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
