---
id: 68b804f3be853a355f055fef
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

`super` 함수는 부모 클래스를 명시적으로 이름을 지정하지 않고 참조할 때 사용합니다. 예를 들어, `super().__init__(arg1, arg2, ...)`은 부모 클래스에 정의된 `__init__` 메서드를 호출합니다.

현재는 `__init__`에서 상속받은 `Movie` 메서드를 사실상 재정의한 상태입니다. 이는 `TVSeries` 객체가 `seasons` 속성과 `total_episodes` 속성만 가지게 된다는 뜻입니다. 이것은 원하는 바가 아닙니다.

`Movie` 클래스의 `__init__`에 정의된 모든 코드를 유지하려면, 부모의 `super`을 호출하기 위해 `__init__`을 사용하세요. 이때 메서드에 필요한 인수들을 전달해야 합니다. 이 경우에는 `title`, `year`, `director`, 그리고 `duration`입니다.

# --hints--

`super().__init__(title, year, director, duration)`을 `__init__` 메서드 안에서 호출해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("super().__init__(title, year, director, duration)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        
        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
