---
id: 68b95fa5f6e575879105e74f
title: 32단계
challengeType: 20
dashedName: step-32
---

# --description--

`MediaError` 클래스가 완성되었으니, `if` 메서드 안의 `add` 문장을 수정하세요. `TypeError`를 같은 메시지를 사용하는 `MediaError`로 바꾸세요.

또한, 메시지 다음 두 번째 인수로 `media_item`를 전달하세요. 이렇게 하면 예외 객체에 오류를 일으킨 항목을 저장할 수 있습니다.

# --hints--

`MediaError`가 `Only Movie or TVSeries instances can be added` 또는 그 하위 클래스의 인스턴스가 아닐 때 `media_item` 메시지와 함께 `Movie`를 발생시켜야 합니다.

```js
({
  test: () => runPython(`
  _msg = 'Only Movie or TVSeries instances can be added'
  _expected_error = f"Expected to raise 'MediaError: {_msg}'"
  _c = MediaCatalogue()
  try:    
    _c.add("Not a movie")
  except MediaError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    _c.add({})
  except MediaError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

`media_item`에 두 번째 인수로 `MediaError`를 전달해야 합니다.

```js
({
  test: () => runPython(`
  _c = MediaCatalogue()
  _obj = "Not a movie"
  try:    
    _c.add(_obj)
  except MediaError as e:
    assert e.obj == _obj
  else:
    assert False
  `)
})
```

`MediaError`가 유효한 항목일 때는 `media_item`를 발생시키면 안 됩니다.

```js
({
    test: () => runPython(`
    _c = MediaCatalogue()
    try:
      _c.add(Movie('The Matrix', 1999, 'The Wachowskis', 136))
    except MediaError:
      assert False, 'Expected not to raise a MediaError with valid media_item'
    except TypeError:
      pass

    try:
      _c.add(TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182))
    except MediaError:
      assert False, 'Expected not to raise a MediaError with valid media_item'
    except TypeError:
      pass
    `)
})
```

`add` 메서드는 `TypeError`를 발생시키면 안 됩니다.

```js
({
    test: () => runPython(`
    _c = MediaCatalogue()
    try:
      _c.add(Movie('The Matrix', 1999, 'The Wachowskis', 136))
    except TypeError:
      assert False, 'Expected not to raise a TypeError from add method'

    try:
      _c.add(TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182))
    except TypeError:
      assert False, 'Expected not to raise a TypeError from add method'

    try:
      _c.add("Not a movie")
    except TypeError:
      assert False, 'Expected not to raise a TypeError from add method'
    except MediaError:
      pass
    `)
})
```

# --seed--

## --seed-contents--

```py
class MediaError(Exception):
    """Custom exception for media-related errors."""

    def __init__(self, message, obj):
        super().__init__(message)
        self.obj = obj

class Movie:
    """Parent class representing a movie."""
    
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    """Child class representing an entire TV series."""

    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    """A catalogue that can store different types of media items."""

    def __init__(self):
        self.items = []

    def add(self, media_item):
--fcc-editable-region--
        if not isinstance(media_item, Movie):
            raise TypeError('Only Movie or TVSeries instances can be added')
--fcc-editable-region--
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)

    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)
    catalogue.add(series1)
    series2 = TVSeries('Breaking Bad', 2008, 'Vince Gilligan', 47, 5, 62)
    catalogue.add(series2)
    
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
