---
id: 68b95956da014b1a8e32ffb3
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

이전 모듈에서 내장 클래스인 `Exception`를 상속하여 자신만의 사용자 정의 예외를 만들 수 있다는 것을 배웠습니다.

```py
class CustomError(Exception):
    ...
```

`MediaError`라는 클래스를 만드세요. 이 클래스는 `Exception` 클래스를 상속받고, `Custom exception for media-related errors.`라는 문서 문자열을 추가하세요.

# --hints--

`MediaError`라는 클래스를 만들어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("MediaError")`)
})
```

`MediaError` 클래스는 `Exception`를 상속받아야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaError").inherits_from("Exception")`)
})
```

`Custom exception for media-related errors.` 클래스에 `MediaError`라는 문서 문자열을 추가해야 합니다.

```js
({
    test: () => runPython(`
    assert MediaError.__doc__ == 'Custom exception for media-related errors.'
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Movie:
    """Parent class representing a movie."""
    
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    """Child class representing an entire TV series."""

    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    """A catalogue that can store different types of media items."""

    def __init__(self):
        self.items = []

    def add(self, media_item):
        if not isinstance(media_item, Movie):
            raise TypeError('Only Movie or TVSeries instances can be added')
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)

    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)
    catalogue.add(series1)
    series2 = TVSeries('Breaking Bad', 2008, 'Vince Gilligan', 47, 5, 62)
    catalogue.add(series2)

    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
