---
id: 68b95c952d50ed52abe02ed3
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

`MediaError` 클래스에 `__init__` 메서드를 `self`, `message`, `obj` 매개변수와 함께 만드세요.

그 안에서 `super().__init__(message)`를 호출하여 `Exception` 클래스에 이미 있는 오류 메시지를 저장하고 표시하는 메커니즘이 제대로 작동하는지 확인하세요.

그 후 `obj`를 `self.obj`에 할당하세요.

# --hints--

`MediaError` 클래스에 `__init__` 메서드가 `self`, `message`, `obj` 매개변수를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaError").find_function("__init__").has_args("self, message, obj")`)
})
```

`super().__init__(message)` 메서드 안에서 `__init__`를 호출해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaError").find_function("__init__").has_stmt("super().__init__(message)")`)
})
```

`obj` 메서드 안에서 `self.obj`를 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaError").find_function("__init__").has_stmt("self.obj = obj")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MediaError(Exception):
    """Custom exception for media-related errors."""


--fcc-editable-region--
class Movie:
    """Parent class representing a movie."""
    
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class TVSeries(Movie):
    """Child class representing an entire TV series."""

    def __init__(self, title, year, director, duration, seasons, total_episodes):
        super().__init__(title, year, director, duration)

        if seasons < 1:
            raise ValueError('Seasons must be 1 or greater')
        if total_episodes < 1:
            raise ValueError('Total episodes must be 1 or greater')
        
        self.seasons = seasons
        self.total_episodes = total_episodes

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.seasons} seasons, {self.total_episodes} episodes, {self.duration} min avg, {self.director}'

class MediaCatalogue:
    """A catalogue that can store different types of media items."""

    def __init__(self):
        self.items = []

    def add(self, media_item):
        if not isinstance(media_item, Movie):
            raise TypeError('Only Movie or TVSeries instances can be added')
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)

    series1 = TVSeries('Scrubs', 2001, 'Bill Lawrence', 24, 9, 182)
    catalogue.add(series1)
    series2 = TVSeries('Breaking Bad', 2008, 'Vince Gilligan', 47, 5, 62)
    catalogue.add(series2)

    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
