---
id: acda2fb1324d9b0fa741e6b5
title: Conferma la fine
challengeType: 1
forumTopicId: 16006
dashedName: confirm-the-ending
---

# --description--

Controlla se una stringa (primo argomento, `str`) termina con la stringa obiettivo data (secondo argomento, `target`).

Questa sfida _può_ essere risolta con il metodo `.endsWith()`, introdotto in ES2015. Ma per questa sfida, vogliamo che usi invece uno dei metodi substring di JavaScript.

# --hints--

`confirmEnding("Bastian", "n")` dovrebbe restituire `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` dovrebbe restituire `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` dovrebbe restituire `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` dovrebbe restituire `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` dovrebbe restituire `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` dovrebbe restituire `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` dovrebbe restituire `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` dovrebbe restituire `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing", "mountain")` dovrebbe restituire `false`.

```js
assert.isFalse(
  confirmEnding(
    'If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing',
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` dovrebbe restituire `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Il tuo codice non dovrebbe usare il metodo integrato `.endsWith()` per risolvere la sfida.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js
function confirmEnding(str, target) {
  return str;
}

confirmEnding('Bastian', 'n');
```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
