---
id: 587d78af367417b2b2512b04
title: 제품 랜딩 페이지 만들기
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301144
dashedName: build-a-product-landing-page
---

# --description--

**목표:** <a href="https://product-landing-page.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://product-landing-page.freecodecamp.rocks</a>와 기능적으로 유사한 앱을 만드세요. **이 데모 프로젝트를 복사하지 마세요**.

**사용자 스토리:**

1. 제품 랜딩 페이지에는 `header` 요소와 해당 `id="header"`가 있어야 합니다.
1. `header` 요소 안에 이미지가 보이고, 해당 `id="header-img"`가 있어야 합니다 (여기에 로고 이미지가 적합합니다).
1. `#header` 요소 안에는 `nav` 요소와 해당 `id="nav-bar"`가 있어야 합니다.
1. `nav` 요소 안에는 클래스 `nav-link`를 가진 클릭 가능한 요소가 최소 세 개 있어야 합니다.
1. `.nav-link` 요소 안의 `nav` 버튼을 클릭하면 랜딩 페이지의 해당 섹션으로 이동해야 합니다.
1. `id="video"`를 사용해 임베드된 제품 동영상을 볼 수 있어야 합니다.
1. 랜딩 페이지에는 `form` 요소와 해당 `id="form"`가 있어야 합니다.
1. 폼 안에는 이메일 주소를 입력할 수 있는 `input` 필드와 `id="email"`가 있어야 합니다.
1. `#email` 입력 필드에는 사용자가 필드 용도를 알 수 있도록 자리 표시자 텍스트가 있어야 합니다.
1. `#email` 입력 필드는 HTML5 유효성 검사를 사용해 입력된 텍스트가 이메일 주소인지 확인해야 합니다.
1. 폼 안에는 제출 `input`와 해당 `id="submit"`가 있어야 합니다.
1. `#submit` 요소를 클릭하면 이메일이 정적 페이지로 제출되어야 합니다 (이 모의 URL 사용: `https://www.freecodecamp.com/email-submit`).
1. 내비게이션 바는 항상 뷰포트 상단에 있어야 합니다.
1. 제품 랜딩 페이지에는 적어도 하나의 미디어 쿼리가 있어야 합니다.
1. 제품 랜딩 페이지는 CSS flexbox를 적어도 한 번 사용해야 합니다.

사용자 스토리를 모두 충족하고 아래 모든 테스트를 통과해야 이 프로젝트를 완료할 수 있습니다. 자신만의 스타일로 만드세요. 코딩을 즐기세요!

**참고:** HTML에 `<link rel="stylesheet" href="styles.css">`를 추가해 스타일시트를 링크하고 CSS를 적용하세요.

# --hints--

`header` 요소에 `id`가 `header`인 속성이 있어야 합니다.

```js
const el = document.getElementById('header')
assert(!!el && el.tagName === 'HEADER')
```

`img` 요소에 `id`가 `header-img`인 속성이 있어야 합니다.

```js
const el = document.getElementById('header-img')
assert(!!el && el.tagName === 'IMG')
```

`#header-img`는 `#header`의 자손 요소여야 합니다.

```js
const els = document.querySelectorAll('#header #header-img')
assert(els.length > 0)
```

`#header-img`에는 `src` 속성이 있어야 합니다.

```js
const el = document.getElementById('header-img')
assert(!!el && !!el.src)
```

`#header-img`의 `src` 값은 유효한 URL이어야 합니다 (`http`로 시작).

```js
const el = document.getElementById('header-img')
assert(!!el && /^http/.test(el.src))
```

`nav` 요소에 `id`가 `nav-bar`인 속성이 있어야 합니다.

```js
const el = document.getElementById('nav-bar')
assert(!!el && el.tagName === 'NAV')
```

`#nav-bar`는 `#header`의 자손 요소여야 합니다.

```js
const els = document.querySelectorAll('#header #nav-bar')
assert(els.length > 0)
```

`.nav-link` 안에 최소 3개의 `#nav-bar` 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('#nav-bar .nav-link')
assert(els.length >= 3)
```

각 `.nav-link` 요소에는 `href` 속성이 있어야 합니다.

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  if (!el.href) assert(false)
})
assert(els.length > 0)
```

각 `.nav-link` 요소는 랜딩 페이지의 해당 요소에 링크해야 합니다 (다른 요소의 id 값을 가진 `href` 속성이 있어야 함, 예: `#footer`).

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  const linkDestination = el.getAttribute('href').slice(1)
  if (!document.getElementById(linkDestination)) assert(false)
})
assert(els.length > 0)
```

`video` 또는 `iframe` 요소에 `id`가 `video`인 속성이 있어야 합니다.

```js
const el = document.getElementById('video')
assert(!!el && (el.tagName === 'VIDEO' || el.tagName === 'IFRAME'))
```

`#video`에는 `src` 속성이 있어야 합니다.

```js
let el = document.getElementById('video')
const sourceNode = el.children;
let sourceElement = null;
if (sourceNode.length) {
  sourceElement = [...video.children].filter(el => el.localName === 'source')[0];
}
if (sourceElement) {
  el = sourceElement;
}
assert(el.hasAttribute('src'));
```

`form` 요소에 `id`가 `form`인 속성이 있어야 합니다.

```js
const el = document.getElementById('form')
assert(!!el && el.tagName === 'FORM')
```

`input` 요소에 `id`가 `email`인 속성이 있어야 합니다.

```js
const el = document.getElementById('email')
assert(!!el && el.tagName === 'INPUT')
```

`#email`는 `#form`의 자손 요소여야 합니다.

```js
const els = document.querySelectorAll('#form #email')
assert(els.length > 0)
```

`#email`에는 자리 표시자 텍스트가 있는 `placeholder` 속성이 있어야 합니다.

```js
const el = document.getElementById('email')
assert(!!el && !!el.placeholder && el.placeholder.length > 0)
```

`#email`는 `type`를 `email`로 설정해 HTML5 유효성 검사를 사용해야 합니다.

```js
const el = document.getElementById('email')
assert(!!el && el.type === 'email')
```

`input` 요소에 `id`가 `submit`인 속성이 있어야 합니다.

```js
const el = document.getElementById('submit')
assert(!!el && el.tagName === 'INPUT')
```

`#submit`는 `#form`의 자손 요소여야 합니다.

```js
const els = document.querySelectorAll('#form #submit')
assert(els.length > 0)
```

`#submit`에는 `type`가 `submit`인 속성이 있어야 합니다.

```js
const el = document.getElementById('submit')
assert(!!el && el.type === 'submit')
```

`#form`에는 `action` 속성이 `https://www.freecodecamp.com/email-submit`이어야 합니다.

```js
const el = document.getElementById('form')
assert(!!el && el.action === 'https://www.freecodecamp.com/email-submit')
```

`#email`에는 `name` 속성이 `email`이어야 합니다.

```js
const el = document.getElementById('email')
assert(!!el && el.name === 'email')
```

`#nav-bar`는 항상 뷰포트 상단에 있어야 합니다.

```js
  const timeout = (milliseconds) => new Promise((resolve) => setTimeout(resolve, milliseconds));

  const header = document.getElementById('header');
  const headerChildren = header.children;
  const navbarCandidates = [header, ...headerChildren];
  
  // Return smallest top position of all navbar candidates
  const getNavbarPosition = (candidates = []) => {
    return candidates.reduce(
      (min, candidate) =>
        Math.min(min, Math.abs(candidate?.getBoundingClientRect().top)),
      Infinity
    );
  };
  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the viewport '
  );

  window.scroll(0, 500);
  await timeout(1);

  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the ' +
      'viewport even after scrolling '
  );
    
  window.scroll(0, 0);
```

제품 랜딩 페이지는 적어도 하나의 미디어 쿼리를 사용해야 합니다.

```js
const htmlSourceAttr = Array.from(document.querySelectorAll('source')).map(el => el.getAttribute('media'))
const cssCheck = new __helpers.CSSHelp(document).getCSSRules('media')
assert(cssCheck.length > 0 || htmlSourceAttr.length > 0);
```

제품 랜딩 페이지는 CSS flexbox를 적어도 한 번 사용해야 합니다.

```js
const hasFlex = (rule) => ["flex", "inline-flex"].includes(rule.style?.display)
const stylesheet = new __helpers.CSSHelp(document).getStyleSheet()
const cssRules = new __helpers.CSSHelp(document).styleSheetToCssRulesArray(stylesheet)
const mediaRules = new __helpers.CSSHelp(document).getCSSRules('media')
const usesFlex = cssRules.find(rule => hasFlex(rule))
const usesFlexMedia = mediaRules.find(mediaRule => {
  return [...mediaRule.cssRules].find(rule => hasFlex(rule))
})
assert(usesFlex || usesFlexMedia)
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Product Landing Page</title>
  </head>
  <body>
    <header id="header">
      <nav id="nav-bar">
        <img
          id="header-img"
          src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
          max-height="50px"
        />
        <a href="#Features" class="nav-link">Features</a> |
        <a href="#Video" class="nav-link">See our facility!</a> |
        <a href="#Pricing" class="nav-link">Purchase</a>
        <hr />
      </nav>
    </header>
    <main>
      <h1>
        Pokemon Daycare Service
      </h1>
      <section id="Features">
        <h2>What we offer</h2>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">Guaranteed friendly and loving staff!</div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Comfortable environment for Pokemon to explore and play!
          </div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Multiple membership plans to fit your lifestyle!
          </div>
        </div>
      </section>
      <section id="Video">
        <h2>Check us out!</h2>
        A sneak peek into our facility:
        <br />
        <iframe
          id="video"
          width="520"
          height="347"
          src="https://www.youtube.com/embed/Nw-ksH2r6AQ"
          frameborder="0"
          allowfullscreen
          alt="A video tour of our facility"
        >
        </iframe>
      </section>
      <section id="Pricing">
        <h2>Membership Plans</h2>
        <div class="flex-mem">
          <div class="flex-mem-box">
            <font size="+2">Basic Membership</font><br />
            <ul>
              <li>One Pokemon</li>
              <li>Food and berries provided</li>
            </ul>
            <em>$9.99/month</em>
          </div>
          <div class="flex-mem-box">
            <font size="+2">Silver Membership</font><br />
            <ul>
              <li>Up to Three Pokemon</li>
              <li>Food and berries provided</li>
              <li>Grooming and accessories included</li>
            </ul>
            <em>$19.99/month</em>
          </div>
          <div class="flex-mem-box">
            <font size="+2">Gold Membership</font><br />
            <ul>
              <li>Up to six Pokemon!</li>
              <li>Food and berries provided</li>
              <li>Grooming and accessories included</li>
              <li>Personal training for each Pokemon</li>
              <li>Breeding and egg hatching</li>
            </ul>
            <em>$29.99/month</em>
          </div>
        </div>
      </section>
      <form id="form" action="https://www.freecodecamp.com/email-submit">
        <p>Sign up for our newsletter!</p>
        <label for="email"><p>Email:</p><input name="email" id="email" type="email" placeholder="johnsmith@email.com" required></label>
        <input type="submit" id="submit">
      </form>
      <footer>
        <a href="../">Return to Project List</a> |
        <a href="https://www.nhcarrigan.com">Return to HomePage</a>
      </footer>
    </main>
  </body>
</html>
```

```css
body {
  background-color: #3a3240;
  color: white;
}
main {
  max-width: 750px;
  margin: 50px auto;
}
input {
  background-color: #92869c;
}
a:not(.nav-link) {
  color: white;
}
#header-img {
  max-height: 25px;
}
#nav-bar {
  position: fixed;
  width: 100%;
  text-align: center;
  top: 0%;
  background-color: #92869c;
}
h1 {
  text-align: center;
}
body {
  text-align: center;
}
footer {
  text-align: center;
}
#bullet {
  max-height: 25px;
}
.flex-here {
  display: flex;
  justify-content: center;
}
.flex-left {
  height: 25px;
}
.flex-mem {
  display: flex;
  justify-content: center;
}
.flex-mem-box {
  background-color: #92869c;
  border-color: black;
  border-width: 5px;
  border-style: solid;
  margin: 10px;
  padding: 10px;
  color: black;
}
@media (max-width: 350px) {
  #video {
    width: 300;
    height: 200;
  }
}
```
