---
id: 587d825a367417b2b2512c89
title: Реалізувати алгоритм швидкого сортування
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

У цій лабораторній роботі ви реалізуєте алгоритм сортування: швидке сортування.

Швидке сортування — це ефективний рекурсивний підхід «розділяй і володарюй» для сортування масиву. У цьому методі у вихідному масиві вибирається опорне значення. Потім масив розбивається на два підмасиви: зі значеннями меншими та більшими за опорне. Далі результат рекурсивного виклику алгоритму швидкого сортування для обох підмасивів об’єднується. Це триває, доки не буде досягнуто базового випадку — порожнього або однеелементного масиву, який повертається. Розгортання рекурсивних викликів повертає відсортований масив.

Хоча вибір опорного значення важливий, для наших цілей підійде будь-яке опорне. Для простоти можна використати перший або останній елемент.

Швидке сортування — дуже ефективний метод сортування, що забезпечує середню продуктивність *O(nlog(n))*. Також його відносно легко реалізувати. Ці характеристики роблять його популярним і корисним методом сортування.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати `quicksort` функцію, яка приймає масив цілих чисел як вхідні дані.
1. Функція `quicksort` повинна повертати масив, що містить ті ж цілі числа, що й вхідний, але впорядковані від найменшого до найбільшого.
1. Функція `quicksort` повинна рекурсивно викликати сама себе для сортування масиву.

# --hints--

Ви повинні мати `quicksort` функцію.

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` має повертати масив, який не змінений, окрім порядку.

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` має повертати відсортований масив (від найменшого до найбільшого).

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` не має використовувати вбудований метод `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
