---
id: 68e5293bd00d2fe134f58985
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

由於 `PercentageDiscount` 繼承自抽象的類別，你必須實作所有抽象的行為。從實作 `is_applicable` 方法開始。`is_applicable` 方法應該接受 `self`、`product: Product` 和 `user_tier: str` 作為參數。它應該有一個返回型別提示為 `bool`。

此方法應該判斷百分比折扣是否有效。若 `self.percent` 小於或相等於 `70`，則傳回 `True`，否則傳回 `False`。這可防止像 90% 折扣這類不切實際的折扣。

# --hints--

你的 `PercentageDiscount` 類別應該有一個 `is_applicable` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

你的 `is_applicable` 方法應該有參數 `self`、`product: Product` 和 `user_tier: str`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

當 `self.percent` 小於或相等於 `70` 時，你的 `is_applicable` 方法應該傳回 `True`。

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

當 `self.percent` 大於 70 時，你的 `is_applicable` 方法應該傳回 `False`。

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

你的 `is_applicable` 方法應該有 `bool` 的返回型別提示。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
