---
id: 68e52994bc7ea2e3dec5309d
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

現在讓我們為高級使用者創建第三種折扣策略。創建一個名為 `PremiumUserDiscount` 的類別，繼承自 `DiscountStrategy`。

與先前的策略不同，這個策略不需要 `__init__` 方法，因為它不儲存任何設定。它將對所有高級用戶套用固定的 20% 折扣。

目前，只需使用 `pass` 創建繼承自 `DiscountStrategy` 的空類別。

# --hints--

你應該創建一個名為 `PremiumUserDiscount` 的類別。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PremiumUserDiscount")`)
})
```

你的 `PremiumUserDiscount` 類別應該繼承自 `DiscountStrategy`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").inherits_from("DiscountStrategy")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
