---
id: 68e5293bd00d2fe134f5898a
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

讓我們測試你的新折扣策略。在現有的程式碼之後，使用金額為 `5` 創建一個 `FixedAmountDiscount` 的實體，並將它指定給名為 `fixed_discount` 的變數。然後將此折扣套用到產品上並列印結果。

# --hints--

你應該創建一個名為 `fixed_discount` 的變數，並使用 `FixedAmountDiscount` 實體。

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("fixed_discount")
    assert isinstance(fixed_discount, FixedAmountDiscount)
  `)
})
```

你的 `FixedAmountDiscount` 實體應該以金額 `5` 初始化。

```js
({
  test: () => runPython(`assert fixed_discount.amount == 5`)
})
```

你應該呼叫 `fixed_discount.apply_discount(product)` 並列印結果。

```js
({
  test: () => runPython(`
    assert _Node(_code).has_call("print(fixed_discount.apply_discount(product))")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
--fcc-editable-region--

--fcc-editable-region--
```
