---
id: 68e5293bd00d2fe134f58981
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

創建一個名為 `DiscountStrategy` 的類別，繼承自 `ABC`。這將是一個抽象的基礎類別，定義所有折扣策略的介面。

目前，只需在類別本體中使用 `pass` 即可創建一個空類別。

# --hints--

你應該創建一個名為 `DiscountStrategy` 的類別。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

你的 `DiscountStrategy` 類別應該繼承自 `ABC`。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
