---
id: 68e52994bc7ea2e3dec530a1
title: 第 24 步
challengeType: 20
dashedName: step-24
---

# --description--

創建一個名為 `DiscountEngine` 的類別。這個類別將管理所有折扣策略並計算產品的最佳價格。

為它提供一個帶有 `self` 和 `strategies` 參數的 `__init__` 方法，其中 `strategies` 的型別提示為 `List[DiscountStrategy]`。該方法的返回型別應為 `None`。在 `__init__` 方法內，將 `strategies` 儲存為實體屬性。

# --hints--

你應該創建一個名為 `DiscountEngine` 的類別。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountEngine")`)
})
```

你的 `DiscountEngine` 類別應該有一個 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("__init__")`)
})
```

你的 `__init__` 方法應該有參數 `self` 和 `strategies: List[DiscountStrategy]`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_args("self, strategies: List[DiscountStrategy]")`)
})
```

你的 `__init__` 方法應該有 `None` 的返回型別。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_returns("None")`)
})
```

你應該在你的 `__init__` 方法中將 `strategies` 指定到 `self.strategies`。

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_stmt("self.strategies = strategies")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
from typing import List

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
