---
id: 68e5293bd00d2fe134f58984
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

現在是時候為你的抽象的類別創建一個具體的實作體。創建一個名為 `PercentageDiscount` 的類別，繼承自 `DiscountStrategy`。

為它提供一個帶有 `self` 和 `percent` 參數的 `__init__` 方法。為 `percent` 添加 `int` 的型別提示，並將該方法的返回型別設為 `None`。在 `__init__` 方法內，將 `percent` 儲存為實體屬性。

# --hints--

你應該創建一個名為 `PercentageDiscount` 的類別。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

你的 `PercentageDiscount` 類別應該繼承自 `DiscountStrategy`。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

你的 `PercentageDiscount` 類別應該有一個 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

你的 `__init__` 方法應該有參數 `self` 和 `percent: int`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

你的 `__init__` 方法應該有一個 `None` 的返回類型提示。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

你應該在你的 `__init__` 方法中將 `percent` 指定給 `self.percent`。

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
