---
id: 68e529226b0defdff26c4a8a
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

現在，使用名稱為 `Wireless Mouse` 且價格為 `50.0` 的 `Product` 創建一個實體。將它指定到名為 `product` 的變數。然後，將 `product` 列印到主控台以查看它的樣子。

# --hints--

你應該有一個名為 `product` 的變數。

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

你應該將一個 `Product` 的實體指定給 `product` 變數。

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

你的 `Product` 實體應該有名稱 `Wireless Mouse` 和價格 `50.0`。

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

你應該將 `product` 列印到主控台。

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

當你列印名稱為 `'Wireless Mouse'` 且價格為 `50.0` 的 `Product` 時，應該顯示 `Wireless Mouse - $50.0`。

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
