---
id: 68e5293bd00d2fe134f58989
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

現在為 `FixedAmountDiscount` 實作 `apply_discount` 方法。此方法應該從產品價格中扣除固定金額並傳回新的價格。

# --hints--

你的 `FixedAmountDiscount` 類別應該有一個 `apply_discount` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("apply_discount")`)
})
```

你的 `apply_discount` 方法應該有參數 `self` 和 `product: Product`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("apply_discount").has_args("self, product: Product")`)
})
```

你的 `apply_discount` 方法應該有一個 `-> float` 的返回型別提示。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("apply_discount").has_returns("float")`)
})
```

你的 `apply_discount` 方法應該傳回 `product.price - self.amount`。

```js
({
  test: () => runPython(`
    test_product = Product("Test", 100.0)
    discount = FixedAmountDiscount(15)
    result = discount.apply_discount(test_product)
    assert result == 85.0
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
