---
id: 68e524c94e2e4ab7b4bfc7e2
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將建置一個折扣計算器，能夠對產品套用不同的折扣策略。系統將根據多個折扣規則為顧客決定最佳價格。

首先定義一個名為 `Product` 的類別。為你的類別新增一個帶有以下參數的 `__init__` 方法：`self`、`name` 和 `price`。在 `__init__` 方法內，將 `name` 和 `price` 參數指定給同名的實體屬性。

# --hints--

你應該創建一個名為 `Product` 的類別。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

你的 `Product` 類別應該有一個 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

你的 `__init__` 方法應該有以下參數：`self`、`name` 和 `price`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

你應該在你的 `__init__` 方法中將 `name` 指定到 `self.name`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

你應該在你的 `__init__` 方法中將 `price` 指定到 `self.price`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
