---
id: 68e529226b0defdff26c4a8b
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

現在你將開始建置折扣系統介面。Python 提供了一種定義抽象的類別的方法，作為其他類別的藍圖。這是使用 `abc` 模組的 `ABC` 類別（抽象基礎類別）來完成的。

在你的程式碼頂部，為你的程式碼添加以下匯入述語：

```py
from abc import ABC, abstractmethod
```

這會匯入 `ABC` 類別和 `abstractmethod` 裝飾器，你將用它們來定義任何繼承自你的抽象類別的類別必須實作的方法。

# --hints--

你應該在你的程式碼頂部有提到的匯入述語。

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
