---
id: 68e5293bd00d2fe134f58983
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

現在為你的 `DiscountStrategy` 類別添加第二個名為 `apply_discount` 的抽象方法。使用 `@abstractmethod` 裝飾器並給它 `self` 和 `product` 參數。

為你的 `product` 添加型別提示，`product` 應該是 `Product` 型別，且該方法應傳回 `float`。在方法主體中使用 `pass`。

# --hints--

你的 `DiscountStrategy` 類別應該有一個名為 `apply_discount` 的方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

你的 `apply_discount` 方法應該有 `@abstractmethod` 裝飾器。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

你的 `apply_discount` 方法應該有參數 `self` 和 `product: Product`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

你的 `apply_discount` 方法應該有一個 `-> float` 的返回型別提示。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
