---
id: 68e529226b0defdff26c4a89
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

現在，如果你創建一個 `Product` 的實體並列印它，你會看到一個不太具資訊性的預設物件表述。舉例來說，它可能看起來像 `<__main__.Product object at 0x7f9b8c2d1d30>`。

為了將預設物件表述轉換成可讀的字串，請在你的 `Product` 類別中創建一個 `__str__` 方法，該方法傳回格式為 `name - $price` 的字串（其中 `name` 和 `price` 應該被對應屬性的值所取代）。

# --hints--

你的 `Product` 類別應該有一個 `__str__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

你的 `__str__` 方法應該傳回格式為 `name - $price` 的字串。

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
