---
id: 68e5293bd00d2fe134f58987
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

現在你將創建另一個折扣策略。創建一個名為 `FixedAmountDiscount` 的類別，繼承自 `DiscountStrategy`。

為它提供一個帶有 `self` 和 `amount` 參數的 `__init__` 方法。為 `amount` 參數添加 `int` 型別提示，並為該方法添加返回類型提示 `None`。在 `__init__` 方法內，將 `amount` 儲存為實體屬性。

# --hints--

你應該創建一個名為 `FixedAmountDiscount` 的類別。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

你的 `FixedAmountDiscount` 類別應該繼承自 `DiscountStrategy`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

你的 `FixedAmountDiscount` 類別應該有一個 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

你的 `__init__` 方法應該有參數 `self` 和 `amount: int`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

你的 `__init__` 方法應該有 `None` 的返回型別。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

你應該在你的 `__init__` 方法中將 `amount` 指定到 `self.amount`。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
