---
id: 68e79d91daf8fb4808f93c47
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

在實作 `apply_discount` 方法後，測試你的折扣策略。使用 `10` 百分比創建一個 `PercentageDiscount` 實體，並將它指定給名為 `discount` 的變數。然後使用 `discount.apply_discount(product)` 套用折扣，並列印結果以查看折扣效果。

# --hints--

你應該創建一個名為 `discount` 的變數，並使用 `PercentageDiscount` 實體。

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

你的 `PercentageDiscount` 實體應該以 `10` 百分比初始化。

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

你應該呼叫 `discount.apply_discount(product)` 並列印結果。

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
