---
id: 68e9207ddcba053c5cb24413
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

在 Python 中，型別提示是選用的信號，用來告訴其他開發者變數或函式的資料型別預期為何。以下是一個為 `name` 參數添加型別提示 `str` 的範例。

```py
def greet(name: str):
    return 'Hello, ' + name
```

對於 `__init__` 方法參數，將 `name` 參數的型別提示更新為 `str`，並將 `price` 參數的型別提示更新為 `float`。

# --hints--

你應該為 `Product` 類別的 `__init__` 方法中的 `name` 參數添加型別提示 `str`，並為 `price` 參數添加型別提示 `float`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
