---
id: 68f119ad905b29ff59d858c8
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

就像前一步驟一樣，在你的 `__str__` 方法中，為返回類型添加 `str` 的提示，因為此方法傳回物件的字串表述。

# --hints--

你的 `__str__` 方法應該有 `str` 的返回型別提示。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
