---
id: aaa48de84e1ecc7c742e1124
title: 回文檢查器
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16004
dashedName: palindrome-checker
---

# --description--

如果給定的字串是迴文，則傳回 `true`。否則，傳回 `false`。

<dfn>回文</dfn>是指忽略標點符號、大小寫和空格後，正讀和反讀都相同的字或句子。

**注意：** 你需要移除 **所有非英數字元**（標點符號、空格和符號），並將所有字元轉成相同大小寫（小寫或大寫），才能檢查迴文。

我們會傳遞具有不同格式的字串，例如 `racecar`、`RaceCar` 和 `race CAR` 等。

我們也會傳遞包含特殊符號的字串，例如 `2A3*3a2`、`2A3 3a2` 和 `2_A3*3#A2`。

# --hints--

`palindrome("eye")` 應該傳回一個布林值。

```js
assert(typeof palindrome('eye') === 'boolean');
```

`palindrome("eye")` 應該傳回 `true`。

```js
assert(palindrome('eye') === true);
```

`palindrome("_eye")` 應該傳回 `true`。

```js
assert(palindrome('_eye') === true);
```

`palindrome("race car")` 應該傳回 `true`。

```js
assert(palindrome('race car') === true);
```

`palindrome("not a palindrome")` 應該傳回 `false`。

```js
assert(palindrome('not a palindrome') === false);
```

`palindrome("A man, a plan, a canal. Panama")` 應該傳回 `true`。

```js
assert(palindrome('A man, a plan, a canal. Panama') === true);
```

`palindrome("never odd or even")` 應該傳回 `true`。

```js
assert(palindrome('never odd or even') === true);
```

`palindrome("nope")` 應該傳回 `false`。

```js
assert(palindrome('nope') === false);
```

`palindrome("almostomla")` 應該傳回 `false`。

```js
assert(palindrome('almostomla') === false);
```

`palindrome("My age is 0, 0 si ega ym.")` 應該傳回 `true`。

```js
assert(palindrome('My age is 0, 0 si ega ym.') === true);
```

`palindrome("1 eye for of 1 eye.")` 應該傳回 `false`。

```js
assert(palindrome('1 eye for of 1 eye.') === false);
```

`palindrome("0_0 (: /-\ :) 0-0")` 應該傳回 `true`。

```js
assert(palindrome('0_0 (: /- :) 0-0') === true);
```

`palindrome("five|\_/|four")` 應該傳回 `false`。

```js
assert(palindrome('five|_/|four') === false);
```

# --seed--

## --seed-contents--

```js
function palindrome(str) {
  return true;
}

palindrome("eye");
```

# --solutions--

```js
function palindrome(str) {
  var string = str.toLowerCase().split(/[^A-Za-z0-9]/gi).join('');
  var aux = string.split('');
  if (aux.join('') === aux.reverse().join('')){
    return true;
  }

  return false;
}
```
