---
id: 680a431692a5f863f726c438
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

У цьому практичному занятті ви створите *шифр Цезаря*. Це одна з найпростіших технік для шифрування тексту, яка полягає у заміні кожної літери відкритого тексту на літеру, що знаходиться на фіксованій кількості позицій далі в алфавіті. Наприклад, зі зсувом на 5, `a` буде замінений на `f`, `b` — на `g` і так далі.

Щоб реалізувати цей шифр, вам потрібно створити нову версію свого алфавіту, яка починається з позиції, вказаної зсувом. Як ви дізналися в попередньому уроці, можна витягти частину рядка, використовуючи зрізи рядка:

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

Створіть змінну з ім’ям `shifted_alphabet` і використайте синтаксис зрізу, щоб призначити їй ту частину `alphabet`, яка починається з індексу `shift`. Потім викличте вбудовану функцію `print()`, щоб вивести `shifted_alphabet` у термінал і подивитися результат.

# --hints--

Ви повинні оголосити нову змінну з ім’ям `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

Ви повинні призначити `alphabet[shift:]` змінній `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

Ви повинні викликати функцію `print()`, передавши їй `shifted_alphabet` як аргумент.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
