---
id: 680a82806c4c6bc3567f2708
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Оголосіть нову змінну з іменем `text` і призначте їй рядок `hello world`. Це буде повідомлення для шифрування.

# --hints--

Ви повинні оголосити змінну з іменем `text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

Ви повинні призначити рядок `hello world` змінній `text`. Не забудьте взяти рядок у одинарні або подвійні лапки та звернути увагу на регістр літер.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
