---
id: 680a836235461fce1ff53b3f
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Метод `translate()` приймає як аргумент таблицю перекладу, створену `str.maketrans()`. Він викликається на рядку і повертає копію оригінального рядка, де символи були замінені на основі таблиці перекладу:

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

Викличте метод `translate()` на `text`, передавши `translation_table` як аргумент, і призначте результат змінній з ім’ям `encrypted_text`.

# --hints--

У вас має бути змінна з ім’ям `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Ви маєте призначити `text.translate(translation_table)` змінній `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
