---
id: 680a549f60f5a3fedfc0edd9
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Як видно з вихідних даних, зсунутий алфавіт починається з літери `f`, тому що `shift` має значення `5`. Але тепер перші п’ять літер алфавіту – `a`, `b`, `c`, `d` і `e` – відсутні у зсунутому алфавіті, тож вам потрібно додати їх у кінець зсунутого алфавіту.

Оператор `+` використовується для об’єднання двох або більше рядків у процесі, який називається конкатенацією, ось так:

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

Змініть існуюче присвоєння змінної `shifted_alphabet`: використайте синтаксис зрізу, щоб витягти відсутню першу частину `alphabet` і конкатенувати її з `alphabet[shift:]`.

Нагадаємо, що `sentence[start:stop]` повертає символи `sentence` від позиції `start` (включно) до `stop` (не включно).

# --hints--

Ви не повинні безпосередньо змінювати змінну `alphabet`.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

Ви повинні призначити конкатенацію `alphabet[shift:]` і відсутньої першої частини `alphabet` змінній `shifted_alphabet` на тій самій лінії, де вона була оголошена. Використайте `shift`, щоб вказати, де зупинити зріз.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
