---
id: 680a7e200400089967b8b1f4
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Метод `str.maketrans()` приймає два рядки однакової довжини і повертає таблицю трансляції, яка зіставляє кожен символ першого рядка з відповідним символом другого рядка. Кожен символ у таблиці трансляції зберігається як Unicode-ординал — число, яке унікально ідентифікує символ.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

Ви дізнаєтеся більше про тип структури, яку повертає `str.maketrans()`, пізніше в навчальній програмі. А поки створіть таблицю трансляції, яка зіставляє символи в `alphabet` із символами в `shifted_alphabet`, і призначте її змінній з ім’ям `translation_table`.

# --hints--

Ви повинні оголосити змінну з ім’ям `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

Ви повинні викликати `str.maketrans()`, передавши `alphabet` і `shifted_alphabet` як аргументи, і призначити результат `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
