---
id: 680a5e7f3d47715196b12ca0
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Тепер, коли у вас є весь зсунутий алфавіт, видаліть виклик `print(shifted_alphabet)`.

# --hints--

Ви маєте видалити `print(shifted_alphabet)` зі свого коду.

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
