---
id: 680a3d4c065557260de2ff33
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Як ви, можливо, пам’ятаєте з попередніх уроків, у Python ви оголошуєте змінну, записуючи ім’я змінної зліва від оператора присвоєння (`=`) і значення, яке потрібно призначити, справа:

```py
variable_name = value
```

Створіть змінну з ім’ям `shift` і призначте новій змінній значення `5`.

# --hints--

Ви повинні оголосити змінну з ім’ям `shift`. Зверніть увагу, що ім’я змінної має бути на початку лінії.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Ви повинні призначити значення `5` змінній `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
