---
id: 680a41ef5bfdc059cfd9d4b5
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

У попередніх уроках ви дізналися про різні типи даних, які можна зберігати у змінній. Ви щойно призначили ціле число. Тепер потрібно призначити рядок — послідовність символів, укладену або в одинарні, або в подвійні лапки:

```py
string_1 = 'I am a string'
string_2 = "I am also a string"
```

Оголосіть іншу змінну з ім’ям `alphabet` і призначте їй рядок `abcdefghijklmnopqrstuvwxyz`.

# --hints--

Ви маєте оголосити змінну з ім’ям `alphabet`. Зверніть увагу, що ім’я змінної має стояти на початку лінії.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Ви маєте призначити рядок `abcdefghijklmnopqrstuvwxyz` змінній `alphabet`. Пам’ятайте, що рядок потрібно укладати або в одинарні, або в подвійні лапки.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
