---
id: 6810891bd6de8a87a833df42
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Як видно з вихідних даних, повідомлення було зашифровано. Наступним кроком буде зробити ваш код багаторазовим, щоб ви могли шифрувати різні повідомлення.

Для цього потрібно створити функцію. Нагадаємо, як створити функцію з іменем `spam`, яка виводить `Spam!` у термінал:

```py
def spam():
    print('Spam!')
```

Створіть функцію з іменем `caesar`. Помістіть увесь ваш існуючий код у тіло функції. Зверніть увагу, щоб усі рядки в тілі функції мали однаковий рівень відступу.

# --hints--

У вас має бути функція з іменем `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

Весь код, який ви написали досі, слід перемістити в тіло функції `caesar`. Переконайтеся, що всі рядки в тілі функції мають однаковий рівень відступу...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
