---
id: 6818ed1adc249bb91d373729
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

У попередніх уроках ви дізналися, що інструкція `return` використовується для повернення значення з функції, щоб ви могли використовувати його в іншому місці вашого коду:

```py
def spam():
    return 'Spam!'
```

Видаліть виклик `print(encrypted_text)` з вашої функції. Потім видаліть оголошення змінної `encrypted_text` і повертайте безпосередньо `text.translate(translation_table)` з вашої функції.

# --hints--

У функції `print(encrypted_text)` не повинно бути `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

У функції `encrypted_text` не повинно бути змінної `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

Ваша функція `caesar` має повертати `text.translate(translation_table)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
