---
id: 6818e9c7c8dc5694e8f64fc3
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Тепер протестуйте функцію `caesar`, викликавши її з рядком `freeCodeCamp` та числом `3` як аргументами. Призначте виклик функції змінній з ім’ям `encrypted_text`.

# --hints--

У вас має бути змінна з ім’ям `encrypted_text` в глобальній області видимості.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Ви маєте призначити `caesar('freeCodeCamp', 3)` змінній `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
