---
id: 6818b80519e6c6c47a8bb6e8
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Як ви дізналися в попередньому уроці, функції можуть мати параметри — це змінні, до яких можна звертатися всередині функції. Ось функція `sum` з двома параметрами, `num1` та `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

Повідомлення для шифрування та зсув досі жорстко закодовані у вашій функції. Додайте вашій функції два параметри з іменами `text` та `shift`, а також видаліть оголошення змінних `text` і `shift` з тіла функції `caesar`.

# --hints--

Ваша функція `caesar` має мати два параметри, `text` і `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

У тілі функції `text` не повинна бути оголошена змінна `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

У тілі функції `shift` не повинна бути оголошена змінна `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
