---
id: 6818eb3f98c0bca5af7c8d03
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Тепер ваш код можна використовувати повторно. Однак функція `caesar` виводить повідомлення в термінал і за замовчуванням повертає `None`. Щоб це довести, виведіть `encrypted_text` наприкінці свого коду.

# --hints--

Ви повинні вивести `encrypted_text` поза межами тіла функції.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
