---
id: 6818fcb250f34e62cab32c39
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Ви могли помітити, що, хоча повідомлення і зашифроване, символи у верхньому регістрі залишилися без змін. Це відбувається тому, що таблиця перекладу не містить літер у верхньому регістрі.

Щоб це виправити, вам потрібно змінити рядки, які передаються у ваш виклик `str.maketrans()`, використовуючи метод `upper()`:

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

Оновіть свій виклик `str.maketrans()`, додавши до кожного аргументу версію цього аргументу у верхньому регістрі.

# --hints--

Вам слід змінити аргументи вашого виклику `str.maketrans()`, додавши до кожного аргументу версію цього аргументу у верхньому регістрі.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
