---
id: 6819c7beb4a8289c2b4f24d5
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Від’ємне або нульове зміщення не повинно прийматися вашою функцією. Тому після вашої першої `if` інструкції створіть ще одну `if` інструкцію, яка перевіряє, чи `shift` менше за `1`, і повертає рядок `Shift must be a positive integer.`

# --hints--

У вас має бути друга `if` інструкція, яка перевіряє, чи `shift` менше за `1`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift < 1", "1 > shift", "shift <= 0", "0 => shift"]
assert any(_cond.is_equivalent(option) for option in _options)
`) })
```

Друга `if` інструкція має повертати `Shift must be a positive integer.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be a positive integer.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
