---
id: 6819ca464db219ab86f37dc4
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Як ви вже перевірили, зсув, переданий для шифрування тексту, має бути додатним. Проте він не може перевищувати `25` (останній індекс `alphabet`).

Додайте другу умову до інструкції `if`, яка перевіряє, що `shift` більший за `25`. Пам’ятайте, що логічна операція АБО в Python реалізується через оператор `or`.

Також оновіть повідомлення, яке повертається, на `Shift must be an integer between 1 and 25.`

# --hints--

Ви маєте використати оператор `or`, щоб додати другу умову після `shift < 1`, яка перевіряє, чи `shift` більший за `25`. Існуючу умову змінювати не потрібно.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

Ваша друга інструкція `if` має повертати `Shift must be an integer between 1 and 25.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
