---
id: 6819c5c9faf2548e556163d5
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Як ви пам’ятаєте з попередніх уроків, оператор `not` — це унарний логічний оператор, який заперечує вираз:

```py
print(not True) # False
print(not False) # True
```

Використайте оператор `not`, щоб виправити умову інструкції `if`, щоб ваша функція повертала `Shift must be an integer value.`, коли `shift` *не* є цілим числом.

# --hints--

Ви повинні використати оператор `not`, щоб заперечити `isinstance(shift, int)` у вашій інструкції `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
