---
id: 6819cf38880d8cc9b9f6af7a
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Python дозволяє вказувати значення за замовчуванням для параметрів у функції, створюючи функцію, яку можна викликати з меншою кількістю або без аргументів. Ось як створити функцію з параметром `name`, який має значення за замовчуванням:

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

Додайте до своєї функції третій параметр з іменем `encrypt` і надайте йому значення за замовчуванням `True`. 

# --hints--

Ви повинні додати до своєї функції третій параметр з іменем `encrypt` і значенням за замовчуванням `True`. Не змінюйте існуючі параметри.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
