---
id: 6819be9055828d65c813eec5
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Функція `isinstance()` повертає `True`, якщо її перший аргумент є екземпляром другого аргументу, інакше повертає `False`:

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

Замініть поточну умову вашої інструкції `if` викликом `isinstance()`. Передайте `shift` як перший аргумент і `int` як другий аргумент.

# --hints--

Ваша інструкція `if` має використовувати `isinstance(shift, int)` як умову.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
