---
id: 6819bc0637b80256a33adccc
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Тепер, коли ви реалізували базові функції шифру, настав час додати валідацію. Для цього вам знадобиться інструкція `if`. Ось нагадування синтаксису інструкції `if`:

```py
if condition:
    # code to run when condition is true
```

На початку тіла вашої функції створіть інструкцію `if`. Поки що використовуйте `True` як умову, а в тілі інструкції `if` поверніть рядок `Shift must be an integer value.`.

# --hints--

Усередині вашої функції `if` має бути інструкція `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

Ваша інструкція `if` має використовувати `True` як умову.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Ви маєте повертати рядок `Shift must be an integer value.` з вашої інструкції `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
