---
id: 681a2a6cf95e82ca30866364
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Оголосіть дві функції з іменами `encrypt` та `decrypt`, обидві з параметрами `text` і `shift`.

Ви будете використовувати `encrypt` для процесу шифрування, а `decrypt` — для дешифрування, позначаючи ці дві дії описовою назвою.

Поверніть виклик `caesar`, передаючи `text` і `shift` з обох ваших нових функцій, але переконайтеся, що також передали `False` як третій аргумент для функції `decrypt`.

# --hints--

Ви повинні мати функцію з іменем `encrypt` з двома параметрами, `text` і `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

Ваша функція `encrypt` має повертати `caesar(text, shift)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

Ви повинні мати функцію з іменем `decrypt` з двома параметрами, `text` і `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

Ваша функція `decrypt` має повертати `caesar(text, shift, False)`.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
