---
id: 681a1d60e037127e7c8e0230
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Ви збираєтеся використати параметр, доданий на попередньому кроці, щоб визначити, чи функція має шифрувати переданий їй текст (поведінка за замовчуванням, `encrypt=True`), чи розшифровувати зашифроване повідомлення.

Створіть інструкцію `if`, яка перевіряє, чи `encrypt` не є істинним. Усередині нової інструкції `if` встановіть `shift` у значення `- shift`. Це необхідно, щоб зсув відбувався у протилежному напрямку щодо процесу шифрування.

# --hints--

Ви повинні встановити `shift` у `- shift`, коли `encrypt` не є істинним.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
