---
id: 5a24c314108439a4d4036143
title: 상태 로직을 Redux로 분리하기
challengeType: 6
forumTopicId: 301429
dashedName: extract-state-logic-to-redux
---

# --description--

React 컴포넌트를 완성했으니, 이제 `state`에서 로컬로 수행하던 로직을 Redux로 옮겨야 합니다. 이것이 간단한 React 앱을 Redux와 연결하는 첫 단계입니다. 앱의 유일한 기능은 사용자가 새 메시지를 순서 없는 목록에 추가하는 것입니다. 이 예제는 React와 Redux가 어떻게 함께 작동하는지 보여주기 위해 단순하게 구성되어 있습니다.

# --instructions--

먼저, 액션 타입 `ADD`을 정의하고 const `ADD`에 할당하세요. 다음으로, 메시지를 추가하는 액션을 생성하는 액션 생성자 `addMessage()`를 정의하세요. 이 액션 생성자에 `message`를 전달하고, 반환되는 `action`에 메시지를 포함해야 합니다.

그다음, 메시지 상태를 처리하는 리듀서 `messageReducer()`를 만드세요. 초기 상태는 빈 배열이어야 합니다. 이 리듀서는 상태에 저장된 메시지 배열에 메시지를 추가하거나 현재 상태를 반환해야 합니다. 마지막으로, 리듀서를 전달하여 Redux 저장소를 만드세요.

# --hints--

const `ADD`가 존재하며 문자열 `ADD`와 같은 값을 가져야 합니다.

```js
assert(ADD === 'ADD');
```

액션 생성자 `addMessage`는 `type`이 `ADD`이고 `message`가 전달된 메시지인 객체를 반환해야 합니다.

```js
assert(
  (function () {
    const addAction = addMessage('__TEST__MESSAGE__');
    return addAction.type === ADD && addAction.message === '__TEST__MESSAGE__';
  })()
);
```

`messageReducer`는 함수여야 합니다.

```js
assert(typeof messageReducer === 'function');
```

저장소는 존재하며 초기 상태가 빈 배열로 설정되어야 합니다.

```js
assert(
  (function () {
    const initialState = store.getState();
    return typeof store === 'object' && initialState.length === 0;
  })()
);
```

저장소에 `addMessage`를 디스패치하면 상태에 저장된 메시지 배열에 불변성을 유지하며 새 메시지가 추가되어야 합니다.

```js
assert(
  (function () {
    const initialState = store.getState();
    const isFrozen = DeepFreeze(initialState);
    store.dispatch(addMessage('__A__TEST__MESSAGE'));
    const addState = store.getState();
    return isFrozen && addState[0] === '__A__TEST__MESSAGE';
  })()
);
```

`messageReducer`는 다른 액션이 호출되면 현재 상태를 반환해야 합니다.

```js
assert(
  (function () {
    const addState = store.getState();
    store.dispatch({ type: 'FAKE_ACTION' });
    const testState = store.getState();
    return addState === testState;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Define ADD, addMessage(), messageReducer(), and store here:
```

# --solutions--

```jsx
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);
```
