---
id: 5a24c314108439a4d4036141
title: React Redux 시작하기
challengeType: 6
forumTopicId: 301430
dashedName: getting-started-with-react-redux
---

# --description--

이 도전 과제 시리즈에서는 React와 Redux를 함께 사용하는 방법을 소개합니다. 먼저, 각 기술의 주요 원칙을 검토해 보겠습니다. React는 데이터를 제공하면 효율적이고 예측 가능한 방식으로 뷰를 렌더링하는 뷰 라이브러리입니다. Redux는 애플리케이션 상태 관리를 단순화하는 데 사용할 수 있는 상태 관리 프레임워크입니다. 일반적으로 React Redux 애플리케이션에서는 전체 앱 상태를 관리하는 단일 Redux 저장소를 만듭니다. React 컴포넌트는 자신의 역할과 관련된 저장소 내 데이터 조각에만 구독합니다. 그런 다음 React 컴포넌트에서 직접 액션을 디스패치하면 저장소 업데이트가 트리거됩니다.

React 컴포넌트가 자체 상태를 로컬에서 관리할 수 있지만, 복잡한 애플리케이션에서는 Redux를 사용해 앱 상태를 한 곳에 모아두는 것이 일반적으로 더 좋습니다. 개별 컴포넌트가 자신만의 로컬 상태를 가질 수 있는 예외도 있습니다. 마지막으로 Redux는 기본적으로 React와 함께 작동하도록 설계되지 않았기 때문에 `react-redux` 패키지를 사용해야 합니다. 이 패키지는 Redux `state`와 `dispatch`를 React 컴포넌트에 `props`로 전달하는 방법을 제공합니다.

다음 도전 과제들에서는 먼저 새 텍스트 메시지를 입력할 수 있는 간단한 React 컴포넌트를 만듭니다. 이 메시지들은 뷰에 표시되는 배열에 추가됩니다. 이는 React 수업에서 배운 내용을 잘 복습하는 기회가 될 것입니다. 다음으로 메시지 배열 상태를 관리하는 Redux 저장소와 액션을 만듭니다. 마지막으로 `react-redux`를 사용해 Redux 저장소를 컴포넌트와 연결하여 로컬 상태를 Redux 저장소로 추출합니다.

# --instructions--

`DisplayMessages` 컴포넌트로 시작하세요. 이 컴포넌트에 생성자를 추가하고 두 개의 프로퍼티가 있는 상태로 초기화하세요: 빈 문자열로 설정된 `input`와 빈 배열로 설정된 `messages`입니다.

# --hints--

`DisplayMessages` 컴포넌트는 빈 `div` 요소를 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(DisplayMessages));
    return mockedComponent.find('div').text() === '';
  })()
);
```

`DisplayMessages` 생성자는 `super`와 함께 `props`를 전달하여 올바르게 호출해야 합니다.

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        noWhiteSpace.includes('constructor(props)') &&
        noWhiteSpace.includes('super(props')
      );
    })()
  );
```

`DisplayMessages` 컴포넌트는 `{input: "", messages: []}`와 같은 초기 상태를 가져야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(DisplayMessages));
    const initialState = mockedComponent.state();
    return (
      typeof initialState === 'object' &&
      initialState.input === '' &&
      Array.isArray(initialState.messages) &&
      initialState.messages.length === 0
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class DisplayMessages extends React.Component {
  // Change code below this line

  // Change code above this line
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
  }
  render() {
    return <div/>
  }
};
```
