---
id: 5a24c314108439a4d4036144
title: Redux를 React에 연결하기 위해 Provider 사용하기
challengeType: 6
forumTopicId: 301435
dashedName: use-provider-to-connect-redux-to-react
---

# --description--

지난 도전 과제에서 메시지 배열을 처리하는 Redux 저장소를 만들고 새 메시지를 추가하는 액션을 생성했습니다. 다음 단계는 React가 Redux 저장소와 업데이트를 디스패치하는 데 필요한 액션에 접근할 수 있도록 하는 것입니다. React Redux는 이러한 작업을 돕기 위해 `react-redux` 패키지를 제공합니다.

React Redux는 두 가지 핵심 기능인 `Provider`과 `connect`를 포함한 작은 API를 제공합니다. `connect`에 관한 내용은 다른 도전 과제에서 다룹니다. `Provider`는 React Redux에서 제공하는 래퍼 컴포넌트로, React 앱을 감쌉니다. 이 래퍼는 컴포넌트 트리 전체에서 Redux `store`와 `dispatch` 함수를 접근할 수 있게 합니다. `Provider`는 Redux 저장소와 앱의 자식 컴포넌트를 두 개의 프로프로 받습니다. App 컴포넌트에 대한 `Provider` 정의는 다음과 같을 수 있습니다:

```jsx
<Provider store={store}>
  <App/>
</Provider>
```

# --instructions--

코드 편집기에는 지난 여러 도전 과제에서 작성한 모든 Redux와 React 코드가 표시됩니다. 여기에는 Redux 저장소, 액션, 그리고 `DisplayMessages` 컴포넌트가 포함되어 있습니다. 유일하게 새로 추가된 부분은 하단의 `AppWrapper` 컴포넌트입니다. 이 최상위 컴포넌트를 사용해 `Provider`에서 `ReactRedux`를 렌더링하고 Redux 저장소를 프로프로 전달하세요. 그런 다음 `DisplayMessages` 컴포넌트를 자식으로 렌더링하세요. 완료하면 React 컴포넌트가 페이지에 렌더링된 것을 볼 수 있습니다.

**참고:** 여기서는 React Redux가 전역 변수로 제공되므로 점 표기법으로 Provider에 접근할 수 있습니다. 편집기 내 코드는 이를 활용해 `Provider`라는 상수에 Provider를 할당하여 `AppWrapper` 렌더 메서드에서 사용할 수 있게 했습니다.

# --hints--

`AppWrapper`가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('AppWrapper').length === 1;
  })()
);
```

`Provider` 래퍼 컴포넌트는 Redux 저장소와 같은 `store` 프로프를 전달받아야 합니다.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
      return __helpers
        .removeWhiteSpace(code)
        .includes('<Providerstore={store}>');
    })()
  );
```

`DisplayMessages`는 `AppWrapper`의 자식으로 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('AppWrapper').find('DisplayMessages').length === 1
    );
  })()
);
```

`DisplayMessages` 컴포넌트는 `h2`, `input`, `button`, `ul` 요소를 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h2').length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('ul').length === 1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};



const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
    this.handleChange = this.handleChange.bind(this);
    this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {  
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Render the Provider below this line

  // Change code above this line
};
```

# --solutions--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
 this.handleChange = this.handleChange.bind(this);
 this.submitMessage = this.submitMessage.bind(this);
 }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };  
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Change code below this line
  render() {
    return (
      <Provider store = {store}>
        <DisplayMessages/>
      </Provider>
    );
  }
  // Change code above this line
};
```
