---
id: 5a24c314108439a4d4036146
title: 프로프스에 디스패치 매핑하기
challengeType: 6
forumTopicId: 301432
dashedName: map-dispatch-to-props
---

# --description--

`mapDispatchToProps()` 함수는 React 컴포넌트에 특정 액션 생성자를 제공하여 Redux 저장소에 액션을 디스패치할 수 있게 합니다. 이전 도전 과제에서 작성한 `mapStateToProps()` 함수와 구조가 비슷합니다. 이 함수는 디스패치 액션을 속성 이름에 매핑하는 객체를 반환하며, 이 속성들은 컴포넌트 `props`가 됩니다. 하지만 각 속성이 `state`의 일부를 반환하는 대신, 액션 생성자와 관련 액션 데이터를 인수로 받아 `dispatch`를 호출하는 함수를 반환합니다. 이 `dispatch`는 함수 정의 시 `mapDispatchToProps()`에 매개변수로 전달되기 때문에 접근할 수 있습니다. 이는 `state`를 `mapStateToProps()`에 전달한 것과 같습니다. 내부적으로 React Redux는 Redux의 `store.dispatch()`를 사용해 `mapDispatchToProps()`와 함께 이 디스패치를 수행합니다. 이는 컴포넌트가 `store.subscribe()`에 매핑될 때 `state`를 사용하는 방식과 유사합니다.

예를 들어, 액션 페이로드로 `loginUser()`를 받는 `username` 액션 생성자가 있습니다. 이 액션 생성자에 대해 `mapDispatchToProps()`가 반환하는 객체는 다음과 비슷할 것입니다:

```jsx
{
  submitLoginUser: function(username) {
    dispatch(loginUser(username));
  }
}
```

# --instructions--

코드 편집기에는 `addMessage()`라는 액션 생성자가 제공됩니다. `mapDispatchToProps()`를 인수로 받아 객체를 반환하는 `dispatch` 함수를 작성하세요. 이 객체는 디스패치 함수에 새 메시지를 매개변수로 전달하여 `submitNewMessage`를 디스패치할 때 사용하는 `addMessage()` 속성을 가져야 합니다.

# --hints--

`addMessage`는 `type`와 `message` 키를 가진 객체를 반환해야 합니다.

```js
assert(
  (function () {
    const addMessageTest = addMessage();
    return (
      addMessageTest.hasOwnProperty('type') &&
      addMessageTest.hasOwnProperty('message')
    );
  })()
);
```

`mapDispatchToProps`는 함수여야 합니다.

```js
assert(typeof mapDispatchToProps === 'function');
```

`mapDispatchToProps`는 객체를 반환해야 합니다.

```js
assert(typeof mapDispatchToProps() === 'object');
```

`addMessage`에서 `submitNewMessage`와 함께 `mapDispatchToProps`를 디스패치하면 디스패치 함수에 메시지를 반환해야 합니다.

```js
assert(
  (function () {
    let testAction;
    const dispatch = (fn) => {
      testAction = fn;
    };
    let dispatchFn = mapDispatchToProps(dispatch);
    dispatchFn.submitNewMessage('__TEST__MESSAGE__');
    return (
      testAction.type === 'ADD' && testAction.message === '__TEST__MESSAGE__'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

// Change code below this line

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: function(message) {
      dispatch(addMessage(message));
    }
  }
};
```
