---
id: 5a24c314108439a4d4036147
title: Redux를 React에 연결하기
challengeType: 6
forumTopicId: 301426
dashedName: connect-redux-to-react
---

# --description--

이제 `mapStateToProps()` 함수와 `mapDispatchToProps()` 함수를 작성했으니, 이를 사용해 `state`와 `dispatch`를 React 컴포넌트 중 하나의 `props`에 매핑할 수 있습니다. React Redux의 `connect` 메서드가 이 작업을 처리할 수 있습니다. 이 메서드는 두 개의 선택적 인수인 `mapStateToProps()`와 `mapDispatchToProps()`를 받습니다. 선택적인 이유는, 어떤 컴포넌트는 `state`에만 접근하면 되고 액션을 디스패치할 필요가 없거나 그 반대일 수 있기 때문입니다.

이 메서드를 사용하려면, 함수를 인수로 전달하고 결과를 즉시 컴포넌트에 호출하세요. 이 구문은 다소 특이하며 다음과 같습니다:

```js
connect(mapStateToProps, mapDispatchToProps)(MyComponent)
```

**참고:** `connect` 메서드의 인수 중 하나를 생략하려면, 그 자리에 `null`를 전달하세요.

# --instructions--

코드 편집기에는 `mapStateToProps()`와 `mapDispatchToProps()` 함수와 `Presentational`라는 새로운 React 컴포넌트가 있습니다. 이 컴포넌트를 `connect` 전역 객체의 `ReactRedux` 메서드로 Redux에 연결하고, `Presentational` 컴포넌트에 즉시 호출하세요. 결과를 연결된 컴포넌트를 나타내는 `const`라는 새 `ConnectedComponent`에 할당하세요. 이제 Redux에 연결되었습니다! `connect`의 인수 중 하나를 `null`로 변경해 보고 테스트 결과를 확인하세요.

# --hints--

`Presentational` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('Presentational').length === 1;
  })()
);
```

`Presentational` 컴포넌트는 `messages`를 통해 `connect` 프로프를 받아야 합니다.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return props.messages === '__INITIAL__STATE__';
  })()
);
```

`Presentational` 컴포넌트는 `submitNewMessage`를 통해 `connect` 프로프를 받아야 합니다.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return typeof props.submitNewMessage === 'function';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line

const ConnectedComponent = connect(mapStateToProps, mapDispatchToProps)(Presentational);
```
