---
id: 5a24c314108439a4d403614a
title: 여기서부터 앞으로 나아가기
challengeType: 6
forumTopicId: 301434
dashedName: moving-forward-from-here
---

# --description--

축하합니다! React와 Redux 수업을 모두 마쳤습니다. 다음 단계로 넘어가기 전에 한 가지 더 짚고 넘어갈 부분이 있습니다. 보통 이렇게 코드 편집기에서 React 앱을 작성하지 않습니다. 이 도전 과제는 자신의 컴퓨터 파일 시스템에서 작업할 때 구문이 어떻게 보이는지 살짝 보여줍니다. 코드는 비슷하게 보일 텐데, `import` 문장 사용 부분만 다릅니다(이 문장들은 도전 과제에서 제공된 모든 의존성을 불러옵니다).

마지막으로, React와 Redux 코드를 작성하려면 일반적으로 약간의 설정이 필요합니다. 이 과정은 금방 복잡해질 수 있습니다. 자신의 컴퓨터에서 직접 실험해보고 싶다면 <a href="https://www.freecodecamp.org/news/install-react-with-create-react-app/" target="_blank" rel="noopener noreferrer nofollow">Create React App</a>이 미리 설정되어 바로 사용할 수 있게 준비되어 있습니다.

또는 CodePen에서 JavaScript 전처리기로 Babel을 활성화하고, React와 ReactDOM을 외부 JavaScript 리소스로 추가해 그곳에서 작업할 수도 있습니다.

# --instructions--

메시지 `'Now I know React and Redux!'`를 콘솔에 출력하세요.

# --hints--

메시지 `Now I know React and Redux!`가 콘솔에 출력되어야 합니다.

```js
assert(
    /console\s*\.\s*log\s*\(\s*('|"|`)Now I know React and Redux!\1\s*\)/.test(
      code)
  );
```

# --seed--

## --seed-contents--

```jsx
/*
import React from 'react'
import ReactDOM from 'react-dom'
import { Provider, connect } from 'react-redux'
import { createStore, combineReducers, applyMiddleware } from 'redux'
import thunk from 'redux-thunk'

import rootReducer from './redux/reducers'
import App from './components/App'

const store = createStore(
  rootReducer,
  applyMiddleware(thunk)
);

ReactDOM.render(
  <Provider store={store}>
    <App/>
  </Provider>,
  document.getElementById('root')
);
*/

// Only change code below this line
```

# --solutions--

```jsx
console.log('Now I know React and Redux!');
```
