---
id: 5a24c314108439a4d4036145
title: "상태를 프로프스에 매핑하기"
challengeType: 6
forumTopicId: 301433
dashedName: map-state-to-props
---

# --description--

`Provider` 컴포넌트는 React 컴포넌트에 `state`와 `dispatch`를 제공할 수 있게 해주지만, 어떤 상태와 액션을 원하는지 정확히 지정해야 합니다. 이렇게 하면 각 컴포넌트가 필요한 상태에만 접근할 수 있게 됩니다. 이를 위해 두 개의 함수를 만듭니다: `mapStateToProps()`와 `mapDispatchToProps()`입니다.

이 함수들에서 접근하고 싶은 상태 조각과 디스패치할 수 있어야 하는 액션 생성자를 선언합니다. 이 함수들이 준비되면, 다른 도전 과제에서 React Redux의 `connect` 메서드를 사용해 컴포넌트에 연결하는 방법을 배우게 됩니다.

**참고:** 내부적으로 React Redux는 `store.subscribe()` 메서드를 사용해 `mapStateToProps()`를 구현합니다.

# --instructions--

`mapStateToProps()` 함수를 만드세요. 이 함수는 `state`를 인수로 받고, 그 상태를 특정 속성 이름에 매핑하는 객체를 반환해야 합니다. 이 속성들은 `props`를 통해 컴포넌트에서 접근할 수 있게 됩니다. 이 예제에서는 앱의 전체 상태를 하나의 배열에 보관하므로, 그 전체 상태를 컴포넌트에 전달할 수 있습니다. 반환하는 객체에 `messages` 속성을 만들고, `state`로 설정하세요.

# --hints--

const `state`는 빈 배열이어야 합니다.

```js
assert(Array.isArray(state) && state.length === 0);
```

`mapStateToProps`는 함수여야 합니다.

```js
assert(typeof mapStateToProps === 'function');
```

`mapStateToProps`는 객체를 반환해야 합니다.

```js
assert(typeof mapStateToProps() === 'object');
```

배열을 상태로 `mapStateToProps`에 전달하면, 이 배열이 `messages` 키에 할당된 상태로 반환되어야 합니다.

```js
assert(mapStateToProps(['messages']).messages.pop() === 'messages');
```

# --seed--

## --seed-contents--

```jsx
const state = [];

// Change code below this line
```

# --solutions--

```jsx
const state = [];

// Change code below this line

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};
```
