---
id: 68a23004e96a105df3a21fc4
title: Construye un Asistente para Contar Cartas
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

En el juego de casino Blackjack, un jugador puede determinar si tiene una ventaja en la siguiente mano sobre la casa llevando la cuenta del número relativo de cartas altas y bajas que quedan en la baraja. Esto se llama cuenta de tarjetas.

Tener más cartas altas en la baraja es una ventaja para el jugador. Cuando el conteo es positivo, el jugador debería apostar alto. Cuando el conteo da 0 o negativo, el jugador debería apostar bajo.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes usar `let` para declarar una variable global llamada `count` y asignarle el valor `0`.
1. Deberías tener una función llamada `cardCounter`.
1. La función `cardCounter` debe recibir un parámetro `card` que puede ser un número o una cadena.
    - Para valores entre `2` y `10`, el parámetro `card` será un número.
    - Para todos los demás valores, el parámetro `card` será una cadena.
1. La función `cardCounter` debería modificar la variable global `count` según ciertos criterios.
1. La variable global `count` debe incrementarse en `1` para las cartas `2`, `3`, `4`, `5` o `6`
1. La variable global `count` debe permanecer sin cambios para las tarjetas `7`, `8`, `9`.
1. La variable global `count` debe disminuirse en `1` para las cartas `10`, `"J"`, `"Q"`, `"K"`, `"A"`
1. La función `cardCounter` debe devolver una cadena con el conteo actual y la cadena `Bet` si el conteo es positivo.
1. La función `cardCounter` debe devolver una cadena con el conteo actual y la cadena `Hold` si el conteo es menor o igual a `0`.
1. En la salida de la función, el conteo actual y la decisión del jugador (`Bet` o `Hold`) deben estar separados por un espacio. Por ejemplo, `-3 Hold`.

# --hints--

Debes usar `let` para declarar una variable global llamada `count` y asignarle el valor `0`. 

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

Debes tener una función llamada `cardCounter`. 

```js
assert.isFunction(cardCounter);
```

Tu función debe devolver el valor de `count` y el texto (`Bet` o `Hold`) con un espacio entre ellos.

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

Después de las cartas `2`, `3`, `4`, `5`, entonces llamar a `cardCounter(6)` debería devolver la cadena `5 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

Después de las cartas `7`, `8`, llamar a `cardCounter(9)` debería devolver la cadena `0 Hold`.

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

Después de las cartas `10`, `"J"`, `"Q"`, `"K"`, entonces llamar a `cardCounter("A")` debería devolver la cadena `-5 Hold`.

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

Después de las cartas `3`, `7`, `"Q"`, `8`, entonces llamar a `cardCounter("A")` debería devolver la cadena `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

Después de las cartas `2`, `"J"`, `9`, `2`, entonces llamar a `cardCounter(7)` debería devolver la cadena `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

Después de las cartas `2`, `2`, entonces llamar a `cardCounter(10)` debería devolver la cadena `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

Después de las cartas `3`, `2`, `"A"`, `10`, entonces llamar a `cardCounter("K")` debería devolver la cadena `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
