---
id: 66eaddd04a9e533fba689001
title: Progetta una scheda per un post del blog
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

In questo laboratorio, praticherai come definire lo stile di sfondi e bordi creando una scheda per un post del blog.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `div` con una classe `blog-post-card` che contenga tutti gli elementi della tua scheda.
2. All’interno dell’elemento `.blog-post-card`, dovresti avere un elemento immagine con un attributo `alt` valido e del testo, e una classe `post-img`. Puoi usare `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` per l’attributo `src` dell’immagine.
3. Dovresti avere un `div` con una classe `post-content` all’interno dell’elemento `.blog-post-card` con quanto segue:
    - Un elemento `h2` con una classe `post-title` e del testo per il titolo del tuo post del blog.
    - Un elemento `p` con una classe `post-excerpt` e del testo per riassumere il tuo post del blog.
    - Un elemento `a` con una classe `read-more` con il testo `Read More`.
4. Dovresti impostare il colore di sfondo dell’elemento `body` su un valore diverso dal bianco.
5. Dovresti applicare i seguenti stili all’elemento `.blog-post-card`:
    - Uno sfondo bianco.
    - Angoli arrotondati.
    - Una larghezza a tua scelta.
    - Un allineamento del testo a tua scelta.
6. L’elemento `.post-img` dovrebbe essere stilizzato in modo che l’immagine riempia tutta la larghezza della scheda e abbia un bordo inferiore.
7. L’elemento `.post-content` dovrebbe essere stilizzato in modo che ci sia spaziatura interna all’interno della scheda.
8. Gli elementi `.post-title` e `.post-excerpt` dovrebbero avere un colore del testo diverso dal predefinito e margini su tutti i lati.
9. L’elemento `.read-more` dovrebbe essere stilizzato come un pulsante e avere:
    - Un colore del testo diverso dal predefinito.
    - Un colore di sfondo.
    - Margini su tutti i lati.
    - Spaziatura interna su tutti i lati.
    - Angoli arrotondati.
    - Una proprietà `display` impostata su `inline-block`.
    - Un effetto al passaggio del mouse che cambia il colore di sfondo.

**Nota:** Assicurati di collegare il tuo foglio di stile nel tuo HTML e di applicare il CSS.

# --hints--

Dovresti avere un `div` con una classe `blog-post-card`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

Dovresti avere un elemento `img` con un `class` di `post-img`. Assicurati che la tua immagine abbia un attributo `alt` con del testo e un attributo `src` con un valore.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

Dovresti avere un `div` con classe `post-content`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

Dovresti avere un elemento `h2` con una classe `post-title`. Assicurati che abbia del contenuto testuale per il titolo del tuo post del blog.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

Dovresti avere un `p` con una classe `post-excerpt`. Assicurati che abbia del contenuto testuale per riassumere il tuo post del blog.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

Dovresti avere un elemento `a` con un `class` di `read-more`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

Il tuo elemento `.read-more` dovrebbe avere il testo `Read More`.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

Il tuo elemento `.blog-post-card` dovrebbe avere una proprietà `border-radius` con un valore (non dovrebbe essere `0` o un valore negativo).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

Il tuo elemento `.blog-post-card` dovrebbe avere uno sfondo bianco.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

Il tuo elemento `body` dovrebbe avere un colore di sfondo diverso dal bianco.

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

Dovresti mirare a `.blog-post-card` e impostare la sua proprietà `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

Dovresti mirare a `.blog-post-card` e impostare la sua proprietà `text-align`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

Dovresti mirare a `.post-content` e impostare la sua proprietà `padding`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

Il tuo elemento `.read-more` dovrebbe avere un effetto al passaggio del mouse.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

Dovresti mirare a `.read-more` e impostare la sua proprietà `color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

Dovresti mirare a `.read-more` e impostare la sua proprietà `background-color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

Dovresti mirare a `.read-more` e impostare la sua proprietà `margin`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

Dovresti mirare a `.read-more` e impostare la sua proprietà `display`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

Dovresti mirare a `.read-more` e impostare la sua proprietà `border-radius`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

Dovresti mirare a `.read-more` e impostare la sua proprietà `padding`. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

Il tuo elemento `.post-img` dovrebbe riempire il `width` della scheda.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

Il tuo elemento `.post-img` dovrebbe avere un valore `border-bottom`.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

I tuoi elementi `.post-title` e `.post-excerpt` dovrebbero avere margini e colori del testo non predefiniti.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
