---
id: 66db57ad34c7089b9b41bfd6
title: 建置 Lightbox Viewer
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

Lightbox 用於網站上，透過在頁面主要內容上方彈出或模態視窗，以更具吸引力的方式展示多媒體內容。

在本實驗中，你將創建一個燈箱畫廊，當點擊縮圖時會顯示全尺寸影像。每個影像都提供兩個版本：縮圖和全尺寸影像。全尺寸影像與縮圖相同，但沒有 `-thumbnail` 後綴。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該在你的 `body` 中有一個類別為 `gallery` 的 `div`。
1. 在 `.gallery` 元素中，你應該有三個影像縮圖，每個都帶有 `gallery-item` 類別。你應該使用以下聯結作為縮圖影像：
   
    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. 你應該在你的 `body` 中有一個 `class` 為 `lightbox` 的 `div`。
1. 你應該在你的 `.lightbox` 元素中有一個 `button`，其 `id` 為 `close-btn`。如果你想，可以使用 `&times;` 作為它的文字。
1. 你應該在你的 `.lightbox` 元素中放置一個 `id` 為 `lightbox-image` 的 `img`。
1. 你的 `.lightbox` 元素應該具有固定位置，以便預覽能夠在目前的影像上方開啟。
1. 你的 `.lightbox` 元素應該透過將高度和寬度設定為容器的 100% 來涵蓋整個視窗。你應該確保 `.lightbox` 元素從容器的左上角開始。
1. `.lightbox` 應該有一個背景顏色。最初，其 `display` 屬性應設為 `none` 以隱藏它。
1. 當你點擊你的 `.gallery-item` 項目之一時，`.lightbox` 項目的 `display` 屬性應該設為 `flex`，以使 `.lightbox` 項目及其內的兩個項目可見。
1. 當你點擊你的 `.gallery-item` 項目其中一個時，`#lightbox-image` 元素的 `src` 應該被設定為點擊的影像的全尺寸版本，方法是從影像的 `src` 屬性中移除 `-thumbnail`。全尺寸影像位於以下連結：

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`
  
1. 當你的 `.lightbox` 元素可見且你點擊 `#close-btn` 或 `.lightbox` 元素時，`.lightbox` 元素的 `display` 應該被設回 `none`。

**注意：** 請務必在你的 HTML 中連結你的樣式表和 JavaScript 檔案。

# --hints--

你應該在你的 `body` 元素內有一個類別為 `gallery` 的 `div`。

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

在 `.gallery` 元素中，你應該有三個帶有 `gallery-item` 類別的 `img` 元素。

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

在 `.gallery` 元素內，你應該有一個 `img` 元素，其 `src` 設定為 `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

在 `.gallery` 元素內，你應該有一個 `img` 元素，其 `src` 設定為 `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

在 `.gallery` 元素內，你應該有一個 `img` 元素，其 `src` 設定為 `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

你應該在你的 `body` 元素內有一個類別為 `lightbox` 的 `div` 元素。

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

在你的 `.lightbox` 元素中，應該有一個 `button` 元素，其 `id` 設定為 `close-btn`。

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

在你的 `.lightbox` 元素中，應該有一個 `img` 元素，其 `id` 設定為 `lightbox-image`。

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

你的 `.lightbox` 元素應該具有固定定位。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

你的 `.lightbox` 元素應該涵蓋整個視窗。

```js
const style = new __helpers.CSSHelp(document).getStyle(".lightbox");

const width = style?.width;
const height = style?.height;

assert.oneOf(
  width,
  ["100%", "100vw"]
);

assert.oneOf(
  height,
  ["100%", "100vh"]
);
```

你的 `.lightbox` 元素應該與容器的左上角對齊。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

你的 `.lightbox` 元素應該有背景顏色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

你的 `.lightbox` 元素應該一開始隱藏。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

當你點擊你的 `.gallery-item` 項目之一時，應將 `.lightbox` 元素的 `display` 屬性設為 `flex`，以使 `.lightbox` 及其內的兩個元素可見。

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

當你點擊你的 `.gallery-item` 項目之一時，`#lightbox-image` 元素的 `src` 應該設定為點擊的影像的全尺寸版本，方法是從影像的 `src` 屬性中移除 `-thumbnail`。

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

當你的 `.lightbox` 元素可見且你點擊 `#close-btn` 按鈕時，`.lightbox` 元素的 `display` 應該被設回 `none`。

```js
const lightbox = document.querySelector(".lightbox");
const closeBtn = document.getElementById("close-btn"); 
const galleryItem = document.querySelector(".gallery-item");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

closeBtn.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

當你的 `.lightbox` 元素可見且你點擊 `.lightbox` 元素時，應將 `.lightbox` 元素的 `display` 設回 `none`。

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

lightbox.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```
