---
id: 66db529d37ad966480ebb633
title: Crie um programa seletor de almoços
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

Neste laboratório, você construirá um programa que ajuda a gerenciar opções de almoço.Você trabalhará com um array de almoços, adicionará e removerá itens do array e selecionará aleatoriamente uma opção de almoço.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma variável `lunches` e atribuir a ela um array vazio que será usado para armazenar os itens de almoço.

2. Você deve criar uma função `addLunchToEnd` que recebe um array como primeiro argumento e uma string como segundo argumento. A função deve:
    - Adicione a string ao final do array.
    - Registre a string `"[Lunch Item] added to the end of the lunch menu."` no console, onde `[Lunch Item]` é a string passada para a função.
    - Retorne o array atualizado.

3. Você deve criar uma função `addLunchToStart` que recebe um array como primeiro argumento e uma string como segundo argumento. A função deve:
    - Adicione a string ao início do array.
    - Registre a string `"[Lunch Item] added to the start of the lunch menu."` no console, onde `[Lunch Item]` é a string passada para a função.
    - Retorne o array atualizado.

4. Você deve criar uma função `removeLastLunch` que recebe um array como argumento. A função deve:
   - Remova o último elemento do array.
   - Se a remoção for bem-sucedida, registre a string `"[Lunch Item] removed from the end of the lunch menu."` no console, onde `[Lunch Item]` é o elemento removido do array.
   - Se o array estiver vazio, registre a string `"No lunches to remove."` no console.
   - Retorne o array atualizado.

5. Você deve criar uma função `removeFirstLunch` que recebe um array como argumento. A função deve:
   - Remova o primeiro elemento do array.
   - Se a remoção for bem-sucedida, registre a string `"[Lunch Item] removed from the start of the lunch menu."` no console, onde `[Lunch Item]` é o elemento removido do array.
   - Se o array estiver vazio, registre a string `"No lunches to remove."` no console.
   - Retorne o array atualizado.

6. Você deve criar uma função `getRandomLunch` que recebe um array como argumento. A função deve:
   - Selecione um elemento aleatório do array.
   - Se for bem-sucedido, registre a string `"Randomly selected lunch: [Lunch Item]"` no console, onde `[Lunch Item]` é um elemento aleatório no array.
   - Se o array estiver vazio, registre a string `"No lunches available."` no console.

7. Você deve criar uma função `showLunchMenu` que recebe um array como argumento e:
   - Se houver elementos no array, registra a string `"Menu items: [Lunch Item], [Lunch Item]..."` no console, onde cada `[Lunch item]` é um dos elementos no array, na ordem.
   - Se o array estiver vazio, registra a string `"The menu is empty."` no console.

# --hints--

Você deve declarar uma variável `lunches` e atribuir a ela um array vazio para armazenar os itens do almoço.

```js
assert.isArray(lunches);
```

Você deve definir uma função `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

A função `addLunchToEnd` deve ter dois parâmetros.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` deve registrar a string `"Tacos adicionados ao final do menu de almoço."` no console.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` deve retornar `["Pizza", "Tacos", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

Você deve definir uma função `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

A função `addLunchToStart` deve ter dois parâmetros.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` deve registrar a string `"Sushi adicionado ao início do menu de almoço."` no console.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` deve retornar `["Pizza", "Burger", "Sushi"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

Você deve definir uma função `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

A função `removeLastLunch` deve ter um parâmetro.

```js
assert.lengthOf(removeLastLunch, 1);
```

Quando o array de entrada estiver vazio, a função `removeLastLunch` deve registrar a string `"Nenhum almoço para remover."` no console.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Ensopado", "Sopa", "Torrada"])` deve registrar a string `"Torrada removida do final do menu de almoço."` no console.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` deve retornar `["Sushi", "Pizza"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

Você deve definir uma função `removeFirstLunch`.

```js
assert.isFunction(removeFirstLunch);
```

A função `removeFirstLunch` deve ter um único parâmetro.

```js
assert.lengthOf(removeFirstLunch, 1);
```

Quando o array de entrada estiver vazio, a função `removeFirstLunch` deve registrar a string `"Nenhum almoço para remover."` no console.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salada", "Ovos", "Queijo"])` deve registrar a string `"Salada removida do início do menu de almoço."` no console.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` deve retornar `["Pizza", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`, `addLunchToStart`, `removeLastLunch` e `removeFirstLunch` devem retornar o mesmo array passado como argumento após atualizá-lo.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

Você deve definir uma função `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

A função `getRandomLunch` deve ter um único parâmetro.

```js
assert.lengthOf(getRandomLunch, 1);
```

Quando o array de entrada estiver vazio, a função `getRandomLunch` deve registrar a string `"Nenhum almoço disponível."` no console.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

Quando o array de entrada não estiver vazio, a função `getRandomLunch` deve registrar uma string no formato `"Almoço selecionado aleatoriamente: [Item do Almoço]"` no console.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

A função `getRandomLunch` não deve alterar o array passado para ela como argumento.

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

Você deve definir uma função `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

A função `showLunchMenu` deve ter um único parâmetro.

```js
assert.lengthOf(showLunchMenu, 1);
```

Quando o array de entrada estiver vazio, a função `showLunchMenu` deve registrar a string `"O menu está vazio."` no console.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Verduras", "Milhos", "Feijões"])` deve registrar `"Itens do menu: Verduras, Milhos, Feijões"` no console.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Hambúrguer", "Batatas Fritas", "Salada"])` deve registrar `"Itens do menu: Pizza, Hambúrguer, Batatas Fritas, Salada"` no console.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
