---
id: 6619240f46cec8e04d77e03a
title: Базове завдання з функцій A
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

Створіть функцію, яка приймає ціле число. Ця функція має повертати заданий `integer + 7`, якщо ціле число менше за `10`. Якщо ціле число більше або дорівнює `10`, вона має повертати заданий `integer - 3`.

Назва функції має бути `addOrSubtract`.

# --hints--

У вас має бути функція з назвою `addOrSubtract`.

```js
assert.isFunction(addOrSubtract);
```

Ваша функція має приймати ціле число як аргумент.

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

Ви маєте повертати задане ціле число + 7, якщо ціле число менше за 10.

```js
assert.strictEqual(addOrSubtract(5), 12);
```

Ви маєте повертати задане ціле число - 3, якщо ціле число більше або дорівнює 10.

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
