---
id: 661e151f068359c3ccf2f4d7
title: Базове завдання з функцій C
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

Напишіть функцію з назвою `capitalize`, яка приймає рядок як параметр і повертає новий рядок із першою літерою, написаною з великої.

# --hints--

У вас має бути функція з назвою `capitalize`.

```js
assert.isFunction(capitalize);
```

Ваша функція має приймати рядок як параметр.

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

Ваша функція має повертати новий рядок із першою літерою, написаною з великої.

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
