---
id: 661e17c6068359c3ccf2f4d8
title: Базове завдання з функцій D
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

Напишіть функцію з іменем `lastLetter`, яка приймає рядок як параметр і повертає останню літеру цього рядка.

# --hints--

У вас має бути функція з іменем `lastLetter`.

```js
assert.isFunction(lastLetter);
```

Ваша функція повинна приймати рядок як параметр.

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

Ви маєте повернути останню літеру рядка.

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
