---
id: 661e131f068359c3ccf2f4d6
title: Основні функції. Завдання B
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

Напишіть функцію з назвою `multiply`, яка приймає два параметри і повертає їхній добуток.

# --hints--

У вас має бути функція з назвою `multiply`.

```js
assert.isFunction(multiply);
```

Ваша функція має приймати два цілі числа як аргументи.

```js
assert.match(multiply.toString(), /\s*multiply\(\s*\w+\s*,\s*\w+\s*\)/);
```

Ви повинні повернути добуток цих двох цілих чисел.

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js 
function multiply(a, b) {
  return a * b;
}
```
