---
id: 587d7da9367417b2b2512b69
title: sort メソッドを使って配列をアルファベット順に並べ替える
challengeType: 1
forumTopicId: 18303
dashedName: sort-an-array-alphabetically-using-the-sort-method
---

# --description--

`sort` メソッドは、配列の要素をコールバック関数に従って並べ替えます。

例えば：

```js
function ascendingOrder(arr) {
  return arr.sort(function(a, b) {
    return a - b;
  });
}

ascendingOrder([1, 5, 2, 3, 4]);
```

これは `[1, 2, 3, 4, 5]` の値を返します。

```js
function reverseAlpha(arr) {
  return arr.sort(function(a, b) {
    return a === b ? 0 : a < b ? 1 : -1;
  });
}

reverseAlpha(['l', 'h', 'z', 'b', 's']);
```

これは `['z', 's', 'l', 'h', 'b']` の値を返します。

JavaScript のデフォルトの並べ替え方法は文字列の Unicode ポイント値によるため、予期しない結果になる場合があります。したがって、配列の要素をどのように並べ替えるかを指定するコールバック関数を提供することをお勧めします。通常 `compareFunction` と呼ばれるそのようなコールバック関数が渡されると、配列の要素は `compareFunction` の戻り値に従って並べ替えられます。もし `compareFunction(a,b)` が 0 未満の値を 2 つの要素 `a` と `b` に対して返した場合、`a` は `b` より前に来ます。もし `compareFunction(a,b)` が 0 より大きい値を 2 つの要素 `a` と `b` に対して返した場合、`b` は `a` より前に来ます。もし `compareFunction(a,b)` が 0 と等しい値を 2 つの要素 `a` と `b` に対して返した場合、`a` と `b` は変更されません。

# --instructions--

`sort` 関数内で `alphabeticalOrder` メソッドを使い、`arr` の要素をアルファベット順に並べ替えてください。関数は並べ替えられた配列を返す必要があります。

# --hints--

コードは `sort` メソッドを使う必要があります。

```js
assert(__helpers.removeJSComments(code).match(/\.sort/g));
```

`alphabeticalOrder(["a", "d", "c", "a", "z", "g"])` は `["a", "a", "c", "d", "g", "z"]` を返すべきです。

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'd', 'c', 'a', 'z', 'g'])) ===
    JSON.stringify(['a', 'a', 'c', 'd', 'g', 'z'])
);
```

`alphabeticalOrder(["x", "h", "a", "m", "n", "m"])` は `["a", "h", "m", "m", "n", "x"]` を返すべきです。

```js
assert(
  JSON.stringify(alphabeticalOrder(['x', 'h', 'a', 'm', 'n', 'm'])) ===
    JSON.stringify(['a', 'h', 'm', 'm', 'n', 'x'])
);
```

`alphabeticalOrder(["a", "a", "a", "a", "x", "t"])` は `["a", "a", "a", "a", "t", "x"]` を返すべきです。

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'a', 'a', 'a', 'x', 't'])) ===
    JSON.stringify(['a', 'a', 'a', 'a', 't', 'x'])
);
```

# --seed--

## --seed-contents--

```js
function alphabeticalOrder(arr) {
  // Only change code below this line

  return arr
  // Only change code above this line
}

alphabeticalOrder(["a", "d", "c", "a", "z", "g"]);
```

# --solutions--

```js
function alphabeticalOrder(arr) {
  return arr.sort();
}
```
