---
id: 587d7dab367417b2b2512b70
title: カリー化と部分適用の紹介
challengeType: 1
forumTopicId: 301232
dashedName: introduction-to-currying-and-partial-application
---

# --description--

<dfn>arity</dfn>（関数のアリティ）とは、関数が必要とする引数の数のことです。<dfn>カリー化</dfn>とは、N個の引数を取る関数を、アリティ1の関数をN個返す関数に変換することを指します。

つまり、関数を1つの引数を受け取り、次の引数を受け取る別の関数を返す形に再構成します。

例を示します：

```js
function unCurried(x, y) {
  return x + y;
}

function curried(x) {
  return function(y) {
    return x + y;
  }
}

const curried = x => y => x + y

curried(1)(2)
```

`curried(1)(2)` は `3` を返します。

これは、関数に一度にすべての引数を渡せない場合に便利です。各関数呼び出しを変数に保存しておけば、次の引数が利用可能になったときに受け取る関数の参照を保持できます。上記のカリー化された関数を使った例を示します：

```js
const funcForY = curried(1);
console.log(funcForY(2)); // 3
```

同様に、<dfn>部分適用</dfn>は、関数にいくつかの引数を渡して別の関数を返し、その関数にさらに引数を渡す操作として説明できます。例を示します：

```js
function impartial(x, y, z) {
  return x + y + z;
}

const partialFn = impartial.bind(this, 1, 2);
partialFn(10); // 13
```

# --instructions--

`add` 関数の本体を完成させて、`x`、`y`、`z` のパラメータをカリー化で加えるようにしてください。

# --hints--

`add(10)(20)(30)` は `60` を返す必要があります。

```js
assert(add(10)(20)(30) === 60);
```

`add(1)(2)(3)` は `6` を返す必要があります。

```js
assert(add(1)(2)(3) === 6);
```

`add(11)(22)(33)` は `66` を返す必要があります。

```js
assert(add(11)(22)(33) === 66);
```

コードの最後には `x + y + z` を返す文が含まれている必要があります。

```js
assert(__helpers.removeJSComments(code).match(/[xyz]\s*?\+\s*?[xyz]\s*?\+\s*?[xyz]/g));
```

# --seed--

## --seed-contents--

```js
function add(x) {
  // Only change code below this line


  // Only change code above this line
}

add(10)(20)(30);
```

# --solutions--

```js
const add = x => y => z => x + y + z
```
