---
id: 587d7b8d367417b2b2512b5b
title: 関数型プログラミングについて学ぶ
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

関数型プログラミングは、解決策がシンプルで独立した関数であり、関数のスコープ外に副作用がないプログラミングスタイルです：`INPUT -> PROCESS -> OUTPUT`

関数型プログラミングのポイントは次の通りです：

1) 独立した関数 - プログラムの状態、つまり変更される可能性のあるグローバル変数に依存しません

2) 純粋な関数 - 同じ入力に対して常に同じ出力を返します

3) 副作用が限定された関数 - 関数外のプログラムの状態に対する変更やミューテーションは慎重に管理されます

# --instructions--

freeCodeCampのメンバーはお茶が大好きです。

コードエディターでは、`prepareTea` と `getTea` 関数がすでに定義されています。`getTea` 関数を呼び出してチームのために40杯のお茶を用意し、`tea4TeamFCC` 変数に保存してください。

# --hints--

`tea4TeamFCC` 変数にはチームのための40杯のお茶が入っている必要があります。

```js
assert(tea4TeamFCC.length === 40);
```

`tea4TeamFCC` 変数には緑茶のカップが入っている必要があります。

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40); 
```
