---
id: 587d7da9367417b2b2512b66
title: concat メソッドを使って2つの配列を結合する
challengeType: 1
forumTopicId: 301229
dashedName: combine-two-arrays-using-the-concat-method
---

# --description--

<dfn>連結</dfn>とは、アイテムを端から端へつなげることを意味します。JavaScript は文字列と配列の両方に同じように動作する `concat` メソッドを提供しています。配列の場合、このメソッドは一方の配列で呼び出され、もう一方の配列が `concat` の引数として渡され、最初の配列の末尾に追加されます。このメソッドは新しい配列を返し、元のどちらの配列も変更しません。例を示します。

```js
[1, 2, 3].concat([4, 5, 6]);
```

返される配列は `[1, 2, 3, 4, 5, 6]` になります。

# --instructions--

`concat` 関数内で `nonMutatingConcat` メソッドを使い、`attach` の末尾に `original` を連結してください。関数は連結された配列を返す必要があります。

# --hints--

コードは `concat` メソッドを使う必要があります。

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

`first` 配列は変更されてはいけません。

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

`second` 配列は変更されてはいけません。

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingConcat([1, 2, 3], [4, 5])` は `[1, 2, 3, 4, 5]` を返すべきです。

```js
assert(
  JSON.stringify(nonMutatingConcat([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingConcat(original, attach) {
  // Only change code below this line


  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingConcat(first, second);
```

# --solutions--

```js
function nonMutatingConcat(original, attach) {
  return original.concat(attach);
}
const first = [1, 2, 3];
const second = [4, 5];
```
