---
id: 587d7dab367417b2b2512b6f
title: 配列の要素が条件を満たすかどうかを some メソッドで確認する
challengeType: 1
forumTopicId: 301314
dashedName: use-the-some-method-to-check-that-any-elements-in-an-array-meet-a-criteria
---

# --description--

`some` メソッドは配列で動作し、*いずれかの*要素が特定のテストに合格するかどうかを確認します。条件を満たす値があればブール値の `true` を返し、なければ `false` を返します。

例えば、次のコードは `numbers` 配列の中に 10 未満の要素があるかどうかをチェックします。

```js
const numbers = [10, 50, 8, 220, 110, 11];

numbers.some(function(currentValue) {
  return currentValue < 10;
});
```

`some` メソッドは `true` を返します。

# --instructions--

`some` 関数内で `checkPositive` メソッドを使い、`arr` の中に正の値があるかどうかをチェックしてください。関数はブール値を返す必要があります。

# --hints--

コードは `some` メソッドを使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/\.some/g));
```

`checkPositive([1, 2, 3, -4, 5])` は `true` を返すべきです。

```js
assert(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` は `true` を返すべきです。

```js
assert(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([-1, -2, -3, -4, -5])` は `false` を返すべきです。

```js
assert(!checkPositive([-1, -2, -3, -4, -5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.some(elem => elem > 0);
}
```
