---
id: 587d7b8f367417b2b2512b64
title: プロトタイプ上に filter メソッドを実装する
challengeType: 1
forumTopicId: 301231
dashedName: implement-the-filter-method-on-a-prototype
---

# --description--

独自の `filter` メソッドを実装すると、その動作について多く学べるでしょう。`for` ループか `Array.prototype.forEach()` を使うことをお勧めします。

# --instructions--

`Array.prototype.myFilter()` とまったく同じ動作をする独自の `Array.prototype.filter()` を作成してください。組み込みの `filter` メソッドは使わないでください。`Array` インスタンスは `myFilter` メソッド内で `this` を使ってアクセスできます。

# --hints--

`[23, 65, 98, 5, 13].myFilter(item => item % 2)` は `[23, 65, 5, 13]` と等しくなるはずです。

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item % 2;
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

`["naomi", "quincy", "camperbot"].myFilter(element => element === "naomi")` は `["naomi"]` を返すはずです。

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element === "naomi";
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

`[1, 1, 2, 5, 2].myFilter((element, index, array) => array.indexOf(element) === index)` は `[1, 2, 5]` を返すはずです。

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array.indexOf(element) === index;
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

コード内で `filter` メソッドを使わないでください。

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?filter/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myFilter = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myFilter = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    if (callback(this[i], i, this)) newArray.push(this[i]);
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const odd_s = s.myFilter(item => item % 2 === 1);
```
