---
id: 587d7b88367417b2b2512b45
title: '高階関数 map、filter、reduce を使って複雑な問題を解く'
challengeType: 1
forumTopicId: 301311
dashedName: use-higher-order-functions-map-filter-or-reduce-to-solve-a-complex-problem
---

# --description--

`map()`、`filter()`、`reduce()` のような高階関数を使ったいくつかのチャレンジを終えたので、これらを使ってより複雑なチャレンジを解いてみましょう。

# --instructions--

`squareList` 関数のコードを `map()`、`filter()`、`reduce()` のいずれかの組み合わせを使って完成させてください。この関数は、実数の配列が渡されたときに、*正の整数のみ*（小数は整数ではありません）の二乗を含む新しい配列を返す必要があります。実数の配列の例は `[-3, 4.8, 5, 3, -3.2]` です。

**注意:** 関数内で `for` や `while` のようなループや `forEach()` 関数を使わないでください。

# --hints--

`squareList` は `function` であるべきです。

```js
assert.typeOf(squareList, 'function'),
  '<code>squareList</code> should be a <code>function</code>';
```

`for`、`while`、`forEach` は使わないでください。

```js
assert(!__helpers.removeJSComments(code).match(/for|while|forEach/g));
```

`map`、`filter`、または `reduce` を使うべきです。

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/\.(map|filter|reduce)\(/g)
);
```

関数は `array` を返すべきです。

```js
assert(Array.isArray(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])));
```

`squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])` は `[16, 1764, 36]` を返すべきです。

```js
assert.deepStrictEqual(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2]), [
  16,
  1764,
  36
]);
```

`squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3])` は `[9, 100, 49]` を返すべきです。

```js
assert.deepStrictEqual(squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3]), [
  9,
  100,
  49
]);
```

# --seed--

## --seed-contents--

```js
const squareList = arr => {
  // Only change code below this line
  return arr;
  // Only change code above this line
};

const squaredIntegers = squareList([-3, 4.8, 5, 3, -3.2]);
console.log(squaredIntegers);
```

# --solutions--

```js
const squareList = arr => {
  const positiveIntegers = arr.filter(num => {
    return num >= 0 && Number.isInteger(num);
  });
  const squaredIntegers = positiveIntegers.map(num => {
    return num ** 2;
  });
  return squaredIntegers;
};
```
