---
id: 587d7da9367417b2b2512b6a
title: 元の配列を変更せずにソートされた配列を返す
challengeType: 1
forumTopicId: 301237
dashedName: return-a-sorted-array-without-changing-the-original-array
---

# --description--

`sort` メソッドの副作用として、元の配列の要素の順序が変わってしまいます。つまり、その配列をその場で変更（ミューテート）します。これを避ける方法の一つは、まず空の配列をソート対象の配列に連結することです（`slice` と `concat` は新しい配列を返すことを覚えておいてください）。その後、`sort` メソッドを実行します。

# --instructions--

`sort` 関数内で `nonMutatingSort` メソッドを使い、配列の要素を昇順にソートしてください。この関数は新しい配列を返し、`globalArray` 変数を変更してはいけません。

# --hints--

コードには `sort` メソッドを使うべきです。

```js
assert(nonMutatingSort.toString().match(/\.sort/g));
```

`globalArray` 変数は変更されてはいけません。

```js
assert(JSON.stringify(globalArray) === JSON.stringify([5, 6, 3, 2, 9]));
```

`nonMutatingSort(globalArray)` は `[2, 3, 5, 6, 9]` を返すべきです。

```js
assert(
  JSON.stringify(nonMutatingSort(globalArray)) ===
    JSON.stringify([2, 3, 5, 6, 9])
);
```

`nonMutatingSort(globalArray)` はハードコーディングしてはいけません。

```js
assert(!nonMutatingSort.toString().match(/\[.*?[23569].*?\]/gs));
```

関数は渡された配列ではなく、新しい配列を返すべきです。

```js
assert(nonMutatingSort(globalArray) !== globalArray);
```

`nonMutatingSort([1, 30, 4, 21, 100000])` は `[1, 4, 21, 30, 100000]` を返すべきです。

```js
assert(JSON.stringify(nonMutatingSort([1, 30, 4, 21, 100000])) ===
    JSON.stringify([1, 4, 21, 30, 100000]))
```

`nonMutatingSort([140000, 104, 99])` は `[99, 104, 140000]` を返すべきです。

```js
assert(JSON.stringify(nonMutatingSort([140000, 104, 99])) ===
    JSON.stringify([99, 104, 140000]))
```

# --seed--

## --seed-contents--

```js
const globalArray = [5, 6, 3, 2, 9];

function nonMutatingSort(arr) {
  // Only change code below this line


  // Only change code above this line
}

nonMutatingSort(globalArray);
```

# --solutions--

```js
const globalArray = [5, 6, 3, 2, 9];
function nonMutatingSort(arr) {
  return [].concat(arr).sort((a,b) => a-b);
}
```
