---
id: 587d7b8e367417b2b2512b5f
title: 関数の外部依存を避けるために引数を渡す
challengeType: 1
forumTopicId: 301234
dashedName: pass-arguments-to-avoid-external-dependence-in-a-function
---

# --description--

前回のチャレンジは関数型プログラミングの原則に一歩近づきましたが、まだ足りない部分があります。

グローバル変数の値は変更しませんでしたが、関数 `incrementer` はグローバル変数 `fixedValue` が存在しなければ動作しません。

関数型プログラミングのもう一つの原則は、依存関係を常に明示的に宣言することです。つまり、関数が変数やオブジェクトに依存している場合、その変数やオブジェクトを引数として直接関数に渡す必要があります。

この原則にはいくつかの良い効果があります。関数のテストがしやすくなり、どの入力を受け取るかが明確になり、プログラムの他の部分に依存しなくなります。

これにより、コードを変更、削除、追加するときに自信を持てます。何を変更できて何を変更できないかがわかり、潜在的な落とし穴も見えやすくなります。

最後に、関数は同じ入力に対して常に同じ出力を返すようになります。コードのどの部分から実行されても同じ結果が得られます。

# --instructions--

`incrementer` 関数を更新して、依存関係を明確に宣言しましょう。

引数を受け取り、その値を1増やして結果を返す `incrementer` 関数を書いてください。

# --hints--

`incrementer` 関数は `fixedValue`（`4`）の値を変更してはいけません。

```js
assert(fixedValue === 4);
```

`incrementer` 関数は引数を受け取る必要があります。

```js
assert(incrementer.length === 1);
```

`incrementer` 関数は `fixedValue` の値より1大きい値を返す必要があります。

```js
const __newValue = incrementer(fixedValue);
assert(__newValue === 5);
```

# --seed--

## --seed-contents--

```js
// The global variable
let fixedValue = 4;

// Only change code below this line
function incrementer() {


  // Only change code above this line
}
```

# --solutions--

```js
let fixedValue = 4;

function incrementer(fixedValue) {
  return fixedValue + 1;
}
```
