---
id: 587d7dab367417b2b2512b6e
title: 配列のすべての要素が条件を満たしているかを every メソッドで確認する
challengeType: 1
forumTopicId: 301312
dashedName: use-the-every-method-to-check-that-every-element-in-an-array-meets-a-criteria
---

# --description--

`every` メソッドは配列を操作して、*すべての* 要素が特定のテストに合格するかどうかを確認します。すべての値が条件を満たす場合はブール値の `true` を返し、そうでなければ `false` を返します。

例えば、次のコードは `numbers` 配列のすべての要素が 10 未満かどうかをチェックします。

```js
const numbers = [1, 5, 8, 0, 10, 11];

numbers.every(function(currentValue) {
  return currentValue < 10;
});
```

ここで `every` メソッドは `false` を返します。

# --instructions--

`every` 関数内で `checkPositive` メソッドを使い、`arr` のすべての要素が正の数かどうかを確認してください。関数はブール値を返す必要があります。

# --hints--

`every` メソッドを使う必要があります。

```js
assert(__helpers.removeJSComments(code).match(/\.every/g));
```

`checkPositive([1, 2, 3, -4, 5])` は `false` を返すべきです。

```js
assert.isFalse(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` は `true` を返すべきです。

```js
assert.isTrue(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([1, -2, 3, -4, 5])` は `false` を返すべきです。

```js
assert.isFalse(checkPositive([1, -2, 3, -4, 5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.every(num => num > 0);
}
```
