---
id: 587d7b8f367417b2b2512b62
title: プロトタイプ上に map を実装する
challengeType: 1
forumTopicId: 301230
dashedName: implement-map-on-a-prototype
---

# --description--

`Array.prototype.map()` や単に `map()` を適用したことからわかるように、`map` メソッドは呼び出された配列と同じ長さの配列を返します。また、コールバック関数が変更しない限り、元の配列は変更されません。

つまり、`map` は純粋な関数であり、その出力は入力のみに依存します。さらに、別の関数を引数として受け取ります。

自分で `map` メソッドのバージョンを実装すると、多くのことが学べるでしょう。`for` ループや `Array.prototype.forEach()` を使うことをお勧めします。

# --instructions--

組み込みの `Array.prototype.myMap()` メソッドの代わりに、`Array.prototype.map()` とまったく同じ動作をする `map` を自分で書いてください。`Array` インスタンスは `myMap` メソッド内で `this` を使ってアクセスできます。

# --hints--

`[23, 65, 98, 5, 13].myMap(item => item * 2)` は `[46, 130, 196, 10, 26]` と等しくなるはずです。

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item * 2;
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`["naomi", "quincy", "camperbot"].myMap(element => element.toUpperCase())` は `["NAOMI", "QUINCY", "CAMPERBOT"]` を返すべきです。

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element.toUpperCase();
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`[1, 1, 2, 5, 2].myMap((element, index, array) => array[index + 1] || array[0])` は `[1, 2, 5, 2, 1]` を返すべきです。

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array[index + 1] || array[0];
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

コード内で `map` メソッドを使わないでください。

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?map/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    newArray.push(callback(this[i], i, this));
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const doubled_s = s.myMap(item => item * 2);
```
