---
id: 587d7b90367417b2b2512b65
title: slice メソッドを使って配列の一部を返す
challengeType: 1
forumTopicId: 301239
dashedName: return-part-of-an-array-using-the-slice-method
---

# --description--

`slice` メソッドは配列の特定の要素のコピーを返します。引数を2つ取ることができ、最初の引数はスライスを開始するインデックスを指定し、2つ目はスライスを終了するインデックス（終了インデックスは含まれません）を指定します。引数が指定されない場合は、配列の最初から最後までがデフォルトとなり、配列全体のコピーを簡単に作成できます。`slice` メソッドは元の配列を変更せず、新しい配列を返します。

例を示します：

```js
const arr = ["Cat", "Dog", "Tiger", "Zebra"];
const newArray = arr.slice(1, 3);
```

`newArray` は `["Dog", "Tiger"]` の値を持ちます。

# --instructions--

`slice` 関数内で `sliceArray` メソッドを使い、与えられた `anim` と `beginSlice` のインデックスに基づいて `endSlice` 配列の一部を返してください。関数は配列を返す必要があります。

# --hints--

コードは `slice` メソッドを使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/\.slice/g));
```

`inputAnim` 変数は変更されてはいけません。

```js
assert(
  JSON.stringify(inputAnim) ===
    JSON.stringify(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 3)` は `["Dog", "Tiger"]` を返すべきです。

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 3)) ===
    JSON.stringify(['Dog', 'Tiger'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 0, 1)` は `["Cat"]` を返すべきです。

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 0, 1)) ===
    JSON.stringify(['Cat'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 4)` は `["Dog", "Tiger", "Zebra"]` を返すべきです。

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 4)) ===
    JSON.stringify(['Dog', 'Tiger', 'Zebra'])
);
```

# --seed--

## --seed-contents--

```js
function sliceArray(anim, beginSlice, endSlice) {
  // Only change code below this line


  // Only change code above this line
}

const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
sliceArray(inputAnim, 1, 3);
```

# --solutions--

```js
function sliceArray(anim, beginSlice, endSlice) {
  return anim.slice(beginSlice, endSlice);
}
const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
```
