---
id: 587d7b8e367417b2b2512b5d
title: 命令型コード使用の危険性を理解する
challengeType: 1
forumTopicId: 301241
dashedName: understand-the-hazards-of-using-imperative-code
---

# --description--

関数型プログラミングは良い習慣です。コードを管理しやすくし、見えにくいバグから守ってくれます。でもその前に、問題が起きやすい箇所を明確にするために、命令型のプログラミング手法を見てみましょう。

英語（および多くの他の言語）では、命令形は命令を与えるために使われます。同様に、プログラミングにおける命令型スタイルは、コンピューターにタスクを実行するための一連の文を与えるものです。

多くの場合、その文はプログラムの状態を変更します。例えばグローバル変数を更新するような場合です。古典的な例は、配列のインデックスを繰り返すための`for`ループを書くことです。

これに対して、関数型プログラミングは宣言型プログラミングの一形態です。メソッドや関数を呼び出すことで、コンピューターに何をしてほしいかを伝えます。

JavaScriptは多くの事前定義されたメソッドを提供しており、一般的なタスクを処理するためにコンピューターがどう実行すべきかを自分で書く必要がありません。例えば、上で述べた`for`ループの代わりに、配列の繰り返し処理の詳細を扱う`map`メソッドを呼び出すことができます。これにより、「デバッグ」セクションで扱った「オフ・バイ・ワン・エラー」のような意味的なミスを避けられます。

次のシナリオを考えてみましょう。ブラウザーでウェブを閲覧していて、開いているタブを追跡したいとします。これを簡単なオブジェクト指向コードでモデル化してみましょう。

Windowオブジェクトはタブで構成されており、通常複数のWindowを開いています。各Windowオブジェクトにある開いているサイトのタイトルは配列に保持されています。ブラウザーで作業（新しいタブを開く、ウィンドウを統合する、タブを閉じる）をした後、まだ開いているタブを出力したいです。閉じたタブは配列から削除され、新しいタブは（簡単のため）配列の末尾に追加されます。

コードエディターには、`tabOpen()`、`tabClose()`、`join()`の関数を使ったこの機能の実装が表示されています。配列`tabs`は、開いているページの名前を格納するWindowオブジェクトの一部です。

# --instructions--

エディターのコードを確認してください。プログラムに副作用をもたらすメソッドを使っているため、動作が正しくありません。`finalTabs.tabs`に格納される最終的な開いているタブのリストは`['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']`であるべきですが、コードが生成するリストは少し異なっています。

正しいタブを削除するように`Window.prototype.tabClose`を変更してください。

# --hints--

`finalTabs.tabs`は`['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']`であるべきです

```js
assert.deepEqual(finalTabs.tabs, [
  'FB',
  'Gitter',
  'Reddit',
  'Twitter',
  'Medium',
  'new tab',
  'Netflix',
  'YouTube',
  'Vine',
  'GMail',
  'Work mail',
  'Docs',
  'freeCodeCamp',
  'new tab'
]);
```

# --seed--

## --seed-contents--

```js
// tabs is an array of titles of each site open within the window
const Window = function(tabs) {
  this.tabs = tabs; // We keep a record of the array inside the object
};

// When you join two windows into one window
Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

// When you open a new tab at the end
Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab'); // Let's open a new tab for now
  return this;
};

// When you close a tab
Window.prototype.tabClose = function(index) {

  // Only change code below this line

  const tabsBeforeIndex = this.tabs.splice(0, index); // Get the tabs before the tab
  const tabsAfterIndex = this.tabs.splice(index + 1); // Get the tabs after the tab

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex); // Join them together

  // Only change code above this line

  return this;
 };

// Let's create three browser windows
const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']); // Your mailbox, drive, and other work sites
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']); // Social sites
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']); // Entertainment sites

// Now perform the tab opening, closing, and other operations
const finalTabs = socialWindow
  .tabOpen() // Open a new tab for cat memes
  .join(videoWindow.tabClose(2)) // Close third tab in video window, and join
  .join(workWindow.tabClose(1).tabOpen());
console.log(finalTabs.tabs);
```

# --solutions--

```js
const Window = function(tabs) {
  this.tabs = tabs;
};

Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab');
  return this;
};

Window.prototype.tabClose = function(index) {
  const tabsBeforeIndex = this.tabs.slice(0, index);
  const tabsAfterIndex = this.tabs.slice(index + 1);

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex);
  return this;
 };

const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']);
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']);
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']);

const finalTabs = socialWindow
  .tabOpen()
  .join(videoWindow.tabClose(2))
  .join(workWindow.tabClose(1).tabOpen());
```
