---
id: 587d7da9367417b2b2512b67
title: concat を使って配列の末尾に要素を追加する（push の代わりに）
challengeType: 1
forumTopicId: 301226
dashedName: add-elements-to-the-end-of-an-array-using-concat-instead-of-push
---

# --description--

関数型プログラミングは、変更を加えない関数を作成し使用することに重点を置いています。

前回のチャレンジでは、元の配列を変更せずに配列を新しい配列にマージする方法として `concat` メソッドを紹介しました。`concat` と `push` メソッドを比較してください。`push` は呼び出された同じ配列の末尾に要素を追加するため、その配列を変更します。例を示します。

```js
const arr = [1, 2, 3];
arr.push(4, 5, 6);
```

`arr` は `[1, 2, 3, 4, 5, 6]` に変更された値を持ちますが、これは関数型プログラミングの方法ではありません。

`concat` は、配列を変更する副作用なしに新しい要素を配列の末尾にマージする方法を提供します。

# --instructions--

`nonMutatingPush` 関数を変更して、`concat` を使い、`newItem` を `original` の末尾にマージしてください。ただし、`original` と `newItem` の配列は変更しないようにしてください。関数は配列を返す必要があります。

# --hints--

コードは `concat` メソッドを使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

コードは `push` メソッドを使ってはいけません。

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?push/g));
```

`first` 配列は変更されてはいけません。

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

`second` 配列は変更されてはいけません。

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingPush([1, 2, 3], [4, 5])` は `[1, 2, 3, 4, 5]` を返すべきです。

```js
assert(
  JSON.stringify(nonMutatingPush([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingPush(original, newItem) {
  // Only change code below this line
  return original.push(newItem);

  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingPush(first, second);
```

# --solutions--

```js
function nonMutatingPush(original, newItem) {
  return original.concat(newItem);
}
const first = [1, 2, 3];
const second = [4, 5];
```
