---
id: 587d824a367417b2b2512c45
title: 익명 메시지 게시판
challengeType: 4
forumTopicId: 301568
dashedName: anonymous-message-board
---

# --description--

다음과 기능적으로 유사한 풀스택 JavaScript 앱을 만드세요: <a href="https://anonymous-message-board.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://anonymous-message-board.freecodecamp.rocks/</a>.

이 프로젝트를 진행하면서 다음 방법 중 하나를 사용해 코드를 작성하게 됩니다:

-   <a href="https://github.com/freeCodeCamp/boilerplate-project-messageboard/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 프로젝트를 완성하세요.
-   원하는 사이트 빌더를 사용해 프로젝트를 완성하세요. GitHub 저장소의 모든 파일을 반드시 포함해야 합니다.

# --instructions--

1.  테스트를 작성할 준비가 되면 `NODE_ENV`를 따옴표 없이 설정하고 `.env`에 데이터베이스 연결 문자열을 입력하세요.
2.  `routes/api.js`에서 컨트롤러/핸들러를 만들고 라우팅을 처리하는 것을 권장합니다.
3.  `server.js`에 보안 기능을 추가하세요.

다음 테스트를 `tests/2_functional-tests.js`에 작성하세요:

-   새 스레드 생성: `/api/threads/{board}`에 POST 요청
-   최근 10개 스레드와 각 스레드의 3개 답글 보기: `/api/threads/{board}`에 GET 요청
-   잘못된 비밀번호로 스레드 삭제: 유효하지 않은 `/api/threads/{board}`와 함께 `delete_password`에 DELETE 요청
-   올바른 비밀번호로 스레드 삭제: 유효한 `/api/threads/{board}`와 함께 `delete_password`에 DELETE 요청
-   스레드 신고: `/api/threads/{board}`에 PUT 요청
-   새 답글 생성: `/api/replies/{board}`에 POST 요청
-   모든 답글이 포함된 단일 스레드 보기: `/api/replies/{board}`에 GET 요청
-   잘못된 비밀번호로 답글 삭제: 유효하지 않은 `/api/replies/{board}`와 함께 `delete_password`에 DELETE 요청
-   올바른 비밀번호로 답글 삭제: 유효한 `/api/replies/{board}`와 함께 `delete_password`에 DELETE 요청
-   답글 신고: `/api/replies/{board}`에 PUT 요청

# --hints--

예시 URL이 아닌 자신만의 프로젝트를 제공할 수 있습니다.

```js
  assert(
    !/.*\/anonymous-message-board\.freecodecamp\.rocks/.test(
      code
    )
  );
```

사이트가 자신의 페이지 내 iFrame에서만 로드되도록 허용하세요.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-frame-options']?.includes('SAMEORIGIN'));
```

DNS 프리페칭을 허용하지 마세요.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-dns-prefetch-control']?.includes('off'));
```

사이트가 자신의 페이지에서만 리퍼러를 전송하도록 허용하세요.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['referrer-policy']?.includes('same-origin'));
```

`/api/threads/{board}`에 POST 요청을 보내고 `text`과 `delete_password`을 포함한 폼 데이터를 전송할 수 있습니다. 저장된 데이터베이스 레코드는 최소한 `_id`, `text`, `created_on`(날짜 및 시간), `bumped_on`(날짜 및 시간, `created_on`와 동일하게 시작), `reported` (불리언), `delete_password`, 및 `replies` (배열) 필드를 가집니다.

```js
  const date = new Date();
  const text = `fcc_test_${date}`;
  const deletePassword = 'delete_me';
  const data = { text, delete_password: deletePassword };
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(url + '/api/threads/fcc_test');
    const parsed = await checkData.json();
    try {
      assert.equal(parsed[0].text, text);
      assert.isNotNull(parsed[0]._id);
      assert.equal(new Date(parsed[0].created_on).toDateString(), date.toDateString());
      assert.equal(parsed[0].bumped_on, parsed[0].created_on);
      assert.isArray(parsed[0].replies);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

`/api/replies/{board}`에 POST 요청을 보내고 `text`, `delete_password`, 및 `thread_id`을 포함한 폼 데이터를 전송할 수 있습니다. 이로써 `bumped_on` 날짜가 댓글 날짜로 업데이트됩니다. 스레드의 `replies` 배열에는 최소한 `_id`, `text`, `created_on`, `delete_password`, 및 `reported` 속성을 가진 객체가 저장됩니다.

```js
  const url = code;
  const body = await fetch(url + '/api/threads/fcc_test');
  const thread = await body.json();

  const date = new Date();
  const text = `fcc_test_reply_${date}`;
  const delete_password = 'delete_me';
  const thread_id = thread[0]._id;
  const replyCount = thread[0].replies.length;

  const data = { text, delete_password, thread_id };
  const res = await fetch(url + '/api/replies/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
    const parsed = await checkData.json();
    try {
      assert.equal(parsed.replies.length, replyCount + 1);
      assert.equal(parsed.replies[0].text, text);
      assert.equal(parsed._id, thread_id);
      assert.equal(parsed.bumped_on, parsed.replies[0].created_on);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

`/api/threads/{board}`에 GET 요청을 보내면 게시판에서 최근 10개의 최신 스레드와 각 스레드의 최신 3개 답글이 배열로 반환됩니다. `reported`와 `delete_password` 필드는 클라이언트에 전송되지 않습니다.

```js
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test');

  if (res.ok) {
    const threads = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isAtMost(threads.length, 10);
      for (let i = 0; i < threads.length; i++) {
        assert.containsAllKeys(threads[i], ["_id", "text", "created_on", "bumped_on", "replies"]);
        assert.isAtMost(threads[i].replies.length, 3);
        assert.notExists(threads[i].delete_password);
        assert.notExists(threads[i].reported);
        for (let j = 0; j < threads[i].replies.length; j++) {
          assert.notExists(threads[i].replies[j].delete_password);
          assert.notExists(threads[i].replies[j].reported);
        }
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

`/api/replies/{board}?thread_id={thread_id}`에 GET 요청을 보내면 모든 답글이 포함된 전체 스레드가 반환되며, 이전 테스트와 동일한 필드는 클라이언트에 제외됩니다.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);

  if (res.ok) {
    const thread = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isObject(thread);
      assert.containsAllKeys(thread, ["_id", "text", "created_on", "bumped_on", "replies"]);
      assert.isArray(thread.replies);
      assert.notExists(thread.delete_password);
      assert.notExists(thread.reported);
      for (let i = 0; i < thread.replies.length; i++) {
        assert.notExists(thread.replies[i].delete_password);
        assert.notExists(thread.replies[i].reported);
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

`/api/threads/{board}`에 DELETE 요청을 보내고 `thread_id` 및 `delete_password`을 전달하여 스레드를 삭제할 수 있습니다. 반환값은 문자열 `incorrect password` 또는 `success`입니다.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  let data = { thread_id, delete_password: "wrong_password" };
  const res_invalid = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  data = { thread_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    const not_deleted = await res_invalid.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      assert.equal(not_deleted, "incorrect password");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

`/api/replies/{board}`에 DELETE 요청을 보내고 `thread_id`, `reply_id`, 및 `delete_password`을 전달할 수 있습니다. 반환값은 문자열 `incorrect password` 또는 `success`입니다. 성공 시 `reply_id`의 텍스트가 `[deleted]`로 변경됩니다.

```js
  const url = code;

  const thread_data = {
    text: "fcc_test_thread",
    delete_password: "delete_me",
  };
  await fetch(`${url}/api/threads/fcc_test`, {
    method: "POST",
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(thread_data)
  });
  let res = await fetch(`${url}/api/threads/fcc_test`);
  let threads = await res.json();
  const thread_id = threads[0]._id;
  
  const reply_data = { thread_id, text: "fcc_test_reply", delete_password: "delete_me" };
  await fetch(`${url}/api/replies/fcc_test`, {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(reply_data)
  });
  res = await fetch(`${url}/api/threads/fcc_test`);
  threads = await res.json();
  const reply_id = threads[0].replies[0]._id;

  const data = { thread_id, reply_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/replies/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
      const thread = await res.json();
      assert.equal(thread._id, thread_id);
      assert.equal(thread.replies[0]._id, reply_id);
      assert.equal(thread.replies[0].text, "[deleted]");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

`/api/threads/{board}`에 PUT 요청을 보내고 `thread_id`를 전달할 수 있습니다. 반환값은 문자열 `reported`입니다. `reported`의 `thread_id` 값이 `true`로 변경됩니다.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const data = { thread_id };
  
  res = await fetch(`${url}/api/threads/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

`/api/replies/{board}`에 PUT 요청을 보내고 `thread_id` 및 `reply_id`을 전달할 수 있습니다. 반환값은 문자열 `reported`입니다. `reported`의 `reply_id` 값이 `true`로 변경됩니다.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const reply_id = threads[0].replies[0]._id;
  const data = { thread_id, reply_id };
  
  res = await fetch(`${url}/api/replies/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

모든 10개 기능 테스트가 완료되어 통과했습니다.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 10);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
    assert.isAtLeast(test.assertions.length, 1);
  });
```
