---
id: 587d824a367417b2b2512c44
title: 주식 가격 확인기
challengeType: 4
forumTopicId: 301572
dashedName: stock-price-checker
---

# --description--

다음과 기능적으로 유사한 풀스택 JavaScript 앱을 만드세요: <a href="https://stock-price-checker.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker.freecodecamp.rocks/</a>.

신뢰할 수 있는 모든 주식 가격 API는 API 키가 필요하기 때문에 우회 방법을 만들었습니다. <a href="https://stock-price-checker-proxy.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker-proxy.freecodecamp.rocks/</a>를 사용해 직접 키를 등록하지 않고도 최신 주식 가격 정보를 받아오세요.

이 프로젝트를 진행하려면 다음 방법 중 하나로 코드를 작성해야 합니다:

-   <a href="https://github.com/freeCodeCamp/boilerplate-project-stockchecker/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 프로젝트를 완성하세요.
-   원하는 사이트 빌더를 사용해 프로젝트를 완성하세요. GitHub 저장소의 모든 파일을 반드시 포함해야 합니다.

# --instructions--

1.  `NODE_ENV` 환경 변수를 `test`로, 따옴표 없이 설정하세요
2.  `routes/api.js`에서 프로젝트를 완성하거나 핸들러/컨트롤러를 만드세요
3.  모든 보안 기능은 `server.js`에 추가하세요
4.  모든 기능 테스트는 `tests/2_functional-tests.js`에 작성하세요

**참고** 개인정보 보호 고려사항: IP당 좋아요는 1회만 허용해야 하므로 IP 주소를 저장해야 합니다. 일반 데이터 보호 규정(GDPR) 같은 데이터 개인정보 보호법을 준수하는 것이 중요합니다. 사용자 데이터 저장 허가를 받는 방법도 있지만, 익명화하는 것이 훨씬 간단합니다. 이 도전 과제에서는 IP 주소를 데이터베이스에 저장하기 전에 익명화하세요. 익명화 방법으로는 해시 처리, 잘라내기, IP 주소 일부를 0으로 설정하는 방법 등이 있습니다.

다음 테스트를 `tests/2_functional-tests.js`에 작성하세요:

-   한 주식 보기: `/api/stock-prices/`에 GET 요청
-   한 주식 보고 좋아요 누르기: `/api/stock-prices/`에 GET 요청
-   같은 주식을 다시 보고 좋아요 누르기: `/api/stock-prices/`에 GET 요청
-   두 주식 보기: `/api/stock-prices/`에 GET 요청
-   두 주식을 보고 좋아요 누르기: `/api/stock-prices/`에 GET 요청

# --hints--

예시 URL이 아니라 직접 만든 프로젝트를 사용할 수 있습니다.

```js
  assert(
    !/.*\/stock-price-checker\.freecodecamp\.rocks/.test(code)
  );
```

서버에서만 스크립트와 CSS 로딩을 허용하도록 콘텐츠 보안 정책을 설정하세요.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("script-src 'self'")
  );
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("style-src 'self'")
  );
```

NASDAQ 주식 심볼을 `GET` 쿼리 매개변수로 전달해 `/api/stock-prices`에 `stock` 요청을 보낼 수 있습니다. 반환된 객체에는 `stockData`라는 속성이 포함됩니다.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  assert.property(parsed, 'stockData');
```

`stockData` 속성은 문자열인 `stock` 심볼, 숫자인 `price`, 그리고 숫자인 `likes`을 포함합니다.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker.price, 'number');
  assert.typeOf(ticker.likes, 'number');
  assert.typeOf(ticker.stock, 'string');
```

좋아요를 추가하려면 `like` 필드를 불리언 값 `true`으로 전달할 수 있습니다. IP당 좋아요는 1회만 허용됩니다.

```js

```

두 주식을 전달하면 반환값은 두 주식 정보를 담은 배열입니다. `likes` 대신 두 `rel_likes` 객체 모두에 대해 좋아요 차이인 `stockData`가 표시됩니다.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG&stock=MSFT'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker, 'array');
  assert.property(ticker[0], 'rel_likes');
  assert.property(ticker[1], 'rel_likes');
```

모든 5가지 기능 테스트가 완료되어 통과했습니다.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 5);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
  });
```
