---
id: 5e601c775ac9d0ecd8b94aff
title: 보안 실시간 멀티플레이어 게임
challengeType: 4
forumTopicId: 462375
dashedName: secure-real-time-multiplayer-game
---

# --description--

HTML 캔버스 API와 Socket.io를 사용하여 이와 기능적으로 유사한 2D 실시간 멀티플레이어 게임을 개발하세요: <a href="https://secure-real-time-multiplayer-game.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://secure-real-time-multiplayer-game.freecodecamp.rocks/</a>. 이 프로젝트를 진행하려면 다음 방법 중 하나를 사용하여 코드를 작성해야 합니다:

- <a href="https://github.com/freeCodeCamp/boilerplate-project-secure-real-time-multiplayer-game/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 프로젝트를 완료하세요.
- 원하는 사이트 빌더를 사용하여 프로젝트를 완료하세요. GitHub 저장소의 모든 파일을 반드시 포함하세요.

# --instructions--

각 플레이어가 자신의 아바타를 움직일 수 있고, 최소 하나의 수집 아이템이 있으며, 플레이어의 점수를 기준으로 순위가 계산되는 보안 멀티플레이어 게임을 만드세요.

자세한 내용은 아래 테스트를 참고하세요.

게임이 반드시 보안적이어야 합니다! 다음 보안 조치를 포함하세요:

- 클라이언트가 MIME 타입을 추측하거나 감지하지 못하게 하세요
- XSS 공격을 방지하세요
- 클라이언트에 웹사이트의 어떤 것도 캐시하지 마세요
- 헤더에 사이트가 `PHP 7.4.3`로 구동된다고 표시하세요

**참고**: `helmet@^3.21.3`은 사용자 스토리에 필요합니다. 이는 사용자 스토리를 달성하는 방법에 대한 정보를 얻으려면 이전 버전의 Helmet 문서를 사용해야 함을 의미합니다.

# --hints--

예제 URL이 아닌 자신만의 프로젝트를 제공할 수 있습니다.

```js
  assert(
    !/.*\/secure-real-time-multiplayer-game\.freecodecamp\.rocks/.test(
      code
    )
  );
```

여러 플레이어가 서버에 연결하여 플레이할 수 있습니다.

```js

```

각 플레이어는 아바타를 가집니다.

```js

```

각 플레이어는 `Player`의 `Player.mjs` 클래스로 생성된 객체로 나타냅니다.

```js

```

최소한 각 플레이어 객체는 고유한 `id`, `score`, 그리고 플레이어 현재 위치를 나타내는 `x`와 `y` 좌표를 포함해야 합니다.

```js

```

게임에는 최소 하나의 수집 아이템 유형이 있습니다. 이를 구현하려면 `Collectible`의 `Collectible.mjs` 클래스를 완성하세요.

```js

```

최소한 `Collectible` 클래스로 생성된 각 수집 아이템 객체는 고유한 `id`, `value`, 그리고 아이템 현재 위치를 나타내는 `x`와 `y` 좌표를 포함해야 합니다.

```js

```

플레이어는 WASD 및/또는 방향키를 사용해 아바타를 움직일 수 있습니다. 이를 구현하려면 `movePlayer`의 `Player.mjs` 메서드를 완성하세요.

```js

```

`movePlayer` 메서드는 "up", "down", "left", 또는 "right" 문자열과 플레이어 위치가 변경될 픽셀 수를 나타내는 숫자 두 개의 인수를 받아야 합니다. `movePlayer`은 호출된 플레이어 객체의 `x`와 `y` 좌표를 조정해야 합니다.

```js

```

플레이어 점수는 다른 플레이어 중에서 순위를 계산하는 데 사용됩니다. 이를 구현하려면 `calculateRank` 클래스의 `Player` 메서드를 완성하세요.

```js

```

`calculateRank` 메서드는 연결된 모든 플레이어를 나타내는 객체 배열을 인수로 받아 `Rank: currentRanking/totalPlayers` 문자열을 반환해야 합니다. 예를 들어, 두 명의 플레이어가 있는 게임에서 플레이어 A가 점수 3, 플레이어 B가 점수 5를 가졌다면, 플레이어 A에 대한 `calculateRank`는 `Rank: 2/2`을 반환해야 합니다.

```js

```

플레이어는 수집 아이템과 충돌할 수 있습니다. 이를 구현하려면 `collision`의 `Player.mjs` 메서드를 완성하세요.

```js

```

`collision` 메서드는 수집 아이템 객체를 인수로 받아야 합니다. 플레이어 아바타가 아이템과 겹치면 `collision` 메서드는 `true`을 반환해야 합니다.

```js

```

모든 플레이어는 동기화 상태를 유지합니다.

```js

```

플레이어는 언제든지 게임에서 연결을 끊을 수 있습니다.

```js

```

클라이언트가 MIME 타입을 추측하거나 감지하지 못하게 하세요.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.equal(parsed.headers['x-content-type-options'], 'nosniff');
```

교차 사이트 스크립팅(XSS) 공격을 방지하세요.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.equal(parsed.headers['x-xss-protection'], '1; mode=block');
```

클라이언트에 웹사이트의 어떤 것도 캐시하지 마세요.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.equal(parsed.headers['surrogate-control'], 'no-store');
  assert.equal(
    parsed.headers['cache-control'],
    'no-store, no-cache, must-revalidate, proxy-revalidate'
  );
  assert.equal(parsed.headers['pragma'], 'no-cache');
  assert.equal(parsed.headers['expires'], '0');
```

헤더에 사이트가 실제로는 아니지만 보안 조치로 "PHP 7.4.3"로 구동된다고 표시하세요.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.equal(parsed.headers['x-powered-by'], 'PHP 7.4.3');
```
