---
id: 612e6afc009b450a437940a1
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

基本的なHTML構造から始めましょう。`DOCTYPE` 宣言と`html`、`head`、`body`、`title` 要素を追加してください。

このページの言語を英語に設定してください。`title` を `Piano` に設定してください。

# --hints--

コードには`DOCTYPE` 参照が含まれている必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れてください。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプを`html` に定義してください。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

タイプの後に`DOCTYPE` で`>` 宣言を閉じてください。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 要素は`lang` 属性が`en` の開始タグを持つ必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素は終了タグを持つ必要があります。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 宣言はHTMLの最初にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

開始`head` タグが必要です。

```js
assert(code.match(/<head\s*>/i));
```

終了`head` タグが必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

開始`body` タグが必要です。

```js
assert(code.match(/<body\s*>/i));
```

終了`body` タグが必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` と`body` 要素は兄弟要素である必要があります。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 要素は`html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は`html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

コードには`title` 要素が必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

プロジェクトには`Piano` のタイトルが必要です。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

タイトルの大文字・小文字と綴りに注意してください。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
