---
id: 612e89562043183c86df287c
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

ブラウザーは特定の要素にデフォルトのマージンやパディングの値を適用する場合があります。ピアノの見た目を正しくするために、ボックスモデルをリセットする必要があります。

CSSファイルに`html`ルールセレクターを追加し、`box-sizing`プロパティを`border-box`に設定してください。

# --hints--

`html`セレクターがあるはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

`html`セレクターには`box-sizing`プロパティが`border-box`に設定されているはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
