---
id: 612e77aba7ca691f598feb02
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

モバイルデバイス向けにページを最適化するための `meta` タグを1つと、ページで受け入れられる `charset` を指定するための `meta` タグを1つ追加してください。

# --hints--

ページに `meta` 要素を2つ追加してください。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

`meta` 要素が2つあるはずです。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

1つの `meta` 要素には `name` が `viewport` に設定され、`content` が `width=device-width, initial-scale=1.0` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

もう1つの `meta` 要素には `charset` 属性が `UTF-8` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
