---
id: 612e8279827a28352ce83a72
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

7つの `.key` 要素のセットをコピーして、`.keys` の div にさらに2セット貼り付けてください。

# --hints--

`.key` 要素が21個あるはずです。

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

`.black--key` 要素が15個あるはずです。

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

すべての `.key` 要素は `.keys` 要素の中にある必要があります。

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

すべての `.black--key` 要素は `.keys` 要素の中にある必要があります。

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
