---
id: 612ea4c4993aba52ab4aa32e
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

これまでに準備した疑似セレクターを使う時が来ました。黒鍵を作るために、新しい`.key.black--key::after`セレクターを追加してください。これは`key black--key`クラスを持つ要素を対象にし、HTML内のこれらの要素の後の疑似要素を選択します。

新しいセレクターでは、`background-color`を`#1d1e22`に設定してください。また、`content`プロパティを`""`に設定してください。これにより疑似要素は空になります。

`content`プロパティは要素の内容を設定またはオーバーライドするために使われます。デフォルトでは、`::before`および`::after`疑似セレクターで作成された疑似要素は空で、ページにレンダリングされません。`content`プロパティを空文字列`""`に設定すると、要素は空のままページにレンダリングされるようになります。

もし試してみたい場合は、`background-color`プロパティを削除し、`content`プロパティに`"♥"`のような異なる値を設定してみてください。テストに合格するために、終わったらこれらの変更は元に戻すことを忘れないでください。

# --hints--

`.key.black--key::after`セレクターがあるはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

`.key.black--key::after`セレクターには`background-color`プロパティが`#1d1e22`に設定されているはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

`.key.black--key::after`セレクターには`content`プロパティが`""`に設定されているはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
