---
id: 612e78af05201622d4bab8aa
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

ピアノの作業を始めましょう。`div` 要素の中に `body` 要素を作成し、`id` を `piano` に設定してください。

# --hints--

新しい `div` 要素を作成してください。

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div` は `body` 要素の中にある必要があります。

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

`div` は `id` を `piano` に設定してください。

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
