---
id: 612e89d254fe5d3df7d6693d
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`html` のボックスモデルをリセットしたので、内部の要素にも同じ設定を適用する必要があります。これを行うには、`box-sizing` プロパティを `inherit` に設定すると、対象の要素が親要素と同じ値を使うようになります。

また、セレクターに続く特別なキーワードである疑似要素も対象にする必要があります。今回使う疑似要素は `::before` と `::after` です。

`::before` セレクターは選択した要素の最初の子要素として疑似要素を作成し、`::after` セレクターは選択した要素の最後の子要素として疑似要素を作成します。これらの疑似要素は装飾的なコンテンツを作るためによく使われますが、このプロジェクトの後半で詳しく見ていきます。

まずは、`*` を持つすべての要素を対象にする CSS セレクターを作成し、疑似要素の `::before` と `::after` も含めてください。そして `box-sizing` プロパティを `inherit` に設定してください。

# --hints--

`*, ::before, ::after` セレクターがあるはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

`*, ::before, ::after` セレクターには `box-sizing` プロパティを `inherit` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
