---
id: 612e9d142affc44a453655db
title: ステップ16
challengeType: 0
dashedName: step-16
---

# --description--

キー自体のスタイルを設定する時間です。`class`要素に対して`.key`セレクターを作成してください。`background-color`セットを`#ffffff`に、`position`プロパティを`relative`に、`width`プロパティを`41px`に、`height`プロパティを`175px`に設定してください。

# --hints--

`.key`セレクターがあるはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key'));
```

`.key`セレクターには`background-color`プロパティが`#ffffff`に設定されているはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor === 'rgb(255, 255, 255)');
```

`.key`セレクターには`position`プロパティが`relative`に設定されているはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.position === 'relative');
```

`.key`セレクターには`width`プロパティが`41px`に設定されているはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.width === '41px');
```

`.key`セレクターには`height`プロパティが`175px`に設定されているはずです。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.height === '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
