---
id: afd15382cdfb22c9efe8b7de
title: تنفيذ مولد أزواج الحمض النووي
challengeType: 26
dashedName: implement-a-dna-pair-generator
---

# --description--

في اللولب المزدوج للحمض النووي، القواعد دائمًا تكون مزدوجة: إذا كان على أحد الخيوط قاعدة <em>A</em>، فهناك على الخيط الآخر مباشرة أمامها قاعدة <em>T</em>، الزوج الآخر هو <em>C</em> و<em>G</em>.

في هذا المختبر، ستكتب دالة لمطابقة أزواج القواعد المفقودة لسلسلة الحمض النووي المقدمة. لكل حرف في السلسلة المقدمة، ابحث عن الحرف المقابل في زوج القاعدة.

على سبيل المثال، للإدخال `ATCG`، أرجع `[["A", "T"], ["T", "A"], ["C", "G"], ["G", "C"]]`

القاعدة <em>A</em> تُزاوج مع القاعدة <em>T</em>، والقاعدة <em>T</em> تُزاوج مع القاعدة <em>A</em>، والقاعدة <em>C</em> تُزاوج مع القاعدة <em>G</em>، وأخيرًا القاعدة <em>G</em> تُزاوج مع القاعدة <em>C</em>.

**الهدف**: نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة `pairElement` تأخذ سلسلة نصية بأي طول كمعلمة.
1. يجب أن تُرجع دالة `pairElement` مصفوفة ثنائية الأبعاد، حيث تحتوي كل مصفوفة داخلية على سلسلتين نصيتين، السلسلة الأولى هي قاعدة من المُدخَل، والسلسلة الثانية هي القاعدة المزاوجة.
1. عند إعطاء `A`، يجب أن تُزاوجها الدالة مع `T`.
1. عند إعطاء `T`، يجب أن تُزاوجها الدالة مع `A`.
1. عند إعطاء `C`، يجب أن تُزاوجها الدالة مع `G`.
1. عند إعطاء `G`، يجب أن تُزاوجها الدالة مع `C`.

# --hints--

يجب أن تنشئ دالة باسم `pairElement`.

```js
assert.isFunction(pairElement);
```

يجب أن تأخذ دالة `pairElement` معلمة واحدة.

```js
assert.lengthOf(pairElement, 1);
```

يجب أن تُرجع `pairElement("ATCGA")` القيمة `[["A","T"],["T","A"],["C","G"],["G","C"],["A","T"]]`.

```js
assert.deepEqual(pairElement('ATCGA'), [
  ['A', 'T'],
  ['T', 'A'],
  ['C', 'G'],
  ['G', 'C'],
  ['A', 'T']
]);
```

يجب أن تُرجع `pairElement("TTGAG")` القيمة `[["T","A"],["T","A"],["G","C"],["A","T"],["G","C"]]`.

```js
assert.deepEqual(pairElement('TTGAG'), [
  ['T', 'A'],
  ['T', 'A'],
  ['G', 'C'],
  ['A', 'T'],
  ['G', 'C']
]);
```

يجب أن تُرجع `pairElement("CTCTA")` القيمة `[["C","G"],["T","A"],["C","G"],["T","A"],["A","T"]]`.

```js
assert.deepEqual(pairElement('CTCTA'), [
  ['C', 'G'],
  ['T', 'A'],
  ['C', 'G'],
  ['T', 'A'],
  ['A', 'T']
]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var lookup = Object.create(null);
lookup.A = 'T';
lookup.T = 'A';
lookup.C = 'G';
lookup.G = 'C';

function pairElement(str) {
 return str.split('').map(function(p) {return [p, lookup[p]];});
}
```
