---
id: 587d825a367417b2b2512c89
title: "implementare l'algoritmo quicksort"
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

In questo laboratorio implementerai l'algoritmo di ordinamento: quicksort.

Quicksort è un metodo efficiente, ricorsivo e basato sulla strategia divide et impera per ordinare un array. In questo metodo, si sceglie un valore pivot nell'array originale. L'array viene quindi suddiviso in due sotto-array di valori minori e maggiori del valore pivot. Poi si combina il risultato della chiamata ricorsiva dell'algoritmo quicksort su entrambi i sotto-array. Questo processo continua fino a raggiungere il caso base di un array vuoto o con un solo elemento, che viene restituito. Lo srotolamento delle chiamate ricorsive restituisce l'array ordinato.

Sebbene la scelta del valore pivot sia importante, per i nostri scopi qui qualsiasi pivot va bene. Per semplicità, si può usare il primo o l'ultimo elemento.

Quicksort è un metodo di ordinamento molto efficiente, che offre una prestazione media di *O(nlog(n))*. È anche relativamente facile da implementare. Queste caratteristiche lo rendono un metodo di ordinamento popolare e utile.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione `quicksort` che prende come input un array di numeri interi.
1. La funzione `quicksort` dovrebbe restituire un array che contiene gli stessi numeri interi dell'input, ma in ordine dal più basso al più alto.
1. La funzione `quicksort` dovrebbe chiamare ricorsivamente se stessa per ordinare l'array.

# --hints--

Dovresti avere una funzione `quicksort`.

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` dovrebbe restituire un array che non cambia se non per l'ordine.

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` dovrebbe restituire un array ordinato (dal più piccolo al più grande).

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` non dovrebbe usare il metodo integrato `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
