---
id: 68ed661e9f463c155bb677a1
title: Passaggio 1
challengeType: 0
dashedName: step-1
---

# --description--

In questo workshop eserciterai l'uso del CSS flexbox progettando un insieme di riquadri colorati.

Per cominciare, all'interno dell'elemento `body` aggiungi un elemento `header`. Annida un elemento `h1` dentro l'elemento `header` con il testo `Colored Boxes Layout`.

# --hints--

Dovresti avere un elemento `header` dentro l'elemento `body`.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

Il tuo elemento `header` dovrebbe avere un elemento `h1` al suo interno.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Il tuo elemento `h1` dovrebbe contenere il testo `Colored Boxes Layout`.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
