---
id: 68ef0b3237fef991c92d9c19
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Ora aggiungi un elemento `div` con la classe `flex-container` sotto il tuo `header`. Poi crea sei elementi `div` con la classe `box` all'interno del `div` che hai appena creato.

Ricorda che un attributo di classe viene spesso usato per riferirsi a un nome di classe in un foglio di stile. In questo caso, tutti gli elementi `div` saranno stilizzati allo stesso modo secondo la definizione di stile `.flex-container` e `.box`.

# --hints--

Dovresti avere un elemento `div` sotto l'elemento `header`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

Il tuo elemento `div` sotto l'elemento `header` dovrebbe avere la classe `flex-container`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

Dovresti avere sei elementi `div` annidati all'interno dell'elemento `.flex-container`.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

Tutti e sei i tuoi elementi `div` dentro l'elemento `.flex-container` dovrebbero avere la classe `box`.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
