---
id: 68f47093c85c7da27c7ee5cb
title: Passaggio 39
challengeType: 0
dashedName: step-39
---

# --description--

Ora crea un selettore `.box4` e aggiungi la proprietà `background` con il valore `#4caf50`.

Poi aggiungi la proprietà `order` con valore `3` al selettore `.box4`. Vedrai questo riquadro spostarsi dopo quello con la classe `box3`.

# --hints--

Dovresti avere un selettore `.box4`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box4'));
```

Il tuo selettore `.box4` dovrebbe avere la proprietà `background` con il valore `#4caf50`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box4')?.getPropVal('background'), 'rgb(76, 175, 80)');
```

Il tuo selettore `.box4` dovrebbe avere la proprietà `order` con il valore `3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box4')?.getPropVal('order'), '3');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box box4">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2;
  flex-shrink: 3; 
}

--fcc-editable-region--

--fcc-editable-region--
```
