---
id: 68f932c803f4d3d9f511b436
title: Passaggio 20
challengeType: 0
dashedName: step-20
---

# --description--

La proprietà `flex-wrap` definisce se gli elementi flex sono costretti a stare su una sola riga oppure possono andare a capo su più righe.

Per impostazione predefinita, gli elementi flex hanno un valore `nowrap`, il che significa che rimangono su una sola riga anche se superano la larghezza del contenitore. Il valore `wrap` invece fa andare gli elementi flex a capo su nuove righe quando non entrano nella larghezza del contenitore, creando un layout più flessibile e responsivo.

Per continuare, aggiungi la proprietà `flex-wrap` con il valore `wrap` al selettore di classe `.flex-container`.

# --hints--

Il tuo selettore `.flex-container` dovrebbe avere la proprietà `flex-wrap` con il valore `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('flex-wrap'), 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
