---
id: 68f91df5776e693049fb8942
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Ora è il momento di aggiungere alcune regole CSS al file `styles.css`. Inizia creando un selettore per l'elemento `h1`.

Centra il contenuto dell'elemento `h1` impostando la sua proprietà `text-align` su `center`. Crea anche una proprietà `margin-bottom` con valore `10px` per aggiungere margine tra il fondo dell'elemento `h1` e qualsiasi elemento HTML che si trovi sotto.

# --hints--

Dovresti avere un selettore `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Il tuo selettore `h1` dovrebbe impostare la proprietà `text-align` su `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

Il tuo selettore `h1` dovrebbe impostare la proprietà `margin-bottom` su `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
