---
id: 692f22933ba1cb86cd75a70c
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Ora è il momento di creare il nome del colore per i riquadri. Per farlo, aggiungi un elemento `p` sotto ogni elemento `h2` con questi nomi di colore in ordine: `Red`, `Orange`, `Yellow`, `Green`, `Blue` e `Indigo`. Un colore per ogni riquadro.

# --hints--

Dovresti aggiungere un elemento `p` con il testo `Red` sotto l'elemento `h2` all'interno del primo elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

Dovresti aggiungere un elemento `p` con il testo `Orange` sotto l'elemento `h2` all'interno del secondo elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

Dovresti aggiungere un elemento `p` con il testo `Yellow` sotto l'elemento `h2` all'interno del terzo elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

Dovresti aggiungere un elemento `p` con il testo `Green` sotto l'elemento `h2` all'interno del quarto elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

Dovresti aggiungere un elemento `p` con il testo `Blue` sotto l'elemento `h2` all'interno del quinto elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

Dovresti aggiungere un elemento `p` con il testo `Indigo` sotto l'elemento `h2` all'interno del sesto elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
