---
id: 692f30e86027a4b253d0a42d
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Aggiungi la proprietà `font-size` con il valore `1.125rem` al selettore `.box`.

La proprietà `font-size` specifica la dimensione del testo. Può essere impostata usando unità diverse come pixel, em, rem, percentuali e così via, rendendola flessibile per design responsivi.

In questo caso, il valore usa unità `rem`, che sono relative all'elemento radice `html`, non all'elemento genitore. Questo significa che la dimensione del carattere tipografico rimane costante e non è influenzata dalle dimensioni dell'elemento genitore, rendendo il layout più prevedibile.

# --hints--

Il tuo selettore `.box` dovrebbe avere la proprietà `font-size` con il valore `1.125rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('font-size'), '1.125rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
