---
id: 692f2c20cb7f9e5050b935a7
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Aggiungi la proprietà `border` con il valore `1px solid #000` al selettore `.box`. La proprietà `border` serve a creare un contorno visibile attorno a un elemento. Può essere personalizzata in termini di `border-width`, `border-style` e `border-color`.

In questo caso, `1px` è `border-width`, che definisce lo spessore del bordo. `solid` è `border-style` e imposta l’aspetto del bordo (solido, tratteggiato, puntinato e così via). E `#000` è `border-color`, che definisce il colore del bordo.

# --hints--

Il tuo selettore `.box` dovrebbe avere la proprietà `border` con il valore `1px solid #000`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border'), '1px solid rgb(0, 0, 0)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
