---
id: 692f318859d5e4bf8fd77ba1
title: Passaggio 18
challengeType: 0
dashedName: step-18
---

# --description--

La proprietà `order` specifica l'ordine di un elemento flex rispetto agli altri elementi flex all'interno dello stesso contenitore. Per impostazione predefinita, tutti gli elementi flex hanno un valore `order` pari a `0`, il che significa che appaiono nell'ordine in cui sono scritti nell'HTML. Gli elementi con valori `order` più bassi appaiono prima, mentre quelli con valori più alti appaiono per ultimi.

Aggiungi la proprietà `order` con il valore `0` al selettore `.box`.

# --hints--

Il tuo selettore `.box` dovrebbe avere la proprietà `order` con il valore `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('order'), '0');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
