---
id: 692f22a2ad44ae887a24c8f9
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

All'interno di ogni elemento `.box`, aggiungi un elemento `h2` con il testo `Box` seguito da uno spazio e da un numero sequenziale che parte da `1`.

# --hints--

Dovresti avere un elemento `h2` con il testo `Box 1` all'interno del primo `div` con la classe `box`.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

Dovresti avere un elemento `h2` con il testo `Box 2` all'interno del secondo `div` con la classe `box`.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

Dovresti avere un elemento `h2` con il testo `Box 3` all'interno del terzo `div` con la classe `box`.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

Dovresti avere un elemento `h2` con il testo `Box 4` all'interno del quarto `div` con la classe `box`.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

Dovresti avere un elemento `h2` con il testo `Box 5` all'interno del quinto `div` con la classe `box`.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

Dovresti avere un elemento `h2` con il testo `Box 6` all'interno dell'ultimo `div` con la classe `box`.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
