---
id: 694b46a414bb2f5079868ed1
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

La proprietà `flex` controlla la dimensione e il comportamento degli elementi all'interno di un contenitore flessibile. È composta da tre proprietà: `flex-grow`, `flex-shrink` e `flex-basis`.

Ora aggiungi la proprietà `flex` con il valore `1 1 100px` al selettore `.box`. Questo farà sì che ogni riquadro inizi da `100px`, permettendo loro di crescere per riempire lo spazio disponibile o ridursi quando necessario.

# --hints--

Il tuo selettore `.box` dovrebbe avere una proprietà `flex` con il valore `1 1 100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex'), '1 1 100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
