---
id: 5d822fd413a79914d39e98cb
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Ifuatayo, ongeza lebo za kufungua na kufunga `head` na `body` ndani ya kipengele cha `html`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `head`.

```js
assert(code.match(/<head\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Unapaswa kuwa na lebo ya kufungua ya `body`.

```js
assert(code.match(/<body\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Vipengele vya `head` na `body` vinapaswa kuwa ndugu.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

Kipengele cha `head` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

Kipengele cha `body` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
  
</html>
--fcc-editable-region--

```
