---
id: 5d822fd413a79914d39e98d6
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Sasa una kitu kinachofanana na jengo. Uko tayari kuunda variable yako ya kwanza. Tamko za variable huanza na alama mbili za dash (`-`) na hupewa jina na thamani kama hii: `--variable-name: value;`. Katika sheria ya darasa la `bb1`, unda variable iitwayo `--building-color1` na umpe thamani ya `#999`.

# --hints--

Unapaswa kuunda variable mpya iitwayo `--building-color1`.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

Unapaswa kufafanua variable ya `--building-color1` ndani ya `.bb1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

Unapaswa kumpa `--building-color1` thamani ya `#999`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(),'#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
