---
id: 564944c91be2204b269d51e3
title: змінити текст всередині елемента за допомогою jQuery
challengeType: 6
forumTopicId: 16773
dashedName: change-text-inside-an-element-using-jquery
---

# --description--

За допомогою jQuery ви можете змінити текст між початковим і кінцевим тегами елемента. Ви навіть можете змінити HTML-розмітку.

У jQuery є функція `.html()`, яка дозволяє додавати HTML-теги та текст всередині елемента. Будь-який вміст, що раніше був у цьому елементі, буде повністю замінений на вміст, який ви надасте за допомогою цієї функції.

Ось як ви могли б переписати та виділити текст нашого заголовка:

```js
$("h3").html("<em>jQuery Playground</em>");
```

У jQuery також є схожа функція `.text()`, яка змінює лише текст без додавання тегів. Іншими словами, ця функція не буде обробляти HTML-теги, які їй передають, а замість цього сприйматиме їх як текст, яким ви хочете замінити існуючий вміст.

Змініть кнопку з ідентифікатором `target4`, виділивши її текст.

<a href="https://www.freecodecamp.org/news/html-elements-explained-what-are-html-tags/#em-element" target="_blank" rel="noopener noreferrer nofollow">Перегляньте нашу новину про &lt;em&gt;</a>, щоб дізнатися різницю між `<i>` і `<em>` та їх використанням.

Зверніть увагу, що хоча тег `<i>` традиційно використовувався для виділення тексту, тепер його застосовують як тег для іконок. Тег `<em>` зараз широко прийнятий як тег для виділення. Для цього завдання підійде будь-який із них.

# --hints--

Ви повинні виділити текст у вашій кнопці `target4`, додавши HTML-теги.

```js
assert.isTrue(
  /<em>|<i>\s*#target4\s*<\/em>|<\/i>/gi.test($('#target4').html())
);
```

Інший текст має залишатися без змін.

```js
assert($('#target4') && $('#target4').text().trim() === '#target4');
```

Ви не повинні змінювати жоден інший текст.

```js
assert.isFalse(/<em>|<i>/gi.test($('h3').html()));
```

Використовуйте `.html()`, а не `.text()`.

```js
assert(code.match(/\.html\(/g));
```

Ви повинні вибрати `button id="target4"` за допомогою jQuery.

```js
assert(code.match(/\$\(\s*?(\"|\')#target4(\"|\')\s*?\)\.html\(/));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target4").html('<em>#target4</em>');
  });
</script>

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
