---
id: bad87fee1348bd9acdd08826
title: "Дізнайтеся, як працюють теги script і document ready"
challengeType: 6
forumTopicId: 18224
dashedName: learn-how-script-tags-and-document-ready-work
---

# --description--

Тепер ми готові вивчати jQuery — найпопулярніший інструмент JavaScript усіх часів.

Перш ніж почати використовувати jQuery, потрібно додати кілька речей до нашого HTML.

Спочатку додайте елемент `script` у верхній частині вашої сторінки. Обов’язково закрийте його на наступній лінії.

Ваш браузер виконає будь-який JavaScript всередині елемента `script`, включно з jQuery.

Всередині вашого елемента `script` додайте цей код: `$(document).ready(function() {` до вашого `script`. Потім закрийте його на наступній лінії (ще всередині вашого елемента `script`) за допомогою: `});`

Ми пізніше дізнаємося більше про `functions`. Важливо знати, що код, який ви помістите всередину цього `function`, буде виконаний одразу після того, як браузер завантажить вашу сторінку.

Це важливо, бо без вашого `document ready function` код може виконатися раніше, ніж буде відтворено HTML, що спричинить помилки.

# --hints--

Ви маєте створити елемент `script`, переконавшись, що він валідний і має кінцевий тег.

```js
assert(
  code.match(/<\/script\s*>/g) &&
    code.match(
      /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
    ) &&
    code.match(/<\/script\s*>/g).length ===
      code.match(
        /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
      ).length
);
```

Ви маєте додати `$(document).ready(function() {` на початок вашого елемента `script`.

```js
assert(
  code.match(
    /\$\s*?\(\s*?document\s*?\)\.ready\s*?\(\s*?function\s*?\(\s*?\)\s*?\{/g
  )
);
```

Ви маєте закрити вашу функцію `$(document).ready(function() {` за допомогою `});`

```js
assert(code.match(/\n*?\s*?\}\s*?\);/g));
```

# --seed--

## --seed-contents--

```html
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
  });
</script>
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
