---
id: 68e4276483b0c71658aab1ce
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Haverá momentos em que você vai querer remover espaços em branco apenas no início da string. É aí que o método `trimStart()` é útil.

Aqui está um exemplo:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

Lembre-se de que strings são imutáveis. Este método criará uma nova string com o resultado. Não modificará a original.

Crie uma variável `trimmedStart` e atribua a ela o resultado de usar o método `trimStart()` na variável `userInput`.

# --hints--

Você deve ter uma variável `trimmedStart`.

```js
assert.exists(trimmedStart);
```

Sua variável `trimmedStart` deve ser uma string.

```js
assert.isString(trimmedStart);
```

Você deve atribuir `userInput.trimStart()` à sua variável `trimmedStart`.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
