---
id: 68e4276483b0c71658aab1d3
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Como você se lembra das lições e workshops anteriores, você pode extrair partes de uma string usando o método `slice()`.

Aqui está um lembrete:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

Altere o valor de `camelCasedVersion` de uma string vazia para o resultado do uso do método `slice()` na variável `lowercaseWord`. Passe os números `0` e `5` que representam os índices inicial e final para o método `slice()`.

Agora você deve ver a palavra `"camel"` no console.

# --hints--

Sua variável `camelCasedVersion` não deve ser uma string vazia.

```js
assert.isNotEmpty(camelCasedVersion);
```

Você deve atribuir `lowercaseWord.slice(0, 5)` à variável `camelCasedVersion`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

Você NÃO deve atribuir a string literal `"camel"` ao `camelCasedVersion`. Por favor, releia as instruções para a abordagem correta.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method — leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method — trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
