---
id: 68e4276483b0c71658aab1d2
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Haverá momentos em que você precisará usar múltiplos métodos para criar uma nova string. Nas próximas etapas, você usará alguns dos métodos que aprendeu até agora para criar uma nova string que utiliza a convenção de nomenclatura camelCase.

Comece criando uma variável chamada `lowercaseWord` e atribua a ela a string `"camelcase"`. 

Abaixo daquela variável `lowercaseWord`, crie outra variável chamada `camelCasedVersion` e atribua a ela uma string vazia por enquanto.

Abaixo da variável `camelCasedVersion`, adicione um `console.log()` com a string `"Camel cased version:"` seguido de outro `console.log()` para a variável `camelCasedVersion`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve ter uma variável chamada `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

Seu `lowercaseWord` deve ser uma string.

```js
assert.isString(lowercaseWord);
```

Você deve atribuir a string `"camelcase"` à sua variável `lowercaseWord`.

```js
assert.equal(lowercaseWord, "camelcase");
```

Você deve ter uma variável `camelCasedVersion`.

```js
assert.exists(camelCasedVersion);
```

Sua variável `camelCasedVersion` deve ser uma string.

```js
assert.isString(camelCasedVersion);
```

Você deve atribuir uma string vazia à sua variável `camelCasedVersion`.

```js
assert.isEmpty(camelCasedVersion);
```

Você deve registrar a string `"Camel cased version:"` no console.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

Você deve registrar a variável `camelCasedVersion` no console.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method — leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method — trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
