---
id: 68e4276483b0c71658aab1d0
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Haverá momentos em que você precisará criar uma nova string com uma capitalização diferente. 

Para criar uma nova string com todas as letras convertidas para maiúsculas, você pode usar o método `toUpperCase()`.

Aqui está um exemplo:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Crie uma variável chamada `upperCaseInput` e atribua a ela o resultado do uso do método `toUpperCase()` na variável `cleanedInput`.

# --hints--

Você deve ter uma variável `upperCaseInput`.

```js
assert.exists(upperCaseInput);
```

Sua variável `upperCaseInput` deve ser uma string.

```js
assert.isString(upperCaseInput);
```

Você deve atribuir `cleanedInput.toUpperCase()` à sua variável `upperCaseInput`.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method — leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method — trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
