---
id: 68e4276483b0c71658aab1d1
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Se você precisar criar uma nova string onde todas as letras estejam em minúsculas, você pode usar o método `toLowerCase()`.  

Comece criando uma variável chamada `lowerCaseInput` e atribua a ela o resultado de usar o método `toLowerCase()` na variável `cleanedInput`. 

Então registre a string `"Result of using the toLowerCase() method:"`.

Finalmente, adicione outro `console.log()` para a variável `lowerCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve ter uma variável `lowerCaseInput`.

```js
assert.exists(lowerCaseInput);
```

Sua variável `lowerCaseInput` deve ser uma string.

```js
assert.isString(lowerCaseInput);
```

Você deve atribuir `cleanedInput.toLowerCase()` à sua variável `lowerCaseInput`.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

Você deve registrar a string `"Result of using the toLowerCase() method:"` no console.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

Você deve registrar a variável `lowerCaseInput` no console.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
