---
id: 68e4276483b0c71658aab1cc
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

O valor de `userInput` tem alguns espaços extras nas duas extremidades da string. Para criar uma nova string com esses espaços em branco removidos, você pode usar o método `trim()`.  

Aqui está um exemplo:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

Crie uma variável chamada `cleanedInput` e atribua a ela o resultado de usar o método `trim()` na variável `userInput`.

# --hints--

Você deve ter uma variável `cleanedInput`.

```js
assert.exists(cleanedInput);
```

Sua variável `cleanedInput` deve ser uma string.

```js
assert.isString(cleanedInput);
```

Você deve atribuir `userInput.trim()` à sua variável `cleanedInput`.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
