---
id: 68e4276483b0c71658aab1cd
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora é hora de ver o resultado.

Comece adicionando um `console.log()` com a string `"Result of trimming whitespace from both ends:"`.

Então, abaixo daquele `console.log()`, adicione outro `console.log()` com a variável `cleanedInput`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar a string `"Result of trimming whitespace from both ends:"` no console.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

Você deve registrar a variável `cleanedInput` no console.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
