---
id: 68e4276483b0c71658aab1ca
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você praticará o uso de vários métodos de string usados para formatar strings.

Comece criando uma variável chamada `userInput` e atribua a ela o seguinte valor: `"   Hello World!   "`.

# --hints--

Você deve ter uma variável `userInput`.

```js
assert.exists(userInput);
```

Sua variável `userInput` deve ser uma string.

```js
assert.isString(userInput);
```

Você deve atribuir `"   Hello World!   "` à sua variável `userInput`. Certifique-se de que há três espaços antes do `H` e três espaços depois do `!`.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
