---
id: 68ed6079fe0e3e98b94a1b75
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Em seguida, adicione um `console.log()` com a string `"Original input:"`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar a string `"Original input:"` no console.

```js
assert.equal(getLogs()[0], "Original input:");
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
--fcc-editable-region--

--fcc-editable-region--
```
