---
id: 68f1179eaea3bb59ba434386
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Para remover espaços em branco no final de uma string, você pode usar o método `trimEnd()`.

Aqui está um exemplo:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Crie uma variável chamada `trimmedEnd` e atribua a ela o resultado do uso do método `trimEnd()` na variável `userInput`. 

# --hints--

Você deve ter uma variável `trimmedEnd`.

```js
assert.exists(trimmedEnd);
```

Sua variável `trimmedEnd` deve ser uma string.

```js
assert.isString(trimmedEnd);
```

Você deve atribuir `userInput.trimEnd()` à sua variável `trimmedEnd`.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method — leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
