---
id: 68f3eeaa65614107ecd5f1f2
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Abaixo do seu primeiro `console.log()`, adicione outro `console.log()` para registrar a variável `userInput` no console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar a variável `userInput` no console.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
