---
id: 68f3f5ca1d2a490c97f14f03
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Agora é hora de ver os resultados de remover espaços em branco no final da string.

Abaixo da sua variável `trimmedEnd`, adicione um `console.log()` com a string `"After using the trimEnd() method, trailing spaces removed:"`. 

Então, abaixo daquele `console.log()`, adicione outro `console.log()` com a variável `trimmedEnd`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar a string `"After using the trimEnd() method, trailing spaces removed:"` no console.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

Você deve registrar a variável `trimmedEnd` no console.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--
const trimmedEnd = userInput.trimEnd();

--fcc-editable-region--
```
