---
id: 68f3f890010c870e17bbe754
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Agora é hora de ver os resultados.

Abaixo da sua variável `upperCaseInput`, adicione um `console.log()` com a string `"Result of using the toUpperCase() method:"`. 

Então abaixo daquele `console.log()`, adicione outro `console.log()` com a variável `upperCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar a string `"Result of using the toUpperCase() method:"` no console.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

Você deve registrar a variável `upperCaseInput` no console.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
