---
id: 68f3f3d65d51d80ac0b412fc
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Agora é hora de ver os resultados de remover espaços em branco no início da string.

Abaixo da sua variável `trimmedStart`, adicione um `console.log()` com a string `"After using the trimStart() method, leading spaces removed:"`. 

Então, abaixo daquele `console.log()`, adicione outro `console.log()` com a variável `trimmedStart`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar a string `"After using the trimStart() method, leading spaces removed:"` no console.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

Você deve registrar a variável `trimmedStart` no console.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
