---
id: 68f40925d810ae14b2d5a169
title: Passo 16
challengeType: 1
dashedName: step-16
---

# --description--

No momento sua variável `camelCasedVersion` contém o valor `camelC`. Mas o resultado desejado é `camelCase`. 

Para obter as letras restantes `ase`, você pode usar o método `slice()` com um valor de índice negativo. Lembre-se de que usar valores de índice negativos indica ao método `slice()` para extrair a partir do final. 

Aqui está um exemplo de lembrete:

```js
const str = "abcd";
str.slice(-2); // cd
```

Use o operador `+` para concatenar `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` com o resultado de usar o método `slice()` em `lowercaseWord`. Passe o número `-3` para o método `slice()`.

E com essa última alteração, seu workshop está completo!

# --hints--

Você deve concatenar `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` com `lowercaseWord.slice(-3)`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

Você NÃO deve atribuir a string literal `"camelCase"` ao `camelCasedVersion`. Por favor, releia as instruções para a abordagem correta.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method — leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method — trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method — leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method — trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
