---
id: 68f405c8c9df44132c8fe1bb
title: Passo 15
challengeType: 1
dashedName: step-15
---

# --description--

Como você se lembra das lições anteriores, a convenção de nomenclatura camelCase começa com uma letra minúscula para a primeira palavra e cada palavra subsequente começa com uma letra maiúscula seguida de letras minúsculas.

Aqui estão alguns exemplos de lembretes: 

```md
freeCodeCamp
camelCase
learningIsFun
```

A segunda palavra na variável `lowercaseWord` é `"case"`. Para acessar o `c` nessa palavra, você pode usar `lowercaseWord[5]`. 

Use o operador `+` para concatenar `lowercaseWord.slice(0, 5)` com o resultado de usar o método correto para converter strings para maiúsculas em `lowercaseWord[5]`.

Agora você deve ver `camelC` no console.

# --hints--

Você deve concatenar `lowercaseWord.slice(0, 5)` com `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

Você NÃO deve atribuir a string literal `"camelC"` ao `camelCasedVersion`. Por favor, releia as instruções para a abordagem correta.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method — leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method — trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

