---
id: 6718d2d59337c822ecb697ff
title: بناء برنامج إدارة حساب بنكي
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب تعريف صنف باسم `BankAccount` مع دالة البناء. يجب أن تضبط دالة البناء القيمة الافتراضية لـ `balance` على `0` وتُهيئ مصفوفة فارغة باسم `transactions` لتخزين سجلات المعاملات ككائنات.

2. يجب أن تكون كل معاملة مخزنة في مصفوفة `transactions` كائنًا يحتوي على خاصيتين: `type` و `amount`. يجب أن تكون الخاصية `type` إما `deposit` أو `withdraw`، والخاصية `amount` هي المبلغ المودع أو المسحوب.

3. يجب تعريف طريقة باسم `deposit` تأخذ مبلغ الإيداع كمعلمة. عندما يكون مبلغ الإيداع أكبر من `0`، يجب أن:

- تضيف كائنًا جديدًا إلى مصفوفة المعاملات بنوع `deposit` والمبلغ المودع.
   - تحدّث الرصيد.
   - تُرجع `"Successfully deposited $[amount]. New balance: $[balance]"`.

4. إذا كان المبلغ أقل من أو يساوي `0`، يجب أن تُرجع طريقة `deposit` الرسالة `"Deposit amount must be greater than zero."`.

5. يجب تعريف طريقة باسم `withdraw` تأخذ مبلغًا كمعلمة. يجب أن تحدّث هذه الطريقة الرصيد الحالي حسب عمليات السحب. عندما يكون المبلغ المراد سحبه أكبر من `0` وأقل من أو يساوي الرصيد الحالي، يجب أن:

- تضيف كائنًا جديدًا إلى مصفوفة المعاملات بنوع `withdraw` والمبلغ المسحوب.
   - تحدّث الرصيد.
   - تُرجع `"Successfully withdrew $[amount]. New balance: $[balance]"`.

6. إذا كان المبلغ المراد سحبه أقل من أو يساوي `0` أو أكبر من الرصيد الحالي، يجب أن تُرجع طريقة `withdraw` الرسالة `"Insufficient balance or invalid amount."`.


7. يجب تعريف طريقة باسم `checkBalance` تُرجع الرصيد الحالي بصيغة `"Current balance: $[balance]"`.

8. يجب تعريف طريقة باسم `listAllDeposits` تقوم بتكرار مصفوفة المعاملات وتُرجع جميع الإيداعات بصيغة `"Deposits: amount,amount,..."`.

9. يجب تعريف طريقة باسم `listAllWithdrawals` تقوم بتكرار مصفوفة المعاملات وتُرجع جميع عمليات السحب بصيغة `"Withdrawals: amount,amount,..."`.

10. يجب إنشاء نسخة جديدة من `BankAccount` باسم `myAccount`.

11. يجب أن يحتوي حسابك البنكي `myAccount` على خمسة معاملات على الأقل.

12. يجب أن يحتوي حسابك البنكي `myAccount` على إيداعين على الأقل.

13. يجب أن يحتوي حسابك البنكي `myAccount` على عمليتي سحب على الأقل.

14. يجب أن يكون رصيد حسابك البنكي `myAccount` أكبر من `$100`.

# --hints--

يجب تعريف صنف باسم `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

يجب أن يحتوي كائن `BankAccount` في البداية على `balance` بقيمة `0` ومصفوفة فارغة `transactions` لتخزين سجلات المعاملات.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

يجب أن تحتوي على طريقة `deposit` تأخذ مبلغ الإيداع كمعلمة.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

يجب أن تحتوي على طريقة `withdraw` تأخذ مبلغ السحب كمعلمة.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

يجب أن تحتوي على طريقة `checkBalance` تتحقق من الرصيد الحالي.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

يجب أن تحتوي على طريقة `listAllDeposits` تعرض جميع الإيداعات.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

يجب أن تحتوي على طريقة `listAllWithdrawals` تعرض جميع عمليات السحب.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` يجب أن تُرجع `"Successfully deposited $100. New balance: $100"`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` يجب أن تُرجع `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` يجب أن تُرجع `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

عندما يكون رصيد الحساب `100`، يجب أن تُرجع `BankAccount.withdraw(150)` الرسالة `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` يجب أن تُرجع `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` يجب أن تُرجع `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

عندما يكون رصيد الحساب `200`، يجب أن تُرجع `BankAccount.withdraw(150)` الرسالة `"Successfully withdrew $150. New balance: $50"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

عندما يكون رصيد الحساب `200`، يجب أن تُرجع `BankAccount.checkBalance()` الرسالة `"Current balance: $200"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

عند إيداع `10`، `35`، `90`، يجب أن تُرجع طريقة `listAllDeposits` الرسالة `"Deposits: 10,35,90"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

عند سحب `20`، `50`، `100`، يجب أن تُرجع طريقة `listAllWithdrawals` الرسالة `"Withdrawals: 20,50,100"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

يجب أن يكون لديك نسخة من `BankAccount` باسم `myAccount`.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

يجب أن يحتوي حسابك البنكي `myAccount` على خمسة معاملات على الأقل.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

يجب أن يحتوي حسابك البنكي `myAccount` على إيداعين على الأقل.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

يجب أن يحتوي حسابك البنكي `myAccount` على عمليتي سحب على الأقل.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```


يجب أن يكون رصيد حسابك البنكي `myAccount` أكبر من `$100`.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
