---
id: 60f0286404aefb0562a4fdf9
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Füge ein `title` und ein `meta` Element innerhalb des `head`-Elementes ein. Gib deinem Projekt den Titel `Registration Form` und füge das `charset` Attribut mit dem Wert `utf-8` deinem `meta` Element ein.

# --hints--

Dein Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your `title` element should be within the `head` element.

```js
assert.exists(document.querySelector('head > title'));
```

Your project should have a title of `Registration Form`. Denke daran, dass Schreibweise und Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

You should create a `meta` element within the `head` element.

```js
assert.exists(document.querySelector('head > meta'));
```

You should give the `meta` element a `charset` attribute with the value of `utf-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Your `meta` element should be a void element, it does not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
