---
id: 60f0286404aefb0562a4fdf9
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Add a `title` and `meta` element inside the `head` element. Give your project a title of `Registration Form`, and add the `charset` attribute with a value of `utf-8` to your `meta` element.

# --hints--

Dein Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Dein `title`-Element sollte sich innerhalb des `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > title'));
```

Dein Projekt sollte den Titel `Registration Form` haben. Remember, the casing and spelling matters for the title.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Du solltest ein `meta`-Element innerhalb des `head`-Elements erstellen.

```js
assert.exists(document.querySelector('head > meta'));
```

Du solltest dem `meta`-Element ein `charset`-Attribut mit dem Wert `utf-8` geben.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Das `meta` Element soll ein leeres Element sein, es hat kein schließendes Tag `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
