---
id: 62cc5b1779e4d313466f73c5
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Nest a `link` element within the `head` element. Give it a `rel` attribute with a value of `stylesheet` and an `href` attribute with a value of `styles.css`.

# --hints--

Dein Code sollte ein `link`-Element enthalten.

```js
assert.isNotNull(document.querySelector('link'));
```

Du solltest deine vorhandenen `head`-Tags nicht ändern. Stelle sicher, dass du das abschließende Tag nicht gelöscht hast.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

Du solltest ein `link`-Element haben.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Dein`link`-Element sollte ein leeres Element sein, es sollte kein schließendes Tag besitzen `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Dein`link`-Element sollte innerhalb des `head`-Elements liegen.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
