---
id: 66be24cb4144f955b6bcc550
title: 构建一页扑克牌
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该构建一个网页，显示至少三张扑克牌。
1. 你应该有一个 `main` 元素，ID 为 `playing-cards`。
1. 在你的 `#playing-cards` 元素内，应该至少有三个 `div` 元素，每个都有一个 `card` 类。
1. 在每个 `.card` 元素内，应该有三个 `div` 元素，第一个带有 `left` 类，第二个带有 `middle` 类，第三个带有 `right` 类。
1. 你的 `#playing-cards` 元素应该使用 flexbox 水平居中其子元素，允许它们换行，并在它们之间设置 `20px` 的间距。
1. 你的每个 `.card` 元素应该使用 flexbox 以 `space-between` 方式排列其子元素。
1. 你的每个 `.left` 元素应该使用 flexbox 项属性，使其在父容器的起始位置对齐。
1. 你的每个 `.middle` 元素应该使用 flexbox 项属性，使其在父容器的中间对齐。
1. 你的每个 `.right` 元素应该使用 flexbox 项属性，使其在父容器的末尾对齐。
1. 你的每个 `.middle` 元素应该使用 flexbox 以列方向显示其子元素。

以下是一些你可以拷贝并粘贴到你的卡片中使用的字符：`♠`、`♣`、`♥` 和 `♦`。

# --hints--

你应该有一个 `main` 元素，其 `id` 为 `playing-cards`。

```js
assert.exists(document.querySelector('main#playing-cards'));
```

你的 `#playing-cards` 元素内应至少有三个带有 `card` 类的 `div` 元素。

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

你的每个 `.card` 元素都应有宽度和高度。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

你的每个 `.card` 元素应有且仅有三个 `div` 元素作为子元素。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

你应该在每个 `.card` 元素内有一个类为 `left` 的 `div`。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

你应该在每个 `.card` 元素内有一个类为 `middle` 的 `div`。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

你应该在每个 `.card` 元素内有一个类为 `right` 的 `div`。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

你应该有一个 `#playing-cards` 选择器，将其元素的 `display` 设置为 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

你应该有一个 `#playing-cards` 选择器，将其元素的 `justify-content` 设置为 `center`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

你应该有一个 `#playing-cards` 选择器，将其元素的 `flex-wrap` 设置为 `wrap`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

你应该有一个 `#playing-cards` 选择器，将其元素的 `gap` 设置为 `20px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

你应该有一个 `.card` 选择器，将其元素的 `display` 设置为 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

你应该有一个 `.card` 选择器，将其元素的 `justify-content` 设置为 `space-between`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

你应该有一个 `.left` 选择器，将其元素的 `align-self` 设置为 `flex-start`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

你应该有一个 `.middle` 选择器，将其元素的 `align-self` 设置为 `center`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

你应该有一个 `.right` 选择器，将其元素的 `align-self` 设置为 `flex-end`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

你应该有一个 `.middle` 选择器，将其元素的 `display` 设置为 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

你应该在你的 `.middle` 类中将 `flex-direction` 设置为 `column`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
