---
id: 564944c91be2204b269d51e3
title: Badilisha maandishi ndani ya kipengee kwa kutumia jQuery
challengeType: 6
forumTopicId: 16773
dashedName: change-text-inside-an-element-using-jquery
---

# --description--

Kwa kutumia jQuery, unaweza kubadilisha maandishi kati ya lebo za mwanzo na mwisho za kipengee. Hata unaweza kubadilisha alama za uandishi za HTML.

jQuery ina kitendakazi kinachoitwa `.html()` kinachokuwezesha kuongeza lebo za HTML na maandishi ndani ya kipengee. Yoyote yaliyomo awali ndani ya kipengee hicho yatabadilishwa kabisa na yaliyomo unayotoa kwa kutumia kitendakazi hiki.

Hapa ni jinsi unavyoweza kuandika upya na kusisitiza maandishi ya kichwa chetu cha sehemu:

```js
$("h3").html("<em>jQuery Playground</em>");
```

jQuery pia ina kitendakazi kinachofanana kinachoitwa `.text()` ambacho hubadilisha maandishi tu bila kuongeza lebo. Kwa maneno mengine, kitendakazi hiki hakitatathmini lebo zozote za HTML zinazotumwa kwacho, bali kitachukulia kama maandishi unayotaka kubadilisha yaliyomo yaliyopo.

Badilisha kitufe chenye kitambulisho `target4` kwa kusisitiza maandishi yake.

<a href="https://www.freecodecamp.org/news/html-elements-explained-what-are-html-tags/#em-element" target="_blank" rel="noopener noreferrer nofollow">Tazama makala yetu ya habari kuhusu &lt;em&gt;</a> kujifunza tofauti kati ya `<i>` na `<em>` na matumizi yao.

Kumbuka kwamba ingawa lebo ya `<i>` kwa kawaida imekuwa ikitumika kusisitiza maandishi, sasa imetumika kama lebo ya alama za picha. Lebo ya `<em>` sasa inakubalika sana kama lebo ya kusisitiza. Moja kati ya hizo zitafanya kazi kwa changamoto hii.

# --hints--

Unapaswa kusisitiza maandishi kwenye kitufe chako cha `target4` kwa kuongeza lebo za HTML.

```js
assert.isTrue(
  /<em>|<i>\s*#target4\s*<\/em>|<\/i>/gi.test($('#target4').html())
);
```

Maandishi mengine yanapaswa kubaki bila kubadilika.

```js
assert($('#target4') && $('#target4').text().trim() === '#target4');
```

Haupaswi kubadilisha maandishi mengine yoyote.

```js
assert.isFalse(/<em>|<i>/gi.test($('h3').html()));
```

Unapaswa kutumia `.html()` na si `.text()`.

```js
assert(code.match(/\.html\(/g));
```

Unapaswa kuchagua `button id="target4"` kwa kutumia jQuery.

```js
assert(code.match(/\$\(\s*?(\"|\')#target4(\"|\')\s*?\)\.html\(/));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target4").html('<em>#target4</em>');
  });
</script>

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
