---
id: bad87fee1348bd9aed508826
title: Nakili kipengee kwa kutumia jQuery
challengeType: 6
forumTopicId: 16780
dashedName: clone-an-element-using-jquery
---

# --description--

Mbali na kusogeza vipengee, unaweza pia kuvinakili kutoka mahali pamoja hadi pengine.

jQuery ina kitendakazi kinachoitwa `clone()` ambacho hufanya nakala ya kipengee.

Kwa mfano, kama tungetaka kunakili `target2` kutoka `left-well` yetu hadi `right-well`, tungetumia:

```js
$("#target2").clone().appendTo("#right-well");
```

Je, umeona hili linahusisha kuunganishwa kwa vitendakazi viwili vya jQuery? Hii huitwa <dfn>uunganishaji wa mfululizo wa vitendakazi</dfn> na ni njia rahisi ya kufanya mambo kwa jQuery.

Nakilisha kipengee chako cha `target5` na ukiongeze kwenye `left-well` yako.

# --hints--

Kipengee chako cha `target5` kinapaswa kuwa ndani ya `right-well` yako.

```js
assert($('#right-well').children('#target5').length > 0);
```

Nakala ya kipengee chako cha `target5` inapaswa pia kuwa ndani ya `left-well` yako.

```js
assert($('#left-well').children('#target5').length > 0);
```

Unapaswa kutumia jQuery tu kusogeza vipengee hivi.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
