---
id: bad87fee1348bd9aed108826
title: Lenga Mtoto Maalum wa Kipengele Kutumia jQuery
challengeType: 6
forumTopicId: 18315
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-a-specific-child-of-an-element-using-jquery
---

# --description--

Umeona kwa nini sifa za id ni rahisi sana kwa ajili ya kulenga kwa kutumia vichaguzi vya jQuery. Lakini hutakuwa na id zilizo rahisi kila wakati kufanya kazi nazo.

Kwa bahati nzuri, jQuery ina mbinu nyingine za kulenga vipengele sahihi.

jQuery hutumia Vichaguzi vya CSS kulenga vipengele. Kichaguzi cha `target:nth-child(n)` cha CSS hukuruhusu kuchagua vipengele vyote vya nambari ya nth vilivyo na darasa la vitu au aina ya kipengele kinacholengwa.

Hapa ni jinsi unavyoweza kumpa kipengele cha tatu katika kila well darasa la bounce:

```js
$(".target:nth-child(3)").addClass("animated bounce");
```

Fanya mtoto wa pili katika kila kipengele chako cha well apige bounce. Lazima uchague watoto wa vipengele vyenye darasa la `target`.

# --hints--

Kipengele cha pili katika vipengele vyako vya `target` kinapaswa kupiga bounce.

```js
assert(
  $('.target:nth-child(2)').hasClass('animated') &&
    $('.target:nth-child(2)').hasClass('bounce')
);
```

Vipengele viwili tu ndizo zinapaswa kupiga bounce.

```js
assert($('.animated.bounce').length === 2);
```

Unapaswa kutumia kichaguzi cha `:nth-child()` kubadilisha vipengele hivi.

```js
assert(code.match(/\:nth-child\(/g));
```

Unapaswa kutumia jQuery tu kuongeza madarasa haya kwa kipengele.

```js
assert(
  code.match(/\$\(".target:nth-child\(2\)"\)/g) ||
    code.match(/\$\('.target:nth-child\(2\)'\)/g) ||
    code.match(/\$\(".target"\).filter\(":nth-child\(2\)"\)/g) ||
    code.match(/\$\('.target'\).filter\(':nth-child\(2\)'\)/g)
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
    $("#right-well").children().css("color", "orange");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
    $("#right-well").children().css("color", "orange");
    $(".target:nth-child(2)").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
