---
id: bad87fee1348bd9aedc08826
title: Lenga vitu kwa darasa la vitu ukitumia jQuery
challengeType: 6
forumTopicId: 18316
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-class-using-jquery
---

# --description--

Unaona jinsi tulivyofanya vitu vyote vya `button` vyarukaruke? Tulivichagua kwa `$("button")`, kisha tukawaongeza madarasa ya CSS kwa kutumia `.addClass("animated bounce");`.

Umetumia kitendakazi cha `.addClass()` cha jQuery, ambacho hukuruhusu kuongeza madarasa kwa vitu.

Kwanza, tuelekee vitu vyako vya `div` vilivyo na darasa la `well` kwa kutumia kichaguzi cha `$(".well")`.

Kumbuka kwamba, kama ilivyo kwa tamko za CSS, unaandika `.` kabla ya jina la darasa.

Kisha tumia kitendakazi cha `.addClass()` cha jQuery kuongeza madarasa `animated` na `shake`.

Kwa mfano, unaweza kufanya vitu vyote vilivyo na darasa la `text-primary` kutetemeka kwa kuongeza yafuatayo kwenye `document ready function` yako:

```js
$(".text-primary").addClass("animated shake");
```

# --hints--

Unapaswa kutumia kitendakazi cha `addClass()` cha jQuery kutoa madarasa `animated` na `shake` kwa vitu vyote vyako vilivyo na darasa la `well`.

```js
assert($('.well').hasClass('animated') && $('.well').hasClass('shake'));
```

Unapaswa kutumia jQuery tu kuongeza madarasa haya kwa kitu.

```js
assert(!code.match(/class\.\*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
