---
id: bad87fee1348bd9aeda08826
title: Lenga vipengele kwa id ukitumia jQuery
challengeType: 6
forumTopicId: 18317
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-id-using-jquery
---

# --description--

Unaweza pia kulenga vipengele kwa sifa zao za id.

Kwanza lenga kipengele chako cha `button` chenye id `target3` kwa kutumia kichaguzi cha `$("#target3")`.

Kumbuka kwamba, kama ilivyo kwa tamko za CSS, unaandika `#` kabla ya jina la id.

Kisha tumia kitendakazi cha jQuery cha `.addClass()` kuongeza madarasa ya `animated` na `fadeOut`.

Hapa ni jinsi unavyofanya kipengele cha `button` chenye id `target6` kififie:

```js
$("#target6").addClass("animated fadeOut");
```

# --hints--

Unapaswa kuchagua kipengele cha `button` kwa `id` ya `target3` na utumie kitendakazi cha jQuery cha `addClass()` kumpa darasa la `animated`.

```js
assert($('#target3').hasClass('animated'));
```

Unapaswa kulenga kipengele chenye id `target3` na utumie kitendakazi cha jQuery cha `addClass()` kumpa darasa la `fadeOut`.

```js
assert(
  ($('#target3').hasClass('fadeOut') || $('#target3').hasClass('fadeout')) &&
    code.match(/\$\(\s*.#target3.\s*\)/g)
);
```

Unapaswa kutumia jQuery tu kuongeza madarasa haya kwa kipengele.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
    $("#target3").addClass("animated fadeOut");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
