---
id: bad87fee1348bd9aeda08726
title: Futa kitendakazi chako cha jQuery
challengeType: 6
forumTopicId: 17561
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: delete-your-jquery-functions
---

# --description--

Miondoko hii ilikuwa nzuri mwanzoni, lakini sasa inakuwa ya kuingiza akili.

Futa vitendakazi vyote vitatu vya jQuery kutoka kwenye `document ready function`, lakini acha `document ready function` yenyewe isiharibike.

# --hints--

Vitendakazi vyote vitatu vya jQuery vinapaswa kufutwa kutoka kwenye `document ready function`.

```js
assert(code.match(/\{\s*\}\);/g));
```

Unapaswa kuacha kipengele chako cha `script` kisiharibike.

```js
assert(code.match(/<script>/g));
```

Unapaswa kuacha `$(document).ready(function() {` mwanzoni mwa kipengele chako cha `script`.

```js
assert(code.match(/\$\(document\)\.ready\(function\(\)\s?\{/g));
```

Unapaswa kuacha lebo ya kufunga ya kitendakazi cha `document.ready` isiharibike `})`.

```js
assert(code.match(/.*\s*\}\);/g));
```

Unapaswa kuacha lebo ya kufunga ya kipengele chako cha `script` isiharibike.

```js
assert(
  code.match(/<\/script>/g) &&
    code.match(/<script/g) &&
    code.match(/<\/script>/g).length === code.match(/<script/g).length
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
    $("#target3").addClass("animated fadeOut");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
