---
id: bad87fee1348bd9aed308826
title: Lenga mzazi wa kipengele kwa kutumia jQuery
challengeType: 6
forumTopicId: 18321
dashedName: target-the-parent-of-an-element-using-jquery
---

# --description--

Kipengele chochote cha HTML kina kipengele cha `parent` ambacho hutumia `inherits` vigezo.

Kwa mfano, kipengele cha `h3` katika `jQuery Playground` yako kina kipengele mzazi cha `<div class="container-fluid">`, ambacho pia kina kipengele mzazi cha `body`.

jQuery ina kitendakazi kinachoitwa `parent()` kinachokuwezesha kupata upatikanaji wa mzazi wa kipengele chochote ulichochagua.

Hapa kuna mfano wa jinsi unavyotumia kitendakazi cha `parent()` ikiwa ungependa kumpa kipengele mzazi cha kipengele cha `left-well` rangi ya nyuma ya bluu:

```js
$("#left-well").parent().css("background-color", "blue")
```

Mpatie mzazi wa kipengele cha `#target1` rangi ya nyuma nyekundu.

# --hints--

Kipengele chako cha `left-well` kinapaswa kuwa na rangi ya nyuma nyekundu.

```js
assert(
  $('#left-well').css('background-color') === 'red' ||
    $('#left-well').css('background-color') === 'rgb(255, 0, 0)' ||
    $('#left-well').css('background-color').toLowerCase() === '#ff0000' ||
    $('#left-well').css('background-color').toLowerCase() === '#f00'
);
```

Unapaswa kutumia kitendakazi cha `.parent()` kubadilisha kipengele hiki.

```js
assert(code.match(/\.parent\s*\(\s*\)\s*\.css/g));
```

Njia ya `.parent()` inapaswa kuitwa kwenye kipengele cha `#target1`.

```js
assert(
  code.match(/\$\s*?\(\s*?(?:'|")\s*?#target1\s*?(?:'|")\s*?\)\s*?\.parent/gi)
);
```

Unapaswa kutumia jQuery tu kuongeza madarasa haya kwenye kipengele.

```js
assert(code.match(/<div class="well" id="left-well">/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");

  });
</script>

<!-- Only change code above this line -->

<body>
  <div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
    <div class="row">
      <div class="col-xs-6">
        <h4>#left-well</h4>
        <div class="well" id="left-well">
          <button class="btn btn-default target" id="target1">#target1</button>
          <button class="btn btn-default target" id="target2">#target2</button>
          <button class="btn btn-default target" id="target3">#target3</button>
        </div>
      </div>
      <div class="col-xs-6">
        <h4>#right-well</h4>
        <div class="well" id="right-well">
          <button class="btn btn-default target" id="target4">#target4</button>
          <button class="btn btn-default target" id="target5">#target5</button>
          <button class="btn btn-default target" id="target6">#target6</button>
        </div>
      </div>
    </div>
  </div>
</body>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
  });
</script>

<!-- Only change code above this line -->

<body>
  <div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
    <div class="row">
      <div class="col-xs-6">
        <h4>#left-well</h4>
        <div class="well" id="left-well">
          <button class="btn btn-default target" id="target1">#target1</button>
          <button class="btn btn-default target" id="target2">#target2</button>
          <button class="btn btn-default target" id="target3">#target3</button>
        </div>
      </div>
      <div class="col-xs-6">
        <h4>#right-well</h4>
        <div class="well" id="right-well">
          <button class="btn btn-default target" id="target4">#target4</button>
          <button class="btn btn-default target" id="target5">#target5</button>
          <button class="btn btn-default target" id="target6">#target6</button>
        </div>
      </div>
    </div>
  </div>
</body>
```
