---
id: bad87fee1348bd9acdd08826
title: Jifunze jinsi lebo za script na Document Ready zinavyofanya kazi
challengeType: 6
forumTopicId: 18224
dashedName: learn-how-script-tags-and-document-ready-work
---

# --description--

Sasa tuko tayari kujifunza jQuery, chombo maarufu zaidi cha JavaScript kwa wakati wote.

Kabla hatujaanza kutumia jQuery, tunahitaji kuongeza vitu fulani kwenye HTML yetu.

Kwanza, ongeza kipengele cha `script` juu kabisa ya ukurasa wako. Hakikisha unafunga kwenye mstari unaofuata.

Kivinjari chako kitaendesha JavaScript yoyote ndani ya kipengele cha `script`, ikiwa ni pamoja na jQuery.

Ndani ya kipengele chako cha `script`, ongeza msimbo huu: `$(document).ready(function() {` kwenye `script` yako. Kisha uifunge kwenye mstari unaofuata (bado ndani ya kipengele chako cha `script`) kwa: `});`

Tutajifunza zaidi kuhusu `functions` baadaye. Kitu muhimu kujua ni kwamba msimbo unaoweka ndani ya `function` utaendeshwa mara tu kivinjari chako kitakapopakua ukurasa wako.

Hii ni muhimu kwa sababu bila `document ready function` yako, msimbo wako unaweza kuendeshwa kabla ya HTML yako kuonyeshwa, jambo ambalo litasababisha hitilafu.

# --hints--

Unapaswa kuunda kipengele cha `script` ukihakikisha ni halali na kina lebo ya kufunga.

```js
assert(
  code.match(/<\/script\s*>/g) &&
    code.match(
      /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
    ) &&
    code.match(/<\/script\s*>/g).length ===
      code.match(
        /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
      ).length
);
```

Unapaswa kuongeza `$(document).ready(function() {` mwanzoni mwa kipengele chako cha `script`.

```js
assert(
  code.match(
    /\$\s*?\(\s*?document\s*?\)\.ready\s*?\(\s*?function\s*?\(\s*?\)\s*?\{/g
  )
);
```

Unapaswa kufunga kitendakazi chako cha `$(document).ready(function() {` kwa `});`

```js
assert(code.match(/\n*?\s*?\}\s*?\);/g));
```

# --seed--

## --seed-contents--

```html
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
  });
</script>
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
