---
id: bad87fee1348bd9aed208826
title: Lenga watoto wa kipengele kwa kutumia jQuery
challengeType: 6
forumTopicId: 18320
dashedName: target-the-children-of-an-element-using-jquery
---

# --description--

Wakati vipengele vya HTML vikiwa katika ngazi moja chini ya kingine, huitwa <dfn>watoto</dfn> wa kipengele hicho. Kwa mfano, vipengele vya kitufe katika changamoto hii vyenye maandishi `#target1`, `#target2`, na `#target3` ni watoto wote wa kipengele cha `<div class="well" id="left-well">`.

jQuery ina kitendakazi kinachoitwa `children()` kinachokuwezesha kupata upatikanaji wa watoto wa kipengele chochote ulichochagua.

Hapa kuna mfano wa jinsi unavyotumia kitendakazi cha `children()` ili kutoa rangi `left-well` kwa watoto wa kipengele chako cha `blue`:

```js
$("#left-well").children().css("color", "blue")
```

# --instructions--

Toa rangi ya chungwa kwa watoto wote wa kipengele chako cha `right-well`.

# --hints--

Watoto wote wa `#right-well` wanapaswa kuwa na maandishi ya rangi ya chungwa.

```js
assert($('#right-well').children().css('color') === 'rgb(255, 165, 0)');
```

Unapaswa kutumia kitendakazi cha `children()` kubadilisha vipengele hivi.

```js
assert(code.match(/\.children\(\)\.css/g));
```

Unapaswa kutumia jQuery tu kuongeza madarasa haya kwa kipengele.

```js
assert(code.match(/<div class="well" id="right-well">/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
    $("#right-well").children().css("color", "orange");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
