---
id: bad87fee1348bd9aed908826
title: Badilisha CSS ya kipengele kwa kutumia jQuery
challengeType: 6
forumTopicId: 16776
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: change-the-css-of-an-element-using-jquery
---

# --description--

Tunaweza pia kubadilisha CSS ya kipengele cha HTML moja kwa moja kwa kutumia jQuery.

jQuery ina kitendakazi kinachoitwa `.css()` kinachokuwezesha kubadilisha CSS ya kipengele.

Hapa ni jinsi tunavyobadilisha rangi yake kuwa buluu:

```js
$("#target1").css("color", "blue");
```

Hii ni tofauti kidogo na tamko la kawaida la CSS, kwa sababu sifa ya CSS na thamani yake ziko ndani ya nukuu, na zimegawanywa kwa koma badala ya nukta mbili.

Futa wachaguzi wako wa jQuery, uache `document ready function` tupu.

Chagua `target1` na ubadilishe rangi yake kuwa nyekundu.

# --hints--

Kipengele chako cha `target1` kinapaswa kuwa na maandishi ya rangi nyekundu.

```js
assert($('#target1').css('color') === 'rgb(255, 0, 0)');
```

Unapaswa kutumia jQuery tu kuongeza madarasa haya kwenye kipengele.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
    $("#target3").addClass("animated fadeOut");
    $("button").removeClass("btn-default");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
    $("#target3").addClass("animated fadeOut");
    $("button").removeClass("btn-default");
    $("#target1").css("color", "red");
  });
</script>

<!-- Only change code above this line -->
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
