---
id: 587d8249367417b2b2512c41
title: محول القياس المتري-الإمبراطوري
challengeType: 4
forumTopicId: 301570
dashedName: metric-imperial-converter
---

# --description--

ابنِ تطبيق جافاسكريبت متكامل الوظائف يشبه هذا: <a href="https://metric-imperial-converter.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://metric-imperial-converter.freecodecamp.rocks/</a>. يتطلب هذا المشروع كتابة شفرتك باستخدام إحدى الطرق التالية:

- استنسخ <a href="https://github.com/freeCodeCamp/boilerplate-project-metricimpconverter/" target="_blank" rel="noopener noreferrer nofollow">مستودع GitHub هذا</a> وأكمل مشروعك محليًا.
- استخدم منشئ مواقع من اختيارك لإكمال المشروع. تأكد من دمج جميع الملفات من مستودع GitHub الخاص بنا.

**ملاحظة:** اختبارات هذا المشروع لا تعمل عند استخدام `glitch.com`.

# --instructions--

- أكمل منطق التحويل اللازم في `/controllers/convertHandler.js`
- أكمل المسارات اللازمة في `/routes/api.js`
- انسخ ملف `sample.env` إلى `.env` واضبط المتغيرات بشكل مناسب
- لتشغيل الاختبارات تلقائيًا، أضف `NODE_ENV=test` في ملف `.env`
- لتشغيل الاختبارات في وحدة التحكم، استخدم الأمر `npm run test`.

اكتب الاختبارات التالية في `tests/1_unit-tests.js`:

- يجب أن يقرأ `convertHandler` بشكل صحيح مُدخل عدد صحيح.
- يجب أن يقرأ `convertHandler` بشكل صحيح مُدخل عدد عشري.
- يجب أن يقرأ `convertHandler` بشكل صحيح مُدخل كسر.
- يجب أن يقرأ `convertHandler` بشكل صحيح مُدخل كسر مع عدد عشري.
- يجب أن يعيد `convertHandler` خطأ عند وجود كسر مزدوج (أي `3/2/3`).
- يجب أن يفترض `convertHandler` القيمة الرقمية `1` إذا لم يُقدم مُدخل رقمي.
- يجب أن يقرأ `convertHandler` كل وحدة مُدخلة صحيحة بشكل صحيح.
- يجب أن يعيد `convertHandler` خطأ لوحدة مُدخلة غير صحيحة.
- يجب أن يعيد `convertHandler` وحدة الإرجاع الصحيحة لكل وحدة مُدخلة صحيحة.
- يجب أن يعيد `convertHandler` اسم الوحدة مكتوبًا بشكل صحيح لكل وحدة مُدخلة صحيحة.
- يجب أن يحول `convertHandler` بشكل صحيح من `gal` إلى `L`.
- يجب أن يحول `convertHandler` بشكل صحيح من `L` إلى `gal`.
- يجب أن يحول `convertHandler` بشكل صحيح من `mi` إلى `km`.
- يجب أن يحول `convertHandler` بشكل صحيح من `km` إلى `mi`.
- يجب أن يحول `convertHandler` بشكل صحيح من `lbs` إلى `kg`.
- يجب أن يحول `convertHandler` بشكل صحيح من `kg` إلى `lbs`.

اكتب الاختبارات التالية في `tests/2_functional-tests.js`:

- تحويل مُدخل صحيح مثل `10L`: طلب `GET` إلى `/api/convert`.
- تحويل مُدخل غير صحيح مثل `32g`: طلب `GET` إلى `/api/convert`.
- تحويل رقم غير صحيح مثل `3/7.2/4kg`: طلب `GET` إلى `/api/convert`.
- تحويل رقم ووحدة غير صحيحتين مثل `3/7.2/4kilomegagram`: طلب `GET` إلى `/api/convert`.
- تحويل بدون رقم مثل `kg`: طلب `GET` إلى `/api/convert`.

# --hints--

يمكنك تقديم مشروعك الخاص، وليس عنوان URL المثال.

```js
  assert(
    !/.*\/metric-imperial-converter\.freecodecamp\.rocks/.test(
      code
    )
  );
```

يمكنك `GET` `/api/convert` مع معلمة واحدة تحتوي على رقم ووحدة مقبولة وتحويلها. (تلميح: قسّم المُدخل بالبحث عن فهرس أول حرف والذي سيحدد بداية الوحدة)

```js

```

يمكنك تحويل `'gal'` إلى `'L'` والعكس. (1 gal يساوي 3.78541 L)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 3.78541);
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10gal');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 37.8541);
    assert.equal(data2.returnUnit, 'L');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.26417);
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10l');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 2.64172);
    assert.equal(data4.returnUnit, 'gal');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

يمكنك تحويل `'lbs'` إلى `'kg'` والعكس. (1 lbs يساوي 0.453592 kg)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1lbs');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 0.45359);
    assert.equal(data1.returnUnit, 'kg');
    const response2 = await fetch(code + '/api/convert?input=10lbs');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 4.53592);
    assert.equal(data2.returnUnit, 'kg');
    const response3 = await fetch(code + '/api/convert?input=1kg');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 2.20462);
    assert.equal(data3.returnUnit, 'lbs');
    const response4 = await fetch(code + '/api/convert?input=10kg');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 22.04624);
    assert.equal(data4.returnUnit, 'lbs');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

يمكنك تحويل `'mi'` إلى `'km'` والعكس. (1 mi يساوي 1.60934 km)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 1.60934);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=10mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 16.0934);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(code + '/api/convert?input=1km');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.62137);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(code + '/api/convert?input=10km');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 6.21373);
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

يجب قبول جميع الوحدات الواردة بحالتها العليا أو السفلى، ولكن يجب إرجاعها في كل من `initUnit` و`returnUnit` بحروف صغيرة، باستثناء اللتر الذي يجب تمثيله بحرف `'L'` كبير.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.initUnit, 'gal');
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10L');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.initUnit, 'L');
    assert.equal(data2.returnUnit, 'gal');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.initUnit, 'L');
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10KM');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.initUnit, 'km');
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

إذا كانت وحدة القياس غير صحيحة، ستُرجع `'invalid unit'`.

```js
  try {
    const response = await fetch(code + '/api/convert?input=1min');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid unit' || data === 'invalid unit');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

إذا كان الرقم غير صحيح، ستُرجع `'invalid number'`.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2gal'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid number' || data === 'invalid number');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

إذا كانت كل من الوحدة والرقم غير صحيحتين، ستُرجع `'invalid number and unit'`.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2min'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(
      data.error === 'invalid number and unit' ||
        data === 'invalid number and unit'
    );
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

يمكنك استخدام الكسور أو الأعداد العشرية أو كلاهما في المعلمة (مثل 5، 1/2، 2.5/6)، ولكن إذا لم يُقدم شيء، فسيُفترض الرقم 1.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.approximately(data1.initNum, 1, 0.001);
    assert.approximately(data1.returnNum, 1.60934, 0.001);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=1/5mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.approximately(data2.initNum, 1 / 5, 0.1);
    assert.approximately(data2.returnNum, 0.32187, 0.001);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(
      code + '/api/convert?input=1.5/7km'
    );
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.approximately(data3.initNum, 1.5 / 7, 0.001);
    assert.approximately(data3.returnNum, 0.13315, 0.001);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(
      code + '/api/convert?input=3/2.7km'
    );
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.approximately(data4.initNum, 3 / 2.7, 0.001);
    assert.approximately(data4.returnNum, 0.69041, 0.001);
    assert.equal(data4.returnUnit, 'mi');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

سيحتوي الإرجاع على `initNum`، `initUnit`، `returnNum`، `returnUnit`، و`string` التي تكتب الوحدات بصيغة `'{initNum} {initUnitString} converts to {returnNum} {returnUnitString}'` مع تقريب النتيجة إلى 5 أرقام عشرية.

```js
  try {
    const response = await fetch(code + '/api/convert?input=2mi');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.json();
    assert.equal(data.initNum, 2);
    assert.equal(data.initUnit, 'mi');
    assert.approximately(data.returnNum, 3.21868, 0.001);
    assert.equal(data.returnUnit, 'km', 'returnUnit did not match');
    assert.equal(data.string, '2 miles converts to 3.21868 kilometers');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

اكتملت جميع اختبارات الوحدة الـ16 ونجحت.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const unitTests = getTests.filter(test => {
      return !!test.context.match(/Unit Tests/gi);
    });
    assert.isAtLeast(unitTests.length, 16, 'At least 16 tests passed');
    unitTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

اكتملت جميع الاختبارات الوظيفية الخمسة ونجحت.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const functTests = getTests.filter(test => {
      return !!test.context.match(/Functional Tests/gi);
    });
    assert.isAtLeast(functTests.length, 5, 'At least 5 tests passed');
    functTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

