---
id: 663d2f8056542a8066fd6cc6
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이제 라운드 결과를 얻을 차례입니다. `getRoundResults` 함수를 완성하세요.

플레이어가 라운드에서 이기면 `playerScore`을 `1`만큼 업데이트하고 메시지 `"Player wins! [player's choice] beats [computer's choice]"`을 반환하세요.

컴퓨터와 플레이어가 같은 선택을 하면 메시지 `"It's a tie! Both chose [player's choice]"`를 반환하세요.

컴퓨터가 라운드에서 이기면 `computerScore`를 `1`만큼 업데이트하고 메시지 `"Computer wins! [computer's choice] beats [player's choice]"`을 반환하세요.

`[computer's choice]`은 `computerResult`로 교체되어야 하며 `[player's choice]`은 `userOption`로 교체되어야 합니다.

**팁**

- 플레이어가 라운드에서 이겼는지 확인하려면 `hasPlayerWonTheRound` 함수를 사용할 수 있습니다.
- 메시지를 만들 때 템플릿 리터럴이나 일반 문자열 연결을 사용할 수 있습니다.


# --hints--

`getRoundResults`은 문자열을 반환해야 합니다.

```js
assert.isString(getRoundResults("Rock"));
```

`getRoundResults` 함수는 누가 라운드에서 이겼는지에 따라 올바른 메시지를 반환해야 합니다. 아무도 이기지 못하면 무승부 메시지를 반환해야 합니다.

```js
const result = getRoundResults("Paper")
if(result.startsWith("Player wins!")){
  assert.strictEqual(result, "Player wins! Paper beats Rock")
  assert.isAtLeast(playerScore, 1)
} else if(result.startsWith("Computer wins!")) {
  assert.strictEqual(result, "Computer wins! Scissors beats Paper")
  assert.isAtLeast(computerScore, 1)
} else if(result.startsWith("It's a tie!")){
  assert.strictEqual(result, "It's a tie! Both chose Paper")
}
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Rock, Paper, Scissors game</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Let's play Rock, Paper, Scissors!</h1>
    <main>
      <details class="rules-container">
        <summary>Rules to the game</summary>

        <p>You will be playing against the computer.</p>
        <p>You can choose between Rock, Paper, and Scissors.</p>
        <p>The first one to three points wins.</p>

        <p>Here are the rules to getting a point in the game:</p>
        <ul>
          <li>Rock beats Scissors</li>
          <li>Scissors beats Paper</li>
          <li>Paper beats Rock</li>
        </ul>
        <p>
          If the player and computer choose the same option (Ex. Paper and
          Paper), then no one gets the point.
        </p>
      </details>

      <div class="score-container">
        <strong
          >Player Score: <span class="score" id="player-score">0</span></strong
        >
        <strong
          >Computer Score:
          <span class="score" id="computer-score">0</span></strong
        >
      </div>

      <section class="options-container">
        <h2>Choose an option:</h2>
        <div class="btn-container">
          <button id="rock-btn" class="btn">Rock</button>
          <button id="paper-btn" class="btn">Paper</button>
          <button id="scissors-btn" class="btn">Scissors</button>
        </div>
      </section>

      <div class="results-container">
        <p id="results-msg"></p>
        <p id="winner-msg"></p>
        <button class="btn" id="reset-game-btn">Play again?</button>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --very-dark-blue: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--very-dark-blue);
  text-align: center;
  color: var(--white);
}

h1 {
  margin: 15px 0 20px;
}

.btn {
  cursor: pointer;
  width: 100px;
  margin: 10px;
  color: var(--very-dark-blue);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

.rules-container {
  padding: 10px 0;
  margin: auto;
  border-radius: 15px;
  border: 5px solid var(--yellow);
  background-color: var(--white);
  color: var(--very-dark-blue);
}

.rules-container ul {
  list-style-type: none;
}

.rules-container p {
  margin: 10px 0;
}

@media (min-width: 760px) {
  .rules-container {
    width: 60%;
  }
}

.score-container {
  display: flex;
  justify-content: space-around;
  margin: 30px 0;
  font-size: 1.2rem;
}

.score {
  font-weight: 500;
}

.results-container {
  font-size: 1.3rem;
  margin: 15px 0;
}

#winner-msg {
  margin-top: 25px;
}

#reset-game-btn {
  display: none;
  margin: 20px auto;
}

```

```js
function getRandomComputerResult() {
  const options = ["Rock", "Paper", "Scissors"];
  const randomIndex = Math.floor(Math.random() * options.length);
  return options[randomIndex];
}

let playerScore = 0;
let computerScore = 0;

function hasPlayerWonTheRound(player, computer) {
  return (
    (player === "Rock" && computer === "Scissors") ||
    (player === "Scissors" && computer === "Paper") ||
    (player === "Paper" && computer === "Rock")
  );
}
--fcc-editable-region--
function getRoundResults(userOption) {
  const computerResult = getRandomComputerResult();
 
}

console.log(getRoundResults("Rock"));
console.log("Player Score: ", playerScore, "Computer Score: ", computerScore);
--fcc-editable-region--
```
