---
id: 587d7daf367417b2b2512b7f
title: Cambia il prototipo con un nuovo oggetto
challengeType: 1
forumTopicId: 301316
dashedName: change-the-prototype-to-a-new-object
---

# --description--

Finora hai aggiunto le proprietà a `prototype` una per una:

```js
Bird.prototype.numLegs = 2;
```

Questo diventa noioso dopo poche proprietà.

```js
Bird.prototype.eat = function() {
  console.log("nom nom nom");
}

Bird.prototype.describe = function() {
  console.log("My name is " + this.name);
}
```

Un modo più efficiente è impostare `prototype` su un nuovo oggetto che contiene già le proprietà. In questo modo, le proprietà vengono aggiunte tutte insieme:

```js
Bird.prototype = {
  numLegs: 2, 
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --instructions--

Aggiungi la proprietà `numLegs` e i due metodi `eat()` e `describe()` al `prototype` di `Dog` impostando `prototype` su un nuovo oggetto.

# --hints--

`Dog.prototype` dovrebbe essere impostato su un nuovo oggetto.

```js
assert(/Dog\.prototype\s*?=\s*?{/.test(__helpers.removeJSComments(code)));
```

`Dog.prototype` dovrebbe avere la proprietà `numLegs`.

```js
assert(Dog.prototype.numLegs !== undefined);
```

`Dog.prototype` dovrebbe avere il metodo `eat()`.

```js
assert(typeof Dog.prototype.eat === 'function');
```

`Dog.prototype` dovrebbe avere il metodo `describe()`.

```js
assert(typeof Dog.prototype.describe === 'function');
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype = {
  // Only change code below this line

};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
numLegs: 4,
  eat () {
    console.log('nom nom nom');
  },
  describe () {
    console.log('My name is ' + this.name);
  }
};
```
