---
id: 587d7db0367417b2b2512b83
title: "Usa l'ereditarietà per non ripeterti"
challengeType: 1
forumTopicId: 301334
dashedName: use-inheritance-so-you-dont-repeat-yourself
---

# --description--

Esiste un principio nella programmazione chiamato <dfn>Don't Repeat Yourself (DRY)</dfn>. Il motivo per cui il codice ripetuto è un problema è che ogni modifica richiede di correggere il codice in più punti. Questo di solito significa più lavoro per i programmatori e più margine di errore.

Nota nell'esempio qui sotto che il metodo `describe` è condiviso da `Bird` e `Dog`:

```js
Bird.prototype = {
  constructor: Bird,
  describe: function() {
    console.log("My name is " + this.name);
  }
};

Dog.prototype = {
  constructor: Dog,
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

Il metodo `describe` è ripetuto in due posti. Il codice può essere modificato per seguire il principio DRY creando un `supertype` (o genitore) chiamato `Animal`:

```js
function Animal() { };

Animal.prototype = {
  constructor: Animal, 
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

Poiché `Animal` include il metodo `describe`, puoi rimuoverlo da `Bird` e `Dog`:

```js
Bird.prototype = {
  constructor: Bird
};

Dog.prototype = {
  constructor: Dog
};
```

# --instructions--

Il metodo `eat` è ripetuto sia in `Cat` che in `Bear`. Modifica il codice nello spirito del DRY spostando il metodo `eat` nel `Animal` `supertype`.

# --hints--

`Animal.prototype` dovrebbe avere la proprietà `eat`.

```js
assert(Animal.prototype.hasOwnProperty('eat'));
```

`Bear.prototype` non dovrebbe avere la proprietà `eat`.

```js
assert(!Bear.prototype.hasOwnProperty('eat'));
```

`Cat.prototype` non dovrebbe avere la proprietà `eat`.

```js
assert(!Cat.prototype.hasOwnProperty('eat'));
```

# --seed--

## --seed-contents--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,

};
```

# --solutions--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
```
