---
id: 587d7daf367417b2b2512b7e
title: Comprendere la proprietà constructor
challengeType: 1
forumTopicId: 301327
dashedName: understand-the-constructor-property
---

# --description--

C’è una proprietà speciale `constructor` situata sulle istanze di oggetto `duck` e `beagle` che sono state create nelle sfide precedenti:

```js
let duck = new Bird();
let beagle = new Dog();

console.log(duck.constructor === Bird); 
console.log(beagle.constructor === Dog);
```

Entrambe queste chiamate `console.log` mostrerebbero `true` nella console.

Nota che la proprietà `constructor` è un riferimento alla funzione costruttore che ha creato l’istanza. Il vantaggio della proprietà `constructor` è che è possibile controllare questa proprietà per scoprire che tipo di oggetto è. Ecco un esempio di come potrebbe essere usata:

```js
function joinBirdFraternity(candidate) {
  if (candidate.constructor === Bird) {
    return true;
  } else {
    return false;
  }
}
```

**Nota:** Poiché la proprietà `constructor` può essere sovrascritta (argomento che tratteremo nelle prossime due sfide), è generalmente meglio usare il metodo `instanceof` per controllare il tipo di un oggetto.

# --instructions--

Scrivi una funzione `joinDogFraternity` che prende un parametro `candidate` e, usando la proprietà `constructor`, restituisce `true` se il candidato è un `Dog`, altrimenti restituisce `false`.

# --hints--

`joinDogFraternity` dovrebbe essere definita come funzione.

```js
assert(typeof joinDogFraternity === 'function');
```

`joinDogFraternity` dovrebbe restituire `true` se `candidate` è un’istanza di `Dog`.

```js
assert(joinDogFraternity(new Dog('')) === true);
```

`joinDogFraternity` dovrebbe usare la proprietà `constructor`.

```js
assert(/\.constructor/.test(__helpers.removeJSComments(code)) && !/instanceof/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
function joinDogFraternity(candidate) {

}
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
function joinDogFraternity(candidate) {
  return candidate.constructor === Dog;
}
```
