---
id: 587d7dae367417b2b2512b7c
title: usa le proprietà del prototype per ridurre il codice duplicato
challengeType: 1
forumTopicId: 301336
dashedName: use-prototype-properties-to-reduce-duplicate-code
---

# --description--

Dato che `numLegs` probabilmente avrà lo stesso valore per tutte le istanze di `Bird`, hai essenzialmente una variabile `numLegs` duplicata dentro ogni istanza di `Bird`.

Questo potrebbe non essere un problema quando ci sono solo due istanze, ma immagina se ce ne fossero milioni. Sarebbe un sacco di variabili duplicate.

Un modo migliore è usare il `prototype` di `Bird`. Le proprietà nel `prototype` sono condivise tra TUTTE le istanze di `Bird`. Ecco come aggiungere `numLegs` al `Bird prototype`:

```js
Bird.prototype.numLegs = 2;
```

Ora tutte le istanze di `Bird` hanno la proprietà `numLegs`.

```js
console.log(duck.numLegs);
console.log(canary.numLegs);
```

Dato che tutte le istanze hanno automaticamente le proprietà sul `prototype`, pensa a un `prototype` come a una "ricetta" per creare oggetti. Nota che il `prototype` per `duck` e `canary` fa parte del costruttore `Bird` come `Bird.prototype`.

# --instructions--

Aggiungi una proprietà `numLegs` al `prototype` di `Dog`

# --hints--

`beagle` dovrebbe avere una proprietà `numLegs`.

```js
assert(beagle.numLegs !== undefined);
```

`beagle.numLegs` dovrebbe essere un numero.

```js
assert(typeof beagle.numLegs === 'number');
```

`numLegs` dovrebbe essere una proprietà `prototype` non una proprietà propria.

```js
assert(beagle.hasOwnProperty('numLegs') === false);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}



// Only change code above this line
let beagle = new Dog("Snoopy");
```

# --solutions--

```js
function Dog (name) {
  this.name = name;
}
Dog.prototype.numLegs = 4;
let beagle = new Dog("Snoopy");
```
