---
id: 587d7db2367417b2b2512b8b
title: "Capire l'espressione di funzione immediatamente invocata (IIFE)"
challengeType: 1
forumTopicId: 301328
dashedName: understand-the-immediately-invoked-function-expression-iife
---

# --description--

Un pattern comune in JavaScript è eseguire una funzione non appena viene dichiarata:

```js
(function () {
  console.log("Chirp, chirp!");
})();
```

Questa è una espressione di funzione anonima che si esegue subito e produce `Chirp, chirp!` immediatamente.

Nota che la funzione non ha nome e non è archiviata in una variabile. Le due parentesi tonde () alla fine dell'espressione della funzione fanno sì che venga eseguita o invocata immediatamente. Questo pattern è noto come <dfn>espressione di funzione immediatamente invocata</dfn> o <dfn>IIFE</dfn>.

# --instructions--

Riscrivi la funzione `makeNest` e rimuovi la sua chiamata in modo che diventi invece una funzione anonima immediatamente invocata (IIFE).

# --hints--

La funzione dovrebbe essere anonima.

```js
assert(/\((function|\(\))(=>|\(\)){?/.test(__helpers.removeJSComments(code).replace(/\s/g, '')));
```

La tua funzione dovrebbe avere parentesi tonde alla fine dell'espressione per chiamarla immediatamente.

```js
assert(/\(.*(\)\(|\}\(\))\)/.test(__helpers.removeJSComments(code).replace(/[\s;]/g, '')));
```

# --seed--

## --seed-contents--

```js
function makeNest() {
  console.log("A cozy nest is ready");
}

makeNest();
```

# --solutions--

```js
(function () {
  console.log("A cozy nest is ready");
})();
```

---

```js
(function () {
  console.log("A cozy nest is ready");
}());
```

---

```js
(() => {
  console.log("A cozy nest is ready");
})();
```

---

```js
(() =>
  console.log("A cozy nest is ready")
)();
```
