---
id: 587d7db1367417b2b2512b87
title: "aggiungere metodi dopo l'ereditarietà"
challengeType: 1
forumTopicId: 301315
dashedName: add-methods-after-inheritance
---

# --description--

Una funzione costruttrice che eredita il suo `prototype` oggetto da una funzione costruttrice supertipo può comunque avere i propri metodi oltre a quelli ereditati.

Per esempio, `Bird` è un costruttore che eredita il suo `prototype` da `Animal`:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
Bird.prototype.constructor = Bird;
```

Oltre a ciò che è ereditato da `Animal`, vuoi aggiungere un comportamento unico agli oggetti `Bird`. Qui, `Bird` otterrà una funzione `fly()`. Le funzioni si aggiungono al `Bird's` `prototype` nello stesso modo di qualsiasi funzione costruttrice:

```js
Bird.prototype.fly = function() {
  console.log("I'm flying!");
};
```

Ora le istanze di `Bird` avranno sia i metodi `eat()` sia `fly()`:

```js
let duck = new Bird();
duck.eat();
duck.fly();
```

`duck.eat()` mostrerebbe la stringa `nom nom nom` nella console, e `duck.fly()` mostrerebbe la stringa `I'm flying!`.

# --instructions--

Aggiungi tutto il codice necessario affinché l'oggetto `Dog` erediti da `Animal` e il costruttore `Dog` di `prototype` sia impostato su `Dog`. Poi aggiungi un metodo `bark()` all'oggetto `Dog` in modo che `beagle` possa sia `eat()` sia `bark()`. Il metodo `bark()` dovrebbe stampare `Woof!` nella console.

# --hints--

`Animal` non dovrebbe rispondere al metodo `bark()`.

```js
assert(typeof Animal.prototype.bark == 'undefined');
```

`Dog` dovrebbe ereditare il metodo `eat()` da `Animal`.

```js
assert(typeof Dog.prototype.eat == 'function');
```

Il prototipo di `Dog` dovrebbe avere un metodo `bark()`.

```js
assert('bark' in Dog.prototype);
```

`beagle` dovrebbe essere un `instanceof` `Animal`.

```js
assert(beagle instanceof Animal);
```

Il costruttore di `beagle` dovrebbe essere impostato su `Dog`.

```js
assert(beagle.constructor === Dog);
```

`beagle.eat()` dovrebbe stampare la stringa `nom nom nom`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.eat();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'nom nom nom');
```

`beagle.bark()` dovrebbe stampare la stringa `Woof!`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.bark();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'Woof!');
```

# --seed--

## --seed-contents--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }

// Only change code below this line




// Only change code above this line

let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Dog.prototype.bark = function () {
  console.log('Woof!');
};
let beagle = new Dog();

beagle.eat();
beagle.bark();
```
