---
id: 587d7db0367417b2b2512b82
title: Comprendere la catena del prototipo
challengeType: 1
forumTopicId: 301329
dashedName: understand-the-prototype-chain
---

# --description--

Tutti gli oggetti in JavaScript (con poche eccezioni) hanno un `prototype`. Inoltre, il `prototype` di un oggetto è esso stesso un oggetto.

```js
function Bird(name) {
  this.name = name;
}

typeof Bird.prototype;
```

Poiché un `prototype` è un oggetto, un `prototype` può avere il proprio `prototype`! In questo caso, il `prototype` di `Bird.prototype` è `Object.prototype`:

```js
Object.prototype.isPrototypeOf(Bird.prototype);
```

A cosa serve questo? Potresti ricordare il metodo `hasOwnProperty` da una sfida precedente:

```js
let duck = new Bird("Donald");
duck.hasOwnProperty("name");
```

Il metodo `hasOwnProperty` è definito in `Object.prototype`, che può essere consultato da `Bird.prototype`, che a sua volta può essere consultato da `duck`. Questo è un esempio della catena `prototype`. In questa catena `prototype`, `Bird` è il `supertype` per `duck`, mentre `duck` è il `subtype`. `Object` è un `supertype` sia per `Bird` sia per `duck`. `Object` è un `supertype` per tutti gli oggetti in JavaScript. Perciò, qualsiasi oggetto può usare il metodo `hasOwnProperty`.

# --instructions--

Modifica il codice per mostrare la catena del prototipo corretta.

# --hints--

Il tuo codice dovrebbe mostrare che `Object.prototype` è il prototipo di `Dog.prototype`

```js
assert(/Object\.prototype\.isPrototypeOf/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

Dog.prototype.isPrototypeOf(beagle);  // yields true

// Fix the code below so that it evaluates to true
???.isPrototypeOf(Dog.prototype);
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
Object.prototype.isPrototypeOf(Dog.prototype);
```
