---
id: 587d7dae367417b2b2512b7a
title: Verifica il costruttore di un oggetto con instanceof
challengeType: 1
forumTopicId: 301337
dashedName: verify-an-objects-constructor-with-instanceof
---

# --description--

Ogni volta che una funzione costruttrice crea un nuovo oggetto, si dice che quell'oggetto sia un’<dfn>istanza</dfn> del suo costruttore. JavaScript offre un modo comodo per verificarlo con l’operatore `instanceof`. `instanceof` ti permette di confrontare un oggetto con un costruttore, restituendo `true` o `false` a seconda che quell’oggetto sia stato creato con quel costruttore oppure no. Ecco un esempio:

```js
let Bird = function(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}

let crow = new Bird("Alexis", "black");

crow instanceof Bird;
```

Questo metodo `instanceof` restituirebbe `true`.

Se un oggetto viene creato senza usare un costruttore, `instanceof` verificherà che non sia un’istanza di quel costruttore:

```js
let canary = {
  name: "Mildred",
  color: "Yellow",
  numLegs: 2
};

canary instanceof Bird;
```

Questo metodo `instanceof` restituirebbe `false`.

# --instructions--

Crea una nuova istanza del costruttore `House`, chiamandola `myHouse` e passando un numero di camere da letto. Poi, usa `instanceof` per verificare che sia un’istanza di `House`.

# --hints--

`myHouse` dovrebbe avere un attributo `numBedrooms` impostato su un numero.

```js
assert(typeof myHouse.numBedrooms === 'number');
```

Dovresti verificare che `myHouse` sia un’istanza di `House` usando l’operatore `instanceof`.

```js
assert(/myHouse\s*instanceof\s*House/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}

// Only change code below this line
```

# --solutions--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}
const myHouse = new House(4);
console.log(myHouse instanceof House);
```
