---
id: 587d7dac367417b2b2512b74
title: Usa la notazione a punto per consultare le proprietà di un oggetto
challengeType: 1
forumTopicId: 301333
dashedName: use-dot-notation-to-access-the-properties-of-an-object
---

# --description--

L'ultima sfida ha creato un oggetto con varie proprietà. Ora vedrai come consultare i valori di quelle proprietà. Ecco un esempio:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
console.log(duck.name);
```

La notazione a punto si usa sul nome dell'oggetto, `duck`, seguito dal nome della proprietà, `name`, per consultare il valore di `Aflac`.

# --instructions--

Stampa entrambe le proprietà dell'oggetto `dog` nella tua console.

# --hints--

Il tuo codice dovrebbe usare `console.log` per stampare il valore della proprietà `name` dell'oggetto `dog`.

```js
assert(/console.log\(.*dog\.name.*\)/g.test(__helpers.removeJSComments(code)));
```

Il tuo codice dovrebbe usare `console.log` per stampare il valore della proprietà `numLegs` dell'oggetto `dog`.

```js
assert(/console.log\(.*dog\.numLegs.*\)/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
// Only change code below this line
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
console.log(dog.name);
console.log(dog.numLegs);
```
