---
id: 587d7daf367417b2b2512b80
title: Ricorda di impostare la proprietà constructor quando cambi il prototype
challengeType: 1
forumTopicId: 301323
dashedName: remember-to-set-the-constructor-property-when-changing-the-prototype
---

# --description--

C’è un effetto collaterale cruciale nel impostare manualmente il prototype su un nuovo oggetto. Viene cancellata la proprietà `constructor`! Questa proprietà può essere usata per verificare quale funzione costruttore ha creato l’istanza, ma dato che la proprietà è stata sovrascritta, ora fornisce risultati falsi:

```js
duck.constructor === Bird;
duck.constructor === Object;
duck instanceof Bird;
```

In ordine, queste espressioni valuterebbero `false`, `true` e `true`.

Per risolvere questo, ogni volta che un prototype viene impostato manualmente su un nuovo oggetto, ricorda di definire la proprietà `constructor`:

```js
Bird.prototype = {
  constructor: Bird,
  numLegs: 2,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name); 
  }
};
```

# --instructions--

Definisci la proprietà `constructor` sul `Dog` `prototype`.

# --hints--

`Dog.prototype` dovrebbe impostare la proprietà `constructor`.

```js
assert(Dog.prototype.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
Dog.prototype = {

  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
  constructor: Dog,
  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```
