---
id: 66b62d0ad68488dd76228d6c
title: Eseguire il debug di un cambiacolori di sfondo casuale
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot sta imparando JavaScript e ha provato a creare un proprio cambiacolori di sfondo casuale. Tuttavia, ha commesso alcuni errori lungo il percorso.

**Obiettivo:** Soddisfa le storie utente qui sotto e fai passare tutti i test affinché il laboratorio funzioni correttamente.

**Storie utente:**

1. Dovresti correggere gli errori di sintassi nell'array `darkColorsArr`.
1. Dovresti correggere il `ReferenceError` lanciato quando si fa riferimento all'oggetto `math` all'interno della funzione `getRandomIndex`.
1. Dovresti arrotondare `randomIndex` per difetto al numero intero più vicino.
1. Dovresti aggiornare il valore della variabile `body` per usare il nome del metodo corretto e correggere il `TypeError` correlato.
1. Dovresti correggere `bgHexCodeSpanElement` in modo che selezioni correttamente l'elemento con l'id `id` di `bg-hex-code`.
1. Dovresti modificare la variabile `color` all'interno di `changeBackgroundColor` in modo che sia impostata su un colore casuale dall'array `darkColorsArr`.
1. Dovresti aggiornare la variabile `btn` in modo che punti all'elemento button. Apri `index.html` per vedere il nome corretto del `id` per quell'elemento button.
1. Dovresti correggere la riga che aggiunge un listener per l'evento `click` a `btn` passando un riferimento alla funzione `changeBackgroundColor` al metodo.
1. Il tuo sfondo dovrebbe cambiare casualmente in uno dei colori di `darkColorsArr` ogni volta che si clicca su `#btn`.

# --hints--

Dovresti risolvere gli errori di sintassi nell'array `darkColorsArr`. Ricorda che gli elementi dell'array devono essere separati da virgole.

```js
assert.isArray(darkColorsArr);
```

Dovresti correggere l'errore di maiuscole nella riga `math.random()`.

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

Dovresti arrotondare `darkColorsArr.length * Math.random()` per difetto al numero intero più vicino.

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

Dovresti aggiornare il valore della tua variabile `body` in modo che `queryselector` usi il camel case.

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

Dovresti correggere `"bg-hex-code"` in modo che faccia riferimento correttamente all'id `id` di `bg-hex-code`. Ricorda che i selettori `id` in CSS iniziano con il simbolo cancelletto (`#`).

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

Dovresti chiamare la funzione `getRandomIndex` all'interno della riga `darkColorsArr[getRandomIndex]`.

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

Dovresti correggere il nome `id` della riga `"#click-btn"` per farlo corrispondere al nome `id` corretto nel file `index.html`.

```js
assert.match(btn?.id, /btn/);
```

Non dovresti chiamare la funzione `changeBackgroundColor`. Invece, dovresti passare il riferimento alla funzione al metodo `addEventListener`.

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
