---
id: 69a04466917c66b132323a1b
title: 第 18 步
challengeType: 1
dashedName: step-18
---

# --description--

如果 `previous` 不是 `null`，这意味着要移除的元素不在头部，而是在列表中的某个位置，因此应该通过将前一个节点直接链接到下一个节点来跳过当前节点。

为此，创建另一个 `if` 语句，检查 previous 是否严格不等于 `null`。如果是这样，设置 `previous.next` 为 `current.next`。

# --hints--

你应该创建另一个条件为 `if` 的 `previous !== null` 语句。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

你应该将 `previous.next` 设置为 `current.next`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
