---
id: 69a03ddce82b0626f329b001
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

现在，创建一个 `while` 循环，用来查看 `current` 是否不严格 `null` 且 `current.element` 不严格相等于元素。

这意味着循环应继续遍历列表，直到到达列表末尾或找到要删除的元素。

# --hints--

你应该创建一个条件为 `while` 的 `current !== null && current.element !== element` 循环。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /while\s+\(current\s*!==\s*null\s*&&\s*current\.element\s*!==\s*element\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    --fcc-editable-region--
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
