---
id: 69a013a86c97dac3784b2636
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

现在，日志应该在控制台显示类似 `{ head: { element: 42, next: null }, length: 1 }` 的内容。

但是，这不是处理添加的节点是列表中第一个元素（即头节点）时的正确方法。

要正确处理，使用 if 语句通过调用 `isEmpty` 函数并传入 `list` 来查看列表是否为空。如果为真，则将你的 `list.head = node` 语句移动到 `if` 语句的体内。

# --hints--

你应该有一个条件为 `if` 的 `isEmpty(list)` 语句。

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

你应该将你的 `list.head = node` 语句移动到 `if` 语句内。

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
