---
id: 69a0102e761a245bb1cbd8e7
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

此时，你应该测试链表以便更好地理解。

创建一个 `myList` 变量并将其设置为 `initList()`。在其下方，使用 `isEmpty()` 作为参数，将 `myList` 的调用记录到控制台。然后，在那之后，使用 `add` 和你选择的一个数字调用 `myList`，然后将 `myList` 记录到控制台。在其下方，再次使用 `isEmpty()` 作为参数，将 `myList` 的调用记录到控制台。

# --hints--

你应该创建一个 `myList` 变量并设置为 `initList()`。

```js
assert.deepEqual(myList, initList());
```

你应该将 `isEmpty(myList)` 的调用记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\)/)
```

你应该使用 `add` 和你选择的一个数字作为参数调用 `myList`。

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

你应该将你的 `myList` 变量记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\)/)
```

你应该再次将 `isEmpty(myList)` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\)\s*console\.log\(isEmpty\(myList\)\)/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
