---
id: 69a00f673f171e29159e3ad8
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

这个节点需要两个属性，一个是要添加的元素，另一个是指向下一个节点的指针。

在 add 函数内部，创建一个名为 `node` 的对象，它应该有一个名为 `element` 的属性，以及一个 `next` 属性，设置为 `null`，这就是指向列表下一个元素的指针。

# --hints--

你应该在 `node` 函数内部创建一个名为 `add` 的对象。

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

你的 `node` 对象应该有一个名为 `element` 的属性，设置为 `element`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

你的 `node` 对象应该有一个名为 `next` 的属性，设置为 `null`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
