---
id: 69a045fd236d0d973a59f103
title: 第 19 步
challengeType: 1
dashedName: step-19
---

# --description--

如果 `previous` 是 `null`，这意味着要删除的元素位于列表头部，因此你应该更新 `head` 指向下一个节点，以便删除第一个元素。

为此，添加一个 `else` 子句。在其中，将 `list.head` 设置为 `current.next`。

# --hints--

你应该为你的 `else` 语句添加一个 `if` 子句。

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

你应该在 `list.head` 子句中将 `current.next` 设置为 `else`。

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
