---
id: 69a00eaa43458e618331873c
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

现在你将编写一个函数，用于将节点添加到列表中。

创建一个带有 `list` 和 `element` 参数的 `add` 函数。

# --hints--

你应该创建一个 `add` 函数。

```js
assert.isFunction(add)
```

你的 `add` 函数应包含 `list` 和 `element` 参数。

```js
const addParams = __helpers.getFunctionParams(add.toString());
const addParamsNames = addParams.map(param => param.name)
assert.include(addParamsNames, "list");
assert.include(addParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
