---
id: 69a00cd4ec513345d40e95ec
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

接下来你要完成的是一个函数，用来查看列表是否为空。

为此，创建一个名为 `isEmpty` 的函数，它接受一个 `list` 参数。在函数内部，返回一个检查，判断 `length` 的 `list` 属性是否严格相等于 `0`。

# --hints--

你应该创建一个 `isEmpty` 函数。

```js
assert.isFunction(isEmpty)
```

你的 `isEmpty` 函数应该有一个 `list` 参数。

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

你的 `isEmpty` 函数应该返回 `list.length === 0`。

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
