---
id: 69a00a051b9a18694d00c16b
title: 第一步
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将实现一个链表。

首先，创建一个名为 `initList` 的函数。在函数内部，返回一个对象，该对象有一个 `head` 属性，值为 `null`，还有一个 `length` 属性，值为 `0`。

这是你用来初始化链表的函数。`head` 表现为列表中的第一个元素，初始时它是 `null`，因为列表中没有元素。`length` 用来跟踪列表中节点（元素）的数量，当你添加或删除节点时，需要对它进行递增或递减。 

# --hints--

你应该创建一个 `initList` 函数。

```js
assert.isFunction(initList)
```

你的 `initList` 函数应该返回一个对象。

```js
assert.isObject(initList());
```

你的 `initList` 函数应该返回一个对象，该对象有一个 `head` 属性，值为 `null`。

```js
assert.propertyVal(initList(), "head", null);
```

你的 `initList` 函数应该返回一个对象，该对象有一个 `length` 属性，值为 `0`。

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
