---
id: 69a03aaa1e6fe8152788ce6a
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

在 `remove` 函数内部，使用 `let` 创建一个 `previous` 变量并赋值为 `null`，再使用 `let` 创建一个 `current` 变量并赋值为 `list.head`。这些变量将在遍历列表的循环中使用。

`current` 跟踪当前正在查看的列表节点，从头节点开始，每次向前移动一个节点。另一方面，`previous` 跟踪当前节点之前的那个节点。

这在删除节点时非常重要，因为你必须更新前一个节点的 next 值，使其指向被删除节点之后的节点，以保持列表的连接。

# --hints--

你应该在 `let` 函数内部使用 `previous` 创建一个 `null` 变量并赋值为 `remove`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

你应该在 `let` 函数内部使用 `current` 创建一个 `list.head` 变量并赋值为 `remove`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
