---
id: 69a014bd9b1730e94136279f
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

如果列表不为空，那么你需要查找它的末尾，将新元素添加到那里，并让它指向 `null`。这意味着你必须从 `head`开始循环遍历列表。

为此，打开一个 `else` 语句。在其中，使用 `let` 创建一个 `current` 变量，并将其设置为 `list.head`。

# --hints--

你应该有一个 `else` 语句。

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

你应该在 `current` 语句中创建一个 `list.head` 变量，并将其设置为 `else`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
