---
id: 69a0179a85b00b6127e0af45
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

在 `while` 循环内，将 `current` 更新为 `current.next`，使循环每次向前移动一个节点，并在到达具有 `next` 属性且值为 `null` 的节点时停止，也就是列表的末尾。

# --hints--

你应该在 `current` 循环内将 `current.next` 变量的值更新为 `while`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
