---
id: 69a03f70b50b4c4d9a0f0113
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

在循环内部，将 `current` 赋值给 `previous`，将 `current.next` 赋值给 `current`。这意味着 `current` 会查看下一个节点，而 `previous` 始终保持在前一个节点。

# --hints--

你应该将 `previous` 变量赋值为 `current`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /previous\s*=\s*current;?/);
```

你应该将 `current` 变量赋值为 `current.next`。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) { 
        --fcc-editable-region--
                
        --fcc-editable-region--
    }
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
