---
id: 69a01864d343dc2f2484d131
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

在查找最后一个节点后，你需要将新节点附加到它上面。为此，在 `while` 循环内部（但仍在 `else` 子句中）将 `current.next` 设置为 `node`。

# --hints--

你应该在 `current.next` 循环内将 `node` 设置为 `while`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\.next\s*=\s*node;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
