---
id: 69a016ac6edea7d8c8020c98
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

现在，创建一个 `while` 循环，条件是 `current.next !== null`。这意味着只要当前节点后面还有另一个节点，循环就应该继续运行。

# --hints--

你应该创建一个条件为 `while` 的 `current.next !== null` 循环。

```js
assert.match(__helpers.removeJSComments(add.toString()), /while\s+\(current\.next\s*!==\s*null\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
