---
id: 69a0102e761a245bb1cbd8e7
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

此时，你应该测试链表以便更好地理解。

创建一个名为 `myList` 的变量，并将其赋值为 `initList()`。在其下方，将调用 `isEmpty()` 并以 `myList` 作为参数的结果记录到控制台。接着，调用 `add` 方法，传入 `myList` 和一个你选择的数字，然后将 `myList` 记录到控制台。在其下方，再次调用 `isEmpty()` 并以 `myList` 作为参数，将结果记录到控制台。

# --hints--

你应该创建一个 `myList` 变量并设置为 `initList()`。

```js
assert.deepEqual(myList, initList());
```

你应该将 `isEmpty(myList)` 的调用记录到控制台。

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(isEmpty(myList))');
```

你应该使用 `add` 和你选择的一个数字作为参数调用 `myList`。

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.match(strippedCode, /add\(myList,\d+\)/);
```

你应该将你的 `myList` 变量记录到控制台。

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(myList)');
```

你应该再次将 `isEmpty(myList)` 记录到控制台。

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
const regexMatch = strippedCode.matchAll(/console\.log\(isEmpty\(myList\)\)/g);
assert.isAtLeast([...regexMatch].length, 2);
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
