---
id: 69a012c6224fac85832247dd
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

注意控制台中的 `true`、`{ head: null, length: 0 }` 和 `true` 结果吗？这是因为你还没有为列表 `myList` 赋值，也没有对 `length` 属性进行递增。

继续将 `node` 赋值给 `list.head`，然后递增 `length` 的 `list` 属性。

# --hints--

你应该将 `node` 赋值给 `list.head`。

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.head\s*=\s*node;?/);
```

你应该递增 `length` 的 `list` 属性。

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.length\+\+;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
