---
id: 66a83bdcf425e7446900b7c4
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Nelle lezioni precedenti, hai imparato a lavorare con l'elemento `form` in questo modo:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

L'attributo `action` serve a specificare dove devono essere inviati i dati del modulo quando il modulo viene inviato.

L'attributo `method` serve a specificare il metodo HTTP da usare per inviare i dati del modulo. I metodi più comuni sono `GET` e `POST`.

**NOTA**: Imparerai come funzionano i metodi HTTP in moduli successivi.

All'interno del tuo elemento `main`, aggiungi un elemento `form` con un attributo `action` impostato su `"https://hotel-feedback.freecodecamp.org"` e un attributo `method` impostato su `"POST"`.

# --hints--

Dovresti avere un tag di apertura `form`.

```js
assert.match(code, /<form.*>/i);
```

Dovresti avere un tag di chiusura `form`.

```js
assert.match(code, /<\/form>/i);
```

Il tuo elemento `form` dovrebbe avere un attributo `action` impostato su `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Il tuo elemento `form` dovrebbe avere un attributo `method` impostato su `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
