---
id: 66a83fec026a7a4631e084d2
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Quando lavori con gli elementi `fieldset`, è comune usare un sottotitolo per descrivere il gruppo di input. Puoi usare l'elemento `legend` per questo.

Ecco un esempio di utilizzo di un elemento `legend`:

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

All'interno del tuo elemento `fieldset`, aggiungi un elemento `legend` con il testo `Personal Information`.

# --hints--

Dovresti avere un tag di apertura `legend`.

```js
assert.match(code, /<legend>/);
```

Dovresti avere un tag di chiusura `legend`.

```js
assert.match(code, /<\/legend>/);
```

Il tuo elemento `legend` dovrebbe essere all'interno dell'elemento `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

Il tuo elemento `legend` dovrebbe contenere il testo `Personal Information`.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
