---
id: 66a8380d911e3f4270d5cadc
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Ora è il momento di aggiungere l'elemento `main` che rappresenta il contenuto principale della pagina.

# --hints--

Dovresti avere un tag `main` di apertura.

```js
assert.match(code, /<main>/i);
```

Dovresti avere un tag `main` di chiusura.

```js
assert.match(code, /<\/main>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
