---
id: 66a84111965a0c46df6bbd0a
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Nelle lezioni precedenti, hai imparato come associare un elemento `label` a un `input` in questo modo:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

L'attributo `for` sull'elemento `label` dovrebbe corrispondere all'attributo `id` sull'elemento `input`. Questa è nota come associazione esplicita.

Sotto il tuo elemento `legend`, aggiungi un elemento `label` con il testo `Name (required):`. Imposta il suo attributo `for` al valore `"full-name"`.

Poi, sotto il tuo elemento `label`, aggiungi un elemento `input` senza attributi. Nei passaggi successivi aggiungerai gli attributi necessari.

# --hints--

Dovresti avere un tag di apertura `label`.

```js
assert.match(code, /<label.*>/);
```

Dovresti avere un tag di chiusura `label`.

```js
assert.match(code, /<\/label>/);
```

Il tuo elemento `label` dovrebbe essere posizionato sotto il tuo `legend` all'interno dell'elemento `fieldset`.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

Il tuo elemento `label` dovrebbe avere un attributo `for` con il valore `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

Il tuo elemento `label` dovrebbe avere il testo `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

Dovresti avere un elemento `input` sotto il tuo elemento `label`.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
