---
id: 66a83601cd819e37f0dccd14
title: Passaggio 1
challengeType: 0
dashedName: step-1
---

# --description--

In questo workshop, praticherai l'uso dei moduli HTML costruendo un modulo di feedback per un hotel.

Per il testo introduttivo, dovrai mostrare il titolo principale seguito da una breve nota sul lasciare un feedback.

All'interno del tuo elemento `body`, aggiungi un elemento `header`.

All'interno dell'elemento `header`, aggiungi un elemento `h1`. Il testo dell'elemento `h1` dovrebbe essere `Hotel Feedback Form`.

Sotto il tuo elemento `h1`, aggiungi un elemento `p`. Il testo dell'elemento `p` dovrebbe essere `Thank you for staying with us. Please provide feedback on your recent stay.`.

# --hints--

Dovresti avere un tag di apertura `header`.

```js
assert.match(code, /<header>/i);
```

Dovresti avere un tag di chiusura `header`.

```js
assert.match(code, /<\/header>/i);
```

Dovresti avere un tag di apertura `h1`.

```js
assert.match(code, /<h1>/i);
```

Dovresti avere un tag di chiusura `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Il tuo elemento `h1` dovrebbe avere il testo `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

Il tuo elemento `h1` dovrebbe essere all'interno dell'elemento `header`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

Dovresti avere un tag di apertura `p`.

```js
assert.match(code, /<p>/i);
```

Dovresti avere un tag di chiusura `p`.

```js
assert.match(code, /<\/p>/i);
```

Il tuo elemento paragrafo dovrebbe avere il testo `Thank you for staying with us. Please provide feedback on your recent stay.`.

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Il tuo elemento paragrafo dovrebbe essere all'interno del tuo `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

Il tuo elemento paragrafo dovrebbe venire dopo il tuo `h1`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
