---
id: 66a83e5e491625454b6f62c3
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

I moduli sono composti da `inputs` dove gli utenti possono inserire i loro dati. Puoi raggruppare input correlati insieme usando l'elemento `fieldset`.

Ecco un esempio di utilizzo di un elemento `fieldset`:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

All'interno del tuo elemento `form`, aggiungi un elemento `fieldset`.

# --hints--

Dovresti avere un tag di apertura `fieldset`.

```js
assert.match(code, /<fieldset>/i);
```

Dovresti avere un tag di chiusura `fieldset`.

```js
assert.match(code, /<\/fieldset>/i);
```

Il tuo elemento `fieldset` dovrebbe essere all'interno dell'elemento `form`.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
