---
id: 66a93c95bc58e26a8fe95818
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Successivamente, aggiungi un `input` con il tipo `"email"` sotto il tuo `label` email. Questo `input` dovrebbe avere un attributo `id` impostato al valore `"email"`. Inoltre, imposta l'attributo `name` al valore `"email"`. 

Questo `input` è anche richiesto, quindi assicurati di aggiungere l'attributo `required`.

Infine, aggiungi un attributo `placeholder` impostato su `"example@email.com"`.

# --hints--

Dovresti avere un elemento `input` sotto il tuo `label` email.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

Il tuo elemento `input` dovrebbe avere un `type` di `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

Il tuo elemento `input` dovrebbe avere un attributo `id` impostato su `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

Il tuo elemento `input` dovrebbe avere un attributo `name` impostato su `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

Il tuo elemento `input` dovrebbe avere un attributo `required`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

Il tuo elemento `input` dovrebbe avere un attributo `placeholder` impostato su `"example@email.com"`.
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
