---
id: 66a9419e2d18476c645ce693
title: Passaggio 14
challengeType: 0
dashedName: step-14
---

# --description--

Il numero `input` viene usato per creare un campo di input numerico.

Ecco un esempio di campo di input di tipo number:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

Gli attributi `min` e `max` servono a impostare i valori minimo e massimo che si possono inserire nel campo di input.

Sotto il tuo elemento `label`, aggiungi un input con l'attributo `type` impostato su `"number"` e un `id` di `"age"`.

L'attributo `name` dovrebbe essere impostato su `"age"`, l'attributo `min` dovrebbe essere impostato su `"3"` e l'attributo `max` dovrebbe essere impostato su `"100"`.

# --hints--

Dovresti avere un `input` di tipo number sotto il tuo elemento `label`.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

Il tuo `input` di tipo number dovrebbe avere un `id` di `age`.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

Il tuo `input` di tipo number dovrebbe avere un attributo `name` impostato su `age`.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

Il tuo `input` di tipo number dovrebbe avere un attributo `min` impostato su `"3"`.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

Il tuo `input` di tipo number dovrebbe avere un attributo `max` impostato su `"100"`.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
