---
id: 66a95d0eff8168747805f1f3
title: Passaggio 19
challengeType: 0
dashedName: step-19
---

# --description--

Quando vuoi che un utente selezioni più opzioni da una lista, puoi usare le caselle di spunta.

Ecco un esempio di come lavorare con le caselle di spunta per le opzioni di cibo:

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

L'attributo `value` serve a specificare il valore che verrà inviato al server quando il modulo viene inviato.

Sotto il tuo elemento `legend`, aggiungi una casella di spunta `input` con gli attributi `id` e `value` impostati su `"ads"`, e l'attributo `name` impostato su `"choice"`.

Sotto la tua casella di spunta `input`, aggiungi un elemento `label` con il testo `Social Media Ads`. L'attributo `for` dovrebbe essere impostato su `"ads"`.

# --hints--

Dovresti avere una casella di spunta `input` sotto il tuo elemento `legend`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

La tua casella di spunta `input` dovrebbe avere un attributo `id` impostato su `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

La tua casella di spunta `input` dovrebbe avere un attributo `name` impostato su `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "choice");
```

La tua casella di spunta `input` dovrebbe avere un attributo `value` impostato su `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

Dovresti avere un elemento `label` sotto la tua casella di spunta `input`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

Il tuo elemento `label` dovrebbe avere il testo `Social Media Ads`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

Il tuo elemento `label` dovrebbe avere un attributo `for` impostato su `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
