---
id: 66a937e74920ba68ebe5e86d
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Nelle lezioni precedenti, hai imparato a lavorare con gli attributi `placeholder` e `required` in questo modo:

```html
<input type="text" placeholder="e.g., John Doe" required>
```

Per il tuo elemento `input` esistente, aggiungi un attributo `placeholder` con il valore `"e.g., John Doe"`.

Inoltre, aggiungi l'attributo `required` all'elemento `input`.

# --hints--

Non dovresti rimuovere l'elemento `input` esistente.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Il tuo `input` dovrebbe avere un attributo `placeholder` con il valore `"e.g., John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

Il tuo `input` dovrebbe avere un attributo `required`.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
