---
id: 66a962954f4e0d76223b37ed
title: Passaggio 21
challengeType: 0
dashedName: step-21
---

# --description--

Successivamente, aggiungi un'altra casella di spunta `input` con gli attributi `id` e `value` impostati su `"location"`, e l'attributo `name` impostato su `"choice"`.

Per l'elemento `label`, il testo di `Location` e l'attributo `for` dovrebbero essere impostati su `"location"`.

Sotto quell'elemento `label`, aggiungi un'altra casella di spunta `input` con gli attributi `id` e `value` impostati su `"reputation"`, e l'attributo `name` impostato su `"choice"`.

Per l'elemento `label`, il testo di `Reputation` e l'attributo `for` dovrebbero essere impostati su `"reputation"`.

# --hints--

Dovresti avere una terza casella di spunta `input`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

La tua terza casella di spunta `input` dovrebbe avere un attributo `id` impostato su `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

La tua terza casella di spunta `input` dovrebbe avere un attributo `name` impostato su `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'choice');
```

La tua terza casella di spunta `input` dovrebbe avere un attributo `value` impostato su `"location"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

Dovresti avere un elemento `label` dopo la tua casella di spunta `input` con il testo `Location`.

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

Il tuo `label` dovrebbe avere l'attributo `for` impostato su `"location"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

Dovresti avere una quarta casella di spunta `input` sotto il tuo `label`.

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

Dovresti avere una quarta casella di spunta `input` con un attributo `id` impostato su `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

La tua quarta casella di spunta `input` dovrebbe avere un attributo `name` impostato su `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'choice');
```

La tua quarta casella di spunta `input` dovrebbe avere un attributo `value` impostato su `"reputation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

Dovresti avere un elemento `label` con il testo `Reputation` sotto la casella di spunta della reputazione.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

Il tuo `label` dovrebbe avere l'attributo `for` impostato su `"reputation"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
