---
id: 66a93730719e1f68410cce54
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

L'attributo `name` serve per identificare i dati del modulo dopo che sono stati inviati al server.

Ecco un esempio di come usare l'attributo `name`:

```html
<input type="email" name="email">
```

Aggiungi un attributo `name` all'elemento `input` con il valore `"name"`.

# --hints--

Non dovresti rimuovere l'elemento `input` esistente.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Non dovresti creare un nuovo elemento `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Il tuo elemento `input` dovrebbe avere un attributo `name` con il valore `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
