---
id: 66a930b20f589b6664c51cb0
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Quando un utente fornisce il proprio nome completo, il `input` accetterà testo semplice.

Nelle lezioni precedenti, hai imparato a lavorare con l'attributo `type` in questo modo:

```html
<input type="text">
```

Per il tuo elemento `input` esistente, aggiungi un attributo `type` impostato su `"text"`.

Inoltre, assegna all'elemento `input` un attributo `id` con il valore `"full-name"`.

# --hints--

Non dovresti rimuovere l'elemento `input` esistente.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Non dovresti creare un nuovo elemento `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Il tuo elemento `input` dovrebbe avere un attributo `type` con il valore `"text"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

Il tuo elemento `input` dovrebbe avere un attributo `id` con il valore `"full-name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
