---
id: 66a954b2bcddba72076c1857
title: Passaggio 17
challengeType: 0
dashedName: step-17
---

# --description--

Sotto il tuo elemento `label`, aggiungi un pulsante `radio` con l'insieme `id` impostato su `"no-option"`, l'attributo `name` impostato su `"hotel-stay"` e un attributo `value` impostato su `"no"`.

Sotto il tuo nuovo pulsante `radio`, aggiungi un altro elemento `label` con l'attributo `for` impostato su `"no-option"`. Il testo per il `label` dovrebbe essere `No`.

Quando hai finito, puoi provare i pulsanti di opzione selezionando un'opzione alla volta.

# --hints--

Dovresti avere un elemento `input` sotto il tuo primo elemento `label`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

Il tuo elemento `input` dovrebbe avere un attributo `type` con valore `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

Il tuo elemento `input` dovrebbe avere un attributo `id` con valore `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

Il tuo elemento `input` dovrebbe avere un attributo `name` con valore `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

Il tuo elemento input dovrebbe avere un attributo `value` impostato su `"no"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) label + input')?.value,
  'no'
);
```

Dovresti avere un secondo elemento `label` sotto il tuo secondo elemento `input`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

Il tuo elemento `label` dovrebbe avere un attributo `for` impostato su `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

Il tuo elemento `label` dovrebbe avere il testo `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
