---
id: 66a9521bc70162712caf118d
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Se vuoi che gli utenti selezionino una sola opzione da un elenco di opzioni, puoi usare un insieme di pulsanti di opzione.

Ecco un esempio di due pulsanti di opzione:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

In questo esempio, i pulsanti di opzione sono raggruppati insieme usando lo stesso valore dell'attributo `name`. Questo significa che può essere selezionato solo un pulsante di opzione alla volta.

Sotto il tuo elemento `legend`, aggiungi un pulsante `radio` con la proprietà `id` impostata su `"yes-option"`, l'attributo `name` impostato su `"hotel-stay"` e un attributo `value` impostato su `"yes"`.

Sotto il tuo pulsante `radio`, aggiungi un elemento `label` con il testo `Yes` e un attributo `for` impostato su `"yes-option"`.

# --hints--

Dovresti avere un pulsante `radio` sotto il tuo elemento `legend`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

Il tuo pulsante `radio` dovrebbe avere un attributo `id` impostato su `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

Il tuo pulsante `radio` dovrebbe avere un attributo `name` impostato su `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

Il tuo pulsante di opzione dovrebbe avere un attributo `value` impostato su `"yes"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

Dovresti avere un elemento `label` sotto il tuo pulsante `radio`.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

Il tuo elemento `label` dovrebbe avere un attributo `for` impostato su `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

Il tuo elemento `label` dovrebbe avere il testo `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
