---
id: 655cd899f8de09431eabb40c
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個專案中，你將透過建置合併排序演算法來學習資料結構。 

這是一個使用分而治之原則來排序資料群集的排序演算法。也就是說，它會將一個群集「分割」成較小的子零件，並透過獨立排序這些子零件來「征服」它們，然後合併已排序的子零件。

創建一個名為 `merge_sort` 的函式（程式）。這個函式（程式）將控制代碼來排序一個數字列表。 

在函式（程式）主體中使用 `pass` 關鍵字。

# --hints--

你應該宣告一個名為 `merge_sort` 的函式。不要忘記在函式主體中使用 `pass` 關鍵字。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
