---
id: 655cd899f8de09431eabb40c
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個項目中，你將通過構建合併排序算法學習數據結構。

這是一種使用分而治之原理對數據集合進行排序的排序算法。 也就是說，它將一個集合“劃分”爲更小的子部分，然後通過對子部分進行獨立排序來“征服”它們，然後合併已排序的子部分。

創建一個名爲 `merge_sort` 的函數。 此函數將處理排序數字列表的任務。

在函數體中使用 `pass` 關鍵字。

# --hints--

您應該聲明一個名爲 `merge_sort` 的函數。 別忘記在函數體中使用 `pass` 關鍵字。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
