---
id: 6564aee9c077774ea49c3faf
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

合併排序演算法主要執行三個動作：

- 將未排序的項目序列分割成子零件
- 對子零件中的項目進行排序
- 合併已排序的子零件

上述過程會遞迴發生，直到子零件合併成完整排序的序列。讓我們從將序列分割開始。

首先，將 `pass` 關鍵字取代為變數 `middle_point`。使用整數除法運算子（`//`）將 `array` 陣列的長度除以 2，並將結果指定給你的新變數 `middle_point`。記得縮排你的程式碼。

# --hints--

你不應該在函式（程式）的主體中使用 `pass`。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

你應該宣告一個 `middle_point` 變數並將 `len(array) // 2` 指定給它。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
