---
id: 6564a9fe51964c229d5b7f4c
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

你需要一個表示要排序的資料群集的參數。在 `merge_sort` 函式中創建一個名為 `array` 的參數。

# --hints--

你的 `merge_sort` 函式（程式）應該接受一個單一參數：`array`。

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
