---
id: 656627b47bd2d2a4afbc945d
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

在前一步驟中，你取得了中點。你可以使用它將 `array` 分成兩部分，並將每個部分指定給新的變數。 

使用切片語法擷取 `array` 的左半部，並將其指定到名為 `left_part` 的變數。


# --hints--

你應該在你的 `merge_sort` 函式中有一個名為 `left_part` 的變數。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

你應該將 `array[:middle_point]` 指定給 `left_part` 變數。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
