---
id: 656639771fed09b5c0e8fe71
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

現在你已經將 `array` 列表分成兩個獨立的列表，你會持續將每個列表分割，直到每個元素都獨立存在於自己的列表中。只有一個數字的列表永遠是已排序的。

為了做到這點，在你的函式中遞迴呼叫 `merge_sort`。目前不要傳入任何引數。

# --hints--

你應該在你的函式（程式）主體底部呼叫 `merge_sort`。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
