---
id: 656680b0fc79f2c38a34d90e
title: 步驟 10
challengeType: 20
dashedName: step-10
---

# --description--

現在是時候對列表（`left_part` 和 `right_part`）進行排序併合併到原始 `array` 中了。

你可以通過比較兩個列表中的元素並將較小的元素合併到主列表中來實現。 你將對 `left_part` 和 `right_part` 中的所有索引進行此比較。

創建三個變量：`left_array_index`、`right_array_index` 和 `sorted_index`，並將它們的值設置爲 `0`。 這些變量將幫助您在排序過程中跟蹤每個索引。

# --hints--

你應該在 `merge_sort` 函數中有一個名爲 `left_array_index` 的變量

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index/);
  }
})
```

`left_array_index` 的值應爲 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index\s*=\s*0/);
  }
})
```

你應該在 `merge_sort` 函數中有一個名爲 `right_array_index` 的變量

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index/);
  }
})
```

`right_array_index` 的值應爲 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index\s*=\s*0/);
  }
})

```

你應該在 `merge_sort` 函數中有一個名爲 `sorted_index` 的變量

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index/);
  }
})
```

`sorted_index` 的值應爲 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index\s*=\s*0/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)

--fcc-editable-region--

--fcc-editable-region--
```
