---
id: 656702f8b4cbd8cbf0a433c6
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

在你的函式（程式）內，創建一個 `while` 迴圈，比較 `left_part` 中的一個元素與 `right_part` 中的一個元素，並將較小的元素合併到主要的 `array` 列表中。 

為迴圈創建兩個條件：一個檢查 `left_array_index` 是否小於 `left_part` 的長度，另一個條件檢查 `right_array_index` 是否小於 `right_part` 的長度。

# --hints--

你的 `while` 迴圈應該有這些條件：`left_array_index < len(left_part) and right_array_index < len(right_part)`。別忘了在迴圈主體中使用 `pass` 關鍵字。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

在迴圈主體中使用 `pass` 關鍵字

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
