---
id: 656709e50ed928da35260276
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

在先前的步驟中，你將 `left_part` 中的元素指定給 `array` 列表，因為它較小。但情況不一定總是如此。在某些比較情況下，右邊的元素可能較小。 

創建一個 `else` 語句，以在 `left_part` 索引不小於 `right_part` 索引時執行。

在 `else` 區段中，將 `right_part[right_array_index]` 指定到 `array[sorted_index]`。

# --hints--

你應該在你的 `else` 語句主體中有 `array[sorted_index] = right_part[right_array_index]`。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /else:\s*[^}]*array\s*\[\s*sorted_index\s*\]\s*=\s*right_part\s*\[\s*right_array_index\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
