---
id: 6569b4e0bd29d17d4b53b16c
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

現在，你將為 `right_part` 複製相同的 `while` 迴圈邏輯。

創建一個 `while` 迴圈，當 `right_array_index` 小於 `len(right_part)` 時執行，並在迴圈主體中使用 `pass` 關鍵字。

# --hints--

你的 `while` 迴圈應該有這個條件：`right_array_index < len(right_part)`。別忘了在迴圈主體中使用 `pass` 關鍵字。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/m);
  }
})
```

你應該在你的 `while` 迴圈主體中使用 `pass` 關鍵字

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+right_array_index\s*<\s*len\(\s*right_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--

--fcc-editable-region--
```
