---
id: 65772ef923f922cd720e5008
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

將 `left_part` 作為引數傳遞給你在上一步新增的 `merge_sort()` 呼叫。

# --hints--

你應該將 `left_part` 作為引數傳遞給你的 `merge_sort()` 呼叫。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*left_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort()
--fcc-editable-region--
```
