---
id: 6577254891048c8f2c19e961
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

再次呼叫 `merge_sort()` 函式（程式）。目前不要傳入任何引數。

# --hints--

你應該在先前的函式呼叫之後再次呼叫 `merge_sort()`。

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
