---
id: 657f59751d5a5c9b069ae0f3
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

你可以使用 `__name__` 變量來確定 Python 腳本是否作爲主程序運行，或者是否作爲模塊（在另一個 Python 文件中編寫的代碼）導入。

如果 `__name__` 的值設置爲 `'__main__'`，則表示當前腳本是主程序，而不是模塊。

在這個項目中，你將使用當前腳本作爲主程序。

創建一個 `if` 語句來判斷 `__name__` 的值是否是 `'__main__'`。

在 `if` 語句中使用 `pass` 關鍵字。

# --hints--

`if` 語句中應該檢查 `__name__` 的值是否爲 `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1

    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
