---
id: 657f59751d5a5c9b069ae0f3
title: 步驟 28
challengeType: 20
dashedName: step-28
---

# --description--

你可以使用 `__name__` 變數來判斷 Python 腳本是作為主要程式執行，還是被匯入為模組（寫在另一個 Python 檔案中的程式碼）。

如果 `__name__` 的值設定為 `'__main__'`，表示目前的腳本是主要程式，而不是模組。

在這個專案中，你將使用目前的腳本作為主要程式。

創建一個 `if` 述語來檢查 `__name__` 的值是否為 `'__main__'`。

在 `if` 述語的主體中使用 `pass` 關鍵字。

# --hints--

你的 `if` 述語應該檢查 `__name__` 是否具有 `'__main__'` 的值

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
