---
id: 619691693bc14b0e528f5a20
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

选择 `.ground` 元素，设置其 `width` 占用视口的全部宽度。 然后，将其 `height` 设置为 `400px`。

# --hints--

你应该使用 `.ground` 选择器。

```js
assert.match(code, /\.ground\s*\{/);
```

你应该将 `.ground` 元素的 `width` 设置为 `100vw`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.width, '100vw');
```

你应该将 `.ground` 元素的 `height` 设置为 `400px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.height, '400px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--
```
