---
id: 68e4c81577243e4db1676491
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafanya mazoezi ya kutumia semantic HTML kwa kujenga ukurasa wa mtandao unaojumuisha baadhi ya vidokezo vya Quincy Larson vya kupata kazi ya msanidi programu. Muundo wa msingi wa HTML umeandaliwa tayari kwa ajili yako.

Anza kwa kuunda kipengele cha `h1` chenye maandishi `Quincy's Tips for Getting a Developer Job`.

# --hints--

Unapaswa kuwa na kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `Quincy's Tips for Getting a Developer Job`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
