---
id: 68e65ec5348e40a739903e6d
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Sasa, tengeneza kipengele cha aya chini ya kipengele cha `h1`. Ndani ya aya hii ongeza maandishi ya `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

# --hints--

Unapaswa kuwa na kipengele cha `p`.

```js
assert.exists(document.querySelector('p'));
```

Aya yako inapaswa kuwa na maandishi ya `Learning to code is hard, but as Quincy Larson says, You can become a developer.`. Hakikisha umeangalia tahajia yako kwa makini.

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

Aya yako inapaswa kuwa chini ya kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
