---
id: 68e964b84f954c322376e592
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Ikiwa chanzo cha nukuu ni tovuti, unaweza kuitaja kwa kutumia sifa ya kipengele `cite`. Thamani ya sifa hii ya kipengele inapaswa kuwa URL halali. Ingawa sifa hii ya kipengele haibadilishi muonekano wa nukuu ya kipande, ni msaada mkubwa kwa kisomaji cha skrini na injini za utafutaji kupata taarifa zaidi kuhusu nukuu hiyo.

Hapa kuna mfano wa kipengele cha nukuu ndani ya mstari chenye sifa ya kipengele `cite`:

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

Ongeza sifa ya kipengele `cite` kwenye kipengele cha nukuu ndani ya mstari chenye URL hii:

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

Kipengele chako cha `q` kinapaswa kuwa na sifa ya kipengele `cite`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

Sifa ya kipengele `cite` ya kipengele chako cha `q` inapaswa kuwa na thamani `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
