---
id: 68e90335bd53a901c41dfc13
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Katika funzo la awali ulijifunza kwamba kipengele cha `q` kinatumika kutofautisha maandishi yaliyotajwa kutoka kwa yaliyomo yanayozunguka.

Hapa kuna mfano:

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

Viburawuza vingi vya kisasa vitaongeza alama za nukuu moja kwa moja kuzunguka nukuu ndani ya mstari unapotumia kipengele cha `q`.

Ndani ya aya, zungusha `You can become a developer.` katika kipengele cha nukuu ndani ya mstari, huku ukiacha sehemu nyingine za aya zisibadilike.

# --hints--

Unapaswa kuwa na kipengele cha `q`.

```js
assert.exists(document.querySelector('q'));
```

Kipengele chako cha `q` kinapaswa kupangwa ndani ya aya ambayo iko chini ya kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

Kipengele chako cha `q` kinapaswa kuwa na maandishi ya `You can become a developer.`

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

Ndani ya kipengele chako cha `q`, maandishi hayapaswi kuwa na nafasi za mbele au nyuma (hakuna nafasi, tab au kuvunjika mstari kabla ya herufi ya kwanza au baada ya ya mwisho).

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

Unapaswa kuongeza lebo za `q` tu kuzunguka sehemu inayofaa ya sentensi bila kubadilisha maandishi ya awali. Kipengele cha `p` kinapaswa bado kusoma: `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
