---
id: 68eaa9ef8a0867071f195f59
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Hadi sasa umekuwa ukitumia sifa ya kipengele `cite` kuonyesha chanzo cha nukuu, lakini sifa hiyo haionyeshi chanzo kwa mtumizi kwa njia ya kuona.

Ikiwa unataka kuonyesha chanzo kwa njia ya kuona, unaweza kuongeza kipengele cha rejea, `cite`, nje ya kipengele cha nukuu ya blok. Kipengele cha rejea ni kipengele cha HTML ambacho unaweza kutumia kuonyesha kichwa cha kazi ya ubunifu iliyorejelewa, kama kitabu, makala, wimbo, filamu, tovuti, au karatasi ya utafiti.

Hapa kuna mfano:

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>&mdash;Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

Ndani ya kipengele `p` chini ya kipengele cha nukuu ya blok, zungusha `How to Learn to Code and Get a Developer Job [Full Book]` katika kipengele cha `cite`.

# --hints--

Unapaswa kuwa na kipengele cha `cite` ndani ya kipengele cha `p` chini ya kipengele cha `blockquote`.

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

Kipengele chako cha `cite` kinapaswa kuwa na maandishi `How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

Ukisahau kipengele cha `cite` ulichokiweka tu, maandishi ya kipengele cha `p` yanapaswa bado kuwa `—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
      --fcc-editable-region--
          &mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
      --fcc-editable-region--
        </p>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
