---
id: 68ea7e7b916e8270483689b6
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Chini ya kipengele cha `h2`, utaongeza nukuu nyingine ya Quincy. Wakati huu, nukuu haitakuwa sehemu ya aya kubwa zaidi. Badala yake, aya nzima itakuwa nukuu. Ili kutofautisha maandishi yaliyoniwekwa nukuu kama haya, unapaswa kutumia kipengele cha nukuu ya kipande: `blockquote`. Kwenye kivinjari, utaona kuwa maandishi yamepunguzwa kidogo upande.

Hapa kuna mfano wa kipengele cha nukuu ya kipande chenye maandishi yaliyoniwekwa nukuu:

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Sasa, ndani ya sehemu ya kwanza, ongeza kipengele cha nukuu ya kipande chini ya kipengele cha `h2` chenye maandishi ya `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

# --hints--

Ndani ya kipengele cha kwanza cha `section`, unapaswa kuwa na kipengele cha `blockquote` chini ya kipengele cha `h2`.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

Kipengele chako cha `blockquote` kinapaswa kuwa na maandishi ya `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
