---
id: 68eab6a2b37d2f13602d0c3f
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Unaweza kuandika maandishi yaliyo na nukuu moja kwa moja ndani ya kipengele cha nukuu ya block, kama ulivyofanya katika kipengele cha kwanza cha `section`. Vinginevyo, unaweza kuyaweka ndani ya kipengele kimoja au zaidi cha aya. Hii ni msaada wakati maandishi yana aya nyingi, lakini unataka kuyaweka ndani ya kipengele kimoja cha nukuu ya block.

Hapa kuna mfano wenye aya mbili:

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  <p>So many people ask me each week: is college still worth it? In this 1-hour video I answer this question and other commonly asked questions about university.</p>
  <p>I've been in adult education for two decades at this point, and even though I'm not a labor market economist, I do feel confident enough to answer these questions.</p>
</blockquote>
```

Katika sehemu ya pili, ndani ya kipengele kilichopo cha nukuu ya block, ongeza vipengele vinne vya `p` vyenye maandishi yafuatayo, kwa mpangilio:

- `So much of getting a job is who you know.`
- `It's OK to be an introvert, but you do need to push your boundaries.`
- `Create GitHub, Twitter, LinkedIn, and Discord accounts.`
- `Go to tech meetups and conferences. Travel if you have to.`

# --hints--

Unapaswa kuwa na vipengele vinne vya `p` vilivyopangwa ndani ya kipengele cha `blockquote` ndani ya sehemu ya pili.

```js
const pElements = document.querySelectorAll('main > section:nth-of-type(2) > blockquote > p');
assert.lengthOf(pElements, 4);
```

Kipengele chako cha kwanza cha `p` kinapaswa kuwa na maandishi `So much of getting a job is who you know.`

```js
const firstPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(1)');
assert.equal(firstPEl?.innerText.trim(), 'So much of getting a job is who you know.');
```

Kipengele chako cha pili cha `p` kinapaswa kuwa na maandishi `It's OK to be an introvert, but you do need to push your boundaries.`

```js
const secondPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(2)');
assert.equal(secondPEl?.innerText.trim(), "It's OK to be an introvert, but you do need to push your boundaries.");
```

Kipengele chako cha tatu cha `p` kinapaswa kuwa na maandishi `Create GitHub, Twitter, LinkedIn, and Discord accounts.`

```js
const thirdPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(3)');
assert.equal(thirdPEl?.innerText.trim(), 'Create GitHub, Twitter, LinkedIn, and Discord accounts.');
```

Kipengele chako cha nne cha `p` kinapaswa kuwa na maandishi `Go to tech meetups and conferences. Travel if you have to.`

```js
const fourthPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(4)');
assert.equal(fourthPEl?.innerText.trim(), 'Go to tech meetups and conferences. Travel if you have to.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>
      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
        --fcc-editable-region--
          
        --fcc-editable-region--
        </blockquote>
      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
