---
id: 68ea9d23318b6481ac3de777
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Kama vile kipengele cha nukuu ndani ya mstari, unaweza pia kuongeza sifa ya kipengele ya `cite` kwa kipengele cha nukuu ya kipande.

Hapa kuna mfano wa kipengele cha nukuu ya kipande chenye sifa ya kipengele ya `cite`:

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Sasa, ongeza sifa ya kipengele ya `cite` kwa kipengele cha nukuu ya kipande chenye URL `https://www.freecodecamp.org/news/learn-to-code-book/`.

# --hints--

Kipengele chako cha `blockquote` kinapaswa kuwa na sifa ya kipengele ya `cite`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.exists(firstBlockquoteEl?.getAttribute('cite'));
```

Sifa ya kipengele ya `cite` inapaswa kuwa na thamani ya `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
      --fcc-editable-region--
        <blockquote>
      --fcc-editable-region--
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
