---
id: 692fb9f4046f1d677f1f6e4e
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Crie uma funГ§ГЈo chamada `isSelect` que receba um Гєnico parГўmetro `element` do tipo `EventTarget | null`. Essa funГ§ГЈo vai verificar se o elemento selecionado Г© um elemento HTML `<select>`.

# --hints--

VocГЄ deve ter uma funГ§ГЈo chamada `isSelect`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.isSelect);
```

Sua funГ§ГЈo `isSelect` deve ter um parГўmetro chamado `element` do tipo `EventTarget | null`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.isSelect.parameters[0].matches("element: EventTarget | null"));
```

Sua funГ§ГЈo `isSelect` deve ter um tipo de retorno `element is HTMLSelectElement`.

```js
assert.match(__helpers.removeJSComments(code), /:\s*element\s+is\s+HTMLSelectElement/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bug Emoji Picker</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <form id="bug-form">
        <label for="species">Species:</Label>
        <select id="species">
          <option value="bee" selected>рџђќ Bee</option>
          <option value="spider">рџ•·пёЏ Spider</option>
        </select>
      </form>
      <div>
        <p id="bug-emoji">рџђќ</p>
      </div>
    </main>
  </body>
  <script src="index.ts"></script>
</html>
```

```css
:root {
  --gradient-back: #987284;
  --gradient-left: #F9B5AC;
  --gradient-right: #AA7674;
  --card: #FFFFFF;
  --card-shadow: 0 10px 30px rgba(30, 41, 59, 0.12);
  --font: #222;
}

html {
  width: 100%;
  min-height: 100%;
  background: radial-gradient(100% 100% at 10% 10%, var(--gradient-left), transparent 60%),
              radial-gradient(90% 85% at 90% 0%, var(--gradient-right), transparent 55%),
              linear-gradient(180deg, var(--gradient-back), #ffffff);
}

body {
  margin: 5rem auto;
  padding: 2rem 3rem;
  width: fit-content;
  min-width: 15%;
  color: var(--font);
  font: 500 16px/1.4 system-ui;
  background: rgba(255, 255, 255, 0.3);
  border: 1px solid rgba(255, 255, 255, 0.3);
  border-radius: 15%;
  box-shadow: var(--card-shadow);

  backdrop-filter: blue(12px);
  -webkit-backdrop-filter: blur(12px);
}

select, input {
  appearance: none;
  -webkit-appearance: none;
  -moz-appearance: none;

  background: rgba(255, 255, 255, 0.25);
  border-radius: 12px;
  padding: 0.1rem 2rem 0.1rem 0.8rem;
  font-size: 1rem;
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

form { 
  display: flex;
  flex-direction: column;
}

#bug-emoji {
  margin: 0 auto;
  width: fit-content;
  font-size: 128px;
}

#bug-name {
  margin: 0 auto;
  width: fit-content;
  font-size: 24;
  font-weight: 700;
}
```

```ts
abstract class Bug<T> {
  emoji!: T;
  emojiElement!: HTMLParagraphElement;
  constructor(emojiElement: HTMLParagraphElement) {
    this.emojiElement = emojiElement;
  }

  abstract render(): void;
}

class Bee extends Bug<string> {
  constructor(emojiElement: HTMLParagraphElement) {
    super(emojiElement);
    this.emoji = "рџђќ";
  }

  override render() {
    this.emojiElement.innerText = this.emoji;
  }
}

class Spider extends Bug<string> {
  constructor(emojiElement: HTMLParagraphElement) {
    super(emojiElement);
    this.emoji = "рџ•·пёЏ";
  }

  override render() {
    this.emojiElement.innerText = this.emoji;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
