---
id: 657ef2a86d4e545cec9a85fb
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

在 camel case 和 pascal case 中，大寫字元標示新字的開始。若要將輸入字串轉換為 snake case，你需要檢查輸入字串中的字元是否為大寫。 

你可以使用 `.isupper()` 字串方法來檢查字元是否為大寫。若字元為大寫，該方法會傳回 `True`，否則傳回 `False`。

在 `for` 迴圈內，為你的程式碼添加一個 `if` 述語，以檢查目前的字元是否為大寫。將 `pass` 述語移動到新的 `if` 述語內。

# --hints--

你應該撰寫一個新的 if 述語，以 char.isupper() 作為條件。記得在結尾添加冒號，並使用 pass 關鍵字填充 if 述語主體。

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

你應該將迴圈主體中的 `pass` 述語替換為 `if` 述語。

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
