---
id: 657effaa2a5e0277d71f9cbe
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

到此為止，變數 `snake_cased_char_list` 已經保存了轉換後的字元列表。若要將這些字元合併成單一字串，你可以使用 `.join()` 方法。

`join` 方法的運作方式是將列表中的每個元素串接成一個字串，並以指定的字串作為分隔符。

```py
result_string = ''.join(characters)
```

上面的範例將 `characters` 列表的元素合併成一個字串，其中每個元素使用空字串作為分隔符連接在一起。

現在，在 `for` 迴圈之後，使用 `.join()` 方法以空字串作為分隔符，將 `snake_cased_char_list` 中的元素連接起來。將結果指定到一個名為 `snake_cased_string` 的新變數。


# --hints--

你應該使用 `.join()` 方法將 `snake_cased_char_list` 列表中的字元串接成字串，並在你的 `for` 迴圈之後將該字串指定給名為 `snake_cased_string` 的變數。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
