---
id: 657ed53c19461d4b95c4757a
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

你需要添加一個空的列表，用來保存你將字串轉換為 snake case 後的字元。 

在函式中，透過建立一個名為 `snake_cased_char_list` 的空列表來取代 `pass` 述語。

# --hints--

你應該在 `convert_to_snake_case` 中將 `pass` 取代為名為 `snake_cased_char_list` 的空列表。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const commentless_code = __helpers.python.removeComments(transformedCode)
        const convert_to_snake_case = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;
        assert.notMatch(function_body, /\bpass\b/);
        assert.match(function_body, / +snake_cased_char_list\s*=\s*\[\s*\]\s*/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    pass    
--fcc-editable-region--
```
