---
id: 657efa642593c5746acc5c81
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

在 `if` 述語主體內，你需要將任何大寫字元轉換為小寫，並在此小寫字元前加上底線。

使用 `.lower()` 字串方法將大寫字元轉換為小寫字元。然後，將底線加到該字元前面。將結果指定到名為 `converted_character` 的變數。

# --hints--

你應該將修改後的字元指定到名為 `converted_character` 的變數。

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

你不應該在你的 `if` 述語中使用 `pass`。

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
