---
id: 657ee28cefc4945568287673
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

在空的列表就緒後，現在你可以開始遍歷輸入的字串並將其轉換成 snake case。

在函式（程式）內，在你剛剛創建的列表下方，為 `pascal_or_camel_cased_string` 添加一個 `for` 迴圈以進行迭代。請確保將標的變數命名為 `char`。目前，在迴圈主體中添加一個 `pass` 述語作為佔位符。

# --hints--

你應該撰寫一個新的 `for` 迴圈，標的變數命名為 `char`。別忘了在結尾加上冒號，並使用 `pass` 來填充迴圈主體。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
