---
id: 657effaa2a5e0277d71f9cbe
title: 步驟 8
challengeType: 20
dashedName: step-8
---

# --description--

此時，變量 `snake_cased_char_list` 保存了轉換後的字符列表。 要將這些字符組合成一個字符串，你可以使用 `.join()` 方法。

`join` 方法的工作原理是將列表的每個元素連接成一個字符串，並用指定的字符串（稱爲分隔符）分隔。

```py
result_string = ''.join(characters)
```

上面的例子將 `characters` 列表中的元素連接成一個字符串，其中每個元素使用空字符串作爲分隔符連接在一起。

現在，在 `for` 循環之後，使用 `.join()` 方法以空字符串作爲分隔符連接 `snake_cased_char_list` 中的元素。 將結果分配給名爲 `snake_cased_string` 的新變量。


# --hints--

你應該使用 `.join()` 方法將 `snake_cased_char_list` 列表中的字符連接成一個字符串，並在 `for` 循環之後將該字符串分配給名爲 `snake_cased_string` 的變量。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
