---
id: 657e928716b77b2277980276
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個專案中，你將建置一個程式，該程式會接受以 `camelCase` 或 `PascalCase` 格式的字串作為輸入，並使用兩種方法將其轉換為 `snake_case` 格式的字串。首先，你會使用 for 迴圈，然後使用列表推導式來達成相同的結果。你將看到列表推導式如何讓你的程式碼更簡潔。

開始定義一個名為 `convert_to_snake_case()` 的新函式（程式），該函式接受一個名為 `pascal_or_camel_cased_string` 的字串作為輸入。暫時在函式內添加一個 `pass` 述語。

# --hints--

你應該定義一個名為 `convert_to_snake_case()` 的新函式（程式）。別忘了在結尾加上冒號，並使用 `pass` 來填充函式（程式）主體。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

你的函式（程式）應該接受一個名為 `pascal_or_camel_cased_string` 的參數。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
