---
id: 657efce98e958b75df86b305
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

在 `if` 述語主體內，你將把轉換後的字元添加到你先前創建的列表中。

為此，將使用 `.append()` 方法。此方法會將給定的物件添加到它被喚起的列表末端。

對 `snake_cased_char_list` 使用 `.append()` 方法，將 `converted_character` 添加到列表中。

# --hints--

你應該使用 `.append()` 方法將轉換後的字元添加到 `snake_cased_char_list`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
