---
id: 657f43d341a0dd17120cdb08
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

你需要將大寫字元轉換為小寫，並在它們前面添加底線。 

在開始處理列表推導式之前，你將為你的函式（程式）給予一個返回值。這樣你就能檢查輸出。

使用 `return` 述語從你的函式傳回 `snake_cased_char_list` 列表。

# --hints--

你應該傳回 `snake_cased_char_list` 列表。確保縮排設定正確。

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*snake_cased_char_list/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--

--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
