---
id: 657f0044be09db790b1eb1c5
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

在 Pascal 命名法中，字串以大寫字元開頭。在將所有字元轉換為小寫並為它們添加底線後，你的字串開頭可能會多出一個底線。

解決此問題最簡單的方法是使用 `.strip()` 字串方法，此方法會從字串中移除作為引數傳入的一組字元中，任何開頭或結尾的字元。例如：

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

`strip()` 方法應用於 `original_string`。這會移除任何前置和後置的底線。上述範例的結果將會是字串 `'example_string'`。

宣告一個名為 `clean_snake_cased_string` 的新變數，並將套用於 `snake_cased_string` 的 `.strip()` 方法的結果指定給它，該方法傳入 `'_'` 作為引數。

# --hints--

你應該將 `snake_cased_string.strip('_')` 的結果字串指定給名為 `clean_snake_cased_string` 的變數。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
