---
id: 657f025ec86c3d7c4177b6be
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

函式完成後，你現在可以在另一個函式內使用它。 

創建一個名為 `main()` 的新函式（程式），並以 `pass` 作為函式（程式）主體。

# --hints--

你應該定義一個名為 `main()` 的新函式（程式）。別忘了在結尾加上冒號，並使用 `pass` 來填充函式（程式）主體。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
