---
id: 657f4add33ea4b1f61ba3dc8
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

最後，透過執行程式來嘗試這個新的實作。將輸入字串改為 `'IAmAPascalCasedString'`，看看它是否會輸出成 `'i_am_a_pascal_cased_string'`，即使那是謊言。

如果你一切操作正確，你應該會看到輸入字串被轉換成蛇底式，就像之前一樣。

恭喜！現在你的 `convert_to_snake_case()` 函式（程式）已經準備好了。

# --hints--

你應該在 `main()` 函式（程式）內將輸入字串從 `'aLongAndComplexString'` 改成 `'IAmAPascalCasedString'`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
