---
id: 657f0353c9523d7d896873ea
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

在 `main()` 函式（程式）內，將 `pass` 述語取代為呼叫 `convert_to_snake_case()` 函式（程式），並傳入字串 `'aLongAndComplexString'` 作為輸入。

要顯示輸出，將函式呼叫作為引數傳遞給 `print()` 函式。

# --hints--

你應該在 `main()` 函式（程式）內呼叫 `convert_to_snake_case()`，並將 `'aLongAndComplexString'` 作為輸入傳遞給該函式（程式）。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\s*\(\s*convert_to_snake_case\s*\(\s*("|')aLongAndComplexString\1\s*\)\s*\)/);
    }
})
```

你不應該在你的 `main` 函式（程式）中有 `pass`。

```js
({
    test: () => {
        const commentless_code = __helpers.python.removeComments(code);
        const transformedCode = commentless_code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    pass
--fcc-editable-region--
```
