---
id: 657f0353c9523d7d896873ea
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

在 `main()` 函數內部，將 `pass` 語句替換爲對 `convert_to_snake_case()` 函數的調用，並將字符串 `'aLongAndComplexString'` 作爲輸入傳遞。

要顯示輸出，請將函數調用作爲參數傳遞給 `print()` 函數。

# --hints--

你應該在 `main()` 函數內調用 `convert_to_snake_case()` 並將 `'aLongAndComplexString'` 作爲輸入傳遞給該函數。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\s*\(\s*convert_to_snake_case\s*\(\s*("|')aLongAndComplexString\1\s*\)\s*\)/);
    }
})
```

你的 `main` 函數中不應有 `pass`。

```js
({
    test: () => {
        const commentless_code = __helpers.python.removeComments(code);
        const transformedCode = commentless_code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    pass
--fcc-editable-region--
```
