---
id: 657f47b12c51e41b3149e584
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

不過，最終結果並不完全是你想達成的。你需要對被 `if` 語句篩選掉的字元執行不同的表達式。你將為此使用 `else` 語句：

```py
spam = [i * 2 if i > 0 else -1 for i in iterable]
```

請注意，與 `if` 語句不同，`if`/`else` 構件必須放在表達式與 `for` 關鍵字之間。

修改你的列表推導式，使當字元不是大寫時保持不變。

# --hints--

你應該修改你的列表推導式，以便在 `char.isupper()` 時評估表達式 `'_' + char.lower()`，否則評估 char。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower() if char.isupper() else char")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string if char.isupper()]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
