---
id: 657f4345abe7f2161f99f1ad
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

將 `pass` 關鍵字取代為變數 `snake_cased_char_list` 並指定一個空列表。使用方括號表示法來創建列表。

# --hints--

你應該在 `convert_to_snake_case` 中宣告一個名為 `snake_cased_char_list` 的空列表。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, /clean_snake_cased_string\s+    snake_cased_char_list\s*=\s*\[\n*\s*\n*\]/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string
--fcc-editable-region--
    pass


--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
