---
id: 657f04ed0035f47ed04d0f1f
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

為了顯示 `convert_to_snake_case()` 函式的輸出，你需要呼叫 `main()` 函式。 

在與兩個現有函式相同的層級，為你的程式碼添加對 `main()` 函式的呼叫。執行時，你應該會看到給定的駝峰式或帕斯卡式字串被轉換成蛇底式。

# --hints--

你應該為你的程式添加對 `main()` 函式的呼叫。

```js
({
  test: () => {
    const pythonCode = `_Node(_code.replace('\\\\r', '')).has_call("main()")`;
    const result = runPython(pythonCode);
    assert(result);
  }
});
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

--fcc-editable-region--
```
