---
id: 657f456223b8c1187b461987
title: 步驟 19
challengeType: 20
dashedName: step-19
---

# --description--

在 Python 中，<dfn>列表推導式</dfn> 是一種構造，它允許你通過將表達式應用於現有可迭代對象中的每個項目並可選地使用條件過濾項目來生成新列表。 除了更簡潔之外，列表推導通常運行速度更快。

基本列表推導由一個表達式和後跟 `for` 子句組成：

```py
spam = [i * 2 for i in iterable]
```

上面使用變量 `i` 來迭代 `iterable`。 通過在當前迭代中評估表達式 `i * 2` 來獲得結果列表的每個元素。

在此步驟中，你需要使用列表推導語法填充空列表 `snake_cased_char_list`。

將空列表轉換爲列表推導，將 `pascal_or_camel_cased_string` 中的每個字符轉換爲小寫字符，並在其前面添加下劃線（之前註釋掉的代碼可能會幫助你編寫表達式）。 使用 `char` 來迭代 `pascal_or_camel_cased_string`。

# --hints--

你應該將 `snake_cased_char_list` 轉換爲對 `pascal_or_camel_cased_string` 進行迭代的列表推導。

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

你的列表推導應該使用 `char` 來迭代 `pascal_or_camel_cased_string`。

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

你的列表推導應該對 `pascal_or_camel_cased_string` 中的每個 `char` 評估 `'_' + char.lower()`。

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
