---
id: 657f456223b8c1187b461987
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

在 Python 中，<dfn>list comprehension</dfn> 是一種構件，允許你透過對現有可迭代物件中的每個項目套用表達式來產生新的列表，並且選用以條件過濾項目。除了更簡潔之外，list comprehension 通常執行速度也較快。

基本的列表推導式由一個表達式，後接一個 `for` 語句組成：

```py
spam = [i * 2 for i in iterable]
```

上述使用變數 `i` 來迭代 `iterable`。結果列表的每個元素都是在當前迭代中評估表達式 `i * 2` 所得到的。

在此步驟中，你需要使用列表推導語法填充空的列表 `snake_cased_char_list`。

將你的空列表轉換成一個列表推導式，將 `pascal_or_camel_cased_string` 中的每個字元轉換成小寫字元並在前面加上底線（你之前註解掉的程式碼可能有助於你撰寫表達式）。使用 `char` 來迭代 `pascal_or_camel_cased_string`。 

# --hints--

你應該將 `snake_cased_char_list` 轉換成一個列表推導式，該推導式迭代 `pascal_or_camel_cased_string`。

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

你的列表推導式應該使用 `char` 來迭代 `pascal_or_camel_cased_string`。

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

你的列表推導式應該為 `pascal_or_camel_cased_string` 中的每個 `char` 計算 `'_' + char.lower()`。

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
