---
id: 657f025ec86c3d7c4177b6be
title: 步驟11
challengeType: 20
dashedName: step-11
---

# --description--

函數完成後，你現在可以在另一個函數中使用它。

創建一個名爲 `main()` 的新函數，並以 `pass` 作爲函數主體。

# --hints--

你應該定義一個名爲 `main()` 的新函數。 不要忘記最後的冒號，並使用 `pass` 來填充函數體。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
