---
id: 657f465f8e718b19c5105ae5
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

列表推導式接受條件述語，只有在符合特定條件時才會評估所提供的表達式：

```py
spam = [i * 2 for i in iterable if i > 0]
```

如你從輸出中所見，從 `pascal_or_camel_cased_string` 產生的字元列表已被串接。由於列表推導式中的表達式會對每個字元進行計算，結果是一個所有字元皆以底線分隔的全小寫字串。

參考上述範例，為你的列表推導式添加一個 `if` 語句，使該表達式僅在字元為大寫時執行。

# --hints--

你應該為你的列表推導式添加一個帶有條件 `char.isupper()` 的 `if` 語句。

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
