---
id: 657f425dbab54e11993c80f0
title: 步驟14
challengeType: 20
dashedName: step-14
---

# --description--

到目前爲止，在這個項目中，你已經使用 `for` 循環來迭代輸入字符串並將其轉換爲所需的輸出。 現在你將開始從 `for` 循環轉換到列表推導。

首先註釋掉 `convert_to_snake_case()` 函數內的所有代碼行。 不要刪除它們，因爲當你使用列表推導實現邏輯時它們會很有幫助。

記得在函數體中添加 `pass` 關鍵字，以防止代碼在測試期間失敗。

# --hints--

你應該註釋掉 `convert_to_snake_case()` 函數內的所有 10 行代碼，不要註釋掉其他任何內容。 添加 `pass` 臨時填充函數體，避免錯誤。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
