---
id: 657f425dbab54e11993c80f0
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

到目前為止，在這個專案中你已經使用 `for` 迴圈來迭代你的輸入字串並將其轉換成所需的輸出。現在你將開始從 `for` 迴圈轉換到列表推導式。

首先將 `convert_to_snake_case()` 函式內所有程式碼行註解起來。不要刪除它們，因為當你使用列表推導式實作邏輯時，它們會很有幫助。

請記得為你的函式（程式）主體添加 `pass` 關鍵字，以防止程式碼在測試期間失敗。

# --hints--

你應該將 `convert_to_snake_case()` 函式（程式）內的所有 10 行程式碼註解掉，其他部分不動。為了暫時填充函式（程式）主體並避免錯誤，將 `pass` 添加進去。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
