---
id: 657f01ae9aea647b27402d3e
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

為了結束該函式，傳回 `clean_snake_cased_string`。這將完成該函式，並允許你使用它將字串從 Pascal 或 camel case 轉換為 snake case。

在函式的結尾添加一個 `return` 述語以傳回 `clean_snake_cased_string`。


# --hints--

你應該在 `convert_to_snake_case` 函式的結尾添加 `return clean_snake_cased_string`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +return\s+clean_snake_cased_string/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')
--fcc-editable-region--


--fcc-editable-region--
```
