---
id: 657f0044be09db790b1eb1c5
title: 步驟 9
challengeType: 20
dashedName: step-9
---

# --description--

在帕斯卡格式中，字符串以大寫字母開頭。 將所有字符轉換爲小寫並添加下劃線後，字符串開頭可能會出現額外的下劃線。

解決此問題的最簡單方法是使用 `.strip()` 字符串方法，該方法從字符串中刪除作爲其參數傳遞的一組字符中的任何前導或尾隨字符。 例如：

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

`strip()` 方法應用於 `original_string`。 這將刪除所有前導和尾隨下劃線。 上述示例的結果將是字符串 `'example_string'`。

聲明一個名爲 `clean_snake_cased_string` 的新變量，並將應用於 `snake_cased_string` 的 `.strip()` 方法的結果分配給它，同時將 `'_'` 作爲參數傳遞給該方法。

# --hints--

你應該將 `snake_cased_string.strip('_')` 的結果字符串分配給名爲 `clean_snake_cased_string` 的變量。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
