---
id: 663b16e62fee463b4caf46e9
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

在將列表 `snake_cased_char_list` 的元素連接後，你需要從結果字串中移除任何前置或後置的底線。為此，使用帶有底線字元 `_` 作為引數的 `strip` 方法。

方法呼叫可以串鏈在一起，這表示一個方法呼叫的結果可以用作另一個方法呼叫的物件。

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

在上述範例中，`.upper()` 方法串鏈到 `' '.join(words_list)`，因此 `.upper()` 是在 `.join()` 呼叫的結果上被呼叫。

修改 `return` 述語，透過串鏈到 `''.join(snake_cased_char_list)` 的 `.strip()` 呼叫方法，以移除任何前置或後置的底線。

# --hints--

你應該修改傳回述語以包含帶有底線字元 `_` 作為引數的 `strip` 方法。

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
