---
id: 663b10c10a4c0a0e095137ee
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

你需要將 `snake_cased_char_list` 的元素使用空字串 `''` 作為分隔符串接成一個字串，而不是傳回該列表。

修改 `return` 述語以傳回使用空字串作為分隔符連接 `snake_cased_char_list` 的結果。

# --hints--

你應該修改 `return` 述語，使用 `join()` 方法以空字串 `''` 作為分隔符，傳回將 `snake_cased_char_list` 連接起來的結果。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_function('convert_to_snake_case').find_return().is_equivalent('return "".join(snake_cased_char_list)')
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return snake_cased_char_list
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
