---
id: 587d778a367417b2b2512aa5
title: figure要素でチャートのアクセシビリティを向上させる
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMqtE'
forumTopicId: 301015
dashedName: improve-chart-accessibility-with-the-figure-element
---

# --description--

HTML5では`figure`要素と関連する`figcaption`が導入されました。これらは一緒に使うことで、画像や図、チャートのような視覚的表現とそのキャプションをまとめて包みます。これらの要素をまとめることで、関連する内容を意味的にグループ化し、`figure`を説明するテキスト代替を提供するという二重のアクセシビリティ向上が得られます。

チャートのようなデータの可視化では、キャプションを使って視覚障害のあるユーザー向けに傾向や結論を簡潔に示すことができます。別のチャレンジでは、チャートのデータをテーブル形式で画面外に移動させ（CSSを使って）スクリーンリーダー利用者に提供する方法を扱います。

例を示します。`figcaption`は`figure`タグの中に入り、他の要素と組み合わせることもできます。

```html
<figure>
  <img src="roundhouseDestruction.jpeg" alt="Photo of Camper Cat executing a roundhouse kick">
  <br>
  <figcaption>
    Master Camper Cat demonstrates proper form of a roundhouse kick.
  </figcaption>
</figure>
```

# --instructions--

Camper Catはステルス、戦闘、武器の訓練に週あたり費やす時間を示す積み上げ棒グラフを作成中です。彼のページ構造を改善するために、彼が使った`div`タグを`figure`タグに変更し、キャプションを囲む`p`タグを`figcaption`タグに変更してください。

# --hints--

コードには`figure`タグが1つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('figure') , 1);
```

コードには`figcaption`タグが1つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('figcaption') , 1);
```

コードに`div`タグがあってはいけません。

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

コードに`p`タグがあってはいけません。

```js
assert.lengthOf(document.querySelectorAll('p') , 0);
```

`figcaption`は`figure`タグの子要素であるべきです。

```js
const figure = document.querySelector('figure');
const children = figure?.querySelectorAll(`:scope ${'figcaption'}`);
assert.lengthOf(children, 1);
```

`figure`要素には終了タグが必要です。

```js
assert.isTrue(
    code.match(/<\/figure>/g)?.length === code.match(/<figure>/g)?.length
);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>

      <!-- Only change code below this line -->
      <div>
        <!-- Stacked bar chart will go here -->
        <br>
        <p>Breakdown per week of time to spend training in stealth, combat, and weapons.</p>
      </div>
      <!-- Only change code above this line -->

    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>
      <figure>
        <!-- Stacked bar chart will go here -->
        <br>
        <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
      </figure>
    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
