---
id: 587d778b367417b2b2512aa8
title: アクセシブルな日付ピッカーを追加する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cR3bRbCV'
forumTopicId: 301008
dashedName: add-an-accessible-date-picker
---

# --description--

フォームにはよく`input`フィールドが含まれており、これを使ってさまざまなフォームコントロールを作成できます。この要素の`type`属性は、どの種類の`input`要素が作成されるかを示します。

以前のチャレンジで`text`や`submit`の入力タイプを見たことがあるかもしれませんが、HTML5では`date`フィールドを指定するオプションが導入されました。ブラウザの対応状況によっては、`input`フィールドがフォーカスされると日付ピッカーが表示され、すべてのユーザーにとってフォームの入力が簡単になります。

古いブラウザではタイプが`text`にデフォルト設定されるため、念のため`label`や`placeholder`のテキストでユーザーに期待される日付形式を示すのが役立ちます。

例を示します：

```html
<label for="input1">Enter a date:</label>
<input type="date" id="input1" name="input1">
```

# --instructions--

Camper Catはモータルコンバットのトーナメントを設定しており、参加者に都合の良い日付を尋ねたいと思っています。`input`タグを追加し、`type`属性に`date`、`id`属性に`pickdate`、`name`属性に`date`を設定してください。

# --hints--

日付セレクター用の`input`タグを1つ追加してください。

```js
assert.lengthOf(document.querySelectorAll('input'), 2);
```

`input`タグには、値が`type`の`date`属性が必要です。

```js
assert.equal(document.querySelector('input')?.getAttribute('type'), 'date');
```

`input`タグには、値が`id`の`pickdate`属性が必要です。

```js
assert.equal(document.querySelector('input')?.getAttribute('id'),'pickdate');
```

`input`タグには、値が`name`の`date`属性が必要です。

```js
assert.equal(document.querySelector('input')?.getAttribute('name'), 'date');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>

        <!-- Only change code below this line -->



        <!-- Only change code above this line -->

        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>
        <input type="date" id="pickdate" name="date">
        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
