---
id: 587d7790367417b2b2512ab1
title: 複数の要素のキーボードフォーカスの順序を指定するために tabindex を使う
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzRRcb'
forumTopicId: 301028
dashedName: use-tabindex-to-specify-the-order-of-keyboard-focus-for-several-elements
---

# --description--

`tabindex` 属性は要素の正確なタブ順序も指定します。これは属性の値を1以上の正の数に設定すると実現されます。

`tabindex="1"` を設定すると、その要素に最初にキーボードフォーカスが移動します。次に指定された `tabindex` の値（2、3など）のシーケンスを順にたどり、最後にデフォルトおよび `tabindex="0"` の項目に移動します。

このようにタブ順序を設定すると、デフォルトの順序（HTMLソースの順序）をオーバーライドすることに注意してください。ページの上部からナビゲーションを開始することを期待しているユーザーには混乱を招く場合があります。この手法は特定の状況で必要になることがありますが、アクセシビリティの観点からは適用前に注意が必要です。

例を示します：

```html
<div tabindex="1">I get keyboard focus, and I get it first!</div>
```

```html
<div tabindex="2">I get keyboard focus, and I get it second!</div>
```

# --instructions--

Camper Cat は Inspirational Quotes ページに検索フィールドを持っており、CSSで右上に配置する予定です。検索の `input` と提出の `input` フォームコントロールをタブ順序の最初の2つの項目にしたいと考えています。`tabindex` `1` に `search` 属性を `input` に設定し、`tabindex` `2` に `submit` 属性を `input` に設定してください。

もう一つ注意すべきことは、いくつかのブラウザでは要素をクリックしたときにタブ順序の途中から開始される場合があることです。ページには常にタブ順序の最初から開始されるようにする要素が追加されています。

# --hints--

コードは `tabindex` `search` タグに `input` 属性を追加する必要があります。

```js
assert.isNotNull(document.querySelector('#search')?.getAttribute('tabindex'));
```

コードは `tabindex` `submit` タグに `input` 属性を追加する必要があります。

```js
assert.isNotNull(document.querySelector('#submit')?.getAttribute('tabindex'));
```

コードは `tabindex` `search` タグに `input` 属性を値 1 で設定する必要があります。

```js
assert.equal(document.querySelector('#search')?.getAttribute('tabindex'), '1');
```

コードは `tabindex` `submit` タグに `input` 属性を値 2 で設定する必要があります。

```js
assert.equal(document.querySelector('#submit')?.getAttribute('tabindex'), '2');
```

# --seed--

## --seed-contents--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input type="search" name="search" id="search">
    <input type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```

# --solutions--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input tabindex="1" type="search" name="search" id="search">
    <input tabindex="2" type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```
