---
id: 587d7790367417b2b2512aaf
title: HTMLのアクセスキーでリンクをナビゲート可能にする
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQvmaTp'
forumTopicId: 301021
dashedName: make-links-navigable-with-html-access-keys
---

# --description--

HTMLは、要素をアクティブにしたりフォーカスを当てたりするためのショートカットキーを指定する`accesskey`属性を提供しています。`accesskey`属性を追加すると、キーボードのみを使うユーザーにとってナビゲーションがより効率的になります。

HTML5ではこの属性を任意の要素に使えますが、特にインタラクティブな要素に使うと便利です。これにはリンク、ボタン、フォームコントロールが含まれます。

例を示します：

```html
<button accesskey="b">Important Button</button>
```

# --instructions--

Camper Catは、ブログ記事タイトルの2つのリンクにキーボードショートカットを設定して、サイトのユーザーが全文にすばやく移動できるようにしたいと考えています。両方のリンクに`accesskey`属性を追加し、最初のリンクには`g`（Garfield用）、2つ目のリンクには`c`（Chuck Norris用）を設定してください。

# --hints--

コードは`accesskey`タグに`a`属性を追加し、`id`を`first`に設定する必要があります。

```js
assert.isTrue(document.querySelector('#first')?.hasAttribute('accesskey'));
```

コードは`accesskey`タグに`a`属性を追加し、`id`を`second`に設定する必要があります。

```js
assert.isTrue(document.querySelector('#second')?.hasAttribute('accesskey'));
```

コードは`accesskey`タグの`a`属性を`id`の`first`に`g`として設定する必要があります。大文字小文字は区別されます。

```js
assert.equal(document.querySelector('#first')?.getAttribute('accesskey'), 'g');
```

コードは`accesskey`タグの`a`属性を`id`の`second`に`c`として設定する必要があります。大文字小文字は区別されます。

```js
assert.equal(document.querySelector('#second')?.getAttribute('accesskey'), 'c');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" accesskey="g" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" accesskey="c" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
