---
id: 587d778f367417b2b2512aac
title: 色覚障害の問題を避けるために十分なコントラストを使う
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMEUw'
forumTopicId: 301012
dashedName: avoid-colorblindness-issues-by-using-sufficient-contrast
---

# --description--

色は視覚デザインの大きな要素ですが、その使用には2つのアクセシビリティの問題があります。まず、重要な情報を伝える手段として色だけを使うべきではありません。なぜならスクリーンリーダーのユーザーは色を認識できないからです。次に、前景色と背景色は色覚障害のあるユーザーが区別できるように十分なコントラストが必要です。

前のチャレンジでは、最初の問題に対応するためにテキストの代替手段を用意することを扱いました。最後のチャレンジでは、2つ目の問題に役立つコントラストチェックツールを紹介しました。WCAGで推奨される4.5:1のコントラスト比は、色の使用だけでなくグレースケールの組み合わせにも適用されます。

色覚障害のあるユーザーは、色相で区別が難しい色があり、場合によっては明度でも区別が難しいことがあります。コントラスト比は前景色と背景色の相対的な輝度（または明度）値を使って計算されることを思い出してください。

実際には、4.5:1のコントラスト比は、暗い色に黒を加えて陰影をつけ、明るい色に白を加えて色調を明るくすることで達成できます。カラーホイール上の暗い色は青、紫、マゼンタ、赤の陰影とされ、明るい色はオレンジ、黄色、緑、青緑の色調とされます。

# --instructions--

Camper Catはブログのテキストと背景に色を使う実験をしていますが、現在の緑がかった`background-color`とマルーンのテキスト`color`の組み合わせはコントラスト比が2.5:1です。彼はCSSの`hsl()`プロパティ（色相、彩度、明度を表す）を使って色を宣言しているので、3番目の引数を変えることで色の明度を簡単に調整できます。`background-color`の明度を35%から55%に上げ、`color`の明度を20%から15%に下げてください。これでコントラスト比は5.9:1に改善されます。

# --hints--

テキストの`color`プロパティの明度の値だけを15%に変更してください。

```js
assert.match(code ,/color:\s*?hsl\(0,\s*?55%,\s*?15%\)/gi);
```

`background-color`プロパティの明度の値だけを55%に変更してください。

```js
assert.match(code ,/background-color:\s*?hsl\(120,\s*?25%,\s*?55%\)/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 20%);
    background-color: hsl(120, 25%, 35%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 15%);
    background-color: hsl(120, 25%, 55%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
