---
id: 587d774c367417b2b2512a9d
title: Altテキストを空白にすべき場合を知る
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

前のチャレンジでは、`alt`タグを使うときに`img`属性を含めることが必須であることを学びました。しかし、画像がすでにキャプションで説明されている場合や装飾目的で使われている場合があります。このような場合、`alt`テキストは冗長または不要に見えることがあります。

画像がすでにテキストで説明されているか、ページに意味を加えない場合でも、`img`には`alt`属性が必要ですが、空文字列に設定できます。例を示します。

```html
<img src="visualDecoration.jpeg" alt="">
```

背景画像も通常は「装飾的」とみなされます。ただし、背景画像は通常CSSルールで適用されるため、スクリーンリーダーが処理するマークアップの一部ではありません。

**注意:** キャプション付きの画像には、検索エンジンが画像の内容を分類するのに役立つため、`alt`テキストを含めることをお勧めします。

# --instructions--

Camper Catは自身のウェブサイトのブログ部分の骨組みページをコーディングしました。彼は2つの記事の間に装飾的な侍の刀の画像で視覚的な区切りを追加する予定です。`alt`タグに`img`属性を追加し、空文字列に設定してください。（画像の`src`は実際のファイルにリンクしていませんが、表示に刀が見えなくても問題ありません。）

# --hints--

あなたの`img`タグには`alt`属性が必要です。

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

`alt`属性は空文字列に設定してください。

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```
