---
id: 587d774e367417b2b2512aa0
title: article要素でコンテンツを囲む
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp79S3'
forumTopicId: 301029
dashedName: wrap-content-in-the-article-element
---

# --description--

`article` は、マークアップに意味論的な意味を追加する新しいHTML5要素のひとつです。`article` はセクショニング要素で、独立した自己完結型のコンテンツを囲むために使われます。このタグはブログのエントリー、フォーラムの投稿、ニュース記事などに適しています。

コンテンツが単独で成立するかどうかは通常判断が必要ですが、いくつか簡単なテストを使えます。周囲の文脈をすべて取り除いても、そのコンテンツは意味を成すかどうか自問してください。同様に、テキストの場合はRSSフィードに入っても内容が通じるかどうかを考えてみてください。

支援技術を使う人たちは、整理され意味論的な意味を持つマークアップに頼って内容を理解しています。

**注意:** `section` 要素もHTML5で新しく導入されており、`article` とは少し異なる意味論的な意味を持ちます。`article` は単独のコンテンツ用で、`section` はテーマ的に関連するコンテンツをグループ化するためのものです。必要に応じて互いに入れ子にできます。例えば、本が`article` なら、各章は`section` です。コンテンツのグループ間に関係がない場合は`div` を使います。

`<div>` - コンテンツをグループ化する  
`<section>` - 関連するコンテンツをグループ化する  
`<article>` - 独立した自己完結型のコンテンツをグループ化する  

# --instructions--

Camper Catはブログページの投稿を`article` タグで囲みましたが、一番上の投稿だけ囲み忘れました。`div` タグを`article` タグに変更してください。

# --hints--

コードには`article` タグが3つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('article') ,3);
```

コードに`div` タグは含まれていないはずです。

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <div>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </div>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```
