---
id: 587d778f367417b2b2512aae
title: 説明的なリンクテキストを使ってリンクに意味を持たせる
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437DcV'
forumTopicId: 301013
dashedName: give-links-meaning-by-using-descriptive-link-text
---

# --description--

スクリーンリーダーユーザーは、デバイスが読み上げるコンテンツの種類をさまざまに選べます。これには、ランドマーク要素にスキップしたり（または飛ばしたり）、メインコンテンツにジャンプしたり、見出しからページの概要を得たりすることが含まれます。もう一つの選択肢は、ページ上のリンクだけを聞くことです。

スクリーンリーダーは、リンクテキスト、つまりアンカー（`a`）タグの間にある内容を読み上げることでこれを実現します。「ここをクリック」や「続きを読む」といったリンクのリストは役に立ちません。代わりに、`a`タグ内に簡潔で説明的なテキストを使い、これらのユーザーにより多くの意味を提供してください。

# --instructions--

Camper Catが使っているリンクテキストは、周囲の文脈がないとあまり説明的ではありません。アンカー（`a`）タグを移動して、`information about batteries`の代わりにテキスト`Click here`を囲むようにしてください。

# --hints--

コードでは、アンカー`a`タグを単語`Click here`の周りから単語`information about batteries`の周りに移動させる必要があります。

```js
assert.match(document.querySelector('a')?.textContent, /^(information about batteries)$/g);
```

`a`要素は、空文字列`href`の値を持つ`""`属性を持つべきです。

```js
assert.isEmpty(document.querySelector('a')?.getAttribute('href'));
```

`a`要素は終了タグを持つべきです。

```js
assert.isTrue(code.match(/<\/a>/g)?.length === code.match(/<a href=(''|"")>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. <a href="">Click here</a> for information about batteries</p>
  </article>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. Click here for <a href="">information about batteries</a></p>
  </article>
</body>
```
