---
id: 587d778a367417b2b2512aa6
title: label要素でフォームフィールドのアクセシビリティを向上させる
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMMAN'
forumTopicId: 301016
dashedName: improve-form-field-accessibility-with-the-label-element
---

# --description--

セマンティックなHTMLマークアップでアクセシビリティを向上させるには、適切なタグ名と属性の両方を使うことが重要です。次のいくつかのチャレンジでは、フォームで属性を使う重要なシナリオを扱います。

`label`タグは特定のフォームコントロール項目のテキスト、通常は選択肢の名前やラベルを囲みます。これにより項目に意味が付与され、フォームが読みやすくなります。`for`属性は`label`タグ上にあり、その`label`をフォームコントロールに明示的に関連付け、スクリーンリーダーで使われます。

ラジオボタンとそのラベルについては、Basic HTMLセクションのレッスンで学びました。そのレッスンでは、ラジオボタンのinput要素を`label`要素の中にラベルテキストと一緒に入れて、テキストをクリック可能にしました。もう一つの方法は、このレッスンで説明する`for`属性を使うことです。

`for`属性の値は、フォームコントロールの`id`属性の値と同じでなければなりません。例を示します。

```html
<form>
  <label for="name">Name:</label>
  <input type="text" id="name" name="name">
</form>
```

# --instructions--

Camper Catは自身の考え抜かれたブログ投稿に多くの関心が寄せられることを期待しており、メール登録フォームを含めたいと考えています。メールの`for`に、`label`フィールドの`id`と一致する`input`属性を追加してください。

# --hints--

コードには空でない`for`属性が`label`タグにあるはずです。

```js
assert.isNotEmpty(document.querySelector('label')?.getAttribute('for'));
```

`for`属性の値はメールの`id`の`input`値と一致する必要があります。

```js
assert.equal(document.querySelector('label')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label>Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
