---
id: 587d778e367417b2b2512aab
title: コントラストの高いテキストで読みやすさを向上させる
challengeType: 0
videoUrl: 'https://scrimba.com/c/cKb3nCq'
forumTopicId: 301017
dashedName: improve-readability-with-high-contrast-text
---

# --description--

前景色と背景色のコントラストが低いと、テキストが読みづらくなります。十分なコントラストがあればコンテンツの読みやすさが向上しますが、「十分」とは具体的にどのような意味でしょうか？

ウェブコンテンツアクセシビリティガイドライン（WCAG）では、通常のテキストに対して少なくとも4.5対1のコントラスト比を推奨しています。この比率は2色の相対輝度値を比較して計算されます。同じ色の場合は1対1（コントラストなし）、白と黒の場合は21対1（最も大きなコントラスト）となります。オンラインにはこの比率を計算してくれる多くのコントラストチェックツールがあります。

# --instructions--

Camper Catが最近のブログ投稿で選んだ白い背景に薄いグレーのテキストは、コントラスト比が1.5対1で読みづらくなっています。テキストの`color`を現在のグレー（`#D3D3D3`）からより暗いグレー（`#636363`）に変更して、コントラスト比を6対1に改善してください。

# --hints--

テキストの`color`を`body`のより暗いグレーに変更してください。

```js
const body = document.querySelector('body');
const bodyColor = window.getComputedStyle(body).color; 
assert(bodyColor == 'rgb(99, 99, 99)');
```

`background-color`の`body`は変更しないでください。

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body).backgroundColor; 
assert.equal(backgroundColor , 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: #D3D3D3;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: #636363;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
