---
id: 587d7787367417b2b2512aa1
title: ヘッダーランドマークでスクリーンリーダーのナビゲーションを簡単にする
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB76vtv'
forumTopicId: 301023
dashedName: make-screen-reader-navigation-easier-with-the-header-landmark
---

# --description--

次に紹介するHTML5要素は意味論的な意味を追加し、アクセシビリティを向上させる`header`タグです。これは親タグの導入情報やナビゲーションリンクを囲むために使われ、複数ページの上部に繰り返し表示されるコンテンツに適しています。

`header`は`main`で見た埋め込みランドマーク機能を共有しており、支援技術がそのコンテンツに素早くアクセスできるようにします。

**注意:** `header`はHTMLドキュメントの`body`タグ内で使うことを想定しています。ページのタイトルやメタ情報などを含む`head`要素とは異なります。

# --instructions--

Camper Catは忍者トレーニングに関する素晴らしい記事を書いており、それらのためのページをサイトに追加したいと考えています。現在`div`を含むトップの`h1`を`header`タグに変更してください。

# --hints--

コードには`header`タグが1つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('header'),1);
```

`header`タグは`h1`を囲むようにしてください。

```js
const header = document.querySelector('header');
const children = header?.querySelectorAll(`:scope ${'h1'}`);
assert.lengthOf(children , 1);
```

コードに`div`タグが含まれていないことを確認してください。

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

`header`要素には終了タグが必要です。

```js
assert.isTrue(code.match(/<\/header>/g)?.length === code.match(/<header>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>

  <div>
    <h1>Training with Camper Cat</h1>
  </div>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>

  <header>
    <h1>Training with Camper Cat</h1>
  </header>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
