---
id: 587d7790367417b2b2512ab0
title: tabindex を使って要素にキーボードフォーカスを追加する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMDHW'
forumTopicId: 301027
dashedName: use-tabindex-to-add-keyboard-focus-to-an-element
---

# --description--

HTML の `tabindex` 属性には、要素のキーボードフォーカスに関する3つの異なる機能があります。タグに付けると、その要素がフォーカス可能であることを示します。値は正の整数、負の整数、またはゼロで、その動作を決定します。

リンクやフォームコントロールなどの特定の要素は、ユーザーがページ内をタブで移動すると自動的にキーボードフォーカスを受け取ります。これは HTML ソースのマークアップに現れる順序と同じです。この同じ機能を `div`、`span`、`p` のような他の要素にも `tabindex="0"` 属性を付けることで与えることができます。例を示します。

```html
<div tabindex="0">I need keyboard focus!</div>
```

**注意:** 負の `tabindex` 値（通常は -1）は、要素がフォーカス可能であるがキーボードでは到達できないことを示します。この方法は、ポップアップウィンドウ用の `div` が起動されたときのように、プログラム的にフォーカスを移動させる場合に一般的に使われ、これらのチャレンジの範囲外です。

# --instructions--

Camper Cat はユーザー情報を収集する新しいアンケートを作成しました。入力フィールドは自動的にキーボードフォーカスを受け取ることは知っていますが、タブで項目を移動するときにキーボードユーザーが説明文で一時停止するようにしたいです。`tabindex` タグに `p` 属性を追加し、その値を `0` に設定してください。ボーナスとして、`tabindex` を使うと `:focus` タグで CSS の疑似クラス `p` も機能します。

# --hints--

コードはフォームの説明文を含む `tabindex` タグに `p` 属性を追加する必要があります。

```js
assert.isNotNull(document.querySelector('p')?.getAttribute('tabindex'));
```

コードは `tabindex` タグに `p` 属性を値 0 で設定する必要があります。

```js
assert.equal(document.querySelector('p')?.getAttribute('tabindex'), '0');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p>Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p tabindex="0">Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
