---
id: 587d774c367417b2b2512a9c
title: 視覚障害者のアクセシビリティのために画像にテキスト代替を追加する
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

他のチャレンジで `alt` 属性が `img` タグに使われているのを見たことがあるでしょう。`alt` テキストは画像の内容を説明し、そのテキスト代替を提供します。`alt` 属性は、画像が読み込まれなかったりユーザーが見ることができない場合に役立ちます。検索エンジンも画像の内容を理解して検索結果に含めるために使います。例を示します：

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

視覚障害のある人はスクリーンリーダーを使ってウェブコンテンツを音声インターフェースに変換しています。視覚的にしか情報が提示されていないと情報を得られません。画像の場合、スクリーンリーダーは `alt` 属性にアクセスしてその内容を読み上げ、重要な情報を伝えます。

良い `alt` テキストは読者に画像の簡潔な説明を提供します。画像には必ず `alt` 属性を含めるべきです。HTML5の仕様では、これは必須とされています。

# --instructions--

Camper Cat はコーディング忍者であり本物の忍者でもあり、知識を共有するためのウェブサイトを作っています。彼が使いたいプロフィール写真は彼のスキルを示しており、サイトの訪問者全員に評価されるべきです。`alt` タグに `img` 属性を追加して、Camper Cat が空手をしていることを説明してください。（画像の `src` は実際のファイルにリンクしていないので、表示には `alt` テキストが見えるはずです。）

# --hints--

`img` タグには `alt` 属性があり、空であってはいけません。

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">
```
