---
id: 587d774e367417b2b2512a9f
title: main要素を使ってコンテンツへ直接ジャンプする
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7zuE'
forumTopicId: 301018
dashedName: jump-straight-to-the-content-using-the-main-element
---

# --description--

HTML5は、開発者により多くの選択肢を提供すると同時にアクセシビリティ機能も組み込んだ新しい要素をいくつか導入しました。これらのタグには`main`、`header`、`footer`、`nav`、`article`、`section`などがあります。

デフォルトでは、ブラウザーはこれらの要素を控えめな`div`と似たようにレンダリングします。しかし、適切な場所で使うことでマークアップに追加の意味を与えられます。タグ名だけでその中に含まれる情報の種類を示せるため、そのコンテンツに意味論的な意味が加わります。支援技術はこの情報にアクセスして、より良いページの概要やナビゲーションオプションをユーザーに提供できます。

`main`要素は（その名の通り）メインコンテンツを囲むために使われ、ページに1つだけ存在するべきです。ページの中心的な話題に関連する情報を囲むことを目的としています。ナビゲーションリンクやバナーのようにページ間で繰り返される項目を含めるためのものではありません。

`main`タグには、支援技術がメインコンテンツへ素早く移動するために使える埋め込みのランドマーク機能もあります。ページの上部に「メインコンテンツへジャンプ」リンクを見たことがあるなら、`main`タグを使うことで支援機器にその機能が自動的に与えられます。

# --instructions--

Camper Catは忍者武器のページに大きなアイデアを持っています。`header`と`footer`（他のチャレンジで扱います）の間に開く`main`タグと閉じる`main`タグを追加して、マークアップを設定するのを手伝ってください。`main`タグは今のところ空のままにしてください。

# --hints--

コードには`main`タグが1つあるべきです。

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`main`タグは閉じる`header`タグと開く`footer`タグの間にあるべきです。

```js
assert.match(code,/<\/header>\s*?<main>\s*?<\/main>/gi);
```

# --seed--

## --seed-contents--

```html
<header>
  <h1>Weapons of the Ninja</h1>
</header>



<footer></footer>
```

# --solutions--

```html
<header>
  <h1>Weapons of the Ninja</h1>
</header>
<main>

</main>
<footer></footer>
```
