---
id: 63b60aaaa65f8922bfce6b7e
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

في عنصر `span` الخاص بك، أنشئ عنصر `label` لـ `entry-dropdown` وأعطه النص `Add food or exercise:`. ثم أنشئ عنصر `select` مع تعيين الخاصية `id` إلى `entry-dropdown` والخاصية `name` إلى `options`. أسفله، أضف عنصر `button` مع تعيين الخاصية `id` إلى `add-entry` والنص `Add Entry`.

أعطِ عنصر `button` الخاص بك خاصية `type` بقيمة `button` لمنع الإرسال التلقائي للنموذج.

# --hints--

يجب أن تضيف عنصر `label` إلى عنصر `span` الخاص بك.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

يجب أن يحتوي عنصر `label` الجديد على خاصية `for` بقيمة `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

يجب أن يحتوي عنصر `label` الجديد على النص `Add food or exercise:`.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

يجب أن تضيف عنصر `select` إلى عنصر `span` الخاص بك.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

يجب أن يأتي عنصر `select` بعد عنصر `label` الخاص بك.

```js
assert.equal(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName, 'LABEL');
```

يجب أن يحتوي عنصر `select` الجديد على خاصية `id` بقيمة `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

يجب أن يحتوي عنصر `select` الجديد على خاصية `name` بقيمة `options`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

يجب أن يكون عنصر `select` فارغًا ويحتوي على وسم إغلاق `</select>`.

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

يجب أن تضيف عنصر `button` إلى عنصر `span` الخاص بك.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

يجب أن يأتي عنصر `button` بعد عنصر `select` الخاص بك.

```js
assert.equal(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName, 'SELECT');
```

يجب أن يحتوي عنصر `button` الجديد على خاصية `id` بقيمة `add-entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

يجب أن يحتوي عنصر `button` الجديد على النص `Add Entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
