---
id: 63b60821c855d01b1eda3c0b
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

بعد ذلك، أنشئ عنصر `fieldset` مع تعيين `id` إلى `lunch`.

داخل ذلك العنصر، أنشئ عنصر `legend` بالنص `Lunch`، وعنصر `div` فارغ مع تعيين `class` إلى `input-container`.

# --hints--

يجب أن تنشئ عنصر `fieldset` ثاني داخل نموذج `form` الخاص بك.

```js
assert.exists(document.querySelectorAll('form fieldset')[1]);
```

يجب أن يأتي عنصر `fieldset` الثاني بعد عنصر `fieldset` الأول.

```js
assert.equal(document.querySelectorAll('form fieldset')[1]?.previousElementSibling?.tagName, "FIELDSET");
```

يجب أن يحتوي عنصر `fieldset` الثاني على `id` معين إلى `lunch`.

```js
assert.equal(document.querySelectorAll('form fieldset')[1]?.id, "lunch");
```

يجب أن يحتوي عنصر `fieldset` الثاني على عنصر `legend`.

```js
assert.exists(document.querySelectorAll('form fieldset')[1]?.querySelector('legend'));
```

يجب أن يحتوي عنصر `legend` الجديد على النص `Lunch`.

```js
assert.equal(document.querySelectorAll('form fieldset')[1]?.querySelector('legend')?.innerText, "Lunch");
```

يجب أن يحتوي عنصر `fieldset` الثاني على عنصر `div`.

```js
assert.exists(document.querySelectorAll('form fieldset')[1]?.querySelector('div'));
```

يجب أن يحتوي عنصر `div` الجديد على `class` معين إلى `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[1]?.querySelector('div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
