---
id: 63507bcbfe4ede356e624395
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Sasa baada ya kupata thamani ya ingizo, unahitaji kuigawanya kuwa safu ya nambari. Tumia njia ya `.split()` kufanya hivyo.

Njia ya `.split()` inachukua mfuatano wa herufi na kuugawanya kuwa safu ya mfuatano wa herufi. Unaweza kuipatia mfuatano wa herufi au usemi wa kawaida (regex) utakaotumika kama kiganja. Kwa mfano, `string.split(",")` itagawanya mfuatano huo kila mara ambapo kuna koma na kurudisha safu ya mfuatano wa herufi.

Tumia usemi wa kawaida wa `/,\s*/g` kugawanya mfuatano wa herufi wa `value` kwa kutumia koma. Unaweza kuubadilisha kulingana na idadi ya nafasi zinazotenganisha thamani zako. Hifadhi safu hiyo katika kigezo cha `array`.

# --hints--

Kitendakazi chako cha `calculate` kinapaswa kuwa na kigezo cha `array`.

```js
assert.match(code.toString(), /array/);
```

Kitendakazi chako cha `calculate` kinapaswa kuita njia ya `.split()` kwenye kigezo cha `value`.

```js
assert.match(code.toString(), /value\.split()/);
```

Kitendakazi chako cha `calculate` kinapaswa kuweka thamani ya matokeo ya njia ya `.split()` kwenye kigezo cha `array`.

```js
assert.match(code.toString(),  /array\s*=\s*value\.split()/);
```

Unapaswa kutumia `/,\s*/g` kama kiganja cha njia ya `split()`.

```js
assert.match(code.toString(), /value\.split\(\s*\/,\s*\\s*\*\s*\/g\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
}
--fcc-editable-region--
```
