---
id: 63507c4b63731437227b0134
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Njia ya `.map()` inachukua kitendakazi cha mwitiko kama hoja yake ya kwanza. Kitendakazi hiki cha mwitiko kinachukua hoja chache, lakini ya kwanza ni kipengele kinachosindikwa kwa sasa. Hapa kuna mfano:

```js
array.map(el => {

})
```

Kitendakazi cha mwitiko kinahitaji kurudisha thamani. Katika kesi hii, unataka kurudisha thamani ya kila kipengele kilichobadilishwa kuwa nambari. Unaweza kufanya hivi kwa kutumia kijenzi cha `Number()`, ukipitia kipengele kama hoja.

Ongeza kitendakazi cha mwitiko kwenye njia yako ya `.map()` kinachobadilisha kila kipengele kuwa nambari.

# --hints--

Njia yako ya `.map()` inapaswa kuwa na kitendakazi cha mwitiko.

```js
assert.match(calculate.toString(), /array\.map\(\s*(\(\s*\w+\s*\)|\w+)\s*=>|array\.map\(\s*function\s*\(\s*\w+\s*\)\s*\{|array\.map\(\s*Number\s*\)/)
```

Kitendakazi chako cha mwitiko kinapaswa kutumia kijenzi cha `Number` kubadilisha kila kipengele cha safu ya data kuwa nambari.

```js
assert.match(calculate.toString(), /Number\(\s*\w+\s*\)|Number/);
```

Kitendakazi chako cha mwitiko hakipaswi kutumia neno kuu la `new` pamoja na kijenzi cha `Number`.

```js
assert.notMatch(calculate.toString(), /new/);
```

Kitendakazi chako cha mwitiko kinapaswa kurudisha kipengele kilichobadilishwa kuwa nambari.

```js
assert.match(calculate.toString(), /(array\.map\(\s*(\(\s*\w+\s*\)|\w+)\s*=>|array\.map\(\s*function\s*\(\s*\w+\s*\)\s*\{)\s*(return\s+)?Number\(\s*\w+\s*\)|array\.map\(\s*Number\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map();
}
--fcc-editable-region--
```
