---
id: 6352ea3a5b79e614ee2282fd
title: Hatua 38
challengeType: 0
dashedName: step-38
---

# --description--

Kuna baadhi ya kesi maalum za kuzingatia wakati wa kuhesabu mode ya seti ya data. Kwanza, ikiwa kila thamani inaonekana mara moja kwa idadi sawa, hakuna mode.

Ili kuhesabu hili, utatumia `Set`. `Set` ni muundo wa data unaoruhusu tu thamani za kipekee. Ukipitisha safu ya data katika kijenzi cha `Set`, itatoa thamani zote zinazojirudia.

Anza kwa kuunda tamko la `if`. Katika masharti, unda `Set` na `new Set()` na upitishe `Object.values()` ya kitu chako cha `counts`. Ikiwa sifa ya `size` ya `Set` hii ni sawa na `1`, hiyo inaonyesha kila thamani inaonekana mara moja kwa idadi sawa. Katika hali hii, rudisha `null` kutoka kwenye kitendakazi chako.

# --hints--

Kitendakazi chako cha `getMode` kinapaswa kuwa na tamko la `if`.

```js
const functionContents = code.split(/const\s+getMode\s*=\s*/)?.[1]?.split(/[^{]\}[^;)]/)?.[0];
assert.match(functionContents, /if\s*\(/);
```

Tamko lako la `if` linapaswa kuunda `Set` mpya na kupitisha `Object.values()` ya kitu chako cha `counts`.

```js
assert.match(getMode.toString(), /if\s*\(\s*new\s+Set\s*\(\s*Object\.values\s*\(\s*counts\s*\)/);
```

Tamko lako la `if` linapaswa kuangalia kama sifa ya `size` ya `Set` mpya ni sawa na `1`.

```js
assert.match(getMode.toString(), /if\s*\(\s*new\s+Set\s*\(\s*Object\.values\s*\(\s*counts\s*\)\s*\)\s*\.size\s*===\s*1/);
```

Tamko lako la `if` linapaswa kurudisha `null` ikiwa sifa ya `size` ya `Set` mpya ni sawa na `1`.

```js
assert.match(getMode.toString(), /if\s*\(\s*new\s+Set\s*\(\s*Object\.values\s*\(\s*counts\s*\)\s*\)\.size\s*===\s*1\s*\)\s*\{?\s*return\s+null\s*;?\s*\}?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

const getMedian = (array) => {
  const sorted = array.toSorted((a, b) => a - b);
  const median =
    sorted.length % 2 === 0
      ? getMean([sorted[sorted.length / 2], sorted[sorted.length / 2 - 1]])
      : sorted[Math.floor(sorted.length / 2)];
  return median;
}

--fcc-editable-region--
const getMode = (array) => {
  const counts = {};
  array.forEach(el => counts[el] = counts[el] ? counts[el] + 1 : 1);
}
--fcc-editable-region--


const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
  
  const mean = getMean(numbers);
  const median = getMedian(numbers);

  document.querySelector("#mean").textContent = mean;
  document.querySelector("#median").textContent = median;
}
```
