---
id: 616971b3cd990186b66c99a1
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

Now that your external CSS file is set up, you can start styling the page.

Kama ukumbusho, hii ndio jinsi ya kulenga kipengele cha aya na kukipanga kulia:

```css
p {
  text-align: right;
}
```

Unda sheria mpya ya CSS inayolenga class ya `h1` na uweke sifa ya `text-align` kuwa `center`.

# --hints--

Unapaswa kutumia kichaguzi cha `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

Sheria yako ya CSS ya `h1` inapaswa kuwa na sifa ya `text-align` iliyowekwa kuwa `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
