---
id: 681cb1b0dab50c87ddb2e518
title: "挑战 7：目标和"
challengeType: 29
dashedName: challenge-7
---

# --description--

给定一个数字数组和一个整数目标，查找数组中两个相加等于目标值的唯一数字。返回包含这两个数字索引的数组，如果没有两个数字的和等于目标，则返回 `"Target not found"`。

- 返回的数组应按索引升序排列。

# --hints--

`find_target([2, 7, 11, 15], 9)` 应该返回 `[0, 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([2, 7, 11, 15], 9), [0, 1])`)
}})
```

`find_target([3, 2, 4, 5], 6)` 应该返回 `[1, 2]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([3, 2, 4, 5], 6), [1, 2])`)
}})
```

`find_target([1, 3, 5, 6, 7, 8], 15)` 应该返回 `[4, 5]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 6, 7, 8], 15), [4, 5])`)
}})
```

`find_target([1, 3, 5, 7], 14)` 应该返回 `'Target not found'`。

```js

({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 7], 14), "Target not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_target(arr, target):

    return arr
```

# --solutions--

```py
def find_target(arr, target):
    for i in range(len(arr)):
        for j in range(i + 1, len(arr)):
            if arr[i] + arr[j] == target:
                return [i, j]
    return 'Target not found'
```
