---
id: 681cb1b0dab50c87ddb2e519
title: "挑战 8：阶乘器"
challengeType: 29
dashedName: challenge-8
---

# --description--

给定一个从零到 20 的整数，返回该数字的阶乘。一个数字的阶乘是从 1 到该数字之间所有数字的乘积。

- 零的阶乘是 1。

# --hints--

`factorial(0)` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(0), 1)`)
}})
```

`factorial(5)` 应该返回 `120`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(5), 120)`)
}})
```

`factorial(20)` 应该返回 `2432902008176640000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(20), 2432902008176640000)`)
}})
```

# --seed--

## --seed-contents--

```py
def factorial(n):

    return n
```

# --solutions--

```py
def factorial(n):
    result = 1
    for i in range(1, n + 1):
        result *= i
    return result
```
