---
id: 681cb1a2dab50c87ddb2e514
title: "挑战 3：斐波那契序列"
challengeType: 29
dashedName: challenge-3
---

# --description--

斐波那契序列是一串数字，其中每个数字都是前述的两个数字之和。当以 `0` 和 `1` 开始时，序列中的前 10 个数字是 `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34`。

给定一个包含斐波那契序列前两个数字的数组，以及一个表现序列长度的整数，返回一个包含给定长度序列的数组。

- 你的函数应处理长度大于或相等于 0 的任意序列。
- 如果长度为 0，返回一个空数组。
- 请注意，起始数字是序列的一部分。

# --hints--


`fibonacci_sequence([0, 1], 20)` 应该返回 `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181])`)
}})
```

`fibonacci_sequence([21, 32], 1)` 应该返回 `[21]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([21, 32], 1), [21])`)
}})
```

`fibonacci_sequence([0, 1], 0)` 应该返回 `[]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 0), [])`)
}})
```

`fibonacci_sequence([10, 20], 2)` 应该返回 `[10, 20]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([10, 20], 2), [10, 20])`)
}})
```

`fibonacci_sequence([123456789, 987654321], 5)` 应该返回 `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541])`)
}})
```

# --seed--

## --seed-contents--

```py
def fibonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def fibonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2]
        sequence.append(next_value)

    return sequence
```
