---
id: 681cb1afdab50c87ddb2e515
title: "挑战 4：S  P  A  C  E  J  A  M"
challengeType: 29
dashedName: challenge-4
---

# --description--

给定一个字串，移除字串中的所有空格，在每个字符之间插入两个空格，将所有字母转换为大写，并返回结果。

- 非字母字符应保持不变（空格除外）。

# --hints--

`space_jam("freeCodeCamp")` 应该返回 `"F  R  E  E  C  O  D  E  C  A  M  P"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("   free   Code   Camp   ")` 应该返回 `"F  R  E  E  C  O  D  E  C  A  M  P"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("Hello World?!")` 应该返回 `"H  E  L  L  O  W  O  R  L  D  ?  !"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !")`)
}})
```

`space_jam("C@t$ & D0g$")` 应该返回 `"C  @  T  $  &  D  0  G  $"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $")`)
}})
```

`space_jam("allyourbase")` 应该返回 `"A  L  L  Y  O  U  R  B  A  S  E"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E")`)
}})
```

# --seed--

## --seed-contents--

```py
def space_jam(s):

    return s
```

# --solutions--

```py
def space_jam(s):
    s = s.replace(" ", "")
    s = "  ".join(s)
    return s.upper()
```
