---
id: 681cb05adab50c87ddb2e513
title: "挑战 2：基础 `Check`"
challengeType: 29
dashedName: challenge-2
---

# --description--

给定一个表现数字的字串，以及一个从 2 到 36 的整数进制，确定该数字在该进制中是否有效。

- 该字串可能包含整数，以及大写或小写字符。
- 该检查应不区分分支。
- 基数可以是 2-36 之间的任何数字。
- 如果每个字符都是给定进制中的有效数字，则数字是有效的。
- 有效进制的数字示例：
  - 基数 2：0-1
  - 八进制：0-7
  - 十进制：0-9
  - 基数 16：0-9 和 A-F
  - 基数 36：0-9 和 A-Z

# --hints--

`is_valid_number("10101", 2)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("10101", 2), True)`);
}})
```

`is_valid_number("10201", 2)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("10201", 2), False)`)
}})
```

`is_valid_number("76543210", 8)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("76543210", 8), True)`)
}})
```

`is_valid_number("9876543210", 8)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("9876543210", 8), False)`)
}})
```

`is_valid_number("9876543210", 10)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("9876543210", 10), True)`)
}})
```

`is_valid_number("ABC", 10)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 10), False)`)
}})
```

`is_valid_number("ABC", 16)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 16), True)`)
}})
```

`is_valid_number("Z", 36)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("Z", 36), True)`)
}})
```

`is_valid_number("ABC", 20)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 20), True)`)
}})
```

`is_valid_number("4B4BA9", 16)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("4B4BA9", 16), True)`)
}})
```

`is_valid_number("5G3F8F", 16)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("5G3F8F", 16), False)`)
}})
```

`is_valid_number("5G3F8F", 17)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("5G3F8F", 17), True)`)
}})
```

`is_valid_number("abc", 10)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("abc", 10), False)`)
}})
```

`is_valid_number("abc", 16)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("abc", 16), True)`)
}})
```

`is_valid_number("AbC", 16)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("AbC", 16), True)`)
}})
```

`is_valid_number("z", 36)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("z", 36), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_number(n, base):

    return n
```

# --solutions--

```py
def is_valid_number(n, base):
    allChars = "0123456789abcdefghijklmnopqrstuvwxyz"
    newN = n.lower()

    availableChars = allChars[0:base]

    for char in newN:
        if char not in availableChars:
            return False

    return True
```
