---
id: 6821ebdf237de8297eaee793
title: "挑战 17：操作顺序"
challengeType: 29
dashedName: challenge-17
---

# --description--

给定一个整数数组和一个字串操作符数组，将操作从左到右循序地应用到数字上。根据需要重复操作，直到所有数字都被使用。返回最终结果。

例如，给定 `[1, 2, 3, 4, 5]` 和 `['+', '*']`，返回从左到右计算 `1 + 2 * 3 + 4 * 5` 的结果，忽略标准的操作顺序。

- 有效的操作符是 `+`、`-`、`*`、`/` 和 `%`。

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` 应该返回 `3`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3)`)
}})
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` 应该返回 `38`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38)`)
}})
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` 应该返回 `60`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60)`)
}})
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` 应该返回 `30`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30)`)
}})
```

`evaluate([33, 11, 29, 13], ['/', '-'])` 应该返回 `-2`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([33, 11, 29, 13], ['/', '-']), -2)`)
}})
```

# --seed--

## --seed-contents--

```py
def evaluate(numbers, operators):

    return numbers
```

# --solutions--

```py
def do_math(a, b, operator):
    if operator == '+':
        return a + b
    elif operator == '-':
        return a - b
    elif operator == '*':
        return a * b
    elif operator == '/':
        return a / b
    else:
        return a % b

def evaluate(numbers, operators):
    total = numbers[0]

    for i in range(1, len(numbers)):
        operator = operators[(i - 1) % len(operators)]
        total = do_math(total, numbers[i], operator)

    return total
```
