---
id: 6821ebda237de8297eaee792
title: "挑战 16：反转括号"
challengeType: 29
dashedName: challenge-16
---

# --description--

给定一个包含正确嵌套括号的字串，使用以下规则返回解码后的字串：

- 每个括号点对内的所有字符都应反转。
- 最终结果中应删除括号。
- 如果括号是嵌套的，应先反转最内层的点对，然后将其结果包含在外层点对的反转中。
- 假设所有括号都是均衡且正确嵌套的。

# --hints--

`decode("(f(b(dc)e)a)")` 应该返回 `"abcdef"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` 应该返回 `"Can you read this?"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` 应该返回 `"freeCodeCamp"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
