---
id: 6821ec02237de8297eaee79a
title: "挑战 24：全字母句"
challengeType: 29
dashedName: challenge-24
---

# --description--

给定一个单词或句子以及一串小写字母，判断该单词或句子是否至少使用了给定集合中的所有字母一次且没有使用其他字母。

- 忽略单词或句子中的非字母字符。
- 忽略单词或句子的字母大小写。

# --hints--

`is_pangram("hello", "helo")` 应该返回 `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "helo"), True)`)
}})
```

`is_pangram("hello", "hel")` 应该返回 `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "hel"), False)`)
}})
```

`is_pangram("hello", "helow")` 应该返回 `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "helow"), False)`)
}})
```

`is_pangram("hello world", "helowrd")` 应该返回 `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello world", "helowrd"), True)`)
}})
```

`is_pangram("Hello World!", "helowrd")` 应该返回 `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("Hello World!", "helowrd"), True)`)
}})
```

`is_pangram("Hello World!", "heliowrd")` 应该返回 `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("Hello World!", "heliowrd"), False)`)
}})
```

`is_pangram("freeCodeCamp", "frcdmp")` 应该返回 `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("freeCodeCamp", "frcdmp"), False)`)
}})
```

`is_pangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz")` 应该返回 `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_pangram(sentence, letters):

    return sentence
```

# --solutions--

```py
import re
def is_pangram(sentence, letters):
    used_letters = []
    for char in sentence.lower():
        if re.match(r'[a-z]', char) and char not in used_letters:
            used_letters.append(char)

    sorted_letters = ''.join(sorted(letters.lower()))
    sorted_used_letters = ''.join(sorted(used_letters))

    return sorted_letters == sorted_used_letters
```
