---
id: 6821ebce237de8297eaee790
title: "挑战 14：字符战斗"
challengeType: 29
dashedName: challenge-14
---

# --description--

给定两个字串，分别表现你的军队和对方军队，你的军队中的每个字符与对方军队中相同位置的字符按照以下规则进行战斗：

- 字符 `a-z` 的强度分别为 `1-26`。
- 字符 `A-Z` 的强度分别为 `27-52`。
- 数字 `0-9` 的强度等于其面值。
- 所有其他字符的值为零。
- 每个 `character` 只能参加一场战斗。

对于每场战斗，较强的`character`获胜。胜利次数更多的军队赢得战争。`Return`以下值：

- `"Opponent retreated"` 如果你的军队中有比对方军队更多的字符。
- `"We retreated"` 如果敌军的字符数量多于你的。
- `"We won"` 如果你的军队赢得了更多战斗。
- 如果敌军赢得更多战斗，则返回 `"We lost"`。
- `"It was a tie"` 如果双方军队赢得相同数量的战斗。

# --hints--

`battle("Hello", "World")` 应该返回 `"We lost"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` 应该返回 `"We won"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` 应该返回 `"We won"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` 应该返回 `"Opponent retreated"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` 应该返回 `"We retreated"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` 应该返回 `"It was a tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` 应该返回 `"It was a tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
