---
id: 6821ebea237de8297eaee795
title: "挑战 19：烛光"
challengeType: 29
dashedName: challenge-19
---

# --description--

给定一个整数，表现你开始时拥有的蜡烛数量，以及一个整数，表现刻录多少根用过的蜡烛可以创建一根新蜡烛，返回在创建并刻录尽可能多的蜡烛后你将使用的蜡烛数量。

例如，如果给定 7 根蜡烛且需要 2 根刻录的蜡烛来制作一根新建的蜡烛：

1. 刻录 7 根蜡烛以获得 7 个剩余物，
2. 将 6 个剩余物回收成 3 根新建蜡烛（剩余 1 个）
3. 刻录 3 根蜡烛以获得 3 个更多剩余物（共 4 个），
4. 将 4 个剩余物回收成 2 根新建蜡烛，
5. 刻录 2 根蜡烛以获得 2 个剩余物，
6. 将 2 个剩余物回收成 1 个新建蜡烛，
7. 刻录 1 根蜡烛。

在示例中你将总共刻录 13 根蜡烛。

# --hints--

`burn_candles(7, 2)` 应该返回 `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

`burn_candles(10, 5)` 应该返回 `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

`burn_candles(20, 3)` 应该返回 `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

`burn_candles(17, 4)` 应该返回 `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

`burn_candles(2345, 3)` 应该返回 `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```
