---
id: 6821ebe4237de8297eaee794
title: "挑战 18：第二名"
challengeType: 29
dashedName: challenge-18
---

# --description--

给定一个表现不同笔记本电脑价格的整数数组，以及一个表现你预算的整数，返回：

1. 如果它在你的预算内，则为第二贵的笔记本电脑，或者
2. 在你的预算范围内最昂贵的笔记本电脑，或者
3. 如果没有笔记本电脑在你的预算内，则为 `0`。

- 应忽略价格副本。

# --hints--

`get_laptop_cost([1500, 2000, 1800, 1400], 1900)` 应该返回 `1800`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900)` 应该返回 `1800`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 2200)` 应该返回 `1899`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 2200), 1899)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 1000)` 应该返回 `0`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 1000), 0)`)
}})
```

`get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` 应该返回 `1400`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_laptop_cost(laptops, budget):

    return laptops
```

# --solutions--

```py
def get_laptop_cost(laptops, budget):
    laptops = sorted(set(laptops), reverse=True)

    if len(laptops) >= 2 and budget >= laptops[1]:
        return laptops[1]
    if not laptops or budget < laptops[-1]:
        return 0

    for price in laptops[2:]:
        if budget >= price:
            return price

    return 0
```
