---
id: 6821ebfd237de8297eaee799
title: "挑战 23：RGB 转 Hex"
challengeType: 29
dashedName: challenge-23
---

# --description--

给定一个 CSS `rgb(r, g, b)` 颜色字串，返回其十六进制等价物。

以下是给定 `input` 的一些示例输出：

| Input   | Output   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- 将所有字母转换为小写。
- 返回一个 `#`，后跟六个字符。不要使用任何简写值。

# --hints--

`rgb_to_hex("rgb(255, 255, 255)")` 应该返回 `"#ffffff"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(255, 255, 255)"), "#ffffff")`)
}})
```

`rgb_to_hex("rgb(1, 11, 111)")` 应该返回 `"#010b6f"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(1, 11, 111)"), "#010b6f")`)
}})
```

`rgb_to_hex("rgb(173, 216, 230)")` 应该返回 `"#add8e6"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(173, 216, 230)"), "#add8e6")`)
}})
```

`rgb_to_hex("rgb(79, 123, 201)")` 应该返回 `"#4f7bc9"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rgb_to_hex("rgb(79, 123, 201)"), "#4f7bc9")`)
}})
```

# --seed--

## --seed-contents--

```py
def rgb_to_hex(rgb):

    return rgb
```

# --solutions--

```py
def rgb_to_hex(rgb):
    import re
    match = re.findall(r'\d+', rgb)
    r, g, b = [max(0, min(255, int(x))) for x in match[:3]]
    return f'#{r:02x}{g:02x}{b:02x}'
```
