---
id: 6821ebd4237de8297eaee791
title: "挑战 15：驼峰式大小写"
challengeType: 29
dashedName: challenge-15
---

# --description--

给定一个字串，使用以下规则返回其驼峰式写法版本：

- 字串参数中的单词由以下集合中的一个或多个字符分隔：空格（` `）、破折号（`-`）或下划线（`_`）。将这些字符的任意序列视为单词分隔符。
- 第一个单词应全部小写。
- 每个后续单词应以大写字母开头，其余部分为小写。
- 所有空格和分隔符都应删除。

# --hints--

`to_camel_case("hello world")` 应该返回 `"helloWorld"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("hello world"), "helloWorld")`)
}})
```

`to_camel_case("HELLO WORLD")` 应该返回 `"helloWorld"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("HELLO WORLD"), "helloWorld")`)
}})
```

`to_camel_case("secret agent-X")` 应该返回 `"secretAgentX"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("secret agent-X"), "secretAgentX")`)
}})
```

`to_camel_case("FREE cODE cAMP")` 应该返回 `"freeCodeCamp"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("FREE cODE cAMP"), "freeCodeCamp")`)
}})
```

`to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` 应该返回 `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_camel_case(s):

    return s
```

# --solutions--

```py
import re
def to_camel_case(s):
    words = re.split(r'[_\- ]+', s)

    camel = [
        words[0].lower() if words else ''
    ] + [
        word.capitalize() for word in words[1:]
    ]

    return ''.join(camel)
```
