---
id: 6821ebf8237de8297eaee798
title: "挑战 22：Tribonacci 序列"
challengeType: 29
dashedName: challenge-22
---

# --description--

Tribonacci 序列是一串数字，其中每个数字都是前述的三个数字之和。当以 `0`、`0` 和 `1` 开始时，序列中的前 10 个数字是 `0`、`0`、`1`、`1`、`2`、`4`、`7`、`13`、`24`、`44`。

给定一个包含 Tribonacci 序列前三个数字的数组，以及一个表现序列长度的整数，返回一个包含给定长度序列的数组。

- 你的函数应处理长度大于或相等于 0 的任意序列。
- 如果长度为 0，返回一个空数组。
- 请注意，起始数字是序列的一部分。

# --hints--

`tribonacci_sequence([0, 0, 1], 20)` 应该返回 `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513])`)
}})
```

`tribonacci_sequence([21, 32, 43], 1)` 应该返回 `[21]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([21, 32, 43], 1), [21])`)
}})
```

`tribonacci_sequence([0, 0, 1], 0)` 应返回 `[]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 0), [])`)
}})
```

`tribonacci_sequence([10, 20, 30], 2)` 应该返回 `[10, 20]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 2), [10, 20])`)
}})
```

`tribonacci_sequence([10, 20, 30], 3)` 应该返回 `[10, 20, 30]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 3), [10, 20, 30])`)
}})
```

`tribonacci_sequence([123, 456, 789], 8)` 应该返回 `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134])`)
}})
```

# --seed--

## --seed-contents--

```py
def tribonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def tribonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return [start_sequence[0], start_sequence[1]]
    if length == 3:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2] + sequence[-3]
        sequence.append(next_value)

    return sequence
```
