---
id: 68216eb60f957572e7c340c4
title: "挑战 11：英里配速"
challengeType: 29
dashedName: challenge-11
---

# --description--

给定跑过的英里数和跑完这些英里所用的时间（格式为 `"MM:SS"`（分钟:秒）），返回一个表示每英里平均用时的字串，格式为 `"MM:SS"`。

- 在需要时添加前导零。

# --hints--

`mile_pace(3, "24:00")` 应该返回 `"08:00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(3, "24:00"), "08:00")`)
}})
```

`mile_pace(1, "06:45")` 应该返回 `"06:45"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(1, "06:45"), "06:45")`)
}})
```

`mile_pace(2, "07:00")` 应该返回 `"03:30"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(2, "07:00"), "03:30")`)
}})
```

`mile_pace(26.2, "120:35")` 应该返回 `"04:36"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(26.2, "120:35"), "04:36")`)
}})
```

# --seed--

## --seed-contents--

```py
def mile_pace(miles, duration):

    return miles
```

# --solutions--

```py
def mile_pace(miles, duration):
    minutes, seconds = map(int, duration.split(":"))
    total_seconds = minutes * 60 + seconds
    avg_seconds_per_mile = total_seconds / miles

    avg_minutes = int(avg_seconds_per_mile // 60)
    avg_seconds = round(avg_seconds_per_mile % 60)

    return f"{avg_minutes:02d}:{avg_seconds:02d}"
```
