---
id: 68216ef80f957572e7c340c5
title: "挑战 12：消息解码器"
challengeType: 29
dashedName: challenge-12
---

# --description--

给定一个秘密消息字串，以及一个表现用于对消息进行编码的字母位移数量的整数，返回解码后的字串。

- 正数表示消息在字母表中被转发了。
- 负数表示消息在字母表中向后移动。
- 分支区分大小写，解码后的字符应保持其编码对应字符的大小写。
- 非字母字符不应被解码。

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` 应该返回 `"This is a secret message."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.")`)
}})
```

`decode("Byffi Qilfx!", 20)` 应该返回 `"Hello World!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Byffi Qilfx!", 20), "Hello World!")`)
}})
```

`decode("Zqd xnt njzx?", -1)` 应该返回 `"Are you okay?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Zqd xnt njzx?", -1), "Are you okay?")`)
}})
```

`decode("oannLxmnLjvy", 9)` 应该返回 `"freeCodeCamp"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("oannLxmnLjvy", 9), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message, shift):

    return message
```

# --solutions--

```py
def decode(message, shift):
    decoded_message = []
    for char in message:
        if char.isalpha():
            base = ord('a') if char.islower() else ord('A')
            new_char = chr((ord(char) - base - shift) % 26 + base)
            decoded_message.append(new_char)
        else:
            decoded_message.append(char)
    return ''.join(decoded_message)
```
