---
id: 68adce01c0e1144d0a902956
title: "挑战 25：元音重复器"
challengeType: 29
dashedName: challenge-25
---

# --description--

给定一个字串，返回该字串的新建版本，其中每个元音的重复次数比你遇到的前一个元音多一次。例如，句子中的第一个元音保持不变。第二个元音应连续出现两次。第三个元音应连续出现三次，依此类推。

- 字母 `a`、`e`、`i`、`o` 和 `u`，无论是大写还是小写，都被视为元音。
- 原始元音应保持其分支。
- 重复的元音应为小写。
- 所有非元音字符应保持其原始大小写。

# --hints--

`repeat_vowels("hello world")` 应该返回 `"helloo wooorld"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` 应该返回 `"freeeCooodeeeeCaaaaamp"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` 应该返回 `"AEeIiiOoooUuuuu"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` 应该返回 `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
