---
id: 68adce01c0e1144d0a90295c
title: "挑战 28：罗马数字解析器"
challengeType: 29
dashedName: challenge-28
---

# --description--

给定一个表现罗马数字的字串，返回其整数值。

罗马数字由以下符号和数值组成：

| Symbol | Value |
|--------|-------|
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

- 数字从左到右读取。如果较小的数字出现在较大的数字之前，则该值被减去。否则，值被相加。

# --hints--

`parse_roman_numeral("III")` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("III"), 3)`)
}})
```

`parse_roman_numeral("IV")` 应该返回 `4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("IV"), 4)`)
}})
```

`parse_roman_numeral("XXVI")` 应该返回 `26`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XXVI"), 26)`)
}})
```

`parse_roman_numeral("XCIX")` 应该返回 `99`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("XCIX"), 99)`)
}})
```

`parse_roman_numeral("CDLX")` 应该返回 `460`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("CDLX"), 460)`)
}})
```

`parse_roman_numeral("DIV")` 应该返回 `504`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("DIV"), 504)`)
}})
```

`parse_roman_numeral("MMXXV")` 应该返回 `2025`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_roman_numeral("MMXXV"), 2025)`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_roman_numeral(numeral):

    return numeral
```

# --solutions--

```py
def parse_roman_numeral(numeral):
    roman_map = {
        "I": 1,
        "V": 5,
        "X": 10,
        "L": 50,
        "C": 100,
        "D": 500,
        "M": 1000
    }

    total = 0
    for i in range(len(numeral)):
        current = roman_map[numeral[i]]
        next_val = roman_map[numeral[i + 1]] if i + 1 < len(numeral) else 0

        if current < next_val:
            total -= current
        else:
            total += current

    return total
```
