---
id: 68adce01c0e1144d0a90295e
title: "挑战 29：首字母缩略词生成器"
challengeType: 29
dashedName: challenge-29
---

# --description--

给定一个包含一个或多个单词的字串，使用以下约束返回这些单词的首字母缩写：

- 缩写应由每个单词的首字母大写组成，除非另有说明。
- 缩写应忽略这些字的第一个字母，除非它们是给定字串的第一个字：`a`、`for`、`an`、`and`、`by` 和 `of`。
- 缩写字母应按给定顺序返回。
- 缩写不应包含任何空格。

# --hints--

`build_acronym("Search Engine Optimization")` 应该返回 `"SEO"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Search Engine Optimization"), "SEO")`)
}})
```

`build_acronym("Frequently Asked Questions")` 应该返回 `"FAQ"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Frequently Asked Questions"), "FAQ")`)
}})
```

`build_acronym("National Aeronautics and Space Administration")` 应该返回 `"NASA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("National Aeronautics and Space Administration"), "NASA")`)
}})
```

`build_acronym("Federal Bureau of Investigation")` 应该返回 `"FBI"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Federal Bureau of Investigation"), "FBI")`)
}})
```

`build_acronym("For your information")` 应该返回 `"FYI"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("For your information"), "FYI")`)
}})
```

`build_acronym("By the way")` 应该返回 `"BTW"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("By the way"), "BTW")`)
}})
```

`build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` 应该返回 `"AUHWPOTIMSH"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH")`)
}})
```

# --seed--

## --seed-contents--

```py
def build_acronym(s):

    return s
```

# --solutions--

```py
def build_acronym(s):
    small_words = {"a", "for", "an", "and", "by", "of"}
    words = s.split()
    acronym = ""

    for i, word in enumerate(words):
        if i == 0 or word.lower() not in small_words:
            acronym += word[0].upper()

    return acronym
```
