---
id: 68af0687ef34c76c28ffa54f
title: "挑战 34：缺失数字"
challengeType: 29
dashedName: challenge-34
---

# --description--

给定一个从 1 到 `n`（含） 的整数数组，返回一个包含 1 到 `n` 之间所有缺失整数的数组（其中 `n` 是给定数组中的最大数字）。

- 给定的数组可能未排序且可能包含副本。
- 返回的数组应为升序。
- 如果没有缺失的整数，返回一个空数组。

# --hints--

`find_missing_numbers([1, 3, 5])` 应该返回 `[2, 4]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 3, 5]), [2, 4])`)
}})
```

`find_missing_numbers([1, 2, 3, 4, 5])` 应该返回 `[]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 2, 3, 4, 5]), [])`)
}})
```

`find_missing_numbers([1, 10])` 应该返回 `[2, 3, 4, 5, 6, 7, 8, 9]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9])`)
}})
```

`find_missing_numbers([10, 1, 10, 1, 10, 1])` 应该返回 `[2, 3, 4, 5, 6, 7, 8, 9]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9])`)
}})
```

`find_missing_numbers([3, 1, 4, 1, 5, 9])` 应该返回 `[2, 6, 7, 8]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([3, 1, 4, 1, 5, 9]), [2, 6, 7, 8])`)
}})
```

`find_missing_numbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` 应该返回 `[11]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_missing_numbers(arr):

    return arr
```

# --solutions--

```py
def find_missing_numbers(arr):
    if not arr:
        return []
    
    n = max(arr)
    seen = set(arr)
    result = [i for i in range(1, n + 1) if i not in seen]
    return result
```
