---
id: 68af0687ef34c76c28ffa549
title: "挑战 31：数组差异"
challengeType: 29
dashedName: challenge-31
---

# --description--

给定两个包含字串值的数组，返回一个新建数组，其中包含仅出现在其中一个数组中的所有值。

- 返回的数组应按字母顺序排序。

# --hints--

`array_diff(["apple", "banana"], ["apple", "banana", "cherry"])` 应该返回 `["cherry"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["apple", "banana"], ["apple", "banana", "cherry"]), ["cherry"])`)
}})
```

`array_diff(["apple", "banana", "cherry"], ["apple", "banana"])` 应该返回 `["cherry"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["apple", "banana", "cherry"], ["apple", "banana"]), ["cherry"])`)
}})
```

`array_diff(["one", "two", "three", "four", "six"], ["one", "three", "eight"])` 应该返回 `["eight", "four", "six", "two"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["one", "two", "three", "four", "six"], ["one", "three", "eight"]), ["eight", "four", "six", "two"])`)
}})
```

`array_diff(["two", "four", "five", "eight"], ["one", "two", "three", "four", "seven", "eight"])` 应该返回 `["five", "one", "seven", "three"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["two", "four", "five", "eight"], ["one", "two", "three", "four", "seven", "eight"]), ["five", "one", "seven", "three"])`)
}})
```

`array_diff(["I", "like", "freeCodeCamp"], ["I", "like", "rocks"])` 应该返回 `["freeCodeCamp", "rocks"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["I", "like", "freeCodeCamp"], ["I", "like", "rocks"]), ["freeCodeCamp", "rocks"])`)
}})
```

# --seed--

## --seed-contents--

```py
def array_diff(arr1, arr2):

    return arr1
```

# --solutions--

```py
def array_diff(arr1, arr2):
    unique1 = []
    for v in arr1:
        if v not in unique1:
            unique1.append(v)
    unique2 = []
    for v in arr2:
        if v not in unique2:
            unique2.append(v)

    only1 = [v for v in unique1 if v not in unique2]
    only2 = [v for v in unique2 if v not in unique1]

    return sorted(only1 + only2)
```
