---
id: 68b06e589bf2273243814773
title: "挑战 37：句子首字母大写器"
challengeType: 29
dashedName: challenge-37
---

# --description--

给定一个段落，返回一个新建段落，其中每个句子的第一个字母都大写。

- 所有其他字符应保持不变。
- 句子可以以句号（`.`）、一个或多个问号（`?`）或一个或多个感叹号（`!`）结尾。

# --hints--

`capitalize("this is a simple sentence.")` 应该返回 `"This is a simple sentence."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("this is a simple sentence."), "This is a simple sentence.")`)
}})
```

`capitalize("hello world. how are you?")` 应该返回 `"Hello world. How are you?"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("hello world. how are you?"), "Hello world. How are you?")`)
}})
```

`capitalize("i did today's coding challenge... it was fun!!")` 应该返回 `"I did today's coding challenge... It was fun!!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!")`)
}})
```

`capitalize("crazy!!!strange???unconventional...sentences.")` 应该返回 `"Crazy!!!Strange???Unconventional...Sentences."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.")`)
}})
```

`capitalize("there's a space before this period . why is there a space before that period ?")` 应该返回 `"There's a space before this period . Why is there a space before that period ?"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?")`)
}})
```

# --seed--

## --seed-contents--

```py
def capitalize(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def capitalize(paragraph):
    def repl(match):
        return match.group(1) + match.group(2) + match.group(3).upper()

    result = re.sub(r'([.!?]+)(\s*)([a-z])', repl, paragraph)

    if result and result[0].islower():
        result = result[0].upper() + result[1:]

    return result
```
