---
id: 68b06e589bf2273243814771
title: "挑战 36：恒温器调节器"
challengeType: 29
dashedName: challenge-36
---

# --description--

给定房间的当前温度和目标温度，返回一个字串，指示如何根据这些约束调整房间温度：

- 如果当前温度低于目标，则返回 `"heat"`。
- 如果当前温度高于目标，则返回 `"cool"`。
- 如果当前温度与目标相等，则返回 `"hold"`。

# --hints--

`adjust_thermostat(68, 72)` 应该返回 `"heat"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(68, 72), "heat")`)
}})
```

`adjust_thermostat(75, 72)` 应该返回 `"cool"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(75, 72), "cool")`)
}})
```

`adjust_thermostat(72, 72)` 应该返回 `"hold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 72), "hold")`)
}})
```

`adjust_thermostat(-20.5, -10.1)` 应该返回 `"heat"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(-20.5, -10.1), "heat")`)
}})
```

`adjust_thermostat(100, 99.9)` 应该返回 `"cool"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(100, 99.9), "cool")`)
}})
```

`adjust_thermostat(0.0, 0.0)` 应该返回 `"hold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(0.0, 0.0), "hold")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(temp, target):

    return temp
```

# --solutions--

```py
def adjust_thermostat(temp, target):
    if temp < target:
        return "heat"
    elif temp > target:
        return "cool"
    else:
        return "hold"
```
