---
id: 68b06e589bf227324381476f
title: "挑战 35：词频"
challengeType: 29
dashedName: challenge-35
---

# --description--

给定一个段落，返回出现频率最高的三个单词的数组。

- 段落中的单词将被空格分开。
- 忽略给定段落中的分支。例如，将 `Hello` 和 `hello` 视为相同的单词。
- 忽略给定段落中的标点符号。标点符号包括逗号（`,`）、句号（`.`）和感叹号（`!`）。
- 返回的数组应包含所有小写单词。
- 返回的数组应按降序排列，最常出现的单词排在最前面。

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` 应该返回 `["coding", "python", "in"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` 应该返回 `["i", "like", "coding"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` 应该返回 `["debug", "test", "deploy"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
