---
id: 68b1f72371a5ac895ac70a0a
title: "挑战 44：字串 镜射"
challengeType: 29
dashedName: challenge-44
---

# --description--

给定两个字串，判断第二个字串是否是第一个字串的镜射。

- 如果一个字串包含相同的字母且顺序相反，则该字串被视为镜射。
- 将大写字母和小写字母视为不同。
- 忽略所有非字母字符。

# --hints--

`is_mirror("helloworld", "helloworld")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("helloworld", "helloworld"), False)`)
}})
```

`is_mirror("Hello World", "dlroW olleH")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Hello World", "dlroW olleH"), True)`)
}})
```

`is_mirror("RaceCar", "raCecaR")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("RaceCar", "raCecaR"), True)`)
}})
```

`is_mirror("RaceCar", "RaceCar")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("RaceCar", "RaceCar"), False)`)
}})
```

`is_mirror("Mirror", "rorrim")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Mirror", "rorrim"), False)`)
}})
```

`is_mirror("Hello World", "dlroW-olleH")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Hello World", "dlroW-olleH"), True)`)
}})
```

`is_mirror("Hello World", "!dlroW !olleH")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Hello World", "!dlroW !olleH"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_mirror(str1, str2):

    return str1
```

# --solutions--

```py
def is_mirror(str1, str2):
    clean1 = "".join(c for c in str1 if c.isalpha())
    clean2 = "".join(c for c in str2 if c.isalpha())
    return clean1[::-1] == clean2
```
