---
id: 68b1f72371a5ac895ac70a06
title: "挑战 42：视频存储"
challengeType: 29
dashedName: challenge-42
---

# --description--

给定一个视频大小、视频大小的单位、硬盘容量以及硬盘的单位，使用以下约束返回硬盘可以保存的视频数量：

- 视频大小的单位可以是字节（`"B"`）、千字节（`"KB"`）、兆字节（`"MB"`）或千兆字节（`"GB"`）。
- 如果未给出上述视频单位之一，返回 `"Invalid video unit"`。
- 硬盘容量的单位可以是千兆字节（`"GB"`）或太字节（`"TB"`）。
- 如果未给出上述硬盘单位之一，返回 `"Invalid drive unit"`。
- 返回驱动器可以容纳的完整 `video` 数量。
- 使用以下转换：

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |
| 1 TB |   1000 GB  |

例如，给定 `500`、`"MB"`、`100` 和 `"GB"` 作为参数，确定 多少 个 500 MB 视频 可以存储在 100 GB 硬盘上。

# --hints--

`number_of_videos(500, "MB", 100, "GB")` 应该返回 `200`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` 应该返回 `"Invalid video unit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` 应该返回 `"Invalid drive unit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` 应该返回 `4000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` 应该返回 `1466`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```
