---
id: 68b1f72371a5ac895ac70a02
title: "挑战 40：照片存储"
challengeType: 29
dashedName: challenge-40
---

# --description--

给定以兆字节（MB）为单位的照片大小和以千兆字节（GB）为单位的硬盘容量，使用以下约束返回硬盘可以保存的照片数量：

- 1 吉字节等于 1000 兆字节。
- 返回驱动器可以保存的完整照片数量。

# --hints--

`number_of_photos(1, 1)` 应该返回 `1000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(1, 1), 1000)`)
}})
```

`number_of_photos(2, 1)` 应该返回 `500`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(2, 1), 500)`)
}})
```

`number_of_photos(4, 256)` 应该返回 `64000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(4, 256), 64000)`)
}})
```

`number_of_photos(3.5, 750)` 应该返回 `214285`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 750), 214285)`)
}})
```

`number_of_photos(3.5, 5.5)` 应该返回 `1571`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 5.5), 1571)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_photos(photo_size_mb, drive_size_gb):

    return photo_size_mb
```

# --solutions--

```py
def number_of_photos(photo_size_mb, drive_size_gb):
    drive_size_mb = drive_size_gb * 1000
    return drive_size_mb // photo_size_mb
```
