---
id: 68b1f72371a5ac895ac70a04
title: "挑战 41：文件存储"
challengeType: 29
dashedName: challenge-41
---

# --description--

给定一个文件大小、文件大小的单位以及以千兆字节（GB）为单位的硬盘容量，返回硬盘可以保存的文件数量，使用以下约束：

- 文件大小的单位可以是字节（`"B"`）、千字节（`"KB"`）或兆字节（`"MB"`）。
- 返回驱动器可以容纳的完整文件数。
- 使用以下转换：

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

例如，给定 `500`、`"KB"` 和 `1` 作为参数，确定在 1 GB 硬盘上可以容纳多少个 500 KB 文件。

# --hints--

`number_of_files(500, "KB", 1)` 应该返回 `2000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(500, "KB", 1), 2000)`)
}})
```

`number_of_files(50000, "B", 1)` 应该返回 `20000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(50000, "B", 1), 20000)`)
}})
```

`number_of_files(5, "MB", 1)` 应该返回 `200`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(5, "MB", 1), 200)`)
}})
```

`number_of_files(4096, "B", 1.5)` 应该返回 `366210`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4096, "B", 1.5), 366210)`)
}})
```

`number_of_files(220.5, "KB", 100)` 应该返回 `453514`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(220.5, "KB", 100), 453514)`)
}})
```

`number_of_files(4.5, "MB", 750)` 应该返回 `166666`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4.5, "MB", 750), 166666)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_files(file_size, file_unit, drive_size_gb):

    return file_size
```

# --solutions--

```py
def number_of_files(file_size, file_unit, drive_size_gb):
    drive_size_bytes = drive_size_gb * 1000 * 1000 * 1000

    if file_unit == "B":
        file_size_bytes = file_size
    elif file_unit == "KB":
        file_size_bytes = file_size * 1000
    else:
        file_size_bytes = file_size * 1000 * 1000

    return int(drive_size_bytes // file_size_bytes)
```
