---
id: 68b1f72371a5ac895ac70a08
title: "挑战 43：数字 对 字母"
challengeType: 29
dashedName: challenge-43
---

# --description--

给定一个字串，如果字串中数字多于字母，则返回 `"digits"`，如果字母多于数字，则返回 `"letters"`，如果数字和字母数量相同，则返回 `"tie"`。

- 数字由 `0-9` 组成。
- 字母由大写或小写的 `a-z` 组成。
- 忽略任何其他字符。

# --hints--

`digits_or_letters("abc123")` 应该返回 `"tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123"), "tie")`)
}})
```

`digits_or_letters("a1b2c3d")` 应该返回 `"letters"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("a1b2c3d"), "letters")`)
}})
```

`digits_or_letters("1a2b3c4")` 应该返回 `"digits"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("1a2b3c4"), "digits")`)
}})
```

`digits_or_letters("abc123!@#DEF")` 应该返回 `"letters"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123!@#DEF"), "letters")`)
}})
```

`digits_or_letters("H3110 W0R1D")` 应该返回 `"digits"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("H3110 W0R1D"), "digits")`)
}})
```

`digits_or_letters("P455W0RD")` 应该返回 `"tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("P455W0RD"), "tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def digits_or_letters(s):

    return s
```

# --solutions--

```py
def digits_or_letters(s):
    digit_count = 0
    letter_count = 0

    for char in s:
        if char.isdigit():
            digit_count += 1
        elif char.isalpha():
            letter_count += 1

    if digit_count > letter_count:
        return "digits"
    if letter_count > digit_count:
        return "letters"
    return "tie"
```
