---
id: 68b7687dded630607aceccb3
title: "挑战 49：CSV 头部解析器"
challengeType: 29
dashedName: challenge-49
---

# --description--

给定逗号分隔值（CSV）文件的第一行，返回包含标题的数组。

- CSV `file` 的第一行包含由逗号分隔的标题。
- 删除每个标题的任何前导或尾随空白。

# --hints--

`get_headings("name,age,city")` 应该返回 `["name", "age", "city"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("name,age,city"), ["name", "age", "city"])`)
}})
```

`get_headings("first name,last name,phone")` 应该返回 `["first name", "last name", "phone"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("first name,last name,phone"), ["first name", "last name", "phone"])`)
}})
```

`get_headings("username , email , signup date ")` 应该返回 `["username", "email", "signup date"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("username , email , signup date "), ["username", "email", "signup date"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_headings(csv):

    return csv
```

# --solutions--

```py
def get_headings(csv):
    return [h.strip() for h in csv.split(",")]
```
