---
id: 68b7cadffed0e75a517da677
title: "挑战 54：P@ssw0rd Str3ngth！"
challengeType: 29
dashedName: challenge-54
---

# --description--

给定一个密码字串，根据密码强度返回 `"weak"`、`"medium"` 或 `"strong"`。

密码根据以下规则进行评估：

- 它至少有 8 个字符长。
- 它包含大写和小写字母。
- 它至少包含 1 个数字。
- 它至少包含来自此集合的一个特殊字符：`!`、`@`、`#`、`$`、`%`、`^`、`&` 或 `*`。

如果密码满足的规则少于 2 条，则返回 `"weak"`。  
如果密码满足 2 条或 3 条规则，则返回 `"medium"`。  
如果密码满足所有 4 条规则，则返回 `"strong"`。

# --hints--

`check_strength("123456")` 应该返回 `"weak"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("123456"), "weak")`)
}})
```

`check_strength("pass!!!")` 应该返回 `"weak"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("pass!!!"), "weak")`)
}})
```

`check_strength("Qwerty")` 应该返回 `"weak"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("Qwerty"), "weak")`)
}})
```

`check_strength("PASSWORD")` 应该返回 `"weak"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD"), "weak")`)
}})
```

`check_strength("PASSWORD!")` 应该返回 `"medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("PassWord%^!")` 应该返回 `"medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PassWord%^!"), "medium")`)
}})
```

`check_strength("qwerty12345")` 应该返回 `"medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("qwerty12345"), "medium")`)
}})
```

`check_strength("S3cur3P@ssw0rd")` 应该返回 `"strong"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("S3cur3P@ssw0rd"), "strong")`)
}})
```

`check_strength("C0d3&Fun!")` 应该返回 `"strong"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("C0d3&Fun!"), "strong")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_strength(password):

    return password
```

# --solutions--

```py
import re
def check_strength(password):
    rules_met = 0

    if len(password) >= 8:
        rules_met += 1
    if re.search(r'[a-z]', password) and re.search(r'[A-Z]', password):
        rules_met += 1
    if re.search(r'\d', password):
        rules_met += 1
    if re.search(r'[!@#$%^&*]', password):
        rules_met += 1

    if rules_met < 2:
        return "weak"
    elif rules_met < 4:
        return "medium"
    else:
        return "strong"
```
