---
id: 68b7687dded630607aceccaf
title: "挑战 47：超速被抓"
challengeType: 29
dashedName: challenge-47
---

# --description--

给定一个表现车辆行驶速度的数字数组和一个表现限速的数字，返回一个包含两项的数组，第一项是超速车辆的数量，第二项是这些车辆超出限速的平均速度。

- 如果没有车辆超速，返回 `[0, 0]`。

# --hints--

`speeding([50, 60, 55], 60)` 应该返回 `[0, 0]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([50, 60, 55], 60), [0, 0])`)
}})
```

`speeding([58, 50, 60, 55], 55)` 应该返回 `[2, 4]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([58, 50, 60, 55], 55), [2, 4])`)
}})
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` 应该返回 `[4, 8.5]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5])`)
}})
```

`speeding([100, 105, 95, 102], 100)` 应该返回 `[2, 3.5]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([100, 105, 95, 102], 100), [2, 3.5])`)
}})
```

`speeding([40, 45, 44, 50, 112, 39], 55)` 应该返回 `[1, 57]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57])`)
}})
```

# --seed--

## --seed-contents--

```py
def speeding(speeds, limit):

    return speeds
```

# --solutions--

```py
def speeding(speeds, limit):
    speeding = 0
    total_over = 0

    for speed in speeds:
        if speed > limit:
            speeding += 1
            total_over += (speed - limit)
    
    if speeding == 0:
        return [0, 0]
    
    return [speeding, total_over / speeding]
```
