---
id: 68b7687dded630607aceccab
title: "挑战 45：完全平方"
challengeType: 29
dashedName: challenge-45
---

# --description--

给定一个整数，判断它是否是一个完全平方数。

- 如果你可以将一个整数乘以它自身得到该数字，那么该数字就是一个完全平方数。例如，9 是一个完全平方数，因为你可以将 3 乘以它自身得到它。

# --hints--

`is_perfect_square(9)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(9), True)`)
}})
```

`is_perfect_square(49)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(49), True)`)
}})
```

`is_perfect_square(1)` 应返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(1), True)`)
}})
```

`is_perfect_square(2)` 应返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(2), False)`)
}})
```

`is_perfect_square(99)` 应返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(99), False)`)
}})
```

`is_perfect_square(-9)` 应返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(-9), False)`)
}})
```

`is_perfect_square(0)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(0), True)`)
}})
```

`is_perfect_square(25281)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(25281), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_perfect_square(n):

    return n
```

# --solutions--

```py
import math
def is_perfect_square(n):
    if n < 0:
        return False
    root = int(math.sqrt(n))
    return root * root == n
```
