---
id: 68c1a929005bf54d342aa8d4
title: "挑战 57：太空周 第 3 天：回家打电话"
challengeType: 29
dashedName: challenge-57
---

# --description--

在太空周的第三天，你将获得一个表示你自己、卫星和你的母星之间通信路线中距离（以公里为单位）的数字数组。使用以下约束确定通过该路线发送的消息到达目标星球所需的时间：

- 数组中的第一个值是你的位置到第一个卫星的距离。
- 除最后一个外，每个后续值都是到下一个卫星的距离。
- 数组中的最后一个值是从前一个卫星到你的母星的距离。
- 该消息以 300,000 公里/秒 的速度传播。
- 每经过一个卫星，消息 **传递** 会增加 0.5 秒的传输延迟。
- 返回一个四位小数的数字，去除末尾的零。

# --hints--

`send_message([300000, 300000])` 应该返回 `2.5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` 应该返回 `3.0627`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` 应该返回 `364.5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` 应该返回 `1674.3333`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` 应该返回 `2.4086`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` 应该返回 `21.1597`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
