---
id: 68c1a929005bf54d342aa8d6
title: "挑战 59：太空周 第 5 天：适居带"
challengeType: 29
dashedName: challenge-59
---

# --description--

在太空周的第 5 天，你将计算恒星的“Goldilocks zone”——恒星周围条件“恰到好处”以存在液态水的区域。

给定恒星的质量，返回一个包含其金发姑娘区起始和结束距离（以天文单位为单位）的数组。

计算 Goldilocks Zone：

1. 通过将恒星的质量提升到 3.5 次方来查找其光度。
2. 该区域的起点是其亮度的 0.95 倍平方根。
3. 该区域的末端是其亮度的 1.37 倍平方根。

- 返回四舍五入到两位小数的距离。

例如，给定 `1` 作为质量，返回 `[0.95, 1.37]`。

# --hints--

`goldilocks_zone(1)` 应该返回 `[0.95, 1.37]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` 应返回 `[0.28, 0.41]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` 应该返回 `[21.85, 31.51]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` 应返回 `[9.38, 13.52]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` 应该返回 `[179.69, 259.13]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
