---
id: 68c497f3aaefc9fd9f1b0e24
title: "挑战 60：太空周第 6 天：月相"
challengeType: 29
dashedName: challenge-60
---

# --description--

在太空周的第 6 天，你将获得一个格式为 `"YYYY-MM-DD"` 的日期，并需要使用以下规则确定该日期的月相：

使用简化的 28 天月周期，分为四个相等的阶段：

- `"New"`：第 1 天至第 7 天
- `"Waxing"`：第 8 天至第 14 天
- `"Full"`：第 15 天至第 21 天
- `"Waning"`：第 22 天至第 28 天

在第 28 天之后，周期从第 1 天开始重复，新月。

- 使用 `"2000-01-06"` 作为参考新月（周期的第 1 天）来确定给定日期的月相。
- 你不会被给予早于 `reference date` 的任何日期。
- 返回正确的阶段作为字串。

**注意：** 第 1 天表现新月当天，意味着自上一个新月以来已过 0 天。

# --hints--

`moon_phase("2000-01-12")` 应该返回 `"New"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` 应该返回 `"Waxing"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` 应该返回 `"Full"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` 应该返回 `"Waning"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` 应该返回 `"New"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
