---
id: 68c497f3aaefc9fd9f1b0e26
title: "挑战 62：十六进制转十进制"
challengeType: 29
dashedName: challenge-62
---

# --description--

给定一个表现十六进制数（基数 16）的字串，返回其十进制（基数 10）值作为一个整数。

十六进制是一种使用 16 个数字的数字系统：

- `0-9` 表现值 `0` 到 `9`。
- `A-F` 表现值从 `10` 到 `15`。

以下是部分转换表：

| Hexadecimal | Decimal |
|:-----------:|:-------:|
| 0           | 0       |
| 1           | 1       |
| ...         | ...     |
| 9           | 9       |
| A           | 10      |
| ...         | ...     |
| F           | 15      |
| 10          | 16      |
| ...         | ...     |
| 9F          | 159     |
| A0          | 160     |
| ...         | ...     |
| FF          | 255     |
| 100         | 256     |

- 该字串只会包含字符 `0–9` 和 `A–F`。

# --hints--

`hex_to_decimal("A")` 应该返回 `10`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A"), 10)`)
}})
```

`hex_to_decimal("15")` 应该返回 `21`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("15"), 21)`)
}})
```

`hex_to_decimal("2E")` 应该返回 `46`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("2E"), 46)`)
}})
```

`hex_to_decimal("FF")` 应该返回 `255`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("FF"), 255)`)
}})
```

`hex_to_decimal("A3F")` 应该返回 `2623`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A3F"), 2623)`)
}})
```

# --seed--

## --seed-contents--

```py
def hex_to_decimal(hex):

    return hex
```

# --solutions--

```py
def hex_to_decimal(hex):

    return int(hex, 16)
```
