---
id: 68c497f3aaefc9fd9f1b0e25
title: "挑战 61：太空周第 7 天：发射燃料"
challengeType: 29
dashedName: challenge-61
---

# --description--

在太空周的最后一天，你将获得你想送入轨道的有效载荷的质量，单位为千克（kg）。使用以下规则确定将你的有效载荷送入轨道所需的燃料量：

- 火箭每提升 5 kg 质量需要 1 kg 燃料。
- 燃料本身有质量。所以当你添加燃料时，需要提升的质量增加，这需要更多燃料，进而增加质量，如此循环。
- 要计算所需的总燃料：从有效载荷质量开始，计算所需的燃料，将该燃料添加到总质量中，然后再次计算。重复此进程，直到额外所需燃料少于 1 公斤，然后停止。
- 忽略火箭本身的质量。只计算提升有效载荷及其自身燃料所需的燃料。

例如，给定 50 kg 的有效载荷质量，你需要 10 kg 燃料来提升它（payload / 5），这将总质量增加到 60 kg，需要 12 kg 来提升（额外 2 kg），这将总质量增加到 62 kg，需要 12.4 kg 来提升——额外 0.4 kg——少于 1 kg 的额外燃料，所以我们在这里停止。总提升质量为 62.4 kg，其中 50 kg 是初始有效载荷，12.4 kg 是燃料。

- 返回所需燃料量，四舍五入到 1 位小数。

# --hints--

`launch_fuel(50)` 应返回 `12.4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(50), 12.4)`)
}})
```

`launch_fuel(500)` 应该返回 `124.8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(500), 124.8)`)
}})
```

`launch_fuel(243)` 应返回 `60.7`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(243), 60.7)`)
}})
```

`launch_fuel(11000)` 应该返回 `2749.8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(11000), 2749.8)`)
}})
```

`launch_fuel(6214)` 应返回 `1553.4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(6214), 1553.4)`)
}})
```

# --seed--

## --seed-contents--

```py
def launch_fuel(payload):

    return payload
```

# --solutions--

```py
def launch_fuel(payload):
    total_mass = payload
    additional_fuel = total_mass / 5
    total_fuel = additional_fuel

    while additional_fuel >= 1:
        total_mass += additional_fuel
        additional_fuel = (total_mass / 5) - total_fuel
        total_fuel += additional_fuel

    return round(total_fuel, 1)
```
