---
id: 68cae5b538ff798bbd4da002
title: "挑战 64：24 到 12"
challengeType: 29
dashedName: challenge-64
---

# --description--

给定一个表现一天时间的 24 小时格式的字串 `"HHMM"`，返回其等效的 12 小时格式 `"H:MM AM"` 或 `"H:MM PM"`。

- 给定的输入将始终是一个四位数的字串，采用 24 小时制格式，范围从 `"0000"` 到 `"2359"`。

# --hints--

`to_12("1124")` 应该返回 `"11:24 AM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1124"), "11:24 AM")`)
}})
```

`to_12("0900")` 应该返回 `"9:00 AM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0900"), "9:00 AM")`)
}})
```

`to_12("1455")` 应该返回 `"2:55 PM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1455"), "2:55 PM")`)
}})
```

`to_12("2346")` 应该返回 `"11:46 PM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("2346"), "11:46 PM")`)
}})
```

`to_12("0030")` 应该返回 `"12:30 AM"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0030"), "12:30 AM")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_12(time):

    return time
```

# --solutions--

```py
def convert_hours(hours):
    if hours == 0:
        return 12
    elif hours > 12:
        return hours - 12
    else:
        return hours

def to_12(time):
    hours = int(time[:2])
    minutes = time[2:]
    period = "AM" if hours < 12 else "PM"

    return f"{convert_hours(hours)}:{minutes} {period}"
```
