---
id: 68cae5b538ff798bbd4da00a
title: "挑战 72：恒温器调节器 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

给定房间当前的华氏温度和摄氏目标温度，返回一个字串，指示如何根据这些约束调整房间温度：

- 如果当前温度低于目标，则返回 `"Heat: X degrees Fahrenheit"`。其中 `X` 是将房间加热到目标所需的华氏度数，四舍五入到 1 位小数。
- 如果当前温度高于目标，则返回 `"Cool: X degrees Fahrenheit"`。其中 `X` 是将房间冷却到目标所需的华氏度数，四舍五入到 1 位小数。
- 如果当前温度与目标相等，则返回 `"Hold"`。

要将摄氏度转换为华氏度，将摄氏温度乘以 1.8 并将 32 添加到结果中（`F = (C * 1.8) + 32`）。

# --hints--

`adjust_thermostat(32, 0)` 应该返回 `"Hold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` 应该返回 `"Heat: 7.0 degrees Fahrenheit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` 应该返回 `"Cool: 7.6 degrees Fahrenheit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` 应该返回 `"Hold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` 应返回 `"Heat: 12.6 degrees Fahrenheit"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
