---
id: 68cae5b538ff798bbd4da006
title: "挑战 68：信用卡掩码器"
challengeType: 29
dashedName: challenge-68
---

# --description--

给定一串信用卡号码字串，使用以下约束返回其掩码版本：

- 该字串将包含四组四位数字（`0-9`），所有组之间用单个空格或单个连字符（`-`）分隔。
- 将所有数字替换为星号（`*`），除了最后四个数字。
- 保留剩余的字符不变。

例如，给定 `"4012-8888-8888-1881"` 返回 `"****-****-****-1881"`。

# --hints--

`mask("4012-8888-8888-1881")` 应该返回 `"****-****-****-1881"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("4012-8888-8888-1881"), "****-****-****-1881")`)
}})
```

`mask("5105 1051 0510 5100")` 应该返回 `"**** **** **** 5100"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("5105 1051 0510 5100"), "**** **** **** 5100")`)
}})
```

`mask("6011 1111 1111 1117")` 应该返回 `"**** **** **** 1117"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("6011 1111 1111 1117"), "**** **** **** 1117")`)
}})
```

`mask("2223-0000-4845-0010")` 应该返回 `"****-****-****-0010"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("2223-0000-4845-0010"), "****-****-****-0010")`)
}})
```

# --seed--

## --seed-contents--

```py
def mask(card):

    return card
```

# --solutions--

```py
def mask(card):
    if '-' in card:
        split = card.split('-')
        return f"****-****-****-{split[3]}"
    else:
        split = card.split(' ')
        return f"**** **** **** {split[3]}"
```
