---
id: 68cae5b538ff798bbd4da005
title: "挑战 67：电子邮件验证器"
challengeType: 29
dashedName: challenge-67
---

# --description--

给定一个字串，使用以下约束确定它是否是有效的电子邮件地址：

- 它必须包含恰好一个 `@` 符号。
- `@` 之前的局部的部分：
  - 只能包含字母（`a-z`、`A-Z`）、数字（`0-9`）、点号（`.`）、下划线（`_`）或连字符（`-`）。
  - 不能以点号开头或结尾。
- `@` 后面的域部分：
  - 必须包含至少一个点。
  - 必须以一个点号后跟至少两个字母结尾。
- 局部的部分和域部分都不能有两个连续的点。

# --hints--

`validate("a@b.cd")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("a@b.cd"), True)`)
}})
```

`validate("hell.-w.rld@example.com")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hell.-w.rld@example.com"), True)`)
}})
```

`validate(".b@sh.rc")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate(".b@sh.rc"), False)`)
}})
```

`validate("example@test.c0")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("example@test.c0"), False)`)
}})
```

`validate("freecodecamp.org")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("freecodecamp.org"), False)`)
}})
```

`validate("develop.ment_user@c0D!NG.R.CKS")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("develop.ment_user@c0D!NG.R.CKS"), True)`)
}})
```

`validate("hello.@wo.rld")` 应返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hello.@wo.rld"), False)`)
}})
```

`validate("hello@world..com")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hello@world..com"), False)`)
}})
```

`validate("develop..ment_user@c0D!NG.R.CKS")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("develop..ment_user@c0D!NG.R.CKS"), False)`)
}})
```

`validate("git@commit@push.io")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("git@commit@push.io"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def validate(email):

    return email
```

# --solutions--

```py
import re
def validate(email):
    if '..' in email:
        return False

    parts = email.split('@')
    if len(parts) != 2:
        return False

    local, domain = parts

    if local.startswith('.') or local.endswith('.'):
        return False
    if not re.match(r'^[a-zA-Z0-9._-]+$', local):
        return False

    if '.' not in domain:
        return False

    tld = domain.split('.')[-1]
    if len(tld) < 2 or not tld.isalpha():
        return False

    return True
```
