---
id: 68cae5b538ff798bbd4da003
title: "挑战 65：字串计数"
challengeType: 29
dashedName: challenge-65
---

# --description--

给定两个字串，确定第二个字串在第一个字串中出现了多少次。

- 模式字串可以在第一个字串中重叠。例如，`"aaa"` 包含两次 `"aa"`。第一次是前两个 `a`，第二次是后两个。

# --hints--

`count('abcdefg', 'def')` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('abcdefg', 'def'), 1)`)
}})
```

`count('hello', 'world')` 应该返回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('hello', 'world'), 0)`)
}})
```

`count('mississippi', 'iss')` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('mississippi', 'iss'), 2)`)
}})
```

`count('she sells seashells by the seashore', 'sh')` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('she sells seashells by the seashore', 'sh'), 3)`)
}})
```

`count('101010101010101010101', '101')` 应该返回 `10`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('101010101010101010101', '101'), 10)`)
}})
```

# --seed--

## --seed-contents--

```py
def count(text, parameter):

    return text
```

# --solutions--

```py
def count(text, pattern):
    if not pattern:
        return 0
    occurrences = 0

    for i in range(len(text) - len(pattern) + 1):
        if text[i:i+len(pattern)] == pattern:
            occurrences += 1

    return occurrences
```
