---
id: 68cae5b538ff798bbd4da004
title: "挑战 66：超文本标记语言 标签剥离器"
challengeType: 29
dashedName: challenge-66
---

# --description--

给定一段 `HTML` 代码字串，移除标签并返回纯文本内容。

- 输入字串将仅包含有效的 `HTML`。
- `HTML` 标签可以嵌套。
- 删除标签和任何属性。

例如，`'<a href="#">Click here</a>'` 应该返回 `"Click here"`。

# --hints--

`strip_tags('<a href="#">Click here</a>')` 应该返回 `"Click here"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` 应该返回 `"Hello World!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` 应返回空字串（`""`）。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` 应该返回 `sectionsection`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
