---
id: 68cae5b538ff798bbd4da001
title: "挑战 63：文字之战"
challengeType: 29
dashedName: challenge-63
---

# --description--

给定两句话，表现你的团队和对手团队，其中你团队的每个单词与对手团队对应的单词进行对战，使用以下规则确定哪个团队获胜：

- 给定的句子将始终包含相同数量的单词。
- 单词由单个空格分隔且仅包含字母。
- 每个单词的值是其字母的总和。
- 字母 `a` 到 `z` 对应值为 `1` 到 `26`。例如，`a` 是 `1`，`z` 是 `26`。
- 大写字母将字母的值加倍。例如，`A` 是 `2`，`Z` 是 `52`。
- 单词按顺序对战：你队伍的第一个单词与对方队伍的第一个单词对战，依此类推。
- 如果一个单词的值大于对方单词的值，则该单词获胜。
- 获胜词更多的队伍获胜。

如果你的团队获胜，则返回 `"We win"`，如果你的团队失败，则返回 `"We lose"`，如果双方获胜次数相同，则返回 `"Draw"`。

# --hints--

`battle("hello world", "hello word")` 应该返回 `"We win"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` 应该返回 `"We win"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` 应该返回 `"We lose"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` 应该返回 `"Draw"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` 应该返回 `"We win"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` 应该返回 `"We lose"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` 应该返回 `"Draw"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
