---
id: 68cae5b538ff798bbd4da009
title: "挑战 71：小费计算器"
challengeType: 29
dashedName: challenge-71
---

# --description--

给定你的餐费价格和一个自定义小费百分比，返回一个包含三个小费值的数组；15％、20％和自定义金额。

- 价格将以 `"$N.NN"` 形式给出。
- 自定义小费百分比将以这种形式给出：`"25%"`。
- 以相同的 `"$N.NN"` 形式返回金额，四舍五入到两位小数。

例如，给定一个 `"$10.00"` 的餐价和一个 `"25%"` 的自定义小费值，返回 `["$1.50", "$2.00", "$2.50"]`。

# --hints--

`calculate_tips("$10.00", "25%")` 应该返回 `["$1.50", "$2.00", "$2.50"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"])`)
}})
```

`calculate_tips("$89.67", "26%")` 应该返回 `["$13.45", "$17.93", "$23.31"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"])`)
}})
```

`calculate_tips("$19.85", "9%")` 应该返回 `["$2.98", "$3.97", "$1.79"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_tips(meal_price, custom_tip):

    return meal_price
```

# --solutions--

```py
def calculate_tips(meal_price, custom_tip):
    meal = float(meal_price[1:])
    custom_percent = float(custom_tip[:-1])

    tips_percent = [15, 20, custom_percent]
    tips = [f"${meal * p / 100:.2f}" for p in tips_percent]

    return tips
```
