---
id: 68d2ba1468508398389487ce
title: "挑战 73：你必须明智地说话"
challengeType: 29
dashedName: challenge-73
---

# --description--

给定一个句子，返回一个听起来像智慧老师建议的版本，遵循以下规则：

- 单词 由 单个 空格 分隔。
- 查找句子中以下单词之一的首次出现：`"have"`、`"must"`、`"are"`、`"will"`、`"can"`。
- 将该词之前及包括该词的所有单词移到句子末尾，然后：
  - 保持单词顺序不变。
  - 全部转换为小写。
  - 并在它们前面添加逗号和空格。
- 将句子中新第一个单词的首字母大写。
- 所有给定的句子都将以单个标点符号结尾。保持原始标点符号并将其移至新句子的末尾。
- 返回新建句子，确保每个单词之间只有一个空格，且句子开头和结尾没有空格。

例如，给定 `"You must speak wisely."` 返回 `"Speak wisely, you must."`

# --hints--

`wise_speak("You must speak wisely.")` 应该返回 `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` 应该返回 `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` 应该返回 `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` 应该返回 `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` 应该返回 `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
