---
id: 68d30845cc08266018fc46bd
title: "挑战 77：持续时间格式化器"
challengeType: 29
dashedName: challenge-77
---

# --description--

给定一个整数秒数，返回一个表现相同时长的字串，格式为 `"H:MM:SS"`，其中 `"H"` 是小时数，`"MM"` 是分钟数，`"SS"` 是秒数。使用以下规则返回时间：

- 秒：应始终为两位数字。
- 分钟：当不需要时应省略前导零。如果持续时间少于 1 分钟，则使用 `"0"`。
- 小时：仅当其大于零时才应包含。

# --hints--

`format(500)` 应该返回 `"8:20"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(500), "8:20")`)
}})
```

`format(4000)` 应该返回 `"1:06:40"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(4000), "1:06:40")`)
}})
```

`format(1)` 应该返回 `"0:01"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(1), "0:01")`)
}})
```

`format(5555)` 应该返回 `"1:32:35"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(5555), "1:32:35")`)
}})
```

`format(99999)` 应该返回 `"27:46:39"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(99999), "27:46:39")`)
}})
```

# --seed--

## --seed-contents--

```py
def format(seconds):

    return seconds
```

# --solutions--

```py
def format(seconds):
    h = seconds // 3600
    m = (seconds % 3600) // 60
    s = seconds % 60

    seconds_str = f"{s:02d}"
    minutes_str = str(m)

    if h > 0:
        hours_str = str(h)
        return f"{hours_str}:{minutes_str.zfill(2)}:{seconds_str}"
    else:
        return f"{minutes_str}:{seconds_str}"
```
