---
id: 68d30845cc08266018fc46be
title: "挑战 78：整数序列"
challengeType: 29
dashedName: challenge-78
---

# --description--

给定一个正整数，返回一个字串，包含从 `1` 到给定数字（包括该数字）的所有整数，按数字顺序排列。

例如，给定 `5`，返回 `"12345"`。

# --hints--

`sequence(5)` 应该返回 `"12345"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(5), "12345")`)
}})
```

`sequence(10)` 应该返回 `"12345678910"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(10), "12345678910")`)
}})
```

`sequence(1)` 应返回 `"1"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(1), "1")`)
}})
```

`sequence(27)` 应返回 `"123456789101112131415161718192021222324252627"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(27), "123456789101112131415161718192021222324252627")`)
}})
```

# --seed--

## --seed-contents--

```py
def sequence(n):

    return n
```

# --solutions--

```py
def sequence(n):
    result = ""
    for i in range(1, n + 1):
        result += str(i)
    return result
```
