---
id: 68d30845cc08266018fc46bc
title: "挑战 76：互补 DNA"
challengeType: 29
dashedName: challenge-76
---

# --description--

给定一个表现 DNA 序列的字串，使用以下规则返回其补集链：

- DNA 由字母 `"A"`、`"C"`、`"G"` 和 `"T"` 组成。
- 字母 `"A"` 和 `"T"` 互为补集。
- 字母 `"C"` 和 `"G"` 互为补集。

例如，给定 `"ACGT"`，返回 `"TGCA"`。

# --hints--

`complementary_dna("ACGT")` 应该返回 `"TGCA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ACGT"), "TGCA")`)
}})
```

`complementary_dna("ATGCGTACGTTAGC")` 应该返回 `"TACGCATGCAATCG"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ATGCGTACGTTAGC"), "TACGCATGCAATCG")`)
}})
```

`complementary_dna("GGCTTACGATCGAAG")` 应该返回 `"CCGAATGCTAGCTTC"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC")`)
}})
```

`complementary_dna("GATCTAGCTAGGCTAGCTAG")` 应该返回 `"CTAGATCGATCCGATCGATC"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC")`)
}})
```

# --seed--

## --seed-contents--

```py
def complementary_dna(strand):

    return strand
```

# --solutions--

```py
def complementary_dna(strand):
    complements = { "A": "T", "T": "A", "C": "G", "G": "C" }
    return "".join(complements[n] for n in strand)
```
