---
id: 68e39ed6106dac2f0a98fd63
title: "挑战 81：第 N 个素数"
challengeType: 29
dashedName: challenge-81
---

# --description--

素数是大于 1 的正整数，只能被 1 和它本身整除。前五个素数是 `2`、`3`、`5`、`7` 和 `11`。

给定一个正整数 `n`，返回第 `n` 个素数。例如，给定 `5` 返回第 5 个素数：`11`。

# --hints--

`nth_prime(5)` 应该返回 `11`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(5), 11)`)
}})
```

`nth_prime(10)` 应该返回 `29`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(10), 29)`)
}})
```

`nth_prime(16)` 应该返回 `53`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(16), 53)`)
}})
```

`nth_prime(99)` 应该返回 `523`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(99), 523)`)
}})
```

`nth_prime(1000)` 应该返回 `7919`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(1000), 7919)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_prime(n):

    return n
```

# --solutions--

```py
def nth_prime(n):
    primes = []
    num = 2

    while len(primes) < n:
        is_prime = True
        for i in range(2, int(num**0.5) + 1):
            if num % i == 0:
                is_prime = False
                break
        if is_prime:
            primes.append(num)
        num += 1

    return primes[n - 1]
```
