---
id: 68e39ed6106dac2f0a98fd64
title: "挑战 82：SpOoKy~CaSe"
challengeType: 29
dashedName: challenge-82
---

# --description--

给定一个表现变量名的字串，使用以下约束将其转换为“spooky case”：

- 将所有下划线（`_`）和连字符（`-`）替换为波浪号（`~`）。
- 将字串的第一个字母和之后的每隔一个字母大写。计数时忽略波浪号字符。将所有其他字母转换为小写。

例如，给定 `hello_world`，返回 `HeLlO~wOrLd`。

# --hints--

`spookify("hello_world")` 应该返回 `"HeLlO~wOrLd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("hello_world"), "HeLlO~wOrLd")`)
}})
```

`spookify("Spooky_Case")` 应该返回 `"SpOoKy~CaSe"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("Spooky_Case"), "SpOoKy~CaSe")`)
}})
```

`spookify("TRICK-or-TREAT")` 应该返回 `"TrIcK~oR~tReAt"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt")`)
}})
```

`spookify("c_a-n_d-y_-b-o_w_l")` 应该返回 `"C~a~N~d~Y~~b~O~w~L"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L")`)
}})
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` 应该返回 `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS")`)
}})
```

# --seed--

## --seed-contents--

```py
def spookify(boo):

    return boo
```

# --solutions--

```py
def spookify(boo):
    replaced = boo.replace("_", "~").replace("-", "~")
    
    result = []
    capitalize = True

    for char in replaced:
        if char == "~":
            result.append(char)
        else:
            result.append(char.upper() if capitalize else char.lower())
            capitalize = not capitalize

    return "".join(result)
```
