---
id: 68ee9e3066cfd4eb2328e8a6
title: "挑战 87：矩阵构建器"
challengeType: 29
dashedName: challenge-87
---

# --description--

给定两个整数（行数和列数），返回一个指定大小的由零（`0`）填充的矩阵（数组的数组）。

例如，给定 `2` 和 `3`，返回：

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`build_matrix(2, 3)` 应该返回 `[[0, 0, 0], [0, 0, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(2, 3), [[0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(3, 2)` 应该返回 `[[0, 0], [0, 0], [0, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(3, 2), [[0, 0], [0, 0], [0, 0]])`)
}})
```

`build_matrix(4, 3)` 应该返回 `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(9, 1)` 应该返回 `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def build_matrix(rows, cols):

    return rows
```

# --solutions--

```py
def build_matrix(rows, cols):

    return [[0 for _ in range(cols)] for _ in range(rows)]
```
