---
id: 68ee9e3066cfd4eb2328e8a4
title: "挑战 85：单词计数器"
challengeType: 29
dashedName: challenge-85
---

# --description--

给定一个句子字串，返回句子中单词的数量。

- 单词是任何非空格字符的序列，并由单个空格分隔。

# --hints--

`count_words("Hello world")` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("Hello world"), 2)`)
}})
```

`count_words("The quick brown fox jumps over the lazy dog.")` 应该返回 `9`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("The quick brown fox jumps over the lazy dog."), 9)`)
}})
```

`count_words("I like coding challenges!")` 应该返回 `4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("I like coding challenges!"), 4)`)
}})
```

`count_words("Complete the challenge in JavaScript and Python.")` 应该返回 `7`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("Complete the challenge in JavaScript and Python."), 7)`)
}})
```

`count_words("The missing semi-colon crashed the entire internet.")` 应该返回 `7`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_words("The missing semi-colon crashed the entire internet."), 7)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_words(sentence):

    return sentence
```

# --solutions--

```py
def count_words(sentence):

    return len(sentence.split(' '))
```
