---
id: 68ee9e3066cfd4eb2328e8a7
title: "挑战 88：工作日查找器"
challengeType: 29
dashedName: challenge-88
---

# --description--

给定格式为 `YYYY-MM-DD` 的字串日期，返回星期几。

有效的返回天数有：

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

务必忽略时区。

# --hints--

`get_weekday("2025-11-06")` 应该返回 `Thursday`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2025-11-06"), "Thursday")`)
}})
```

`get_weekday("1999-12-31")` 应该返回 `Friday`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1999-12-31"), "Friday")`)
}})
```

`get_weekday("1111-11-11")` 应该返回 `Saturday`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1111-11-11"), "Saturday")`)
}})
```

`get_weekday("2112-12-21")` 应该返回 `Wednesday`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2112-12-21"), "Wednesday")`)
}})
```

`get_weekday("2345-10-01")` 应该返回 `Monday`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2345-10-01"), "Monday")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_weekday(date_string):

    return date_string
```

# --solutions--

```py
import datetime
def get_weekday(date_string):
    days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"]
    year, month, day = map(int, date_string.split("-"))
    date = datetime.date(year, month, day)

    return days[date.weekday() % 7 + 1 if date.weekday() != 6 else 0]
```
