---
id: 68f6587287ad1f4ad39b0c85
title: "挑战 99：指纹测试"
challengeType: 29
dashedName: challenge-99
---

# --description--

给定两个表现指纹的字串，使用以下规则确定它们是否匹配：

- 每个指纹将仅由小写字母组成（`a-z`）。
- 如果两个指纹满足以下条件，则视为匹配：
  - 它们长度相同。
  - 不同的字符数不超过指纹长度的 10%。

# --hints--

`is_match("helloworld", "helloworld")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("helloworld", "helloworld"), True)`)
}})
```

`is_match("helloworld", "helloworlds")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("helloworld", "helloworlds"), False)`)
}})
```

`is_match("helloworld", "jelloworld")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("helloworld", "jelloworld"), True)`)
}})
```

`is_match("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"), True)`)
}})
```

`is_match("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"), True)`)
}})
```

`is_match("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_match("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_match(fingerprint_a, fingerprint_b):

    return fingerprint_a
```

# --solutions--

```py
def is_match(fingerprint_a, fingerprint_b):
    if len(fingerprint_a) != len(fingerprint_b):
        return False

    mismatches = 0

    for c1, c2 in zip(fingerprint_a, fingerprint_b):
        if c1 != c2:
            mismatches += 1
            if mismatches > len(fingerprint_a) // 10:
                return False

    return True
```
