---
id: 68f6587287ad1f4ad39b0c82
title: "挑战 96：今天是周末吗？"
challengeType: 29
dashedName: challenge-96
---

# --description--

给定格式为 `"YYYY-MM-DD"` 的日期，返回距离周末剩余的天数。

- 周末从星期六开始。
- 如果给定的日期是星期六或星期日，返回 `"It's the weekend!"`。
- 否则，返回 `"X days until the weekend."`，其中 `X` 是距离星期六的天数。
- 如果 `X` 是 `1`，使用 `"day"`（单数）而不是 `"days"`（复数）。
- 确保计算忽略你的局部时区。

# --hints--

`days_until_weekend("2025-11-14")` 应该返回 `"1 day until the weekend."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-11-14"), "1 day until the weekend.")`)
}})
```

`days_until_weekend("2025-01-01")` 应该返回 `"3 days until the weekend."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-01-01"), "3 days until the weekend.")`)
}})
```

`days_until_weekend("2025-12-06")` 应该返回 `"It's the weekend!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-12-06"), "It's the weekend!")`)
}})
```

`days_until_weekend("2026-01-27")` 应该返回 `"4 days until the weekend."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-01-27"), "4 days until the weekend.")`)
}})
```

`days_until_weekend("2026-09-07")` 应该返回 `"5 days until the weekend."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-09-07"), "5 days until the weekend.")`)
}})
```

`days_until_weekend("2026-11-29")` 应该返回 `"It's the weekend!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-11-29"), "It's the weekend!")`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_weekend(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime

def days_until_weekend(date_string):
    date = datetime.strptime(date_string, "%Y-%m-%d").date()
    day_of_week = date.weekday()

    if day_of_week == 5 or day_of_week == 6:
        return "It's the weekend!"

    days_until_saturday = 5 - day_of_week
    return f"{days_until_saturday} day{'s' if days_until_saturday != 1 else ''} until the weekend."
```
