---
id: 68f6587287ad1f4ad39b0c80
title: "挑战 94：电子邮件签名生成器"
challengeType: 29
dashedName: challenge-94
---

# --description--

给定一个人的名字、职称和公司字串，返回一个电子邮件签名作为单个字串，使用以下规则：

- 名称应首先出现，前面加上一个前缀，该前缀取决于名称的首字母。对于以（不区分大小写）开头的名称：
  - `A-I`：使用 `>>` 作为前缀。
  - `J-R`：使用 `--` 作为前缀。
  - `S-Z`：使用 `::` 作为前缀。
- 名称后应跟一个逗号和空格（`, `）。
- 标题和公司名称应跟随逗号和空格，用 `" at "`（两边带空格）分隔。

例如，给定 `"Quinn Waverly"`、`"Founder and CEO"` 和 `"TechCo"`，返回 `"--Quinn Waverly, Founder and CEO at TechCo"`。

# --hints--

`generate_signature("Quinn Waverly", "Founder and CEO", "TechCo")` 应该返回 `"--Quinn Waverly, Founder and CEO at TechCo"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo")`)
}})
```

`generate_signature("Alice Reed", "Engineer", "TechCo")` 应该返回 `">>Alice Reed, Engineer at TechCo"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo")`)
}})
```

`generate_signature("Tina Vaughn", "Developer", "example.com")` 应该返回 `"::Tina Vaughn, Developer at example.com"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com")`)
}})
```

`generate_signature("B. B.", "Product Tester", "AcmeCorp")` 应该返回 `">>B. B., Product Tester at AcmeCorp"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp")`)
}})
```

`generate_signature("windstorm", "Cloud Architect", "Atmospheronics")` 应该返回 `"::windstorm, Cloud Architect at Atmospheronics"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_signature(name, title, company):

    return name
```

# --solutions--

```py
def generate_signature(name, title, company):
    first_letter = name[0].upper()
    if first_letter in "ABCDEFGHI":
        prefix = ">>"
    elif first_letter in "JKLMNOPQR":
        prefix = "--"
    else:
        prefix = "::"

    return f"{prefix}{name}, {title} at {company}"
```
