---
id: 68f6587287ad1f4ad39b0c7c
title: "挑战 90：字符限制"
challengeType: 29
dashedName: challenge-90
---

# --description--

在此挑战中，你将获得一个字串，需要判断它是否适合发布在社交媒体上。根据给定规则返回以下字串：

- `"short post"` 如果它符合 40 字符限制。
- `"long post"` 如果它超过 40 个字符且适合在 80 个字符限制内。
- `"invalid post"` 如果它太长而无法符合任一限制。

# --hints--

`can_post("Hello world")` 应该返回 `"short post"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` 应该返回 `"long post"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` 应该返回 `"invalid post"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
