---
id: 68f6587287ad1f4ad39b0c84
title: "挑战 98：矩形计数"
challengeType: 29
dashedName: challenge-98
---

# --description--

给定两个表现矩形宽度和高度的正整数，确定在给定矩形中可以容纳多少个矩形。

- 只计算宽度和高度为整数的矩形。

例如，给定 `1` 和 `3`，返回 `6`。三个 1x1 矩形，两个 1x2 矩形，和一个 1x3 矩形。

# --hints--

`count_rectangles(1, 3)` 应该返回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 3), 6)`)
}})
```

`count_rectangles(3, 2)` 应该返回 `18`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(3, 2), 18)`)
}})
```

`count_rectangles(1, 2)` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 2), 3)`)
}})
```

`count_rectangles(5, 4)` 应该返回 `150`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(5, 4), 150)`)
}})
```

`count_rectangles(11, 19)` 应该返回 `12540`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(11, 19), 12540)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_rectangles(width, height):

    return width
```

# --solutions--

```py
def count_rectangles(width, height):
    total = 0
    for w in range(1, width + 1):
        for h in range(1, height + 1):
            total += (width - w + 1) * (height - h + 1)
    return total
```
