---
id: 68f6587287ad1f4ad39b0c7f
title: "挑战 93：元音和构造"
challengeType: 29
dashedName: challenge-93
---

# --description--

给定一个字串，返回一个数组，其中包含字串中元音和辅音的数量。

- 元音由任意大小写的 `a`、`e`、`i`、`o`、`u` 组成。
- 辅音由任何分支中的所有其他字母组成。
- 忽略任何非字母字符。

例如，给定 `"Hello World"`，返回 `[3, 7]`。

# --hints--

`count("Hello World")` 应该返回 `[3, 7]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("JavaScript")` 应该返回 `[3, 7]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("JavaScript"), [3, 7])`)
}})
```

`count("Python")` 应该返回 `[1, 5]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Python"), [1, 5])`)
}})
```

`count("freeCodeCamp")` 应该返回 `[5, 7]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("freeCodeCamp"), [5, 7])`)
}})
```

`count("Hello, World!")` 应该返回 `[3, 7]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("The quick brown fox jumps over the lazy dog.")` 应该返回 `[11, 24]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

# --seed--

## --seed-contents--

```py
def count(s):

    return s
```

# --solutions--

```py
def count(s):
    vowels = "aeiou"
    consonants = "bcdfghjklmnpqrstvwxyz"
    v = 0
    c = 0

    for char in s:
        ch = char.lower()
        if ch in vowels:
            v += 1
        if ch in consonants:
            c += 1

    return [v, c]
```
