---
id: 68f6587287ad1f4ad39b0c7e
title: "挑战 92：扩展名提取器"
challengeType: 29
dashedName: challenge-92
---

# --description--

给定一个表现文件名的字串，返回该文件的扩展名。

- 扩展名是文件名中最后一个句点（`.`）之后的部分。
- 如果文件名不包含句点或以句点结尾，返回 `"none"`。
- 扩展名应保持原样返回，保留分支。

# --hints--

`get_extension("document.txt")` 应该返回 `"txt"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("document.txt"), "txt")`)
}})
```

`get_extension("README")` 应该返回 `"none"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("README"), "none")`)
}})
```

`get_extension("image.PNG")` 应该返回 `"PNG"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("image.PNG"), "PNG")`)
}})
```

`get_extension(".gitignore")` 应该返回 `"gitignore"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension(".gitignore"), "gitignore")`)
}})
```

`get_extension("archive.tar.gz")` 应该返回 `"gz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("archive.tar.gz"), "gz")`)
}})
```

`get_extension("final.draft.")` 应该返回 `"none"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("final.draft."), "none")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_extension(filename):

    return filename
```

# --solutions--

```py
def get_extension(filename):
    last_dot = filename.rfind('.')
    
    if last_dot == -1 or last_dot == len(filename) - 1:
        return "none"
    
    return filename[last_dot + 1:]
```
