---
id: 68f6587287ad1f4ad39b0c83
title: "挑战 97：最大公约数"
challengeType: 29
dashedName: challenge-97
---

# --description--

给定两个正整数，返回它们的最大公因数（GCD）。

- 两个整数的最大公约数是能够整除这两个数且不留余数的最大数字。

例如，`4` 的除数是 `1`、`2` 和 `4`。`6` 的除数是 `1`、`2`、`3` 和 `6`。所以给定 `4` 和 `6`，返回 `2`，这是两个除数集合中都出现的最大数字。

# --hints--

`gcd(4, 6)` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(4, 6), 2)`)
}})
```

`gcd(20, 15)` 应该返回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(20, 15), 5)`)
}})
```

`gcd(13, 17)` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(13, 17), 1)`)
}})
```

`gcd(654, 456)` 应该返回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(654, 456), 6)`)
}})
```

`gcd(3456, 4320)` 应该返回 `864`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(3456, 4320), 864)`)
}})
```

# --seed--

## --seed-contents--

```py
def gcd(x, y):

    return x
```

# --solutions--

```py
def gcd(x, y):
    while y != 0:
        x, y = y, x % y
    return x
```
