---
id: 68ffb91507a5b645769328c9
title: "挑战 106：消息验证器"
challengeType: 29
dashedName: challenge-106
---

# --description--

给定一个消息字串和一个验证字串，确定该消息是否有效。

- 如果消息中的每个单词都按顺序以验证字串中对应的字母开头，则该消息是有效的。
- 字母不区分大小写。
- 消息中的单词由单个空格分隔。

# --hints--

`is_valid_message("hello world", "hw")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("hello world", "hw"), True)`)
}})
```

`is_valid_message("ALL CAPITAL LETTERS", "acl")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("ALL CAPITAL LETTERS", "acl"), True)`)
}})
```

`is_valid_message("Coding challenge are boring.", "cca")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("Coding challenge are boring.", "cca"), False)`)
}})
```

`is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"), True)`)
}})
```

`is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_message("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_message(message, validation):

    return message
```

# --solutions--

```py
def is_valid_message(message, validation):
    words = message.split()
    
    if len(words) != len(validation):
        return False

    for word, val_letter in zip(words, validation):
        if word[0].lower() != val_letter.lower():
            return False

    return True
```
