---
id: 68ffb91507a5b645769328cc
title: "挑战 109：我到底几岁了？"
challengeType: 29
dashedName: challenge-109
---

# --description--

给定某人生日的日期，格式为 `YYYY-MM-DD`，返回该人在 2025 年 11 月 27 日的年龄。

- 假设所有生日均为 2025 年 11 月 27 日之前的有效日期。
- 将年龄作为整数返回。
- 请确保考虑该人在 2025 年是否已经过生日。

# --hints--

`calculate_age("2000-11-20")` 应该返回 `25`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-11-20"), 25)`)
}})
```

`calculate_age("2000-12-01")` 应该返回 `24`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-12-01"), 24)`)
}})
```

`calculate_age("2014-10-25")` 应该返回 `11`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2014-10-25"), 11)`)
}})
```

`calculate_age("1994-01-06")` 应该返回 `31`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-01-06"), 31)`)
}})
```

`calculate_age("1994-12-14")` 应该返回 `30`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-12-14"), 30)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_age(birthday):

    return birthday
```

# --solutions--

```py
from datetime import datetime
def calculate_age(birthday):
    today = datetime(2025, 11, 27)
    birth_date = datetime.strptime(birthday, "%Y-%m-%d")

    age = today.year - birth_date.year

    if (today.month, today.day) < (birth_date.month, birth_date.day):
        age -= 1

    return age
```
