---
id: 68ffb91507a5b645769328c5
title: "挑战 102：最长单词"
challengeType: 29
dashedName: challenge-102
---

# --description--

给定一个句子字串，返回句子中最长的单词。

- 单词 由 单个 空格 分隔。
- 只有字母（`a-z`，不区分大小写）计入单词长度。
- 如果有多个长度相同的单词，返回第一个出现的单词。
- 返回字串中出现的单词，去除标点符号。

# --hints--

`longest_word("The quick red fox")` 应该返回 `"quick"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("The quick red fox"), "quick")`)
}})
```

`longest_word("Hello coding challenge.")` 应该返回 `"challenge"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Hello coding challenge."), "challenge")`)
}})
```

`longest_word("Do Try This At Home.")` 应该返回 `"This"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Do Try This At Home."), "This")`)
}})
```

`longest_word("This sentence... has commas, ellipses, and an exclamation point!")` 应该返回 `"exclamation"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation")`)
}})
```

`longest_word("A tie? No way!")` 应该返回 `"tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("A tie? No way!"), "tie")`)
}})
```

`longest_word("Wouldn't you like to know.")` 应该返回 `"Wouldnt"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Wouldn't you like to know."), "Wouldnt")`)
}})
```

# --seed--

## --seed-contents--

```py
def longest_word(sentence):

    return sentence
```

# --solutions--

```py
import re
def longest_word(sentence):
    words = sentence.split(" ")
    longest = ""

    for word in words:
        cleaned = re.sub(r'[^a-zA-Z]', '', word)
        if len(cleaned) > len(longest):
            longest = cleaned

    return longest
```
