---
id: 68ffb91507a5b645769328ca
title: "挑战 107：FizzBuzz"
challengeType: 29
dashedName: challenge-107
---

# --description--

给定一个整数（`n`），返回一个从 `1` 到 `n`（含） 的整数数组，替换是以下倍数的数字：

- 3 和 `"Fizz"`。
- 5 和 `"Buzz"`。
- 3 和 5 与 `"FizzBuzz"`。

# --hints--

`fizz_buzz(2)` 应该返回 `[1, 2]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz(2), [1, 2])`)
}})
```

`fizz_buzz(4)` 应该返回 `[1, 2, "Fizz", 4]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz(4), [1, 2, "Fizz", 4])`)
}})
```

`fizz_buzz(8)` 应该返回 `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz(8), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8])`)
}})
```

`fizz_buzz(20)` 应该返回 `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz(20), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])`)
}})
```

`fizz_buzz(50)` 应该返回 `[1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz(50), [1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"])`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz(n):

    return n
```

# --solutions--

```py
def fizz_buzz(n):
    result = []
    for i in range(1, n+1):
        if i % 15 == 0:
            result.append("FizzBuzz")
        elif i % 3 == 0:
            result.append("Fizz")
        elif i % 5 == 0:
            result.append("Buzz")
        else:
            result.append(i)
    return result
```
