---
id: 69162d64f96574d9bb629f03
title: "挑战 118：日期格式化器"
challengeType: 29
dashedName: challenge-118
---

# --description--

给定格式为 `"Month day, year"` 的日期，返回格式为 `"YYYY-MM-DD"` 的日期。

- 给定的月份将是完整的英文月份名称。例如：`"January"`、`"February"` 等。
- 在返回值中，如果需要，使用前导零填充月份和日期以确保两位数。

例如，给定 `"December 6, 2025"`，返回 `"2025-12-06"`。

# --hints--

`format_date("December 6, 2025")` 应该返回 `"2025-12-06"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("December 6, 2025"), "2025-12-06")`)
}})
```

`format_date("January 1, 2000")` 应该返回 `"2000-01-01"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("January 1, 2000"), "2000-01-01")`)
}})
```

`format_date("November 11, 1111")` 应该返回 `"1111-11-11"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("November 11, 1111"), "1111-11-11")`)
}})
```

`format_date("September 7, 512")` 应该返回 `"512-09-07"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("September 7, 512"), "512-09-07")`)
}})
```

`format_date("May 4, 1950")` 应该返回 `"1950-05-04"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("May 4, 1950"), "1950-05-04")`)
}})
```

`format_date("February 29, 1992")` 应该返回 `"1992-02-29"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("February 29, 1992"), "1992-02-29")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_date(date_string):

    return date_string
```

# --solutions--

```py
def format_date(date_string):
    months = {
        "January": "01",
        "February": "02",
        "March": "03",
        "April": "04",
        "May": "05",
        "June": "06",
        "July": "07",
        "August": "08",
        "September": "09",
        "October": "10",
        "November": "11",
        "December": "12"
    }

    month, day, year = date_string.replace(",", "").split(" ")
    day = day.zfill(2)

    return f"{year}-{months[month]}-{day}"
```
