---
id: 69162d64f96574d9bb629f02
title: "挑战 117：对称差集"
challengeType: 29
dashedName: challenge-117
---

# --description--

给定两个数组，返回一个包含它们对称差的新建数组。

- 两个集合的对称差是出现在任一集合中但不同时出现在两个集合中的值的集合。
- 返回值按它们在输入数组中首次出现的顺序。

# --hints--

`difference([1, 2, 3], [3, 4, 5])` 应该返回 `[1, 2, 4, 5]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5])`)
}})
```

`difference(["a", "b"], ["c", "b"])` 应该返回 `["a", "c"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"])`)
}})
```

`difference([1, "a", 2], [2, "b", "a"])` 应该返回 `[1, "b"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"])`)
}})
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` 应该返回 `[2, 4, 6, 8]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8])`)
}})
```

# --seed--

## --seed-contents--

```py
def difference(arr1, arr2):

    return arr1
```

# --solutions--

```py
def difference(arr1, arr2):
    diff = []

    for v in arr1:
        if v not in arr2 and v not in diff:
            diff.append(v)

    for v in arr2:
        if v not in arr1 and v not in diff:
            diff.append(v)

    return diff
```
