---
id: 69162d64f96574d9bb629f00
title: "挑战 115：Markdown 有序列表项转换器"
challengeType: 29
dashedName: challenge-115
---

# --description--

给定一个表示 Markdown 中有序列表项的字串，返回等效的超文本标记语言字串。

Markdown 中有效的有序列表项必须：

- 以零个或多个空格开头，后跟
- 一个数字（1 或更大）和一个句点（`.`），后跟
- 至少一个空格，然后
- 列表项文本。

如果字串没有上述的精确形式，返回 `"Invalid format"`。否则，将列表项文本包裹在 `li` 标签中并返回该字串。

例如，给定 `"1. My item"`，返回 `"<li>My item</li>"`。

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签。查看浏览器控制台以查看包含标签的日志。

# --hints--

`convert_list_item("1. My item")` 应该返回 `"<li>My item</li>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` 应该返回 `"<li>Another item</li>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` 应该返回 `"Invalid format"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` 应该返回 `"<li>list item text</li>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` 应该返回 `"Invalid format"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` 应该返回 `"Invalid format"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
