---
id: 69162d64f96574d9bb629efc
title: "挑战 111：球的轨迹"
challengeType: 29
dashedName: challenge-111
---

# --description--

今天的挑战灵感来自于 1972 年 11 月 29 日发布的游戏 Pong。

给定一个包含球的位置（`2`）和球的先前位置（`1`）的矩阵（数组的数组），返回球的下一个位置的矩阵索引。

- 球总是沿直线移动。
- 移动方向由球从 `1` 移动到 `2` 的方式决定。
- 矩阵的边缘被视为墙壁。如果球撞击：
  - 顶部或底部墙壁时，它通过反转其垂直方向来弹跳。
  - 当碰到左墙或右墙时，它通过反转其水平方向来弹跳。
  - 角落，它通过反转两个方向来弹跳。

# --hints--

`get_next_location([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])` 应该返回 `[2, 3]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3])`)
}})
```

`get_next_location([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])` 应该返回 `[3, 0]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0])`)
}})
```

`get_next_location([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])` 应该返回 `[1, 2]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2])`)
}})
```

`get_next_location([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])` 应该返回 `[1, 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1])`)
}})
```

`get_next_location([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])` 应该返回 `[2, 2]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_next_location(matrix):

    return matrix
```

# --solutions--

```py
def get_next_location(matrix):
    prev = None
    curr = None

    for r in range(len(matrix)):
        for c in range(len(matrix[0])):
            if matrix[r][c] == 1:
                prev = (r, c)
            if matrix[r][c] == 2:
                curr = (r, c)

    prev_row, prev_col = prev
    curr_row, curr_col = curr

    dir_x = curr_col - prev_col
    dir_y = curr_row - prev_row

    next_row = curr_row + dir_y
    next_col = curr_col + dir_x

    max_row = len(matrix) - 1
    max_col = len(matrix[0]) - 1

    if next_col < 0 or next_col > max_col:
        dir_x *= -1
    if next_row < 0 or next_row > max_row:
        dir_y *= -1

    next_row = curr_row + dir_y
    next_col = curr_col + dir_x

    return [next_row, next_col]
```
