---
id: 69162d64f96574d9bb629f04
title: "挑战 119：字串压缩"
challengeType: 29
dashedName: challenge-119
---

# --description--

给定一个字串 sentence，返回该句子的压缩版本，其中连续的副本词被替换为该词，后跟其重复次数的括号。

- 只有连续的副本会被压缩。
- 单词由单个空格分隔。

例如，给定 `"yes yes yes please"`，返回 `"yes(3) please"`。

# --hints--

`compress_string("yes yes yes please")` 应该返回 `"yes(3) please"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("yes yes yes please"), "yes(3) please")`)
}})
```

`compress_string("I have have have apples")` 应该返回 `"I have(3) apples"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("I have have have apples"), "I have(3) apples")`)
}})
```

`compress_string("one one three and to the the the the")` 应该返回 `"one(2) three and to the(4)"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("one one three and to the the the the"), "one(2) three and to the(4)")`)
}})
```

`compress_string("route route route route route route tee tee tee tee tee tee")` 应返回 `"route(6) tee(6)"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compress_string("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)")`)
}})
```

# --seed--

## --seed-contents--

```py
def compress_string(sentence):

    return sentence
```

# --solutions--

```py
def compress_string(sentence):
    words = sentence.split(" ")
    result = []
    count = 1

    for i in range(len(words)):
        if i < len(words) - 1 and words[i] == words[i + 1]:
            count += 1
        else:
            if count > 1:
                result.append(f"{words[i]}({count})")
            else:
                result.append(words[i])
            count = 1

    return " ".join(result)
```
