---
id: 69162d64f96574d9bb629f01
title: "挑战 116：排列计数"
challengeType: 29
dashedName: challenge-116
---

# --description--

给定一个字串，返回可以由其字符组成的不同排列数。

- 排列是字串中字符的任意重新排序。
- 不要计算副本排列。
- 如果字串包含重复的字符，重复的排列只应计数一次。
- 该字串将仅包含字母（`A-Z`、`a-z`）。

例如，给定 `"abb"`，返回 `3`，因为有三种不同的方式排列这些字母：`"abb"`、`"bab"` 和 `"bba"`。

# --hints--

`count_permutations("abb")` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` 应该返回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` 应该返回 `630`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` 应该返回 `39916800`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
