---
id: 69162d64f96574d9bb629efe
title: "挑战 113：英里转公里"
challengeType: 29
dashedName: challenge-113
---

# --description--

给定以英里为单位的距离数字，返回等效的公里数。

- 输入将始终是非负数。
- 1 英里相等 1.60934 公里。
- 将结果四舍五入到两位小数。
- 从四舍五入的结果中移除不必要的尾随零。

# --hints--

`convert_to_km(1)` 应该返回 `1.61`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(1), 1.61)`)
}})
```

`convert_to_km(21)` 应该返回 `33.8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(21), 33.8)`)
}})
```

`convert_to_km(3.5)` 应该返回 `5.63`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(3.5), 5.63)`)
}})
```

`convert_to_km(0)` 应该返回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0), 0)`)
}})
```

`convert_to_km(0.621371)` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0.621371), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_km(miles):

    return miles
```

# --solutions--

```py
def convert_to_km(miles):
    km = miles * 1.60934
    return round(km, 2)
```
