---
id: 691b559495c5cb5a37b9b487
title: "挑战 127：速度查看"
challengeType: 29
dashedName: challenge-127
---

# --description--

根据你以每小时英里数（MPH）行驶的速度和以每小时公里数（KPH）表示的限速，判断你是否超速以及你会收到警告还是罚单。

- 1 英里相等 1.60934 公里。
- 如果你的速度小于或相等于限速，返回 `"Not Speeding"`。
- 如果你超速不超过 5 KPH，返回 `"Warning"`。
- 如果你超速超过限速 5 KPH，返回 `"Ticket"`。

# --hints--

`speed_check(30, 70)` 应该返回 `"Not Speeding"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` 应该返回 `"Warning"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` 应该返回 `"Not Speeding"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` 应该返回 `"Ticket"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` 应该返回 `"Warning"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` 应该返回 `"Ticket"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
