---
id: 691b559495c5cb5a37b9b489
title: "挑战 129：Markdown 块引用解析器"
challengeType: 29
dashedName: challenge-129
---

# --description--

给定一个包含 Markdown 中引用的字串，返回等效的超文本标记语言字串。

Markdown 中的引用是任何以以下内容开头的行：

- 以零个或多个空格开始
- 后面跟着一个大于号（`>`）
- 然后，一个或多个空格
- 最后，引用文本。

返回由开启和关闭超文本标记语言 `blockquote` 标签包围的引用文本。

例如，给定 `"> This is a quote"`，返回 `<blockquote>This is a quote</blockquote>`。

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签。查看浏览器控制台以查看包含标签的日志。

# --hints--

`parse_blockquote("> This is a quote")` 应该返回 `"<blockquote>This is a quote</blockquote>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` 应该返回 `"<blockquote>This is also a quote</blockquote>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` 应该返回 `"<blockquote>So  Is  This</blockquote>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
