---
id: 691b559495c5cb5a37b9b485
title: "挑战 125：生命游戏"
challengeType: 29
dashedName: challenge-125
---

# --description--

给定一个表现康威生命游戏当前状态的矩阵（数组的数组），使用以下规则返回矩阵的下一状态：

- 每个单元格要么是 `1`（存活），要么是 `0`（死亡）。
- 一个单元格的邻居是最多八个周围的单元格（垂直、水平和对角线方向）。
- 边缘的单元格邻居少于八个。

更新每个单元格的规则：

- 任何活细胞如果活邻居少于 2 个则死亡（人口过少）。
- 任何有两个或三个活邻居的活细胞将继续存活。
- 任何有超过 3 个活邻居的活细胞都会死亡（过度繁殖）。
- 任何恰好有三个活邻居的死细胞会变为活细胞（繁殖）。

例如，给定：

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

返回：

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

每个单元格根据存活邻居的数量进行更新。例如，`[0][0]` 保持死亡（2 个存活邻居），`[0][1]` 保持存活（2 个存活邻居），`[0][2]` 死亡（3 个存活邻居），依此类推。

# --hints--

`game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` 应该返回 `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]])`)
}})
```

`game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` 应该返回 `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]])`)
}})
```

`game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` 应该返回 `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]])`)
}})
```

`game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` 应该返回 `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def game_of_life(grid):

    return grid
```

# --solutions--

```py
def game_of_life(grid):
    rows, cols = len(grid), len(grid[0])

    def count_live_neighbors(r, c):
        count = 0
        for i in range(r-1, r+2):
            for j in range(c-1, c+2):
                if 0 <= i < rows and 0 <= j < cols and (i != r or j != c):
                    count += grid[i][j]
        return count

    next_state = [row[:] for row in grid]

    for r in range(rows):
        for c in range(cols):
            live_neighbors = count_live_neighbors(r, c)

            if grid[r][c] == 1:
                if live_neighbors < 2 or live_neighbors > 3:
                    next_state[r][c] = 0
            else:
                if live_neighbors == 3:
                    next_state[r][c] = 1

    return next_state
```
