---
id: 691b559495c5cb5a37b9b488
title: "挑战 128：辅音计数"
challengeType: 29
dashedName: challenge-128
---

# --description--

给定一个字串和一个目标数字，确定该字串是否恰好包含目标数量的构造。

- 辅音是除任何大小写的 `"a"`、`"e"`、`"i"`、`"o"` 和 `"u"` 之外的所有字母字符。
- 忽略数字、标点符号、空格和其他非字母字符时计数。

# --hints--

`has_consonant_count("helloworld", 7)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("helloworld", 7), True)`)
}})
```

`has_consonant_count("eieio", 5)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("eieio", 5), False)`)
}})
```

`has_consonant_count("freeCodeCamp Rocks!", 11)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("freeCodeCamp Rocks!", 11), True)`)
}})
```

`has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24), False)`)
}})
```

`has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_consonant_count(text, target):

    return text
```

# --solutions--

```py
def has_consonant_count(text, target):
    vowels = {"a", "e", "i", "o", "u"}
    count = 0

    for ch in text.lower():
        if ch.isalpha() and ch not in vowels:
            count += 1

    return count == target
```
