---
id: 691b559495c5cb5a37b9b483
title: "挑战 123：罗马数字生成器"
challengeType: 29
dashedName: challenge-123
---

# --description--

给定一个整数，返回其对应的罗马数字值。

罗马数字使用这些符号：

| Symbol | Value |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

罗马数字从最大到最小，从左到右书写，遵循以下规则：

- 当符号后面跟着一个相等或更小的值时，使用加法。例如，`18` 写作 `XVIII`（10 + 5 + 1 + 1 + 1 = 18）。
- 当较小的符号出现在较大符号之前时，使用减法来表现任何位值中的 4 或 9。例如，19 写作 `XIX`（10 + (10 - 1)）。
- 没有符号可以连续重复超过三次。当你需要连续写一个符号超过三次时，使用减法。因此你能写的最大数字是 3999。

这是另一个例子：给定 `1464`，返回 `"MCDLXIV"`（1000 + (500 - 100) + 50 + 10 + (5 - 1)）。

# --hints--

`to_roman(18)` 应该返回 `"XVIII"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` 应该返回 `"XIX"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` 应该返回 `"MCDLXIV"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` 应该返回 `"MMXXV"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` 应返回 `"MMMCMXCIX"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```
