---
id: 691b559495c5cb5a37b9b486
title: "挑战 126：首字母大写"
challengeType: 29
dashedName: challenge-126
---

# --description--

给定一个字串 `title`，返回一个按照以下规则格式化为标题形式的新建字串：

- 将每个单词的首字母大写。
- 将每个单词中的其他所有字母转换为小写。
- 单词总是由单个空格分隔。

# --hints--

`title_case("hello world")` 应该返回 `"Hello World"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("hello world"), "Hello World")`)
}})
```

`title_case("the quick brown fox")` 应该返回 `"The Quick Brown Fox"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("the quick brown fox"), "The Quick Brown Fox")`)
}})
```

`title_case("JAVASCRIPT AND PYTHON")` 应该返回 `"Javascript And Python"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("JAVASCRIPT AND PYTHON"), "Javascript And Python")`)
}})
```

`title_case("AvOcAdO tOAst fOr brEAkfAst")` 应该返回 `"Avocado Toast For Breakfast"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast")`)
}})
```

# --seed--

## --seed-contents--

```py
def title_case(title):

    return title
```

# --solutions--

```py
def title_case(title):
    return " ".join(
        w[:1].upper() + w[1:].lower()
        for w in title.split(" ")
    )
```
