---
id: 691b559495c5cb5a37b9b484
title: "挑战 124：库存更新"
challengeType: 29
dashedName: challenge-124
---

# --description--

给定一个表现你商店库存的二维数组，以及另一个表现你收到的货物的二维数组，返回你更新后的库存。

- 数组中的每个元素将具有以下形式：`[quantity, "item"]`，其中 `quantity` 是一个整数，`"item"` 是一个字串。
- 通过添加货运中任何匹配项的数量来更新库存中的项。
- 如果收到的 `item` 不存在于当前库存中，将其作为新项添加到库存末尾。
- 以给定顺序返回库存，新建项按它们在货物中出现的顺序放在末尾。

例如，给定一个库存为 `[[2, "apples"], [5, "bananas"]]` 和一个货物为 `[[1, "apples"], [3, "bananas"]]`，返回 `[[3, "apples"], [8, "bananas"]]`。

# --hints--

`update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` 应该返回 `[[3, "apples"], [8, "bananas"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]])`)
}})
```

`update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` 应该返回 `[[3, "apples"], [8, "bananas"], [4, "oranges"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]])`)
}})
```

`update_inventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` 应该返回 `[[10, "apples"], [30, "bananas"], [20, "oranges"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]])`)
}})
```

`update_inventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` 应该返回 `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def update_inventory(inventory, shipment):

    return inventory
```

# --solutions--

```py
def update_inventory(inventory, shipment):
    inventory_map = {item: i for i, (qty, item) in enumerate(inventory)}

    for qty, item in shipment:
        if item in inventory_map:
            index = inventory_map[item]
            inventory[index][0] += qty
        else:
            inventory.append([qty, item])
            inventory_map[item] = len(inventory) - 1

    return inventory
```
