---
id: 691b559495c5cb5a37b9b481
title: "挑战 121：最频繁"
challengeType: 29
dashedName: challenge-121
---

# --description--

给定一个元素数组，返回出现频率最高的元素。

- 总是会有一个出现频率最高的元素。

# --hints--

`most_frequent(["a", "b", "a", "c"])` 应该返回 `"a"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(most_frequent(["a", "b", "a", "c"]), "a")`)
}})
```

`most_frequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9])` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(most_frequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9]), 2)`)
}})
```

`most_frequent([True, False, "False", "True", False])` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(most_frequent([True, False, "False", "True", False]), False)`)
}})
```

`most_frequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60])` 应该返回 `40`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(most_frequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60]), 40)`)
}})
```

# --seed--

## --seed-contents--

```py
def most_frequent(arr):

    return arr
```

# --solutions--

```py
def most_frequent(arr):
    freq = {}
    max_count = 0
    most_elem = None

    for el in arr:
        freq[el] = freq.get(el, 0) + 1
        if freq[el] > max_count:
            max_count = freq[el]
            most_elem = el

    return most_elem
```
