---
id: 6925e2068081f40f549ced1b
title: "挑战 137：雪花生成器"
challengeType: 29
dashedName: challenge-137
---

# --description--

给定一个使用换行字符（`\n`）表现换行的多行字串，返回一个新字串，其中每一行都水平镜射并附加到原始行的末尾。

- 通过反转所有字符（包括空格）来镜射一行。

例如，给定 `"* \n *\n* "`，它会记录到控制台为：

```sh
* 
 *
* 
```

返回 `"*  *\n ** \n*  *"`，它会将日志输出到控制台，如下所示：

```sh
*  *
 ** 
*  *
```

请仔细注意给定和返回的字串中的空格。确保不要删除任何空格。

# --hints--

`generate_snowflake("* \n *\n* ")` 应该返回 `"*  *\n ** \n*  *"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` 应该返回 `"X=~~=X"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` 应该返回 `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` 应该返回 `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` 应该返回 `"*  --  *\n * -- * \n*  --  *"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
