---
id: 6925e2068081f40f549ced1d
title: "挑战 139：石头、剪刀、布"
challengeType: 29
dashedName: challenge-139
---

# --description--

给定两个字串，第一个表现玩家 1，第二个表现玩家 2，确定石头剪刀布比赛的获胜者。

- 输入字串将始终是 `"Rock"`、`"Paper"` 或 `"Scissors"`。
- `"Rock"` 胜过 `"Scissors"`。
- `"Paper"` 胜过 `"Rock"`。
- `"Scissors"` 胜过 `"Paper"`。

返回：

- 如果 Player 1 获胜，则显示 `"Player 1 wins"`。
- 如果 Player 2 获胜，则返回 `"Player 2 wins"`。
- `"Tie"` 如果双方玩家选择相同的选项。

# --hints--

`rock_paper_scissors("Rock", "Rock")` 应该返回 `"Tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Rock"), "Tie")`)
}})
```

`rock_paper_scissors("Rock", "Paper")` 应该返回 `"Player 2 wins"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Paper"), "Player 2 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Paper")` 应该返回 `"Player 1 wins"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Paper"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Rock", "Scissors")` 应该返回 `"Player 1 wins"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Scissors"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Scissors")` 应该返回 `"Tie"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Scissors"), "Tie")`)
}})
```

`rock_paper_scissors("Scissors", "Rock")` 应该返回 `"Player 2 wins"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Rock"), "Player 2 wins")`)
}})
```

# --seed--

## --seed-contents--

```py
def rock_paper_scissors(player1, player2):

    return player1
```

# --solutions--

```py
def rock_paper_scissors(player1, player2):
    if player1 == player2:
        return "Tie"
    
    if (player1 == "Rock" and player2 == "Scissors") or \
       (player1 == "Paper" and player2 == "Rock") or \
       (player1 == "Scissors" and player2 == "Paper"):
        return "Player 1 wins"
    else:
        return "Player 2 wins"
```
