---
id: 6925e2068081f40f549ced1a
title: "挑战 136：Markdown 图像解析器"
challengeType: 29
dashedName: challenge-136
---

# --description--

给定 Markdown 中的图像字串，返回等效的超文本标记语言字串。

Markdown 图像具有以下形式：`"![alt text](image_url)"`。其中：

- `alt text` 是图像的描述（`alt` 属性值）。
- `image_url` 是图像的源 URL（`src` 属性值）。

返回一个超文本标记语言 `img` 标签的字串，其中 `src` 设置为图像 URL，`alt` 设置为替代文本。

例如，给定 `"![Cute cat](cat.png)"` 返回 `'<img src="cat.png" alt="Cute cat">'`；

- 确保标签、属性顺序、空格和引用的使用与上述完全相同。

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签——查看浏览器控制台以查看包含标签的日志。

# --hints--

`parse_image("![Cute cat](cat.png)")` 应该 `return` `'<img src="cat.png" alt="Cute cat">'`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` 应该返回 `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` 应该 `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
