---
id: 6925e2068081f40f549ced1c
title: "挑战 138：除数之和"
challengeType: 29
dashedName: challenge-138
---

# --description--

给定一个正整数，返回它所有除数的和。

- 除数是任何能将该数字整除（余数为 `0`）的整数。
- 每个除数只计算一次。

例如，给定 `6`，返回 `12`，因为 `6` 的除数是 `1`、`2`、`3` 和 `6`，它们的和是 `12`。

# --hints--

`sum_divisors(6)` 应该返回 `12`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` 应该返回 `14`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` 应该返回 `56`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` 应该返回 `224`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` 应该返回 `806`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` 应该返回 `23520`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
