---
id: 69272dcf1c24b44fd79137c3
title: "挑战 140：SCREAMING_SNAKE_CASE"
challengeType: 29
dashedName: challenge-140
---

# --description--

给定一个表现变量名的字串，返回转换为 SCREAMING_SNAKE_CASE 的变量名。

给定的变量名将以以下一种形式书写：

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

在上述形式中，单词由下划线（`_`）、连字符（`-`）分隔，或新单词以大写字母开头。

转换为 SCREAMING_SNAKE_CASE：

- 将所有字母转换为大写
- 用下划线（`_`）分隔单词

# --hints--

`to_screaming_snake_case("userEmail")` 应该返回 `"USER_EMAIL"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("userEmail"), "USER_EMAIL")`)
}})
```

`to_screaming_snake_case("UserPassword")` 应该返回 `"USER_PASSWORD"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("UserPassword"), "USER_PASSWORD")`)
}})
```

`to_screaming_snake_case("user_id")` 应该返回 `"USER_ID"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user_id"), "USER_ID")`)
}})
```

`to_screaming_snake_case("user-address")` 应该返回 `"USER_ADDRESS"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user-address"), "USER_ADDRESS")`)
}})
```

`to_screaming_snake_case("username")` 应该返回 `"USERNAME"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("username"), "USERNAME")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_screaming_snake_case(variable_name):

    return variable_name
```

# --solutions--

```py
import re
def to_screaming_snake_case(variable_name):
    temp = re.sub(r'[-_]+', ' ', variable_name)
    temp = re.sub(r'([a-z0-9])([A-Z])', r'\1 \2', temp)
    words = temp.strip().split()
    return '_'.join(words).upper()
```
