---
id: 69272dcf1c24b44fd79137c4
title: "挑战 141：起飞燃油"
challengeType: 29
dashedName: challenge-141
---

# --description--

给定你飞机当前的燃料加仑数和到达目的地所需的燃料升数，确定你应该添加多少额外的燃料加仑。

- 1 加仑相等于 3.78541 升。
- 如果飞机已经有足够的燃料，返回 `0`。
- 你只能添加整加仑。
- 返回数字中不要包含小数。

# --hints--

`fuel_to_add(0, 1)` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(0, 1), 1)`)
}})
```

`fuel_to_add(5, 40)` 应该返回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(5, 40), 6)`)
}})
```

`fuel_to_add(10, 30)` 应该返回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(10, 30), 0)`)
}})
```

`fuel_to_add(896, 20500)` 应该返回 `4520`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(896, 20500), 4520)`)
}})
```

`fuel_to_add(1000, 50000)` 应该返回 `12209`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(1000, 50000), 12209)`)
}})
```

# --seed--

## --seed-contents--

```py
def fuel_to_add(current_gallons, required_liters):

    return current_gallons
```

# --solutions--

```py
import math
def fuel_to_add(current_gallons, required_liters):
    liters_per_gallon = 3.78541
    current_liters = current_gallons * liters_per_gallon
    if current_liters >= required_liters:
        return 0
    liters_needed = required_liters - current_liters
    return math.ceil(liters_needed / liters_per_gallon)
```
