---
id: 69306364df283fcaff2e1ada
title: "挑战 149：vOwElcAsE"
challengeType: 29
dashedName: challenge-149
---

# --description--

给定一个字串，返回一个新建字串，其中所有元音字母转换为大写，所有其他字母字符转换为小写。

- 元音是任何分支中的 `"a"`、`"e"`、`"i"`、`"o"` 和 `"u"`。
- 非字母字符应保持不变。

# --hints--

`vowel_case("vowelcase")` 应该返回 `"vOwElcAsE"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("vowelcase"), "vOwElcAsE")`)
}})
```

`vowel_case("coding is fun")` 应该返回 `"cOdIng Is fUn"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("coding is fun"), "cOdIng Is fUn")`)
}})
```

`vowel_case("HELLO, world!")` 应该返回 `"hEllO, wOrld!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HELLO, world!"), "hEllO, wOrld!")`)
}})
```

`vowel_case("git cherry-pick")` 应该返回 `"gIt chErry-pIck"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("git cherry-pick"), "gIt chErry-pIck")`)
}})
```

`vowel_case("HEAD~1")` 应该返回 `"hEAd~1"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HEAD~1"), "hEAd~1")`)
}})
```

# --seed--

## --seed-contents--

```py
def vowel_case(s):

    return s
```

# --solutions--

```py
def vowel_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char in vowels:
            result += char.upper()
        elif char.isalpha():
            result += char.lower()
        else:
            result += char

    return result
```
