---
id: 69306364df283fcaff2e1ad7
title: "挑战 146：`Left-Handed Seat at the Table`"
challengeType: 29
dashedName: challenge-146
---

# --description--

给定一个 4x2 矩阵（数组的数组）表现用餐的座位安排，确定左撇子可以坐的座位数量。

- 左撇子不能坐在右撇子左侧紧邻的座位上。

在给定的矩阵中：

- `"R"` 是一个由右撇子占据的座位。
- 一个 `"L"` 是一个被左撇子占用的座位。
- `"U"` 是一个空位。
- 只有空闲座位可供坐。
- 最上排的座位面向 `"down"`，最下排的座位面向 `"up"`（像一个表格），所以左右是相对于座位的方向。
- 角落座位旁边只有一个座位。

例如，在给定的矩阵中：

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

左上角的座位不能坐，因为左边有一个右撇子。其他两个空座位可以坐，因为左边没有右撇子。

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` 应该返回 `8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` 应该返回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` 应该返回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
