---
id: 69306364df283fcaff2e1ad6
title: "挑战 145：第 N 个斐波那契数"
challengeType: 29
dashedName: challenge-145
---

# --description--

给定一个整数 `n`，返回斐波那契序列中的第 `n` 个数字。

斐波那契序列是一串数字，其中每个数字都是前述的两个数字之和。序列中的前 10 个数字是 `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34`。

# --hints--

`nth_fibonacci(4)` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(4), 2)`)
}})
```

`nth_fibonacci(10)` 应该返回 `34`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(10), 34)`)
}})
```

`nth_fibonacci(15)` 应该返回 `377`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(15), 377)`)
}})
```

`nth_fibonacci(40)` 应该返回 `63245986`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(40), 63245986)`)
}})
```

`nth_fibonacci(75)` 应该返回 `1304969544928657`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(75), 1304969544928657)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_fibonacci(n):

    return n
```

# --solutions--

```py
def nth_fibonacci(n):
    if n == 1:
        return 0
    if n == 2:
        return 1

    a, b = 0, 1
    for _ in range(3, n + 1):
        a, b = b, a + b
    return b
```
