---
id: 69306364df283fcaff2e1ad8
title: "挑战 147：闰年计算器"
challengeType: 29
dashedName: challenge-147
---

# --description--

给定一个整数年份，判断它是否为闰年。

如果一年满足以下规则，则该年是闰年：

- 该年份能被 4 整除，且
- 该年份不能被 100 整除，除非
- 该年份能被 400 整除。

# --hints--

`is_leap_year(2024)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2024), True)`)
}})
```

`is_leap_year(2023)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2023), False)`)
}})
```

`is_leap_year(2100)` 应返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2100), False)`)
}})
```

`is_leap_year(2000)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2000), True)`)
}})
```

`is_leap_year(1999)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(1999), False)`)
}})
```

`is_leap_year(2040)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2040), True)`)
}})
```

`is_leap_year(2026)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2026), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_leap_year(year):

    return year
```

# --solutions--

```py
def is_leap_year(year):
    if (year % 4 == 0 and year % 100 != 0) or (year % 400 == 0):
        return True
    else:
        return False
```
