---
id: 69373793f5a867f769cde135
title: "挑战 150：Markdown 无序列表解析器"
challengeType: 29
dashedName: challenge-150
---

# --description--

给定 Markdown 中有效无序列表的字串，返回等效的超文本标记语言字串。

无序列表由一个或多个列表项组成。有效的列表项单独占一行且：

- - 以短横线（`"-"`）开头，后跟
- 至少一个空格，然后
- 列表项文本。

该列表作为单个字串给出，换行由换行字符（`"\n"`）分隔。不要在项文本中包含换行字符。

将每个列表项包裹在 HTML `li` 标签中，并将整个列表包裹在 `ul` 标签中。

例如，给定 `"- Item A\n- Item B"`，返回 `"<ul><li>Item A</li><li>Item B</li></ul>"`。

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签。查看浏览器控制台以查看包含标签的日志。

# --hints--

`parse_unordered_list("- Item A\n- Item B")` 应该返回 `"<ul><li>Item A</li><li>Item B</li></ul>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- Item A\\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>")`)
}})
```

`parse_unordered_list("-  JavaScript\n-  Python")` 应该返回 `"<ul><li>JavaScript</li><li>Python</li></ul>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("-  JavaScript\\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>")`)
}})
```

`parse_unordered_list("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` 应该返回 `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- 2 C Flour\\n- 1/2 C Sugar\\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>")`)
}})
```

`parse_unordered_list("- A-1\n- A-2\n- B-1")` 应该返回 `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- A-1\\n- A-2\\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_unordered_list(markdown):

    return markdown
```

# --solutions--

```py
def parse_unordered_list(markdown):
    lines = markdown.split("\n")
    list_items = []

    for line in lines:
        list_items.append(f"<li>{line[1:].strip()}</li>")

    return f"<ul>{''.join(list_items)}</ul>"
```
