---
id: 69373793f5a867f769cde137
title: "挑战 152：循环素数"
challengeType: 29
dashedName: challenge-152
---

# --description--

给定一个整数，判断它是否是循环素数。

循环素数是一个整数，其所有数字的旋转都是素数。

例如，`197` 是一个循环素数，因为其所有数字的旋转：`197`、`971` 和 `719`，都是素数。

# --hints--

`is_circular_prime(197)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(197), True)`)
}})
```

`is_circular_prime(23)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(23), False)`)
}})
```

`is_circular_prime(13)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(13), True)`)
}})
```

`is_circular_prime(89)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(89), False)`)
}})
```

`is_circular_prime(1193)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(1193), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_circular_prime(n):

    return n
```

# --solutions--

```py
import math
def is_prime(n):
    if n < 2:
        return False
    for i in range(2, int(math.isqrt(n)) + 1):
        if n % i == 0:
            return False
    return True

def rotations(n):
    s = str(n)
    return [int(s[i:] + s[:i]) for i in range(len(s))]

def is_circular_prime(n):
    return all(is_prime(rot) for rot in rotations(n))
```
