---
id: 69373793f5a867f769cde136
title: "挑战 151：排序数组？"
challengeType: 29
dashedName: challenge-151
---

# --description--

给定一个数字数组，判断这些数字是按升序排列、降序排列还是都不是。

如果给定的数组是：

- 按升序（从低到高）返回 `"Ascending"`。
- 按降序（从高到低）返回 `"Descending"`。
- 未按升序或降序排序，返回 `"Not sorted"`。

# --hints--

`is_sorted([1, 2, 3, 4, 5])` 应该返回 `"Ascending"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([1, 2, 3, 4, 5]), "Ascending")`)
}})
```

`is_sorted([10, 8, 6, 4, 2])` 应该返回 `"Descending"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([10, 8, 6, 4, 2]), "Descending")`)
}})
```

`is_sorted([1, 3, 2, 4, 5])` 应该返回 `"Not sorted"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([1, 3, 2, 4, 5]), "Not sorted")`)
}})
```

`is_sorted([3.14, 2.71, 1.61, 0.57])` 应该返回 `"Descending"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([3.14, 2.71, 1.61, 0.57]), "Descending")`)
}})
```

`is_sorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])` 应该返回 `"Ascending"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending")`)
}})
```

`is_sorted([0.4, 0.5, 0.3])` 应该返回 `"Not sorted"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(is_sorted([0.4, 0.5, 0.3]), "Not sorted")`)
}})
```

# --seed--

## --seed-contents--

```py
def is_sorted(arr):

    return arr
```

# --solutions--

```py
def is_sorted(arr):
    ascending = True
    descending = True

    for i in range(1, len(arr)):
        if arr[i] < arr[i - 1]:
            ascending = False
        if arr[i] > arr[i - 1]:
            descending = False

    if ascending:
        return "Ascending"
    if descending:
        return "Descending"
    return "Not sorted"
```
