---
id: 69373793f5a867f769cde13a
title: "挑战 155：种植作物"
challengeType: 29
dashedName: challenge-155
---

# --description--

给定一个表现你农场字段大小的整数，以及表现农场字段大小单位的 `"acres"` 或 `"hectares"`，还有一种作物类型，确定你可以在字段中种植多少该类型的植物。

- 1 英亩相等于 4046.86 平方米。
- 1 公顷等于 10,000 平方米。

以下是作为 `input` 提供的作物 `list` 以及单株植物占用的空间：

|Crop|Space per plant|
|-|-|
|`"corn"`|1 square meter|
|`"wheat"`|0.1 square meters|
|`"soybeans"`|0.5 square meters|
|`"tomatoes"`|0.25 square meters|
|`"lettuce"`|0.2 square meters|

返回适合该字段的植物数量，向下取整到最接近的整数植物数。

# --hints--

`get_number_of_plants(1, "acres", "corn")` 应该返回 `4046`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

`get_number_of_plants(2, "hectares", "lettuce")` 应该返回 `100000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

`get_number_of_plants(20, "acres", "soybeans")` 应该返回 `161874`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

`get_number_of_plants(3.75, "hectares", "tomatoes")` 应该返回 `150000`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

`get_number_of_plants(16.75, "acres", "tomatoes")` 应该返回 `271139`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
