---
id: 6939b873185d8e00d453563d
title: "挑战 158：数组交换"
challengeType: 29
dashedName: challenge-158
---

# --description--

给定一个包含两个值的数组，返回一个值已交换的数组。

例如，给定 `["A", "B"]` 返回 `["B", "A"]`。

# --hints--

`array_swap(["A", "B"])` 应该返回 `["B", "A"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["A", "B"]), ["B", "A"])`)
}})
```

`array_swap([25, 20])` 应该返回 `[20, 25]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([25, 20]), [20, 25])`)
}})
```

`array_swap([True, False])` 应该返回 `[False, True]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([True, False]), [False, True])`)
}})
```

`array_swap(["1", 1])` 应该返回 `[1, "1"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["1", 1]), [1, "1"])`)
}})
```

# --seed--

## --seed-contents--

```py
def array_swap(arr):

    return arr
```

# --solutions--

```py
def array_swap(arr):

    return [arr[1], arr[0]]
```
