---
id: 6939b873185d8e00d453563b
title: "挑战 156：奇数还是偶数？"
challengeType: 29
dashedName: challenge-156
---

# --description--

给定一个正整数，如果它是奇数，则返回 `"Odd"`，如果是偶数，则返回 `"Even"`。

# --hints--

`odd_or_even(1)` 应该返回 `"Odd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(1), "Odd")`)
}})
```

`odd_or_even(2)` 应该返回 `"Even"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(2), "Even")`)
}})
```

`odd_or_even(13)` 应该返回 `"Odd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(13), "Odd")`)
}})
```

`odd_or_even(196)` 应该返回 `"Even"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(196), "Even")`)
}})
```

`odd_or_even(123456789)` 应该返回 `"Odd"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(123456789), "Odd")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even(n):

    return n
```

# --solutions--

```py
def odd_or_even(n):
    return "Even" if n % 2 == 0 else "Odd"
```
