---
id: 6939b873185d8e00d453563c
title: "挑战 157：Markdown 链接解析器"
challengeType: 29
dashedName: challenge-157
---

# --description--

给定 Markdown 中的链接字串，返回等效的超文本标记语言字串。

Markdown 图像具有以下形式：`"[link_text](link_url)"`。返回超文本标记语言 `a` 标签的字串，其 `href` 设置为 `link_url`，并且 `link_text` 作为标签内容。

例如，给定 `"[freeCodeCamp](https://freecodecamp.org/)"` 返回 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`；

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签——查看浏览器控制台以查看包含标签的日志。

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` 应该返回 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` 应该 `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` 应该返回 `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
