---
id: 694596b0585c11170ac7c7fa
title: "挑战 162：能量消耗"
challengeType: 29
dashedName: challenge-162
---

# --description--

根据锻炼过程中刻录的卡路里数和锻炼期间你的电子设备使用的瓦时数，确定哪一个消耗的能量更多。

要比较它们，请使用以下换算将两个值都转换为焦耳：

- 1 卡路里等于 4184 焦耳。
- 1 瓦时等于 3600 焦耳。

返回：

- 如果锻炼消耗了更多能量，则为 `"Workout"`。
- 如果 `Devices` 使用了更多能量。
- `"Equal"` 如果两者使用了相同量的能量。

# --hints--

`compare_energy(250, 50)` 应该返回 `"Workout"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` 应该返回 `"Devices"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` 应该返回 `"Equal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` 应该返回 `"Workout"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` 应该返回 `"Devices"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` 应该返回 `"Equal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
