---
id: 694596b0585c11170ac7c7fb
title: "挑战 163：辅音分支"
challengeType: 29
dashedName: challenge-163
---

# --description--

给定一个表现变量名的字串，使用以下规则将其转换为 consonant case：

- 所有辅音字母应转换为大写。
- 所有元音（无论大小写的 `a`、`e`、`i`、`o`、`u`）都应转换为小写。
- 所有连字符（`-`）应转换为下划线（`_`）。

# --hints--

`to_consonant_case("helloworld")` 应该返回 `"HeLLoWoRLD"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("helloworld"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("HELLOWORLD")` 应该返回 `"HeLLoWoRLD"`。


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("HELLOWORLD"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("_hElLO-WOrlD-")` 应该返回 `"_HeLLo_WoRLD_"`。


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_hElLO-WOrlD-"), "_HeLLo_WoRLD_")`)
}})
```

`to_consonant_case("_~-generic_~-variable_~-name_~-here-~_")` 应该返回 `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`。


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_consonant_case(s):

    return s
```

# --solutions--

```py
def to_consonant_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char == "-":
            result += "_"
        elif char.isalpha():
            if char in vowels:
                result += char.lower()
            else:
                result += char.upper()
        else:
            result += char

    return result
```
