---
id: 694596b0585c11170ac7c7f9
title: "挑战 161：免运费"
challengeType: 29
dashedName: challenge-161
---

# --description--

给定一个表现你购物车中项的字串数组和一个表示达到免运费资格的最低订单金额的数字，确定你购物车中的项是否符合免运费资格。

给定的数组将包含以下列表中的项：

| Item | Price |
| - | - |
| `"shirt"`	| 34.25 |
| `"jeans"`	| 48.50 |
| `"shoes"` |	75.00 |
| `"hat"`	| 19.95 |
| `"socks"` | 15.00 |
| `"jacket"` | 109.95 |

# --hints--

`gets_free_shipping(["shoes"], 50)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["shoes"], 50), True)`)
}})
```

`gets_free_shipping(["hat", "socks"], 50)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["hat", "socks"], 50), False)`)
}})
```

`gets_free_shipping(["jeans", "shirt", "jacket"], 75)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["jeans", "shirt", "jacket"], 75), True)`)
}})
```

`gets_free_shipping(["socks", "socks", "hat"], 75)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["socks", "socks", "hat"], 75), False)`)
}})
```

`gets_free_shipping(["shirt", "shirt", "jeans", "socks"], 100)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["shirt", "shirt", "jeans", "socks"], 100), True)`)
}})
```

`gets_free_shipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def gets_free_shipping(cart, minimum):

    return cart
```

# --solutions--

```py
def gets_free_shipping(cart, minimum):
    prices = {
        "shirt": 34.25,
        "jeans": 48.50,
        "shoes": 75.00,
        "hat": 19.95,
        "socks": 15.00,
        "jacket": 109.95
    }

    total = 0

    for item in cart:
        total += prices[item]

    return total >= minimum
```
