---
id: 696655d24b614176d4c9b78c
title: "挑战 169：FizzBuzz 迷你"
challengeType: 29
dashedName: challenge-169
---

# --description--

给定一个整数，返回一个基于以下规则的字串：

- 如果数字能被 3 整除，返回 `"Fizz"`。
- 如果数字能被 5 整除，返回 `"Buzz"`。
- 如果数字能被 3 和 5 同时整除，返回 `"FizzBuzz"`。
- 否则，返回给定的数字作为字串。

# --hints--

`fizz_buzz_mini(3)` 应该返回 `"Fizz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(3), "Fizz")`)
}})
```

`fizz_buzz_mini(4)` 应该返回 `"4"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(4), "4")`)
}})
```

`fizz_buzz_mini(35)` 应该返回 `"Buzz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(35), "Buzz")`)
}})
```

`fizz_buzz_mini(75)` 应该返回 `"FizzBuzz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(75), "FizzBuzz")`)
}})
```

`fizz_buzz_mini(98)` 应该返回 `"98"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(98), "98")`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_mini(n):

    return n
```

# --solutions--

```py
def fizz_buzz_mini(n):
    if n % 3 == 0 and n % 5 == 0:
        return "FizzBuzz"
    if n % 3 == 0:
        return "Fizz"
    if n % 5 == 0:
        return "Buzz"
    return str(n)
```
