---
id: 696655d24b614176d4c9b789
title: "挑战 166：Hex 验证器"
challengeType: 29
dashedName: challenge-166
---

# --description--

给定一个字串，判断它是否是有效的 CSS 十六进制颜色。有效的 CSS 十六进制颜色必须：

- 以 `#` 开头，且
- 后面跟随 3 个或 6 个十六进制字符。

十六进制字符是数字 `0` 到 `9` 和字母 `a` 到 `f`（不区分分支）。

# --hints--

`is_valid_hex("#123")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#123"), True)`)
}})
```

`is_valid_hex("#123abc")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#123abc"), True)`)
}})
```

`is_valid_hex("#ABCDEF")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#ABCDEF"), True)`)
}})
```

`is_valid_hex("#0a1B2c")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#0a1B2c"), True)`)
}})
```

`is_valid_hex("#12G")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#12G"), False)`)
}})
```

`is_valid_hex("#1234567")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#1234567"), False)`)
}})
```

`is_valid_hex("#12 3")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#12 3"), False)`)
}})
```

`is_valid_hex("fff")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("fff"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_hex(s):

    return s
```

# --solutions--

```py
def is_valid_hex(s):
    if not s.startswith("#"):
        return False

    hex_part = s[1:]
    if len(hex_part) not in (3, 6):
        return False

    for char in hex_part:
        if not char.lower() in "0123456789abcdef":
            return False

    return True
```
