---
id: 69738771fb5a7b8b24cca2a4
title: "挑战 178：截断 `text`"
challengeType: 29
dashedName: challenge-178
---

# --description--

给定一个字串，如果它的长度不超过 20 个字符，则原样返回。如果它超过 20 个字符，则截断为前 17 个字符，并在末尾追加 `"..."`（总共 20 个字符），然后返回结果。

# --hints--

`truncate_text("Hello, world!")` 应该返回 `"Hello, world!"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Hello, world!"), "Hello, world!")`)
}})
```

`truncate_text("This string should get truncated.")` 应该返回 `"This string shoul..."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("This string should get truncated."), "This string shoul...")`)
}})
```

`truncate_text("Exactly twenty chars")` 应该返回 `"Exactly twenty chars"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Exactly twenty chars"), "Exactly twenty chars")`)
}})
```

`truncate_text(".....................")` 应该返回 `"...................."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("....................."), "....................")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(text):

    return text
```

# --solutions--

```py
def truncate_text(text):
    if len(text) <= 20:
        return text
    return text[:17] + "..."
```
