---
id: 69738771fb5a7b8b24cca2a2
title: "挑战 176：土拨鼠日"
challengeType: 29
dashedName: challenge-176
---

# --description--

今天是土拨鼠日，土拨鼠根据是否看到它的影子来预测天气。

给定一个表现土拨鼠外观的值，返回正确的预测：

- 如果给定值是布尔值 `true`（土拨鼠看到了它的影子），则返回 `"Looks like we'll have six more weeks of winter."`。
- 如果值是布尔类型的 `false`（土拨鼠没有看到它的影子），则返回 `"It's going to be an early spring."`。
- 如果值是其他任何内容（旱獭没有出现），返回 `"No prediction this year."`。

# --hints--

`groundhog_day_prediction(True)` 应该返回 `"Looks like we'll have six more weeks of winter."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

`groundhog_day_prediction(False)` 应该返回 `"It's going to be an early spring."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

`groundhog_day_prediction(None)` 应该返回 `"No prediction this year."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` 应返回 `"No prediction this year."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

`groundhog_day_prediction("True")` 应该返回 `"No prediction this year."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
