---
id: 69738771fb5a7b8b24cca2a3
title: "挑战 177：字串 镜射"
challengeType: 29
dashedName: challenge-177
---

# --description--

给定一个字串，返回一个新建的字串，该字串由给定字串和其反转拷贝追加到末尾组成。

# --hints--

`mirror("freeCodeCamp")` 应该返回 `"freeCodeCamppmaCedoCeerf"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf")`)
}})
```

`mirror("RaceCar")` 应该返回 `"RaceCarraCecaR"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("RaceCar"), "RaceCarraCecaR")`)
}})
```

`mirror("helloworld")` 应该返回 `"helloworlddlrowolleh"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("helloworld"), "helloworlddlrowolleh")`)
}})
```

`mirror("The quick brown fox...")` 应该返回 `"The quick brown fox......xof nworb kciuq ehT"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT")`)
}})
```

# --seed--

## --seed-contents--

```py
def mirror(s):

    return s
```

# --solutions--

```py
def mirror(s):
    reversed_s = s[::-1]
    return s + reversed_s
```
