---
id: 697a49e6ff50d756c9b69364
title: "挑战 187：2026 冬季游戏第 8 天：雪橇"
challengeType: 29
dashedName: challenge-187
---

# --description--

给定一个包含五个数字的数组，每个数字表现了一个雪橇运动员完成赛道某一段所用的时间（以秒为单位），确定哪一段的速度最快以及该速度是多少。

赛道被划分为以下几个区段：

- 段 1：320 米
- 第 2 段：280 米
- 第 3 段：350 米
- 第 4 段：300 米
- 第 5 段：250 米

给定数组中的第一个值对应于第 1 段的时间，第二个值对应于第 2 段，依此类推。

要计算一个段的速度（以米每秒为单位），请将距离除以时间。

返回 `"The luger's fastest speed was X m/s on segment Y."`。其中 `X` 是最快速度，四舍五入到小数点后两位，`Y` 是最快速度发生的赛段编号。

# --hints--

`get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128])` 应该返回 `"The luger's fastest speed was 35.07 m/s on segment 5."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.")`)
}})
```

`get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284])` 应该返回 `"The luger's fastest speed was 37.75 m/s on segment 2."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.")`)
}})
```

`get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912])` 应该返回 `"The luger's fastest speed was 38.49 m/s on segment 3."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.")`)
}})
```

`get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840])` 应该返回 `"The luger's fastest speed was 37.69 m/s on segment 1."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.")`)
}})
```

`get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508])` 应该返回 `"The luger's fastest speed was 39.24 m/s on segment 4."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_fastest_speed(times):

    return times
```

# --solutions--

```py
def get_fastest_speed(times):
    distances = [320, 280, 350, 300, 250]
    fastest_speed = 0
    fastest_segment = 0

    for i, t in enumerate(times):
        speed = distances[i] / t
        if speed > fastest_speed:
            fastest_speed = speed
            fastest_segment = i + 1  # 1-indexed

    return f"The luger's fastest speed was {fastest_speed:.2f} m/s on segment {fastest_segment}."
```
