---
id: 697a49e6ff50d756c9b6935e
title: "挑战 181：2026 冬季游戏第 2 天：单板滑雪"
challengeType: 29
dashedName: challenge-181
---

# --description--

给定单板滑雪者的起始姿势和旋转角度，确定他们的着陆姿势。

- 滑雪板手的站姿要么是 `"Regular"`，要么是 `"Goofy"`。
- 技巧旋转是 90 度的倍数。正数表示顺时针旋转，负数表示逆时针旋转。
- 着陆姿势每旋转 180 度翻转一次。

例如，给定 `"Regular"` 和 `90`，返回 `"Regular"`。给定 `"Regular"` 和 `180` 度，返回 `"Goofy"`。

# --hints--

`get_landing_stance("Regular", 90)` 应该返回 `"Regular"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` 应该返回 `"Goofy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` 应该返回 `"Regular"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` 应该返回 `"Goofy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` 应该返回 `"Goofy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` 应该返回 `"Regular"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
