---
id: 697a49e6ff50d756c9b6935d
title: "挑战 180：冬季运动会第 1 天：开幕日"
challengeType: 29
dashedName: challenge-180
---

# --description--

接下来的 17 天将带给你受冬季运动会启发的编程挑战。

对于第一个，你会得到一个两字母的国家代码，需要返回该国家的标记表情符号。

使用此列表：

| 国家 | 代码 | 国旗 |
| - | -: | :--: |
| 阿尔巴尼亚                    |            `"AL"` | `"🇦🇱"` |
| 安道尔                    |            `"AD"` | `"🇦🇩"` |
| 阿根廷                  |            `"AR"` | `"🇦🇷"` |
| 亚美尼亚                    |            `"AM"` | `"🇦🇲"` |
| 澳大利亚                  |            `"AU"` | `"🇦🇺"` |
| 奥地利                    |            `"AT"` | `"🇦🇹"` |
| 阿塞拜疆                 |            `"AZ"` | `"🇦🇿"` |
| 比利时                    |            `"BE"` | `"🇧🇪"` |
| 贝宁                      |            `"BJ"` | `"🇧🇯"` |
| 玻利维亚                    |            `"BO"` | `"🇧🇴"` |
| 波斯尼亚和黑塞哥维那     |            `"BA"` | `"🇧🇦"` |
| 巴西                     |            `"BR"` | `"🇧🇷"` |
| 保加利亚                   |            `"BG"` | `"🇧🇬"` |
| 加拿大                     |            `"CA"` | `"🇨🇦"` |
| 智利                      |            `"CL"` | `"🇨🇱"` |
| 中国                      |            `"CN"` | `"🇨🇳"` |
| 哥伦比亚                   |            `"CO"` | `"🇨🇴"` |
| 克罗地亚                    |            `"HR"` | `"🇭🇷"` |
| 塞浦路斯                     |            `"CY"` | `"🇨🇾"` |
| 捷克共和国             |            `"CZ"` | `"🇨🇿"` |
| 丹麦                    |            `"DK"` | `"🇩🇰"` |
| 厄瓜多尔                    |            `"EC"` | `"🇪🇨"` |
| 厄立特里亚                    |            `"ER"` | `"🇪🇷"` |
| 爱沙尼亚                    |            `"EE"` | `"🇪🇪"` |
| 芬兰                    |            `"FI"` | `"🇫🇮"` |
| 法国                     |            `"FR"` | `"🇫🇷"` |
| 格鲁吉亚                    |            `"GE"` | `"🇬🇪"` |
| 德国                    |            `"DE"` | `"🇩🇪"` |
| 大不列颠              |            `"GB"` | `"🇬🇧"` |
| 希腊                     |            `"GR"` | `"🇬🇷"` |
| 几内亚比绍              |            `"GW"` | `"🇬🇼"` |
| 海地                      |            `"HT"` | `"🇭🇹"` |
| 香港                  |            `"HK"` | `"🇭🇰"` |
| 匈牙利                    |            `"HU"` | `"🇭🇺"` |
| 冰岛                    |            `"IS"` | `"🇮🇸"` |
| 印度                      |            `"IN"` | `"🇮🇳"` |
| 伊朗                       |            `"IR"` | `"🇮🇷"` |
| 爱尔兰                    |            `"IE"` | `"🇮🇪"` |
| 以色列                     |            `"IL"` | `"🇮🇱"` |
| 意大利                      |            `"IT"` | `"🇮🇹"` |
| 牙买加                    |            `"JM"` | `"🇯🇲"` |
| 日本                      |            `"JP"` | `"🇯🇵"` |
| 哈萨克斯坦                 |            `"KZ"` | `"🇰🇿"` |
| 肯尼亚                      |            `"KE"` | `"🇰🇪"` |
| 科索沃                     |            `"XK"` | `"🇽🇰"` |
| 吉尔吉斯斯坦                 |            `"KG"` | `"🇰🇬"` |
| 拉脱维亚                     |            `"LV"` | `"🇱🇻"` |
| 黎巴嫩                    |            `"LB"` | `"🇱🇧"` |
| 列支敦士登              |            `"LI"` | `"🇱🇮"` |
| 立陶宛                  |            `"LT"` | `"🇱🇹"` |
| 卢森堡                 |            `"LU"` | `"🇱🇺"` |
| 马达加斯加                 |            `"MG"` | `"🇲🇬"` |
| 马来西亚                   |            `"MY"` | `"🇲🇾"` |
| 马耳他                      |            `"MT"` | `"🇲🇹"` |
| 墨西哥                     |            `"MX"` | `"🇲🇽"` |
| 摩尔多瓦                    |            `"MD"` | `"🇲🇩"` |
| 摩纳哥                     |            `"MC"` | `"🇲🇨"` |
| 蒙古                   |            `"MN"` | `"🇲🇳"` |
| 黑山                 |            `"ME"` | `"🇲🇪"` |
| 摩洛哥                    |            `"MA"` | `"🇲🇦"` |
| 荷兰                |            `"NL"` | `"🇳🇱"` |
| 新西兰                |            `"NZ"` | `"🇳🇿"` |
| 尼日利亚                    |            `"NG"` | `"🇳🇬"` |
| 北马其顿            |            `"MK"` | `"🇲🇰"` |
| 挪威                     |            `"NO"` | `"🇳🇴"` |
| 巴基斯坦                   |            `"PK"` | `"🇵🇰"` |
| 菲律宾                |            `"PH"` | `"🇵🇭"` |
| 波兰                     |            `"PL"` | `"🇵🇱"` |
| 葡萄牙                   |            `"PT"` | `"🇵🇹"` |
| 波多黎各                |            `"PR"` | `"🇵🇷"` |
| 罗马尼亚                    |            `"RO"` | `"🇷🇴"` |
| 圣马力诺                 |            `"SM"` | `"🇸🇲"` |
| 沙特阿拉伯               |            `"SA"` | `"🇸🇦"` |
| 塞尔维亚                     |            `"RS"` | `"🇷🇸"` |
| 新加坡                  |            `"SG"` | `"🇸🇬"` |
| 斯洛伐克                   |            `"SK"` | `"🇸🇰"` |
| 斯洛文尼亚                   |            `"SI"` | `"🇸🇮"` |
| 南非               |            `"ZA"` | `"🇿🇦"` |
| 韩国                |            `"KR"` | `"🇰🇷"` |
| 西班牙                      |            `"ES"` | `"🇪🇸"` |
| 瑞典                     |            `"SE"` | `"🇸🇪"` |
| 瑞士                |            `"CH"` | `"🇨🇭"` |
| 泰国                   |            `"TH"` | `"🇹🇭"` |
| 特立尼达和多巴哥          |            `"TT"` | `"🇹🇹"` |
| 土耳其                     |            `"TR"` | `"🇹🇷"` |
| 乌克兰                    |            `"UA"` | `"🇺🇦"` |
| 阿拉伯联合酋长国       |            `"AE"` | `"🇦🇪"` |
| 美国              |            `"US"` | `"🇺🇸"` |
| 乌拉圭                    |            `"UY"` | `"🇺🇾"` |
| 乌兹别克斯坦                 |            `"UZ"` | `"🇺🇿"` |
| 委内瑞拉                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`get_flag("AL")` 应该返回 `"🇦🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "🇦🇱")`)
}})
```

`get_flag("AD")` 应该返回 `"🇦🇩"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "🇦🇩")`)
}})
```

`get_flag("AR")` 应该返回 `"🇦🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "🇦🇷")`)
}})
```

`get_flag("AM")` 应该返回 `"🇦🇲"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "🇦🇲")`)
}})
```

`get_flag("AU")` 应该返回 `"🇦🇺"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "🇦🇺")`)
}})
```

`get_flag("AT")` 应该返回 `"🇦🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "🇦🇹")`)
}})
```

`get_flag("AZ")` 应该返回 `"🇦🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "🇦🇿")`)
}})
```

`get_flag("BE")` 应该返回 `"🇧🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "🇧🇪")`)
}})
```

`get_flag("BJ")` 应该返回 `"🇧🇯"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "🇧🇯")`)
}})
```

`get_flag("BO")` 应该返回 `"🇧🇴"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "🇧🇴")`)
}})
```

`get_flag("BA")` 应该返回 `"🇧🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "🇧🇦")`)
}})
```

`get_flag("BR")` 应该返回 `"🇧🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "🇧🇷")`)
}})
```

`get_flag("BG")` 应该返回 `"🇧🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "🇧🇬")`)
}})
```

`get_flag("CA")` 应该返回 `"🇨🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "🇨🇦")`)
}})
```

`get_flag("CL")` 应该返回 `"🇨🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "🇨🇱")`)
}})
```

`get_flag("CN")` 应该返回 `"🇨🇳"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "🇨🇳")`)
}})
```

`get_flag("CO")` 应该返回 `"🇨🇴"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "🇨🇴")`)
}})
```

`get_flag("HR")` 应该返回 `"🇭🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "🇭🇷")`)
}})
```

`get_flag("CY")` 应该返回 `"🇨🇾"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "🇨🇾")`)
}})
```

`get_flag("CZ")` 应该返回 `"🇨🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "🇨🇿")`)
}})
```

`get_flag("DK")` 应该返回 `"🇩🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "🇩🇰")`)
}})
```

`get_flag("EC")` 应该返回 `"🇪🇨"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "🇪🇨")`)
}})
```

`get_flag("ER")` 应该返回 `"🇪🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "🇪🇷")`)
}})
```

`get_flag("EE")` 应该返回 `"🇪🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "🇪🇪")`)
}})
```

`get_flag("FI")` 应该返回 `"🇫🇮"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "🇫🇮")`)
}})
```

`get_flag("FR")` 应该返回 `"🇫🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "🇫🇷")`)
}})
```

`get_flag("GE")` 应该返回 `"🇬🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "🇬🇪")`)
}})
```

`get_flag("DE")` 应该返回 `"🇩🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "🇩🇪")`)
}})
```

`get_flag("GB")` 应该返回 `"🇬🇧"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "🇬🇧")`)
}})
```

`get_flag("GR")` 应该返回 `"🇬🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "🇬🇷")`)
}})
```

`get_flag("GW")` 应该返回 `"🇬🇼"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "🇬🇼")`)
}})
```

`get_flag("HT")` 应该返回 `"🇭🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "🇭🇹")`)
}})
```

`get_flag("HK")` 应该返回 `"🇭🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "🇭🇰")`)
}})
```

`get_flag("HU")` 应该返回 `"🇭🇺"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "🇭🇺")`)
}})
```

`get_flag("IS")` 应该返回 `"🇮🇸"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "🇮🇸")`)
}})
```

`get_flag("IN")` 应该返回 `"🇮🇳"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "🇮🇳")`)
}})
```

`get_flag("IR")` 应该返回 `"🇮🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "🇮🇷")`)
}})
```

`get_flag("IE")` 应该返回 `"🇮🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "🇮🇪")`)
}})
```

`get_flag("IL")` 应该返回 `"🇮🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "🇮🇱")`)
}})
```

`get_flag("IT")` 应该返回 `"🇮🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "🇮🇹")`)
}})
```

`get_flag("JM")` 应该返回 `"🇯🇲"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "🇯🇲")`)
}})
```

`get_flag("JP")` 应该返回 `"🇯🇵"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "🇯🇵")`)
}})
```

`get_flag("KZ")` 应该返回 `"🇰🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "🇰🇿")`)
}})
```

`get_flag("KE")` 应该返回 `"🇰🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "🇰🇪")`)
}})
```

`get_flag("XK")` 应该返回 `"🇽🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "🇽🇰")`)
}})
```

`get_flag("KG")` 应该返回 `"🇰🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "🇰🇬")`)
}})
```

`get_flag("LV")` 应该返回 `"🇱🇻"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "🇱🇻")`)
}})
```

`get_flag("LB")` 应该返回 `"🇱🇧"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "🇱🇧")`)
}})
```

`get_flag("LI")` 应该返回 `"🇱🇮"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "🇱🇮")`)
}})
```

`get_flag("LT")` 应该返回 `"🇱🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "🇱🇹")`)
}})
```

`get_flag("LU")` 应该返回 `"🇱🇺"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "🇱🇺")`)
}})
```

`get_flag("MG")` 应该返回 `"🇲🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "🇲🇬")`)
}})
```

`get_flag("MY")` 应该返回 `"🇲🇾"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "🇲🇾")`)
}})
```

`get_flag("MT")` 应该返回 `"🇲🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "🇲🇹")`)
}})
```

`get_flag("MX")` 应该返回 `"🇲🇽"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "🇲🇽")`)
}})
```

`get_flag("MD")` 应该返回 `"🇲🇩"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "🇲🇩")`)
}})
```

`get_flag("MC")` 应该返回 `"🇲🇨"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "🇲🇨")`)
}})
```

`get_flag("MN")` 应该返回 `"🇲🇳"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "🇲🇳")`)
}})
```

`get_flag("ME")` 应该返回 `"🇲🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "🇲🇪")`)
}})
```

`get_flag("MA")` 应该返回 `"🇲🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "🇲🇦")`)
}})
```

`get_flag("NL")` 应该返回 `"🇳🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "🇳🇱")`)
}})
```

`get_flag("NZ")` 应该返回 `"🇳🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "🇳🇿")`)
}})
```

`get_flag("NG")` 应该返回 `"🇳🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "🇳🇬")`)
}})
```

`get_flag("MK")` 应该返回 `"🇲🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "🇲🇰")`)
}})
```

`get_flag("NO")` 应该返回 `"🇳🇴"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "🇳🇴")`)
}})
```

`get_flag("PK")` 应该返回 `"🇵🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "🇵🇰")`)
}})
```

`get_flag("PH")` 应该返回 `"🇵🇭"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "🇵🇭")`)
}})
```

`get_flag("PL")` 应该返回 `"🇵🇱"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "🇵🇱")`)
}})
```

`get_flag("PT")` 应该返回 `"🇵🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "🇵🇹")`)
}})
```

`get_flag("PR")` 应该返回 `"🇵🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "🇵🇷")`)
}})
```

`get_flag("RO")` 应该返回 `"🇷🇴"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "🇷🇴")`)
}})
```

`get_flag("SM")` 应该返回 `"🇸🇲"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "🇸🇲")`)
}})
```

`get_flag("SA")` 应该返回 `"🇸🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "🇸🇦")`)
}})
```

`get_flag("RS")` 应该返回 `"🇷🇸"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "🇷🇸")`)
}})
```

`get_flag("SG")` 应该返回 `"🇸🇬"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "🇸🇬")`)
}})
```

`get_flag("SK")` 应该返回 `"🇸🇰"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "🇸🇰")`)
}})
```

`get_flag("SI")` 应该返回 `"🇸🇮"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "🇸🇮")`)
}})
```

`get_flag("ZA")` 应该返回 `"🇿🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "🇿🇦")`)
}})
```

`get_flag("KR")` 应该返回 `"🇰🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "🇰🇷")`)
}})
```

`get_flag("ES")` 应该返回 `"🇪🇸"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "🇪🇸")`)
}})
```

`get_flag("SE")` 应该返回 `"🇸🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "🇸🇪")`)
}})
```

`get_flag("CH")` 应该返回 `"🇨🇭"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "🇨🇭")`)
}})
```

`get_flag("TH")` 应该返回 `"🇹🇭"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "🇹🇭")`)
}})
```

`get_flag("TT")` 应该返回 `"🇹🇹"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "🇹🇹")`)
}})
```

`get_flag("TR")` 应该返回 `"🇹🇷"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "🇹🇷")`)
}})
```

`get_flag("UA")` 应该返回 `"🇺🇦"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "🇺🇦")`)
}})
```

`get_flag("AE")` 应该返回 `"🇦🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "🇦🇪")`)
}})
```

`get_flag("US")` 应该返回 `"🇺🇸"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "🇺🇸")`)
}})
```

`get_flag("UY")` 应该返回 `"🇺🇾"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "🇺🇾")`)
}})
```

`get_flag("UZ")` 应该返回 `"🇺🇿"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "🇺🇿")`)
}})
```

`get_flag("VE")` 应该返回 `"🇻🇪"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "🇻🇪")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
