---
id: 697a49e9860d24853adef67d
title: "挑战 192：2026 冬季游戏第 13 天：北欧两项"
challengeType: 29
dashedName: challenge-192
---

# --description--

给定运动员跳跃成绩的数组，计算他们越野滑雪部分的起跑延迟时间。

跳跃得分最高的运动员首先开始（延迟 0 秒）。所有其他运动员根据他们的跳跃得分与最佳跳跃得分的差距延迟开始。

要计算每个运动员的延迟，先用最佳整体跳跃分数减去运动员的跳跃分数，然后将结果乘以 1.5。将延迟向上取整到最接近的整数。

# --hints--

`calculate_start_delays([120, 110, 125])` 应该返回 `[8, 23, 0]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([120, 110, 125]), [8, 23, 0])`)
}})
```

`calculate_start_delays([118, 125, 122, 120])` 应该返回 `[11, 0, 5, 8]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([118, 125, 122, 120]), [11, 0, 5, 8])`)
}})
```

`calculate_start_delays([100, 105, 95, 110, 120, 115, 108])` 应该返回 `[30, 23, 38, 15, 0, 8, 18]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18])`)
}})
```

`calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` 应该返回 `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_start_delays(jump_scores):

    return jump_scores
```

# --solutions--

```py
import math
def calculate_start_delays(jump_scores):
    best_jump = max(jump_scores)
    return [math.ceil((best_jump - score) * 1.5) for score in jump_scores]
```
