---
id: 697a49e6ff50d756c9b69365
title: "挑战 188：2026 冬季游戏 第 9 天：Skeleton"
challengeType: 29
dashedName: challenge-188
---

# --description--

给定一个表现骨架赛道弯道的字串，确定赛道的难度。

- 给定的字串将仅包含以下字母：
  - `"L"` 表示左转
  - `"R"` 表示右转
  - `"S"` 表示直线段

- 每次方向改变添加 15 分（一个 `"L"` 后跟一个 `"R"`，或反之亦然）。
- 所有其他曲线各添加 5 点（所有其他 `"L"` 或 `"R"` 字符）。
- 直线段添加 0 分。

赛道的难度基于总分。返回：

- `"Easy"` 如果总数在 0 - 100 之间
- `"Medium"` 如果总数在 101-200 之间
- `"Hard"` 如果总数超过 200

# --hints--

`get_difficulty("SLSLLSRRLSRLRL")` 应该返回 `"Easy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLSLLSRRLSRLRL"), "Easy")`)
}})
```

`get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` 应该返回 `"Hard"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard")`)
}})
```

`get_difficulty("SRRRRLSLLRLRSSRLSRL")` 应该返回 `"Medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRRRRLSLLRLRSSRLSRL"), "Medium")`)
}})
```

`get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` 应该返回 `"Hard"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard")`)
}})
```

`get_difficulty("SLLSSLRLSLSLRSLSSLRL")` 应该返回 `"Medium"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium")`)
}})
```

`get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` 应该返回 `"Easy"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_difficulty(track):

    return track
```

# --solutions--

```py
def get_difficulty(track):
  score = 0

  for i in range(len(track)):
    current = track[i]
    previous = track[i - 1] if i > 0 else None

    if current == 'S':
      score += 0
    elif current == 'L' or current == 'R':
      if previous and previous != 'S' and previous != current:
        score += 15
      else:
        score += 5

  if score <= 100:
    return 'Easy'
  elif score <= 200:
    return 'Medium'
  else:
    return 'Hard'
```
