---
id: 697a49e9860d24853adef67c
title: "挑战 191：2026 冬季游戏第 12 天：雪橇"
challengeType: 29
dashedName: challenge-191
---

# --description--

给定一个表现雪橇队运动员权值的数组和一个表现雪橇权值的数字，确定该队是否有资格比赛。

- 数组的长度决定团队规模：1 人、2 人或 4 人团队。
- 所有给定的权值均以千克（kg）为单位。

雪橇（单独的雪橇）必须具有最低权值：

- 1 人团队的 162 kg
- 2 人团队的 170 kg
- 4 人团队的 210 kg

雪橇（运动员加雪橇）的总权值不得超过：

- 1 人团队的 247 kg
- 2 人团队的 390 kg
- 4 人团队的 630 kg

返回 `"Eligible"` 如果团队满足所有要求，或者返回 `"Not Eligible"` 如果团队未能满足一个或多个要求。

# --hints--

`check_eligibility([78], 165)` 应该返回 `"Eligible"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([78], 165), "Eligible")`)
}})
```

`check_eligibility([80], 160)` 应该返回 `"Not Eligible"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 160), "Not Eligible")`)
}})
```

`check_eligibility([80], 170)` 应该返回 `"Not Eligible"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 170), "Not Eligible")`)
}})
```

`check_eligibility([85, 90], 170)` 应该返回 `"Eligible"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 90], 170), "Eligible")`)
}})
```

`check_eligibility([85, 95], 168)` 应该返回 `"Not Eligible"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 95], 168), "Not Eligible")`)
}})
```

`check_eligibility([112, 97], 185)` 应该返回 `"Not Eligible"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([112, 97], 185), "Not Eligible")`)
}})
```

`check_eligibility([110, 102, 90, 106], 222)` 应该返回 `"Eligible"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([110, 102, 90, 106], 222), "Eligible")`)
}})
```

`check_eligibility([106, 99, 90, 88], 205)` 应该返回 `"Not Eligible"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 90, 88], 205), "Not Eligible")`)
}})
```

`check_eligibility([106, 99, 103, 96], 227)` 应该返回 `"Not Eligible"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 103, 96], 227), "Not Eligible")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_eligibility(athlete_weights, sled_weight):

    return athlete_weights
```

# --solutions--

```py
def check_eligibility(athlete_weights, sled_weight):
    team_size = len(athlete_weights)

    rules = {
        1: {"min_sled": 162, "max_total": 247},
        2: {"min_sled": 170, "max_total": 390},
        4: {"min_sled": 210, "max_total": 630}
    }

    min_sled = rules[team_size]["min_sled"]
    max_total = rules[team_size]["max_total"]

    total_weight = sum(athlete_weights) + sled_weight

    if sled_weight < min_sled:
        return "Not Eligible"

    if total_weight > max_total:
        return "Not Eligible"

    return "Eligible"
```
