---
id: 698a1a863194f1f4e63f645f
title: "挑战 208：路径遍历"
challengeType: 29
dashedName: challenge-208
---

# --description--

给定一个表现你路径地图的字串数组，返回到达目标所需移动的字串。

给定的字串将包含以下值：

- `"C"`：你当前位置
- `"G"`：你的目标
- `"T"`：路径的可通行部分
- `"-"`：地图中不可通行的部分

返回一个包含从你的位置到目标所需路径移动的字串，其中：

- `"R"` 是向右移动
- `"D"` 是向下移动
- `"L"` 是向左移动
- `"U"` 是向上移动

- 从你当前位置到目标，总会有一条连续的路径，没有任何分支。
- 每个路径位置最多有两个可通行的位置与之相邻。

例如，给定：

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

返回 `"RDDRDD"`。

# --hints--

`navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` 应该返回 `"RDDRDD"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD")`)
}})
```

`navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` 应该返回 `"RRUUURR"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR")`)
}})
```

`navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` 应该返回 `"DLDDRRRRD"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD")`)
}})
```

`navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` 应该返回 `"RRRDDL"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL")`)
}})
```

`navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` 应该返回 `"ULLLUUURRRRRRDDDR"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate_trail(map):

    return map
```

# --solutions--

```py
def navigate_trail(map):
    rows = len(map)
    cols = len(map[0])

    directions = [
        (0, 1, "R"),
        (1, 0, "D"),
        (0, -1, "L"),
        (-1, 0, "U"),
    ]

    for r in range(rows):
        for c in range(cols):
            if map[r][c] == "C":
                row, col = r, c

    result = ""
    prev = None

    while map[row][col] != "G":
        for dr, dc, move in directions:
            new_row = row + dr
            new_col = col + dc

            if (
                0 <= new_row < rows and
                0 <= new_col < cols and
                (map[new_row][new_col] == "T" or
                 map[new_row][new_col] == "G") and
                (not prev or (new_row, new_col) != prev)
            ):
                result += move
                prev = (row, col)
                row, col = new_row, new_col
                break

    return result
```
