---
id: 698a1a73ade5ac0e19180fa5
title: "挑战 202：添加标点符号"
challengeType: 29
dashedName: challenge-202
---

# --description--

给定一个缺少句号的句子字串，在以下位置添加句号（`"."`）：

- 在每个紧接大写字母之前的空格前
- 并且在字串的末尾

返回结果字串。

# --hints--

`add_punctuation("Hello world")` 应该返回 `"Hello world."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Hello world"), "Hello world.")`)
}})
```

`add_punctuation("Hello world It's nice today")` 应该返回 `"Hello world. It's nice today."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Hello world It's nice today"), "Hello world. It's nice today.")`)
}})
```

`add_punctuation("JavaScript is great Sometimes")` 应该返回 `"JavaScript is great. Sometimes."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("JavaScript is great Sometimes"), "JavaScript is great. Sometimes.")`)
}})
```

`add_punctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z")` 应该返回 `"A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z"), "A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z.")`)
}})
```

`add_punctuation("Wait.. For it")` 应该返回 `"Wait... For it."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Wait.. For it"), "Wait... For it.")`)
}})
```

# --seed--

## --seed-contents--

```py
def add_punctuation(sentences):

    return sentences
```

# --solutions--

```py
def add_punctuation(sentences):
    result = ""
    length = len(sentences)

    for i, c in enumerate(sentences):
        if c == " " and i + 1 < length and sentences[i + 1].isupper():
            result += "."
        result += c

    if not result.endswith("."):
        result += "."

    return result
```
