---
id: 698a1a863194f1f4e63f645e
title: "挑战 207：最小间隙"
challengeType: 29
dashedName: challenge-207
---

# --description--

给定一个 `string`，返回两个具有最小字符间隔（最小间隙）的相同字符之间的子串。

- 总会至少有一对匹配的字符。
- 返回的子串应排除匹配的字符。
- 如果两个或更多空白长度相同，`return` 第一个空白中的字符。

例如，给定 `"ABCDAC"`，返回 `"DA"`。

- 只有 `"A"` 和 `"C"` 在字串中重复。
- 两个 `"A"` 字符之间的字符数是 3，两个 `"C"` 字符之间的字符数是 2。
- 所以返回两个 `"C"` 字符之间的字串。

# --hints--

`smallest_gap("ABCDAC")` 应该返回 `"DA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` 应该返回 `"e"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` 应该返回 `"#q6e&rkf(p"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` 应该返回 `""`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` 应该返回 `"fox"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
