---
id: 698a1a73ade5ac0e19180fa3
title: "挑战 200：字母和数字计数"
challengeType: 29
dashedName: challenge-200
---

# --description--

给定一个字串，返回一条包含其中字母和数字数量的消息。

- 字母是 `A-Z` 和 `a-z`。
- 数字是 `0-9`。
- 忽略所有其他字符。

返回 `"The string has X letters and Y numbers."`，其中 `"X"` 是字母的数量，`"Y"` 是数字的数量。如果任一数量为 1，则该项使用单数形式。例如：使用 `"1 letter"` 而不是 `"1 letters"`，使用 `"1 number"` 而不是 `"1 numbers"`。

# --hints--

`count_letters_and_numbers("helloworld123")` 应该返回 `"The string has 10 letters and 3 numbers."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("helloworld123"), "The string has 10 letters and 3 numbers.")`)
}})
```

`count_letters_and_numbers("Catch 22")` 应该返回 `"The string has 5 letters and 2 numbers."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("Catch 22"), "The string has 5 letters and 2 numbers.")`)
}})
```

`count_letters_and_numbers("A1!")` 应该返回 `"The string has 1 letter and 1 number."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("A1!"), "The string has 1 letter and 1 number.")`)
}})
```

`count_letters_and_numbers("12345")` 应该返回 `"The string has 0 letters and 5 numbers."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("12345"), "The string has 0 letters and 5 numbers.")`)
}})
```

`count_letters_and_numbers("password")` 应该返回 `"The string has 8 letters and 0 numbers."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("password"), "The string has 8 letters and 0 numbers.")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_letters_and_numbers(s):

    return s
```

# --solutions--

```py
def count_letters_and_numbers(s):
    letter_count = 0
    number_count = 0

    for char in s:
        if ('A' <= char <= 'Z') or ('a' <= char <= 'z'):
            letter_count += 1
        elif '0' <= char <= '9':
            number_count += 1

    letter_word = "letter" if letter_count == 1 else "letters"
    number_word = "number" if number_count == 1 else "numbers"

    return f"The string has {letter_count} {letter_word} and {number_count} {number_word}."
```
