---
id: 698a1a73ade5ac0e19180fa1
title: "挑战 198：工作日计数"
challengeType: 29
dashedName: challenge-198
---

# --description--

给定一个开始日期和结束日期，返回两者之间的工作日天数。

- 给定的日期格式为 `"YYYY-MM-DD"`。
- 工作日是工作日（星期一至星期五）。
- 周末不是工作日（星期六和星期日）。
- 计算时包括开始和结束日期。

# --hints--

`count_business_days("2026-02-24", "2026-02-26")` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2026-02-26"), 3)`)
}})
```

`count_business_days("2026-02-24", "2026-02-28")` 应该返回 `4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2026-02-28"), 4)`)
}})
```

`count_business_days("2026-02-21", "2026-03-01")` 应该返回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-21", "2026-03-01"), 5)`)
}})
```

`count_business_days("2026-03-08", "2026-03-17")` 应该返回 `7`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-03-08", "2026-03-17"), 7)`)
}})
```

`count_business_days("2026-02-24", "2027-02-24")` 应该返回 `262`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2027-02-24"), 262)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_business_days(start, end):

    return start
```

# --solutions--

```py
from datetime import datetime, timedelta

def count_business_days(start, end):
    start_date = datetime.strptime(start, "%Y-%m-%d")
    end_date = datetime.strptime(end, "%Y-%m-%d")
    
    count = 0
    current = start_date
    
    while current <= end_date:
        if current.weekday() < 5:
            count += 1
        current += timedelta(days=1)
    
    return count
```
