---
id: 698a1a73ade5ac0e19180fa4
title: "挑战 201：矩阵移位"
challengeType: 29
dashedName: challenge-201
---

# --description--

给定一个数字的矩阵（数组的数组）和一个整数，将矩阵中的所有值按给定的数量进行位移。

- 正向移位会将数值向右移动。
- 负移位将值向左移动。

值应换行：

- 将矩阵视为一个连续的值序列。
- 当一个值超过一行的末尾时，它会继续从下一行的开头开始。
- 当一个值移动超过矩阵中的最后一个位置时，它会回绕到第一个位置。
- 向左移动时同样适用相反的情况。

例如，给定：

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

将所有数字向右移动 `1` 位：

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shift_matrix([[1, 2, 3], [4, 5, 6]], 1)` 应该返回 `[[6, 1, 2], [3, 4, 5]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6]], -1)` 应该返回 `[[2, 3, 4], [5, 6, 1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` 应该返回 `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` 应该 `return` `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` 应该 `return` `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` 应该 `返回` `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_matrix(matrix, shift):

    return matrix
```

# --solutions--

```py
def shift_matrix(matrix, shift):
    rows = len(matrix)
    cols = len(matrix[0])
    flat = [num for row in matrix for num in row]
    n = len(flat)
    s = shift % n
    shifted = flat[-s:] + flat[:-s] if s != 0 else flat[:]
    result = []
    for i in range(rows):
        result.append(shifted[i*cols:(i+1)*cols])

    return result
```
