---
id: 698a1a73ade5ac0e19180fa7
title: "挑战 204：字母求和"
challengeType: 29
dashedName: challenge-204
---

# --description--

给定一个字串，返回其字母的总和。

- 字母是大写或小写的 A-Z
- 字母值为：`"A"` = 1，`"B"` = 2，...，`"Z"` = 26
- 大写字母和小写字母具有相同的值。
- 忽略所有非字母字符。

# --hints--

`sum_letters("Hello")` 应该返回 `52`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("Hello"), 52)`)
}})
```

`sum_letters("freeCodeCamp")` 应该返回 `94`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("freeCodeCamp"), 94)`)
}})
```

`sum_letters("The quick brown fox jumps over the lazy dog.")` 应该返回 `473`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("The quick brown fox jumps over the lazy dog."), 473)`)
}})
```

`sum_letters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` 应该返回 `1681`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681)`)
}})
```

`sum_letters("</404>"）` 应该返回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("</404>"), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_letters(s):

    return s
```

# --solutions--

```py
def sum_letters(s):
    total = 0
    for char in s:
        upper = char.upper()
        if 'A' <= upper <= 'Z':
            total += ord(upper) - ord('A') + 1
    return total
```
