---
id: 6994cff2290543b3aec9f511
title: "挑战 212：数组插入"
challengeType: 29
dashedName: challenge-212
---

# --description--

给定一个数组、一个要插入到数组中的值以及插入该值的索引，返回一个在指定索引插入该值的新建数组。

# --hints--

`insert_into_array([2, 4, 8, 10], 6, 2)` 应该返回 `[2, 4, 6, 8, 10]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(insert_into_array([2, 4, 8, 10], 6, 2), [2, 4, 6, 8, 10])`)
}})
```

`insert_into_array(["the", "quick", "fox"], "brown", 2)` 应该返回 `["the", "quick", "brown", "fox"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(insert_into_array(["the", "quick", "fox"], "brown", 2), ["the", "quick", "brown", "fox"])`)
}})
```

`insert_into_array([], 0, 0)` 应该返回 `[0]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(insert_into_array([], 0, 0), [0])`)
}})
```

`insert_into_array([0, 1, 1, 2, 3, 8, 13], 5, 5)` 应该返回 `[0, 1, 1, 2, 3, 5, 8, 13]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(insert_into_array([0, 1, 1, 2, 3, 8, 13], 5, 5), [0, 1, 1, 2, 3, 5, 8, 13])`)
}})
```

# --seed--

## --seed-contents--

```py
def insert_into_array(arr, value, index):

    return arr
```

# --solutions--

```py
def insert_into_array(arr, value, index):
    return arr[:index] + [value] + arr[index:]
```
