---
id: 6994cff2290543b3aec9f50f
title: "挑战 210：HSL 验证器"
challengeType: 29
dashedName: challenge-210
---

# --description--

给定一个字串，判断它是否是有效的 CSS `hsl()` 颜色值。

- 有效的 HSL 值必须是 `"hsl(h, s%, l%)"` 形式，其中：
  - `h`（色相）必须是 0 到 360（含）之间的数字。
  - `s`（饱和度）必须是介于 0% 和 100% 之间的百分比。
  - `l`（亮度）必须是介于 0% 和 100% 之间的百分比。

- 只允许空格：
  - 在左括号之后
  - 逗号之前和/或之后
  - 在闭合括号之前和/或之后

- 值可以以分号 (`";"`) 结尾，且这是可选的。

例如，`"hsl(240, 50%, 50%)"` 是一个有效的 HSL 值。

# --hints--

`is_valid_hsl("hsl(240, 50%, 50%)")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(240, 50%, 50%)"), True)`)
}})
```

`is_valid_hsl("hsl( 200 , 50% , 75% )")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl( 200 , 50% , 75% )"), True)`)
}})
```

`is_valid_hsl("hsl(99,60%,80%);")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(99,60%,80%);"), True)`)
}})
```

`is_valid_hsl("hsl(0, 0%, 0%) ;")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(0, 0%, 0%) ;"), True)`)
}})
```

`is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;"), True)`)
}})
```

`is_valid_hsl("hsl(361, 50%, 80%)")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(361, 50%, 80%)"), False)`)
}})
```

`is_valid_hsl("hsl(300, 101%, 70%)")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(300, 101%, 70%)"), False)`)
}})
```

`is_valid_hsl("hsl(200, 55%, 75)")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(200, 55%, 75)"), False)`)
}})
```

`is_valid_hsl("hsl (80, 20%, 10%)")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl (80, 20%, 10%)"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_hsl(hsl):

    return hsl
```

# --solutions--

```py
import re
def is_valid_hsl(hsl):
    pattern = r"^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$"
    
    match = re.match(pattern, hsl)
    if not match:
        return False

    hue = int(match.group(1))
    sat = int(match.group(2))
    light = int(match.group(3))

    if not (0 <= hue <= 360):
        return False
    if not (0 <= sat <= 100):
        return False
    if not (0 <= light <= 100):
        return False

    return True
```
