---
id: 699c8e045ee7cb94ed2322d4
title: "挑战 213：单词长度转换器"
challengeType: 29
dashedName: challenge-213
---

# --description--

给定一个字串，返回一个新建字串，其中每个单词都被其长度替换。

- 给定字串中的单词将由单个空格分隔
- 保持返回字串中的空格。

例如，给定 `"hello world"`，返回 `"5 5"`。

# --hints--

`convert_words("hello world")` 应该返回 `"5 5"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("hello world"), "5 5")`)
}})
```

`convert_words("Thanks and happy coding")` 应该返回 `"6 3 5 6"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Thanks and happy coding"), "6 3 5 6")`)
}})
```

`convert_words("The quick brown fox jumps over the lazy dog")` 应该返回 `"3 5 5 3 5 4 3 4 3"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3")`)
}})
```

`convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` 应该 `return` `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_words(s):

    return s
```

# --solutions--

```py
def convert_words(s):

    return " ".join(str(len(word)) for word in s.split(" "))
```
