---
id: 699c8e045ee7cb94ed2322d9
title: "挑战 218：均匀可整除"
challengeType: 29
dashedName: challenge-218
---

# --description--

给定两个整数，判断你是否可以将第一个整数整除第二个整数。

# --hints--

`is_evenly_divisible(4, 2)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4, 2), True)`)
}})
```

`is_evenly_divisible(7, 3)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(7, 3), False)`)
}})
```

`is_evenly_divisible(5, 10)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(5, 10), False)`)
}})
```

`is_evenly_divisible(48, 6)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(48, 6), True)`)
}})
```

`is_evenly_divisible(3186, 9)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(3186, 9), True)`)
}})
```

`is_evenly_divisible(4192, 11)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4192, 11), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_evenly_divisible(a, b):

    return a
```

# --solutions--

```py
def is_evenly_divisible(a, b):
    return a % b == 0
```
