---
id: 699c8e045ee7cb94ed2322dd
title: "挑战 222：昼夜平分影子"
challengeType: 29
dashedName: challenge-222
---

# --description--

今天是昼夜平分点，太阳正好位于赤道上空，并且在中午时分完全处于头顶。给定一个时间，确定一根 4 英尺高的垂直杆投下的阴影。

- 时间将是 `"HH:MM"` 24 小时制格式的字串（例如，`"15:00"` 是下午 3 点）。
- 你将只获得以 30 分钟为增量的时间。

规则：

- 太阳在早上 6 点正对 `"east"` 升起，晚上 6 点正对 `"west"` 落下。
- 影子总是指向与太阳相反的方向。
- 影子的长度（以英尺为单位）是距离中午小时数的立方。
- 日出前（6 点前）、日落后（6 点或更晚）或中午时没有阴影。

返回：

- 如果存在阴影，返回 `"(length)ft (direction)"`。例如，`"8ft west"`。
- 否则，返回 `"No shadow"`。

例如，给定 `"10:00"`，返回 `"8ft west"`，因为上午 10 点距离中午还有 2 小时，所以 2<sup>3</sup> = 8 英尺，且阴影指向西方，因为上午 10 点太阳在东方。

# --hints--

`get_shadow("10:00")` 应该返回 `"8ft west"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` 应该返回 `"27ft east"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` 应该返回 `"No shadow"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` 应该返回 `"166.375ft east"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` 应该返回 `"No shadow"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` 应该返回 `"216ft west"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` 应该返回 `"No shadow"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` 应该返回 `"91.125ft west"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` 应该返回 `"No shadow"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
