---
id: 699c8e045ee7cb94ed2322da
title: "挑战 219：周年里程碑"
challengeType: 29
dashedName: challenge-219
---

# --description--

给定一个表现夫妻结婚年数的整数，返回他们根据此表最近的周年纪念里程碑：

| Years Married | Milestone    |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- 如果他们还没有达到第一个里程碑，返回 `"Newlyweds"`。

# --hints--

`get_milestone(0)` 应该返回 `"Newlyweds"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(0), "Newlyweds")`)
}})
```

`get_milestone(1)` 应该返回 `"Paper"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(1), "Paper")`)
}})
```

`get_milestone(8)` 应该返回 `"Wood"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(8), "Wood")`)
}})
```

`get_milestone(10)` 应该返回 `"Tin"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(10), "Tin")`)
}})
```

`get_milestone(26)` 应该返回 `"Silver"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(26), "Silver")`)
}})
```

`get_milestone(45)` 应该返回 `"Ruby"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(45), "Ruby")`)
}})
```

`get_milestone(50)` 应该返回 `"Gold"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(50), "Gold")`)
}})
```

`get_milestone(64)` 应该返回 `"Diamond"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(64), "Diamond")`)
}})
```

`get_milestone(71)` 应该返回 `"Platinum"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_milestone(71), "Platinum")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_milestone(years):

    return years
```

# --solutions--

```py
def get_milestone(years):
    if years < 1:
        return "Newlyweds"

    milestones = [
        (1, "Paper"),
        (5, "Wood"),
        (10, "Tin"),
        (25, "Silver"),
        (40, "Ruby"),
        (50, "Gold"),
        (60, "Diamond"),
        (70, "Platinum")
    ]

    for year, milestone in reversed(milestones):
        if years >= year:
            return milestone
```
