---
id: 699c8e045ee7cb94ed2322db
title: "挑战 220：最大数字"
challengeType: 29
dashedName: challenge-220
---

# --description--

给定一个由各种标点分隔的数字字串，返回最大的数字。

- 给定的字串将只包含数字和分隔符。
- 分隔符可以是逗号（`","`）、感叹号（`"!"`）、问号（`"?"`）、冒号（`":"`）或分号（`";"`）。

# --hints--

`largest_number("1,2")` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("1,2"), 2)`)
}})
```

`largest_number("4;15:60,26?52!0")` 应该返回 `60`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("4;15:60,26?52!0"), 60)`)
}})
```

`largest_number("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` 应该返回 `-402`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402)`)
}})
```

`largest_number("12;-50,99.9,49.1!-10.1?88?16")` 应该返回 `99.9`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("12;-50,99.9,49.1!-10.1?88?16"), 99.9)`)
}})
```

# --seed--

## --seed-contents--

```py
def largest_number(s):

    return s
```

# --solutions--

```py
import re
def largest_number(s):
    numbers = [float(n) for n in re.split(r'[,!?:;]', s)]
    return max(numbers)
```
