---
id: 69a890af247de743333bd4d0
title: "挑战 227：冷却时间"
challengeType: 29
dashedName: challenge-227
---

# --description--

给定两个时间戳，第一个表现用户完成考试的时间，第二个表现当前时间，确定用户是否可以再次参加考试。

- 两个时间戳都将采用格式：`"YYYY-MM-DDTHH:MM:SS"`，例如 `"2026-03-25T14:00:00"`。请注意时间为 24 小时制。
- 用户必须等待至少 48 小时后才能重新参加考试。

# --hints--

`can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00"), True)`)
}})
```

`can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00"), False)`)
}})
```

`can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00"), True)`)
}})
```

`can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_retake(finish_time, current_time):

    return finish_time
```

# --solutions--

```py
from datetime import datetime, timedelta

def can_retake(finish_time, current_time):
    fmt = "%Y-%m-%dT%H:%M:%S"
    finished_dt = datetime.strptime(finish_time, fmt)
    current_dt = datetime.strptime(current_time, fmt)
    diff = current_dt - finished_dt
    cooldown = timedelta(hours=48)

    return diff >= cooldown
```
