---
id: 69a890af247de743333bd4d1
title: "挑战 228：电影之夜"
challengeType: 29
dashedName: challenge-228
---

# --description--

给定一个表示星期几的字串、另一个表示放映时间的字串和一个整数票数，返回该场电影票的总费用。

给定的日期将是以下之一：

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

放映时间将以 `"H:MMam"` 或 `"H:MMpm"` 形式给出。例如 `"10:00am"` 或 `"10:00pm"`。

以 `"$D.CC"` 形式返回总费用，使用以下规则：

- 周末（星期五－星期日）：每张票 12.00 美元。
- 工作日（星期一 - 星期四）：每张票 10.00 美元。
- 下午 5:00 之前的场次：每张票减免 $2.00（星期二除外）。
- 星期二：所有票价均为 $5.00 每张。

# --hints--

`get_movie_night_cost("Saturday", "10:00pm", 1)` 应该返回 `"$12.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Saturday", "10:00pm", 1), "$12.00")`)
}})
```

`get_movie_night_cost("Sunday", "10:00am", 1)` 应该返回 `"$10.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Sunday", "10:00am", 1), "$10.00")`)
}})
```

`get_movie_night_cost("Tuesday", "7:20pm", 2)` 应该返回 `"$10.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "7:20pm", 2), "$10.00")`)
}})
```

`get_movie_night_cost("Wednesday", "5:40pm", 3)` 应该返回 `"$30.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Wednesday", "5:40pm", 3), "$30.00")`)
}})
```

`get_movie_night_cost("Monday", "11:50am", 4)` 应该返回 `"$32.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Monday", "11:50am", 4), "$32.00")`)
}})
```

`get_movie_night_cost("Friday", "4:30pm", 5)` 应该返回 `"$50.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Friday", "4:30pm", 5), "$50.00")`)
}})
```

`get_movie_night_cost("Tuesday", "11:30am", 1)` 应该返回 `"$5.00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "11:30am", 1), "$5.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):

    return day
```

# --solutions--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):
    if day == "Tuesday":
        price = 5
    else:
        if day in ["Friday", "Saturday", "Sunday"]:
            price = 12
        else:
            price = 10

        hour = int(showtime.split(":")[0])
        period = showtime[-2:]

        hour24 = hour

        if period == "pm" and hour != 12:
            hour24 += 12
        if period == "am" and hour == 12:
            hour24 = 0

        if hour24 < 17:
            price -= 2

    total = price * number_of_tickets

    return f"${total:.2f}"
```
