---
id: 69a890af247de743333bd4cc
title: "挑战 223：QR 解码器"
challengeType: 29
dashedName: challenge-223
---

# --description--

给定一个 6x6 的矩阵（数组的数组），表现一个 QR 码，返回代码中的二进制数据字串。

- 二维码可以以任意 90 度增量旋转的方式给出。
- 一个正确定位的代码在左下角、左上角和右上角有一个 2x2 的 `1`（定位标记）组。
- 这三个 2x2 方向标记不是二进制 `data` 的一部分。
- 当二维码正确放置时，二进制数据按从左到右、从上到下的顺序（像书本一样）被读取。
- 一个代码总是只有一个有效的方向。

例如，给定：

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

或者给出相同的代码但方向不同：

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

返回 `"000000000000000000000001"`，所有二进制数据，排除三个 2x2 方向标记。

# --hints--

`decode_qr(["110011", "110011", "000000", "000000", "110000", "110001"])` 应该返回 `"000000000000000000000001"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001")`)
}})
```

`decode_qr(["100011", "000011", "000000", "000000", "110011", "110011"])` 应该返回 `"000000000000000000000001"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001")`)
}})
```

`decode_qr(["110011", "111111", "010000", "110000", "110011", "110100"])` 应该返回 `"001101000011000000110100"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100")`)
}})
```

`decode_qr(["011011", "101011", "101000", "100010", "110011", "111011"])` 应该返回 `"010001000100010101010110"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110")`)
}})
```

`decode_qr(["111100", "110001", "100011", "001101", "110011", "110011"])` 应该返回 `"010000100100100101001110"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode_qr(qr_code):

    return qr_code
```

# --solutions--

```py
def is_oriented(code):
    marker_indices = [
        (0, 0),
        (0, 1),
        (0, 4),
        (0, 5),
        (1, 0),
        (1, 1),
        (1, 4),
        (1, 5),
        (4, 0),
        (4, 1),
        (5, 0),
        (5, 1)
    ]

    for row, col in marker_indices:
        if code[row][col] != '1':
            return False

    return True


def rotate(code):
    new_code = ['', '', '', '', '', '']
    for i in range(6):
        for j in range(5, -1, -1):
            new_code[i] += code[j][i]
    return new_code


def get_data(code):
    data = ''
    data += code[0][2:4]
    data += code[1][2:4]
    data += code[2]
    data += code[3]
    data += code[4][2:]
    data += code[5][2:]
    return data


def decode_qr(qr_code):
    code = qr_code[:]

    while not is_oriented(code):
        code = rotate(code)

    return get_data(code)
```
