---
id: 69a890af247de743333bd4cf
title: "挑战 226：通过考试计数"
challengeType: 29
dashedName: challenge-226
---

# --description--

给定一个学生考试成绩的数组和通过考试所需的分数，返回通过考试的学生人数。

# --hints--

`passing_count([90, 85, 75, 60, 50], 70)` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([90, 85, 75, 60, 50], 70), 3)`)
}})
```

`passing_count([100, 80, 75, 88, 72, 74, 79, 71, 60, 92], 75)` 应该返回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([100, 80, 75, 88, 72, 74, 79, 71, 60, 92], 75), 6)`)
}})
```

`passing_count([79, 60, 88, 72, 74, 59, 75, 71, 80, 92], 60)` 应该返回 `9`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([79, 60, 88, 72, 74, 59, 75, 71, 80, 92], 60), 9)`)
}})
```

`passing_count([76, 79, 80, 70, 71, 65, 79, 78, 59, 72], 85)` 应该返回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([76, 79, 80, 70, 71, 65, 79, 78, 59, 72], 85), 0)`)
}})
```

`passing_count([84, 65, 98, 53, 58, 71, 91, 80, 92, 70, 73, 83, 86, 69, 84, 77, 72, 58, 69, 75, 66, 68, 72, 96, 90, 63, 88, 63, 80, 67], 60)` 应该返回 `27`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([84, 65, 98, 53, 58, 71, 91, 80, 92, 70, 73, 83, 86, 69, 84, 77, 72, 58, 69, 75, 66, 68, 72, 96, 90, 63, 88, 63, 80, 67], 60), 27)`)
}})
```

# --seed--

## --seed-contents--

```py
def passing_count(scores, passing_score):

    return scores
```

# --solutions--

```py
def passing_count(scores, passing_score):
    count = 0

    for score in scores:
        if score >= passing_score:
            count += 1

    return count
```
