---
id: 69a890af247de743333bd4d2
title: "挑战 229：截断 `text` 2"
challengeType: 29
dashedName: challenge-229
---

# --description--

给定一个 `string`，返回一个被截断的 `string`，使得字符的总宽度不超过 50 个单位。

每个字符都有特定的宽度：

| Letters | Width |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

上表包含所有大写和小写字母。此外：

- 空格（`" "`）的宽度为 2
- 句号（`"."`）的宽度为 1

- 如果给定的字串长度在 50 个单位或以下，返回该字串本身，否则
- 截断字串并在末尾添加三个句点（`"..."`），使其总宽度（包括三个句点）尽可能接近 50 个单位且不超过。

# --hints--

`truncate_text("The quick brown fox")` 应该返回 `"The quick brown f..."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The quick brown fox"), "The quick brown f...")`)
}})
```

`truncate_text("The silky smooth sloth")` 应该返回 `"The silky smooth s..."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The silky smooth sloth"), "The silky smooth s...")`)
}})
```

`truncate_text("THE LOUD BRIGHT BIRD")` 应该返回 `"THE LOUD BRIG..."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...")`)
}})
```

`truncate_text("The fast striped zebra")` 应该返回 `"The fast striped z..."`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The fast striped zebra"), "The fast striped z...")`)
}})
```

`truncate_text("The big black bear")` 应该返回 `"The big black bear"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The big black bear"), "The big black bear")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(s):

    return s
```

# --solutions--

```py
def truncate_text(s):
    MAX_WIDTH = 50
    ELLIPSIS = "..."
    ELLIPSIS_WIDTH = 3
    TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH

    char_groups = {
        "ilI.": 1,
        "fjrt ": 2,
        "abcdeghkmnopqrstuvwxyzJL": 3,
        "ABCDEFGHKMNOPQRSTUVWXYZ": 4
    }

    def get_char_width(char):
        for key, width in char_groups.items():
            if char in key:
                return width
        return 3

    def string_width(string):
        total = 0
        for char in string:
            total += get_char_width(char)
        return total

    if string_width(s) <= MAX_WIDTH:
        return s

    result = ""
    total_width = 0

    for char in s:
        char_width = get_char_width(char)

        if total_width + char_width > TRUNCATE_LIMIT:
            break

        result += char
        total_width += char_width

    return result + ELLIPSIS
```
