---
id: 69a890af247de743333bd4ce
title: "挑战 225：不允许连续重复"
challengeType: 29
dashedName: challenge-225
---

# --description--

给定一个字串，判断它是否没有重复的字符。

- 如果一个字串中没有相同的字符连续出现两次或更多次，则该字串没有重复。

# --hints--

`has_no_repeats("hi world")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hi world"), True)`)
}})
```

`has_no_repeats("hello world")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("hello world"), False)`)
}})
```

`has_no_repeats("abcdefghijklmnopqrstuvwxyz")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("abcdefghijklmnopqrstuvwxyz"), True)`)
}})
```

`has_no_repeats("freeCodeCamp")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("freeCodeCamp"), False)`)
}})
```

`has_no_repeats("The quick brown fox jumped over the lazy dog.")` 应该 `return` `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("The quick brown fox jumped over the lazy dog."), True)`)
}})
```

`has_no_repeats("Mississippi")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_no_repeats("Mississippi"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_no_repeats(s):

    return s
```

# --solutions--

```py
def has_no_repeats(s):
    for i in range(1, len(s)):
        if s[i] == s[i-1]:
            return False
    return True
```
