---
id: 69b1028d6e265413d0198a2a
title: "挑战 231：ISBN-10 验证器"
challengeType: 29
dashedName: challenge-231
---

# --description--

给定一个字串，判断它是否是有效的 ISBN-10。

ISBN-10 由连字符（`"-"`）和其他 10 个字符组成。去掉连字符（`"-"`）后：

- 前 9 个字符必须是数字，
- 最后一个字符可以是数字或字母 `"X"`，它代表数字 10。

验证方法：

- 将每个数字（或值）乘以它的位置（第一个数字乘以 1，第二个乘以 2，依此类推）。
- 将所有结果相加。
- 如果总和能被 11 整除，则该 ISBN-10 有效。

# --hints--

`is_valid_isbn10("0-306-40615-2")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-306-40615-2"), True)`)
}})
```

`is_valid_isbn10("0-306-40615-1")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-306-40615-1"), False)`)
}})
```

`is_valid_isbn10("0-8044-2957-X")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-8044-2957-X"), True)`)
}})
```

`is_valid_isbn10("X-306-40615-2")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("X-306-40615-2"), False)`)
}})
```

`is_valid_isbn10("0-6822-2589-4")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-6822-2589-4"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn10(s):

    return s
```

# --solutions--

```py
def is_valid_isbn10(s):
    isbn = s.replace("-", "")

    if len(isbn) != 10:
        return False

    total = 0

    for i in range(9):
        if not isbn[i].isdigit():
            return False
        total += int(isbn[i]) * (i + 1)

    last = isbn[9]

    if last == "X":
        total += 10 * 10
    elif last.isdigit():
        total += int(last) * 10
    else:
        return False

    return total % 11 == 0
```
