---
id: 69b1028d6e265413d0198a2f
title: "挑战 236：浏览器历史记录"
challengeType: 29
dashedName: challenge-236
---

# --description--

给定一个浏览器命令的数组，返回一个包含两个值的数组：历史记录作为 URL 的数组，以及当前页面的索引。

有效的命令有：

- `"URL"` —— 其中 URL 是一个网页地址（例如 `"freecodecamp.org"`）。导航到给定的 URL，将其添加到历史记录的下一个位置，并丢弃任何前进的历史。
- `"Back"` —— 移动到历史记录中的上一页，如果没有上一页，则停留在当前页面。
- `"Forward"` —— 移动到历史记录中的下一页，如果没有下一页，则停留在当前页面。

例如，给定 `["freecodecamp.org", "freecodecamp.org/learn", "Back"]`，返回 `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`。

# --hints--

`get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` 应该返回 `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` 应该返回 `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"])` 应该返回 `[["example.com", "example.com/contact", "example.com/blog"], 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` 应该返回 `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "Back"])` 应该返回 `[["example.com", "example.com/about"], 0]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Forward"])` 应该返回 `[["example.com", "example.com/about"], 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_browser_history(commands):

    return commands
```

# --solutions--

```py
def get_browser_history(commands):
    stack = []
    index = -1

    for command in commands:
        if command == "Back":
            if index > 0:
                index -= 1
        elif command == "Forward":
            if index < len(stack) - 1:
                index += 1
        else:
            del stack[index + 1:]
            stack.append(command)
            index += 1

    return [stack, index]
```
