---
id: 69b58ce40693f140c84c855a
title: "挑战 241：FizzBuzz 验证器"
challengeType: 29
dashedName: challenge-241
---

# --description--

给定一个循序的整数数组，其中 3 和 5 的倍数被替换，判断它是否是一个有效的 FizzBuzz 序列。

在有效的 FizzBuzz 序列中：

- 3 的倍数被替换为 "Fizz"。
- 5 的倍数被替换为 "Buzz"。
- 同时是 3 和 5 的倍数被替换为 "FizzBuzz"。
- 其他所有数字保持为整数。

# --hints--

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"])` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([13, 14, "FizzBuzz", 16, 17])` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([13, 14, "FizzBuzz", 16, 17]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, 5])` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, 5]), False)`)
}})
```

`is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", "Buzz", 5])` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", "Buzz", 5]), False)`)
}})
```

`is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]), False)`)
}})
```

`is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_fizz_buzz(arr):

    return arr
```

# --solutions--

```py
def is_fizz_buzz(arr):
    start = None
    for i, val in enumerate(arr):
        if isinstance(val, int):
            start = val - i
            break

    if start is None:
        return False

    for i in range(len(arr)):
        n = start + i
        if n % 15 == 0:
            expected = "FizzBuzz"
        elif n % 3 == 0:
            expected = "Fizz"
        elif n % 5 == 0:
            expected = "Buzz"
        else:
            expected = n
        if arr[i] != expected:
            return False

    return True
```
