---
id: 69b559d2903b9e4afe9075f7
title: "Challenge 237: Equation Validation"
challengeType: 29
dashedName: challenge-237
---

# --description--

给定一个表示数学方程的字串，判断它是否正确。

- 左侧最多包含三个正整数和操作符 `+`、`-`、`*` 和 `/`。
- 方程将以 `"number operator number = number"` 的形式给出（左侧有两个或三个数字）。例如：`"2 + 2 = 4"` 或 `"2 + 3 - 1 = 4"`。
- 右侧始终是一个整数。

遵循标准的运算顺序：乘法和除法先于加法和减法，从左到右依次计算。

# --hints--

`is_valid_equation("2 + 2 = 4")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 2 = 4"), True)`)
}})
```

`is_valid_equation("2 + 3 - 1 = 4")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 3 - 1 = 4"), True)`)
}})
```

`is_valid_equation("8 / 2 = 4")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("8 / 2 = 4"), True)`)
}})
```

`is_valid_equation("10 * 5 = 50")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("10 * 5 = 50"), True)`)
}})
```

`is_valid_equation("2 - 2 = 0")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 - 2 = 0"), True)`)
}})
```

`is_valid_equation("2 + 9 / 3 = 5")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 9 / 3 = 5"), True)`)
}})
```

`is_valid_equation("20 - 2 * 3 = 14")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("20 - 2 * 3 = 14"), True)`)
}})
```

`is_valid_equation("2 + 5 = 6")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 5 = 6"), False)`)
}})
```

`is_valid_equation("10 - 2 * 3 = 24")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("10 - 2 * 3 = 24"), False)`)
}})
```

`is_valid_equation("3 + 9 / 3 = 4")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("3 + 9 / 3 = 4"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_equation(equation):

    return equation
```

# --solutions--

```py
def is_valid_equation(equation):
    left, right = equation.split(" = ")
    tokens = left.split(" ")

    i = 1
    while i < len(tokens) - 1:
        op = tokens[i]
        if op in ("*", "/"):
            a = int(tokens[i - 1])
            b = int(tokens[i + 1])
            computed = a * b if op == "*" else a // b
            tokens[i - 1:i + 2] = [str(computed)]
        else:
            i += 2

    result = int(tokens[0])
    i = 1
    while i < len(tokens) - 1:
        op = tokens[i]
        val = int(tokens[i + 1])
        result = result + val if op == "+" else result - val
        i += 2

    return result == int(right)
```
