---
id: 69b58ce40693f140c84c855b
title: "挑战 242：下一个宾果号码"
challengeType: 29
dashedName: challenge-242
---

# --description--

给定一个宾果号码，返回下一个循序的宾果号码。

宾果号码由一个字母和一个根据下表对应范围内的数字组成：

| 字母 | 数字范围 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

例如，给定 `"B10"`，返回下一个宾果号码 `"B11"`。如果给定的是最后一个宾果号码，则返回 `"B1"`。

# --hints--

`get_next_bingo_number("B10")` 应该返回 `"B11"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("B10"), "B11")`)
}})
```

`get_next_bingo_number("N33")` 应该返回 `"N34"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("N33"), "N34")`)
}})
```

`get_next_bingo_number("I30")` 应该返回 `"N31"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("I30"), "N31")`)
}})
```

`get_next_bingo_number("G60")` 应该返回 `"O61"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("G60"), "O61")`)
}})
```

`get_next_bingo_number("O75")` 应该返回 `"B1"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("O75"), "B1")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_next_bingo_number(n):

    return n
```

# --solutions--

```py
def get_next_bingo_number(n):
    num = int(n[1:])
    next_num = 1 if num == 75 else num + 1

    if next_num <= 15:
        return "B" + str(next_num)
    if next_num <= 30:
        return "I" + str(next_num)
    if next_num <= 45:
        return "N" + str(next_num)
    if next_num <= 60:
        return "G" + str(next_num)
    return "O" + str(next_num)
```
