---
id: 69b58ce40693f140c84c8559
title: "挑战 240：回文字符"
challengeType: 29
dashedName: challenge-240
---

# --description--

给定一个字串，判断它是否是回文，并返回中间的字符（如果长度是奇数）或中间的两个字符（如果长度是偶数）。

- 回文是指正着读和反着读都相同的字串。
- 如果不是回文，返回 `"none"`。

# --hints--

`palindrome_locator("racecar")` 应该返回 `"e"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("racecar"), "e")`)
}})
```

`palindrome_locator("level")` 应该返回 `"v"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("level"), "v")`)
}})
```

`palindrome_locator("freecodecamp")` 应该返回 `"none"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("freecodecamp"), "none")`)
}})
```

`palindrome_locator("noon")` 应该返回 `"oo"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("noon"), "oo")`)
}})
```

`palindrome_locator("11100111")` 应该返回 `"00"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("11100111"), "00")`)
}})
```

# --seed--

## --seed-contents--

```py
def palindrome_locator(s):

    return s
```

# --solutions--

```py
def palindrome_locator(s):
    if s != s[::-1]:
        return "none"
    mid = len(s) // 2
    return s[mid] if len(s) % 2 == 1 else s[mid - 1] + s[mid]
```
