---
id: 69bc6cb30c1d112a2e110a03
title: "Challenge 246: Name Initials"
challengeType: 29
dashedName: challenge-246
---

# --description--

给定一个完整的姓名字串，返回他们的首字母。

- 需要初始化的名字由空格分隔。
- 首字母应大写。
- 首字母之间应以点号分隔。

例如，`"Tommy Millwood"` 返回 `"T.M."`。

# --hints--

`get_initials("Tommy Millwood")` should return `"T.M."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Tommy Millwood"), "T.M.")`)
}})
```

`get_initials("Savanna Puddlesplash")` should return `"S.P."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Savanna Puddlesplash"), "S.P.")`)
}})
```

`get_initials("Frances Cowell Conrad")` should return `"F.C.C."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Frances Cowell Conrad"), "F.C.C.")`)
}})
```

`get_initials("Dragon")` should return `"D."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Dragon"), "D.")`)
}})
```

`get_initials("Dorothy Vera Clump Haverstock Norris")` should return `"D.V.C.H.N."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_initials(name):

    return name
```

# --solutions--

```py
def get_initials(name):
    return ''.join(word[0].upper() + '.' for word in name.split())
```
