---
id: 69bc6cb30c1d112a2e110a07
title: "Challenge 250: Hidden Key"
challengeType: 29
dashedName: challenge-250
---

# --description--

欢迎来到第 250 个每日挑战！

给定一个加密字串，使用加密密钥将其解码并返回。

查找密钥的方法：

- 查看截至今天所有挑战编号是 25 的倍数的每日挑战（包括本次）。
- 取出这些挑战标题的第一个字母并组合成一个字串。如果标题以非字母开头，找到它的第一个字母。

解码消息时，遍历加密消息中的每个字母：

1. 查看密钥中对应的字母（如果消息比密钥长，则重复密钥）。
  2. 将密钥字母转换为对应的数字：`"A"` = 1，`"B"` = 2，…，`"Z"` = 26。
  3. 将加密字母在字母表中向后移动该数字个位置。
  4. 如果移动超出 `"A"` 之前，则从 `"Z"` 重新开始。

例如，如果加密消息以 `"Y"` 开头，且第一个密钥字母是 `"V"`（22），则将 `"Y"` 向后移动 22 位得到 `"C"`。对每个字母重复此过程以解码完整消息。

- 只有字母会被移动，空格保持不变。
- 所有给定和返回的字母均为大写。

# --hints--

`decode("YAVJYNXE")` should return `"CONGRATS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")` should return `"CODING LEGEND"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")` should return `"YOU ARE AWESOME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")` should return `"KEEP GOING"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` should return `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
