---
id: 69bc6cb30c1d112a2e110a06
title: "Challenge 249: String Math"
challengeType: 29
dashedName: challenge-249
---

# --description--

给定一个包含数字和其他字符的字串，根据数字之间非数字字符的数量对数字执行数学运算。

- 如果分隔两个数字的字符数量是偶数，则使用加法。
- 如果是奇数，则使用减法。
- 连续的数字构成一个单独的数字。
- 运算从左到右依次执行。
- 忽略开头和结尾非数字的字符。

例如，给定 `"3ab10c8"`，返回 `5`。将 3 和 10 相加得到 13，因为它们之间有偶数个字符。然后从 13 中减去 8，因为结果和 8 之间有奇数个字符。

# --hints--

`do_math("3ab10c8")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("3ab10c8"), 5)`)
}})
```

`do_math("6MINUS4")` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("6MINUS4"), 2)`)
}})
```

`do_math("9plus3")` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("9plus3"), 12)`)
}})
```

`do_math("5fkwo#10i#%.<>15P=@20!#B/25")` should return `15`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("5fkwo#10i#%.<>15P=@20!#B/25"), 15)`)
}})
```

`do_math("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` should return `67`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67)`)
}})
```

# --seed--

## --seed-contents--

```py
def do_math(s):

    return s
```

# --solutions--

```py
import re

def do_math(s):
    tokens = re.findall(r'\d+|\D+', s)
    result = None
    pending_op = None
    for token in tokens:
        if re.match(r'^\d+$', token):
            if result is None:
                result = int(token)
            elif pending_op == 'add':
                result += int(token)
            else:
                result -= int(token)
        else:
            pending_op = 'add' if len(token) % 2 == 0 else 'subtract'
    return result
```
