---
id: 69c5f3d787b1725d5f00c8b8
title: "Challenge 254: Odd Words"
challengeType: 29
dashedName: challenge-254
---

# --description--

给定一个字串，返回其中字母数为奇数的字词。

- 给定字串中的字词将由单个空格分隔。
- 返回的字词也用单个空格分隔。

# --hints--

`get_odd_words("This is a super good test")` should return `"a super"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odd_words("This is a super good test"), "a super")`)
}})
```

`get_odd_words("one two three four")` should return `"one two three"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odd_words("one two three four"), "one two three")`)
}})
```

`get_odd_words("banana split sundae with rainbow sprinkles on top")` should return `"split rainbow sprinkles top"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odd_words("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top")`)
}})
```

`get_odd_words("The quick brown fox jumped over the lazy river")` should return `"The quick brown fox the river"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odd_words("The quick brown fox jumped over the lazy river"), "The quick brown fox the river")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_odd_words(s):

    return s
```

# --solutions--

```py
def get_odd_words(s):
    return ' '.join(word for word in s.split(' ') if len(word) % 2 == 1)
```
