---
id: 69c5f3d787b1725d5f00c8ba
title: "挑战 256：最近的时间方向"
challengeType: 29
dashedName: challenge-256
---

# --description--

给定两个时间，判断从第一个时间到第二个时间，向前移动还是向后移动更快。

- 时间以 24 小时制形式给出（`"HH:MM"`）
- 时钟是循环的（向前移动时，23:59 会到 00:00，向后移动时，00:00 会到 23:59）

返回：

- 如果向前移动更短，返回 `"forward"`
- 如果向后移动更短，返回 `"backward"`
- 如果两个方向所需时间相同，返回 `"equal"`

# --hints--

`get_direction("10:00", "12:00")` 应该返回 `"forward"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("10:00", "12:00"), "forward")`)
}})
```

`get_direction("11:00", "05:00")` 应该返回 `"backward"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("11:00", "05:00"), "backward")`)
}})
```

`get_direction("00:00", "12:00")` 应该返回 `"equal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("00:00", "12:00"), "equal")`)
}})
```

`get_direction("15:45", "01:10")` 应该返回 `"forward"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("15:45", "01:10"), "forward")`)
}})
```

`get_direction("03:30", "19:50")` 应该返回 `"backward"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("03:30", "19:50"), "backward")`)
}})
```

`get_direction("06:30", "18:30")` 应该返回 `"equal"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("06:30", "18:30"), "equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_direction(time1, time2):

    return time1
```

# --solutions--

```py
def get_direction(time1, time2):
    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    start = to_minutes(time1)
    end = to_minutes(time2)
    forward = (end - start + 1440) % 1440
    backward = (start - end + 1440) % 1440

    if forward < backward:
        return 'forward'
    if backward < forward:
        return 'backward'
    return 'equal'
```
