---
id: 69c5f3d787b1725d5f00c8bd
title: "挑战 259：FizzBuzz 爆炸"
challengeType: 29
dashedName: challenge-259
---

# --description--

给定一个整数，返回将单词 `"fizzbuzz"` 变成包含至少给定数量 `"z"` 的字串所需的步骤数，规则如下：

- 从字串 `"fizzbuzz"` 开始。
- 每一步，使用字串中字母的位置应用标准 FizzBuzz 规则（第一个 `"f"` 的位置是 1）。
  - 如果字母位置能被 3 整除，则将该字母替换为 `"fizz"`
  - 如果能被 5 整除，则将该字母替换为 `"buzz"`
  - 如果能被 3 和 5 同时整除，则将该字母替换为 `"fizzbuzz"`

所以经过 1 步，`"fizzbuzz"` 会变成 `"fifizzzbuzzfizzzz"`，它包含 9 个 `"z"`。

# --hints--

`explode_fizzbuzz(9)` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` 应该返回 `4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` 应该返回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` 应该返回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` 应该返回 `11`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` 应该返回 `14`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
