---
id: 69c5f3d787b1725d5f00c8b9
title: "挑战 255：地球日清理队"
challengeType: 29
dashedName: challenge-255
---

# --description--

今天是地球日。给定一个你清理过的项数组，根据以下规则返回你的总清理得分。

给定的项将是以下之一：

| 项 | 基础值 |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

稀有项表现为 `["rare", value]`。例如，`["rare", 80]`。稀有项不会获得连击奖励。

- 连击奖励：如果同一项连续出现，它会获得递增的奖励分。
  - 第一次连续出现：基础值
  - 第二次：基础值 + 1
  - 第三次：基础值 + 2
  - 依此类推。
- 第五项倍数：每收集到第五项都会获得一个倍数。
  - 第五项：*2
  - 第十项：*3
  - 依此类推。

- 在计算任何奖励后应用倍数。

# --hints--

`get_cleanup_score(["bottle", "straw", "shoe", "battery"])` 应该返回 `44`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["bottle", "straw", "shoe", "battery"]), 44)`)
}})
```

`get_cleanup_score(["electronics", "straw", "newspaper", "bottle", "bag"])` 应该返回 `58`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["electronics", "straw", "newspaper", "bottle", "bag"]), 58)`)
}})
```

`get_cleanup_score(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` 应该返回 `79`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79)`)
}})
```

`get_cleanup_score(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` 应该返回 `358`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358)`)
}})
```

`get_cleanup_score(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` 应该返回 `383`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_cleanup_score(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_cleanup_score(items):

    return items
```

# --solutions--

```py
def get_cleanup_score(items):
    values = {
        'bottle': 10, 'can': 6, 'bag': 8, 'tire': 35,
        'straw': 4, 'cardboard': 3, 'newspaper': 3,
        'shoe': 12, 'electronics': 25, 'battery': 18, 'mattress': 38
    }
    total = 0
    prev_item = None
    streak = 0
    for i, item in enumerate(items):
        is_rare = isinstance(item, list) and item[0] == 'rare'
        if is_rare:
            base = item[1]
            prev_item = None
            streak = 0
        else:
            base = values[item]
            streak = streak + 1 if item == prev_item else 0
            prev_item = item
        points = base if is_rare else base + streak
        if (i + 1) % 5 == 0:
            multiplier = (i + 1) // 5 + 1
            points *= multiplier
        total += points
    return total
```
