---
id: 69c6ff713a52713463aa7928
title: "挑战 262：URL 查询解析器"
challengeType: 29
dashedName: challenge-262
---

# --description--

给定一个包含查询字串的 URL，将查询字串解析成一个包含点对的对象（或字典）。

- 查询字串从 `"?"` 之后开始，
- 每个参数由 `"&"` 分隔，
- 每个点对由 `"="` 分隔

例如，给定 `"https://example.com/search?name=Alice&age=30"`，应返回：

```json
{
  "name": "Alice",
  "age": "30"
}
```

所有值都应作为字串返回。

# --hints--

`parse_url_query("https://example.com/search?name=Alice&age=30")` 应返回 `{"name": "Alice", "age": "30"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"})`)
}})
```

`parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python")` 应返回 `{"skill": "programming", "language": "python"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"})`)
}})
```

`parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2")` 应返回 `{"category": "books", "sort": "asc", "page": "2"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"})`)
}})
```

`parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now")` 应返回 `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_url_query(url):

    return url
```

# --solutions--

```py
def parse_url_query(url):
    query_string = url.split('?')[1]
    return dict(param.split('=') for param in query_string.split('&'))
```
