---
id: 69c6ff713a52713463aa7927
title: "挑战 261：数字单词"
challengeType: 29
dashedName: challenge-261
---

# --description--

给定一个 0 到 99 之间的整数，返回它的英文单词表现形式。

- `0` 返回 `"zero"`。
- 1 到 19 的数字有独特的名称（`"one"`、`"two"`、…、`"ten"`、`"eleven"`、…、`"eighteen"`、`"nineteen"`）。
- 20 到 90 的 10 的倍数有它们自己的名称（`"twenty"`、`"thirty"`、…、`"eighty"`、`"ninety"`）。
- 21 到 99 中不是 10 的倍数的数字写作两个单词，中间用连字符连接。例如 `"forty-two"` 和 `"fifty-three"`。

# --hints--

`get_number_words(0)` 应该返回 `"zero"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(0), "zero")`)
}})
```

`get_number_words(10)` 应该返回 `"ten"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(10), "ten")`)
}})
```

`get_number_words(19)` 应该返回 `"nineteen"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(19), "nineteen")`)
}})
```

`get_number_words(30)` 应该返回 `"thirty"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(30), "thirty")`)
}})
```

`get_number_words(53)` 应该返回 `"fifty-three"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(53), "fifty-three")`)
}})
```

`get_number_words(7)` 应该返回 `"seven"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(7), "seven")`)
}})
```

`get_number_words(12)` 应该返回 `"twelve"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(12), "twelve")`)
}})
```

`get_number_words(60)` 应该返回 `"sixty"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(60), "sixty")`)
}})
```

`get_number_words(67)` 应该返回 `"sixty-seven"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(67), "sixty-seven")`)
}})
```

`get_number_words(98)` 应该返回 `"ninety-eight"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(98), "ninety-eight")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_words(n):

    return n
```

# --solutions--

```py
def get_number_words(n):
    ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
            "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
            "seventeen","eighteen","nineteen"]
    tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"]

    if n < 20:
        return ones[n]
    if n % 10 == 0:
        return tens[n // 10]
    return tens[n // 10] + "-" + ones[n % 10]
```
