---
id: 69c6ff713a52713463aa7929
title: "挑战 263：二进制填字游戏"
challengeType: 29
dashedName: challenge-263
---

# --description--

给定一个字符，判断它的 8 位二进制表现是否可以在下列网格中找到，横向或纵向，任一方向：

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

例如，`"A"` 的二进制表现是 `01000001`，它出现在第一行从左到右的位置。

# --hints--

`is_in_crossword("I")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("I"), True)`)
}})
```

`is_in_crossword("D")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("D"), True)`)
}})
```

`is_in_crossword("0")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("0"), True)`)
}})
```

`is_in_crossword("u")` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("u"), True)`)
}})
```

`is_in_crossword("Y")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("Y"), False)`)
}})
```

`is_in_crossword("p")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("p"), False)`)
}})
```

`is_in_crossword("1")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("1"), False)`)
}})
```

`is_in_crossword("Q")` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("Q"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_in_crossword(char):

    return char
```

# --solutions--

```py
def is_in_crossword(char):
    grid = [
        "01000001",
        "01101111",
        "01000100",
        "01100101",
        "01010010",
        "01010100",
        "01101000",
        "10101110",
    ]
    target = format(ord(char), '08b')
    cols = [''.join(grid[r][c] for r in range(8)) for c in range(8)]
    return any(target in seq or target in seq[::-1] for seq in grid + cols)
```
