---
id: 69cfca90e8a0a6d4d6871c4e
title: "Challenge 264: Anagram Groups"
challengeType: 29
dashedName: challenge-264
---

# --description--

给定一个单词数组，返回一个二维数组，将单词按字谜分组。

- 如果单词包含相同字母且顺序可以不同，则它们是字谜。
- 每个单词只属于一个分组。
- 返回的顺序无关紧要。

例如，给定 `["listen", "silent", "hello", "enlist", "world"]`，返回 `[["listen", "silent", "enlist"], ["hello"], ["world"]]`。

# --hints--

`group_anagrams(["listen", "silent", "hello", "enlist", "world"])` should return `[["listen", "silent", "enlist"], ["hello"], ["world"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["listen", "silent", "hello", "enlist", "world"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["enlist", "listen", "silent"], ["hello"], ["world"]])`)
}})
```

`group_anagrams(["eat", "tea", "tan", "ate", "nat", "bat"])` should return `[["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["eat", "tea", "tan", "ate", "nat", "bat"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["ate", "eat", "tea"], ["bat"], ["nat", "tan"]])`)
}})
```

`group_anagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])` should return `[["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]])`)
}})
```

`group_anagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])` should return `[["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def group_anagrams(words):

    return words
```

# --solutions--

```py
def group_anagrams(words):
    groups = {}
    for word in words:
        key = ''.join(sorted(word))
        if key not in groups:
            groups[key] = []
        groups[key].append(word)
    return list(groups.values())
```
