---
id: 69cfca90e8a0a6d4d6871c50
title: "Challenge 266: Good Day"
challengeType: 29
dashedName: challenge-266
---

# --description--

给定一个 `"HH:MM"` 形式的时间字串（24 小时制），返回：

- `"Good morning"` 对于时间 `05:00` 到 `11:59`
- `"Good afternoon"` 对于时间 `12:00` 到 `17:59`
- `"Good evening"` 对于时间 `18:00` 到 `21:59`
- `"Good night"` 对于时间 `22:00` 到 `04:59`

# --hints--

`get_greeting("06:30")` should return `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("06:30"), "Good morning")`)
}})
```

`get_greeting("12:00")` should return `"Good afternoon"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("12:00"), "Good afternoon")`)
}})
```

`get_greeting("21:59")` should return `"Good evening"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("21:59"), "Good evening")`)
}})
```

`get_greeting("00:01")` should return `"Good night"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("00:01"), "Good night")`)
}})
```

`get_greeting("11:30")` should return `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("11:30"), "Good morning")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_greeting(s):

    return s
```

# --solutions--

```py
def get_greeting(s):
    hours, minutes = map(int, s.split(':'))
    total = hours * 60 + minutes

    if 300 <= total < 720:
        return 'Good morning'
    if 720 <= total < 1080:
        return 'Good afternoon'
    if 1080 <= total < 1320:
        return 'Good evening'
    return 'Good night'
```
