---
id: 69d03e549613fbdbe21d11e0
title: "挑战 271：用药提醒"
challengeType: 29
dashedName: challenge-271
---

# --description--

给定一个药物数组和一个表示当前时间的字串，返回你下一次需要服用的药物以及距离服用的时间。

- 每种药物的形式是 `[name, lastTaken]`，其中 `name` 是药物名称，`lastTaken` 是上次服用的时间。
- 所有给定的时间都将采用 `"HH:MM"`（24 小时）格式。

使用以下用药调度：

| 名称 | 调度 |
| - | - |
| Deployxitrin | 08:00，16:00 |
| Debuggamanizole | 07:00，13:00，21:00 |
| Mergeflictamine | 每 4 小时一次 |

返回格式为 `"{name} in Hh Mm"` 的字串。例如，`"Debuggamanizole in 2h 0m"` 或 `"Deployxitrin in 1h 5m"`。

# --hints--

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00")` 应该返回 `"Debuggamanizole in 2h 0m"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00"), "Debuggamanizole in 2h 0m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55")` 应该返回 `"Deployxitrin in 1h 5m"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55"), "Deployxitrin in 1h 5m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15")` 应该返回 `"Mergeflictamine in 0h 45m"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15"), "Mergeflictamine in 0h 45m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59")` 应该返回 `"Debuggamanizole in 0h 1m"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59"), "Debuggamanizole in 0h 1m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55")` 应该返回 `"Debuggamanizole in 0h 5m"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55"), "Debuggamanizole in 0h 5m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00")` 应该返回 `"Mergeflictamine in 3h 30m"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00"), "Mergeflictamine in 3h 30m")`)
}})
```

# --seed--

## --seed-contents--

```py
def medication_reminder(medications, current_time):

    return medications
```

# --solutions--

```py
def medication_reminder(medications, current_time):
    schedule = {
        'Deployxitrin': {'type': 'fixed', 'times': ['08:00', '16:00']},
        'Debuggamanizole': {'type': 'fixed', 'times': ['07:00', '13:00', '21:00']},
        'Mergeflictamine': {'type': 'interval', 'hours': 4},
    }

    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    current = to_minutes(current_time)
    earliest = float('inf')
    next_med = ''

    for name, last_taken in medications:
        med = schedule[name]
        if med['type'] == 'fixed':
            future_times = [to_minutes(t) for t in med['times'] if to_minutes(t) > current]
            if not future_times:
                continue
            next_due = min(future_times)
        else:
            next_due = to_minutes(last_taken) + med['hours'] * 60

        diff = next_due - current
        if 0 < diff < earliest:
            earliest = diff
            next_med = name

    h = earliest // 60
    m = earliest % 60
    return f"{next_med} in {h}h {m}m"
```
