---
id: 69e2383af7832c8032603b92
title: "挑战 276：错误的元素"
challengeType: 29
dashedName: challenge-276
---

# --description--

给定一个按升序排序但有一个元素位置错误的整数数组，返回该元素的索引。

- 如果有多个元素可能被认为位置错误，返回第一个的索引。

# --hints--

`find_offender([1, 6, 2, 3, 4, 5])` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([1, 6, 2, 3, 4, 5]), 1)`)
}})
```

`find_offender([1, 2, 3, 5, 4, 5])` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([1, 2, 3, 5, 4, 5]), 3)`)
}})
```

`find_offender([2, 1])` 应该返回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([2, 1]), 0)`)
}})
```

`find_offender([2, 4, 1, 6, 8])` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([2, 4, 1, 6, 8]), 2)`)
}})
```

`find_offender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` 应该返回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_offender(arr):

    return arr
```

# --solutions--

```py
def find_offender(arr):
    for i in range(len(arr)):
        without = arr[:i] + arr[i+1:]
        if all(without[j-1] <= without[j] for j in range(1, len(without))):
            return i
```
