---
id: 69e2383af7832c8032603b8f
title: "Challenge 273: ISBN-13 Validator"
challengeType: 29
dashedName: challenge-273
---

# --description--

给定一个字串，判断它是否是有效的 ISBN-13 编号。

A valid ISBN-13:

- 只包含数字和连字符
- 去除连字符后恰好有 13 个数字
- 通过以下验证：
  1. 交替将每个数字乘以 1 或 3（第一个数字乘以 1，第二个乘以 3，第三个乘以 1，依此类推）。
  2. 结果的总和必须能被 10 整除。

# --hints--

`is_valid_isbn_13("9780306406157")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9780306406157"), True)`)
}})
```

`is_valid_isbn_13("97803064061570")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("97803064061570"), False)`)
}})
```

`is_valid_isbn_13("978-0-13-595705-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-0-13-595705-9"), True)`)
}})
```

`is_valid_isbn_13("978-030-64061A-4")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-030-64061A-4"), False)`)
}})
```

`is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn_13(s):

    return s
```

# --solutions--

```py
def is_valid_isbn_13(s):
    digits = s.replace("-", "")
    if not digits.isdigit() or len(digits) != 13:
        return False
    total = sum(int(d) * (1 if i % 2 == 0 else 3) for i, d in enumerate(digits))
    return total % 10 == 0
```
