---
id: 69f35a5bb823ed620fcb7cbc
title: "挑战 283：字串拉链"
challengeType: 29
dashedName: challenge-283
---

# --description--

给定两个字串，返回一个新字串，将它们的字符交错排列。若其中一个字串较长，则将剩余字符追加到末尾。

从第一个字串的第一个字符开始。

# --hints--

`zip_strings("abc", "123")` 应该返回 `"a1b2c3"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("abc", "123"), "a1b2c3")`)
}})
```

`zip_strings("acegikmoqsuwy", "bdfhjlnprtvxz")` 应该返回 `"abcdefghijklmnopqrstuvwxyz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz")`)
}})
```

`zip_strings("day", "night")` 应该返回 `"dnaiyght"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("day", "night"), "dnaiyght")`)
}})
```

`zip_strings("python", "javascript")` 应该返回 `"pjyatvhaosncript"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("python", "javascript"), "pjyatvhaosncript")`)
}})
```

`zip_strings("feCdCm", "reoeap")` 应该返回 `"freeCodeCamp"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("feCdCm", "reoeap"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def zip_strings(a, b):

    return a
```

# --solutions--

```py
def zip_strings(a, b):
    result = ""
    for i in range(max(len(a), len(b))):
        if i < len(a):
            result += a[i]
        if i < len(b):
            result += b[i]
    return result
```
