---
id: 69f35a5bb823ed620fcb7cbe
title: "挑战 285：会议时间"
challengeType: 29
dashedName: challenge-285
---

# --description--

给定一个表示多个人可用时间窗口的三维数组，返回所有人都有一小时空闲的最早时间。如果不存在这样的时间，返回 `"None"`。

- 每个人的可用时间是一个包含 `[start, end]` 个整数点对的数组，时间采用 24 小时制。例如，`[10, 12]` 表示该人在 10 点到 12 点之间有空。开始时间范围是 0-23，结束时间范围是 1-24。

例如，给定：

```json
[
  [[10, 12], [15, 16]], // person 1
  [[11, 14], [15, 16]]  // person 2
]
```

返回 `11`，即他们第一个共同空闲小时的开始时间。

# --hints--

`get_meeting_time([[[10, 12], [15, 16]], [[11, 14], [15, 16]]])` 应该返回 `11`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[10, 12], [15, 16]], [[11, 14], [15, 16]]]), 11)`)
}})
```

`get_meeting_time([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]])` 应该返回 `13`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]]), 13)`)
}})
```

`get_meeting_time([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]])` 应该返回 `9`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]]), 9)`)
}})
```

`get_meeting_time([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]])` 应该返回 `None`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]]), "None")`)
}})
```

`get_meeting_time([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]])` 应该返回 `21`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]]), 21)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_meeting_time(availability):

    return availability
```

# --solutions--

```py
def get_meeting_time(availability):
    for hour in range(24):
        if all(
            any(start <= hour and hour + 1 <= end for start, end in person)
            for person in availability
        ):
            return hour
    return "None"
```
