---
id: 69f35a5bb823ed620fcb7cb9
title: "挑战 280：Mongo ID 日期"
challengeType: 29
dashedName: challenge-280
---

# --description--

给定一个 MongoDB ID 字串，返回它的创建时间，格式为 ISO 8601 字串。

- MongoDB ID 是一个 24 字符的十六进制字串。前 8 个字符表现为一个 Unix 时间戳（以秒为单位），以 16 进制整数编码。

例如，`"6a094b50bcf86cd799439011"` 的时间戳是十六进制的 `"6a094b50"`，对应十进制的 `1778994000`，表现为创建时间 `"2026-05-17T05:00:00.000Z"`。

# --hints--

`mongo_id_to_date("6a094b50bcf86cd799439011")` 应该返回 `"2026-05-17T05:00:00.000Z"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z")`)
}})
```

`mongo_id_to_date("695344eb1f4a4c1123042128")` 应该返回 `"2025-12-30T03:20:11.000Z"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z")`)
}})
```

`mongo_id_to_date("386da62df34123ac54617e56")` 应该返回 `"2000-01-01T07:01:01.000Z"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z")`)
}})
```

`mongo_id_to_date("69f571c3d7711807afd3dd55")` 应该返回 `"2026-05-02T03:38:43.000Z"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z")`)
}})
```

`mongo_id_to_date("68adce01c0e1144d0a90295a")` 应该返回 `"2025-08-26T15:08:49.000Z"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z")`)
}})
```

# --seed--

## --seed-contents--

```py
def mongo_id_to_date(s):

    return s
```

# --solutions--

```py
from datetime import datetime, timezone

def mongo_id_to_date(s):
    timestamp = int(s[:8], 16)
    dt = datetime.fromtimestamp(timestamp, tz=timezone.utc)
    return dt.strftime("%Y-%m-%dT%H:%M:%S.000Z")
```
