---
id: 69f8c998d78ad3171a0713be
title: "Challenge 291: FizzBuzz Count"
challengeType: 29
dashedName: challenge-291
---

# --description--

给定一个起始数字和结束数字，计算在该范围内（包含端点）fizz 和 buzz 出现的次数。

- 能被 3 整除的数字计为 fizz。
- 能被 5 整除的数字计为 buzz。
- 能同时被 3 和 5 整除的数字既计为 fizz 也计为 buzz。

返回一个对象或字典，格式为 `{ fizz, buzz }`，其中包含计数结果。

# --hints--

`fizz_buzz_count(1, 11)` should return `{"fizz": 3, "buzz": 2}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(1, 11), {"fizz": 3, "buzz": 2})`)
}})
```

`fizz_buzz_count(14, 41)` should return `{"fizz": 9, "buzz": 6}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(14, 41), {"fizz": 9, "buzz": 6})`)
}})
```

`fizz_buzz_count(24, 100)` should return `{"fizz": 26, "buzz": 16}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(24, 100), {"fizz": 26, "buzz": 16})`)
}})
```

`fizz_buzz_count(-635, -14)` should return `{"fizz": 207, "buzz": 125}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-635, -14), {"fizz": 207, "buzz": 125})`)
}})
```

`fizz_buzz_count(-5432, 6789)` should return `{"fizz": 4074, "buzz": 2444}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-5432, 6789), {"fizz": 4074, "buzz": 2444})`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_count(start, end):

    return start
```

# --solutions--

```py
def fizz_buzz_count(start, end):
    fizz = sum(1 for i in range(start, end + 1) if i % 3 == 0)
    buzz = sum(1 for i in range(start, end + 1) if i % 5 == 0)
    return {"fizz": fizz, "buzz": buzz}
```
