---
id: 69f8c998d78ad3171a0713bd
title: "挑战 290：披萨派对"
challengeType: 29
dashedName: challenge-290
---

# --description--

给定一个数组，表示每个人今天工作的小时数，返回为披萨派对需要订购的披萨数量。

- 将每个人的工作小时数除以 3，得到他们应得的披萨片数。
- 你不能吃半片披萨，所以将每个人的披萨片数向上取整到最接近的整数。
- 每个人至少得到两片披萨。
- 每个披萨有 8 片。将披萨总数向上取整到最接近的整张披萨。

# --hints--

`get_pizzas_to_order([8, 8, 8])` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8]), 2)`)
}})
```

`get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10])` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10]), 3)`)
}})
```

`get_pizzas_to_order([1, 2, 3, 4, 5])` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([1, 2, 3, 4, 5]), 2)`)
}})
```

`get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8])` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8]), 3)`)
}})
```

`get_pizzas_to_order([9, 9, 6])` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([9, 9, 6]), 1)`)
}})
```

`get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0])` 应该返回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_pizzas_to_order(hours_worked):

    return hours_worked
```

# --solutions--

```py
import math

def get_pizzas_to_order(hours_worked):
    total_slices = sum(max(math.ceil(h / 3), 2) for h in hours_worked)
    return math.ceil(total_slices / 8)
```
