---
id: 69f8c998d78ad3171a0713bb
title: "Challenge 288: Secret Number"
challengeType: 29
dashedName: challenge-288
---

# --description--

给定一个秘密数字和一个猜测，判断猜测是否正确。

Return:

- `"higher"` 如果秘密数字比猜测大。
- `"lower"` 如果秘密数字比猜测小。
- `"you got it!"` 如果猜测正确。

# --hints--

`guess_number(50, 30)` should return `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(50, 30), "higher")`)
}})
```

`guess_number(85, 99)` should return `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(85, 99), "lower")`)
}})
```

`guess_number(2026, 2026)` should return `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(2026, 2026), "you got it!")`)
}})
```

`guess_number(92904, 11283)` should return `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(92904, 11283), "higher")`)
}})
```

`guess_number(230495, 423920)` should return `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(230495, 423920), "lower")`)
}})
```

`guess_number(120349, 120349)` should return `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(120349, 120349), "you got it!")`)
}})
```

# --seed--

## --seed-contents--

```py
def guess_number(secret, guess):

    return secret
```

# --solutions--

```py
def guess_number(secret, guess):
    if guess == secret:
        return "you got it!"
    return "higher" if guess < secret else "lower"
```
