---
id: 69f90c1329a94b37e2a2086c
title: "挑战 293：最佳牌型"
challengeType: 29
dashedName: challenge-293
---

# --description--

给定一个包含五张扑克牌的字串数组，返回最佳牌型的名称。

- 每张牌用一个两个字符的字串表现：先是点数，后是花色，`"2h"` 例如。
  - 点数从低到高依次是：`"2"`、`"3"`、`"4"`、`"5"`、`"6"`、`"7"`、`"8"`、`"9"`、`"T"`、`"J"`、`"Q"`、`"K"` 和 `"A"`。
  - 花色有：`"h"`、`"d"`、`"c"` 和 `"s"`。
- A（`"A"`）在顺子中可以作为最大点数或最小点数使用。

牌型按从差到好排序如下：

| 名称 | 描述 |
| - | - |
| `"High Card"` | 无对子或更好牌型 |
| `"Pair"` | 一对 |
| `"Two Pair"` | 两对 |
| `"Three of a Kind"` | 三条 |
| `"Straight"` | 顺子（五张连续点数） |
| `"Flush"` | 同花（五张同花色） |
| `"Full House"` | 葫芦（三条加一对） |
| `"Four of a Kind"` | 铁支（四条） |
| `"Straight Flush"` | 同花顺（五张同花色的顺子） |
| `"Royal Flush"` | 皇家同花顺（`"A"`、`"K"`、`"Q"`、`"J"`、`"T"` 同花色） |

返回最佳牌型的名称。

# --hints--

`get_best_hand(["7s", "7h", "7d", "2c", "5h"])` 应该返回 `"Three of a Kind"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind")`)
}})
```

`get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"])` 应该返回 `"Full House"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House")`)
}})
```

`get_best_hand(["2h", "5h", "7h", "9h", "Jh"])` 应该返回 `"Flush"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["2h", "5h", "7h", "9h", "Jh"]), "Flush")`)
}})
```

`get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"])` 应该返回 `"Four of a Kind"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind")`)
}})
```

`get_best_hand(["Ts", "Th", "9d", "9c", "8h"])` 应该返回 `"Two Pair"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair")`)
}})
```

`get_best_hand(["9c", "8c", "7c", "6c", "5c"])` 应该返回 `"Straight Flush"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush")`)
}})
```

`get_best_hand(["As", "Kh", "Jd", "8c", "5h"])` 应该返回 `"High Card"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Kh", "Jd", "8c", "5h"]), "High Card")`)
}})
```

`get_best_hand(["As", "2h", "3d", "4c", "5h"])` 应该返回 `"Straight"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "2h", "3d", "4c", "5h"]), "Straight")`)
}})
```

`get_best_hand(["Ts", "Th", "7c", "6d", "5h"])` 应该返回 `"Pair"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "7c", "6d", "5h"]), "Pair")`)
}})
```

`get_best_hand(["As", "Ks", "Qs", "Js", "Ts"])` 应该返回 `"Royal Flush"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_best_hand(cards):

    return cards
```

# --solutions--

```py
def get_best_hand(cards):
    rank_order = "23456789TJQKA"

    ranks = [c[0] for c in cards]
    suits = [c[1] for c in cards]

    is_flush = len(set(suits)) == 1
    rank_indices = sorted(rank_order.index(r) for r in ranks)

    def is_straight(indices):
        is_regular = all(indices[i] == indices[i - 1] + 1 for i in range(1, len(indices)))
        is_low_ace = indices == [0, 1, 2, 3, 12]
        return is_regular or is_low_ace

    straight = is_straight(rank_indices)

    counts = {}
    for r in ranks:
        counts[r] = counts.get(r, 0) + 1
    freq = sorted(counts.values(), reverse=True)

    if is_flush and rank_indices == [8, 9, 10, 11, 12]:
        return "Royal Flush"
    if is_flush and straight:
        return "Straight Flush"
    if freq[0] == 4:
        return "Four of a Kind"
    if freq[0] == 3 and freq[1] == 2:
        return "Full House"
    if is_flush:
        return "Flush"
    if straight:
        return "Straight"
    if freq[0] == 3:
        return "Three of a Kind"
    if freq[0] == 2 and freq[1] == 2:
        return "Two Pair"
    if freq[0] == 2:
        return "Pair"
    return "High Card"
```
