---
id: 6a0dcd03ee4e68698080ef6b
title: "挑战 307：分区规定"
challengeType: 29
dashedName: challenge-307
---

# --description--

给定一个二维网格（数组的数组），表现一个城市的建筑版型，返回所有违反分区规定的建筑的坐标。

网格中的每个格子包含下表中的标签之一。如果建筑的（最多）4 个邻居（水平或垂直方向）中有任何一种类型是它不能相邻的，则该建筑违反规定。

| 标签 | 类型 | 不能相邻的类型 |
| - | - | - |
| `"i"` | 工业 | `"R"`, `"I"` |
| `"A"` | 农业 | `"C"` |
| `"R"` | 住宅 | `"i"`, `"C"` |
| `"I"` | 机构 | `"i"` |
| `"C"` | 商业 | `"R"`, `"A"` |
| `""`（空字串） | 未开发 | 无限制 |

将所有违规格子的坐标作为 `[row, col]` 点对数组返回，顺序不限。如果没有违规，返回空数组。

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` 应该返回 `[[0, 0], [0, 1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` 应该返回 `[[0, 1], [1, 1]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` 应该返回 `[]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` 应该返回 `[[0, 1], [0, 2], [0, 3]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` 应该返回 `[[2, 3], [2, 4], [3, 1], [3, 2]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
