---
id: 6a0dcc730cb92a616f86f0bf
title: "挑战 295：模式验证器 第 1 部分"
challengeType: 29
dashedName: challenge-295
---

# --description--

给定一个对象（JavaScript）或字典（Python），判断它是否符合以下模式：

```json
{
  username: string
}
```

- 允许存在额外的密钥

# --hints--

`is_valid_schema({"username": "bob"})` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bob"}), True)`)
}})
```

`is_valid_schema({"username": "jen", "posts": 30})` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "jen", "posts": 30}), True)`)
}})
```

`is_valid_schema({"username": ""})` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": ""}), True)`)
}})
```

`is_valid_schema({"username": 7})` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 7}), False)`)
}})
```

`is_valid_schema({"posts": 25})` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"posts": 25}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return isinstance(obj.get("username"), str)
```
