---
id: 6a0dcd03ee4e68698080ef66
title: "挑战 302：时差反应"
challengeType: 29
dashedName: challenge-302
---

# --description--

给定出发城市、到达城市、飞行时长（小时）和旅行方向，返回旅客经历的时差小时数。

给定的城市将来自以下包含其 UTC 偏移量的列表：

| 城市 | 偏移量 |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

计算时差小时数的方法：

1. 查找两个城市之间的时区差（小时）。
2. 确定方向乘数。如果是 `"east"` 方向，乘数为 1.5，否则为 1.0。
3. 使用公式计算时差小时数：时区差 +（飞行时长 * 0.1）* 方向乘数

返回四舍五入到一位小数的时差小时数。

# --hints--

`get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east")` 应该返回 `6.5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east"), 6.5)`)
}})
```

`get_jet_lag_hours("London", "New York", 8, "west")` 应该返回 `5.8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("London", "New York", 8, "west"), 5.8)`)
}})
```

`get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east")` 应该返回 `1.6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east"), 1.6)`)
}})
```

`get_jet_lag_hours("Dubai", "London", 7, "west")` 应该返回 `4.7`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Dubai", "London", 7, "west"), 4.7)`)
}})
```

`get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west")` 应该返回 `17.5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west"), 17.5)`)
}})
```

`get_jet_lag_hours("Tokyo", "Dubai", 9, "west")` 应该返回 `5.9`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Tokyo", "Dubai", 9, "west"), 5.9)`)
}})
```

`get_jet_lag_hours("New York", "Istanbul", 10, "east")` 应该返回 `9.5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("New York", "Istanbul", 10, "east"), 9.5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):

    return departure_city
```

# --solutions--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):
    offsets = {
        "Los Angeles": -8,
        "New York": -5,
        "London": 0,
        "Istanbul": 3,
        "Dubai": 4,
        "Hong Kong": 8,
        "Tokyo": 9
    }
    timezone_diff = abs(offsets[arrival_city] - offsets[departure_city])
    multiplier = 1.5 if direction == "east" else 1.0
    score = timezone_diff + (flight_duration * 0.1) * multiplier
    return round(score, 1)
```
