---
id: 6a0dcd03ee4e68698080ef68
title: "挑战 304：行程安排"
challengeType: 29
dashedName: challenge-304
---

# --description--

给定一个包含至少两个可选停靠点的日间旅行数组，返回有效行程安排的数量。

行程总是包括 `"breakfast"`、`"lunch"` 和 `"dinner"`，这些不会作为参数传入。可选停靠点可以放在行程中的任何位置，但必须遵守以下规则：

- `"breakfast"` 总是第一个，且在 `"lunch"` 之前至少有一个停靠点。
- `"lunch"` 必须出现在 `"dinner"` 之前，且两者之间至少有一个停靠点。
- `"dinner"` 之后最多只能有一个可选停靠点。

返回有效安排的数量。

# --hints--

`get_itinerary_count(["library", "park"])` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` 应该返回 `18`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` 应该返回 `120`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` 应该返回 `840`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` 应该返回 `55440`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
