---
id: 6a15cadf5f240d05a264955a
title: "挑战 316：1337 语言"
challengeType: 29
dashedName: challenge-316
---

# --description--

给定一个小写字串，通过将下面的字母替换为它们的 1337 代换，返回转换后的 1337 语言字串：

| 字母 | 1337 |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- 没有代换的字符保持不变。

# --hints--

`make_leet("cool")` 应该返回 `"c001"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("cool"), "c001")`)
}})
```

`make_leet("leet")` 应该返回 `"1337"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("leet"), "1337")`)
}})
```

`make_leet("hacker")` 应该返回 `"h4ck3r"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("hacker"), "h4ck3r")`)
}})
```

`make_leet("satellite")` 应该返回 `"547311173"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("satellite"), "547311173")`)
}})
```

`make_leet("abcdefghijklmnopqrstuvwxyz")` 应该返回 `"4bcd3f9h1jk1mn0pqr57uvwxyz"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz")`)
}})
```

# --seed--

## --seed-contents--

```py
def make_leet(s):

    return s
```

# --solutions--

```py
def make_leet(s):
    leet_map = {"a": "4", "e": "3", "g": "9", "i": "1", "l": "1", "o": "0", "t": "7", "s": "5"}
    return "".join(leet_map.get(c, c) for c in s)
```
