---
id: 6a151d32d772271dd2448c2d
title: "挑战 311：咒语施放者"
challengeType: 29
dashedName: challenge-311
---

# --description--

给定一个你正在施放的咒语代码字串，计算总分。

字串中的每个字符代表一个咒语：

| 代码 | 咒语 | 类别 | 基础分数 |
| - | - | - | - |
| `"f"` | 火焰 | 破坏 | 3 |
| `"l"` | 闪电 | 破坏 | 3 |
| `"i"` | 冰霜 | 控件 | 2 |
| `"w"` | 风 | 控件 | 2 |
| `"h"` | 治疗 | 恢复 | 1 |
| `"s"` | 盾牌 | 恢复 | 1 |

根据连续施放的咒语类别不同，应用连击倍数：

- 第一个咒语总是按基础分数计分。
- 每当连续施放的咒语类别与前一个不同，连击倍数加 1。
- 施放与前一个咒语类别相同的咒语时，连击倍数重置为 1。
- 每个咒语的得分是其基础分数乘以当前连击倍数。

返回这串咒语的总分。

# --hints--

`cast("fihwl")` 应该返回 `33`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` 应该返回 `45`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` 应该返回 `37`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` 应该返回 `50`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` 应该返回 `101`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
