---
id: 6a15cadf5f240d05a264955b
title: "挑战 317：BMI 计算器"
challengeType: 29
dashedName: challenge-317
---

# --description--

给定以磅为单位的体重和以英寸为单位的身高，返回四舍五入到一位小数的 BMI（体质指数）。

计算 BMI：用体重除以身高的平方，然后将结果乘以 703。

# --hints--

`calculate_bmi(180, 70)` 应该返回 `25.8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(180, 70), 25.8)`)
}})
```

`calculate_bmi(140, 64)` 应该返回 `24.0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(140, 64), 24.0)`)
}})
```

`calculate_bmi(160, 76)` 应该返回 `19.5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(160, 76), 19.5)`)
}})
```

`calculate_bmi(200, 60)` 应该返回 `39.1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(200, 60), 39.1)`)
}})
```

`calculate_bmi(150, 68)` 应该返回 `22.8`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(150, 68), 22.8)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_bmi(weight, height):

    return weight
```

# --solutions--

```py
def calculate_bmi(weight, height):
    return round((weight / (height * height)) * 703, 1)
```
