---
id: 6a15cadf5f240d05a2649557
title: "挑战 313：租赁费用"
challengeType: 29
dashedName: challenge-313
---

# --description--

给定一个租赁时间戳、一个归还时间戳和一个租赁等级，返回租赁的总费用，包括任何逾期费用。

- 给定的时间戳是 UTC ISO 字串，例如：`"2026-06-18T18:30:00Z"`。
- 租赁等级是租赁到期归还前的天数：`1`、`3` 或 `7`。
- 租赁必须在租赁期最后一天的 UTC 时间中午 12:00 或更早归还。例如，3 月 15 日任何时间租出的 1 天租赁，必须在 3 月 16 日 UTC 时间中午 12:00 前归还。
- 每超过到期日期和时间一天，将产生逾期费用。

价格如下：

| 等级 | 基础费用 | 每天逾期费用 |
| - | - | - |
| 1 天 | $4.99 | $3.99 |
| 3 天 | $3.99 | $2.99 |
| 7 天 | $2.99 | $0.99 |

返回总费用，四舍五入到小数点后两位，格式为 `"$D.CC"`。

# --hints--

`get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` 应该返回 `"$4.99"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99")`)
}})
```

`get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` 应该返回 `"$12.97"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97")`)
}})
```

`get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` 应该返回 `"$3.99"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99")`)
}})
```

`get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` 应该返回 `"$9.97"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97")`)
}})
```

`get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` 应该返回 `"$2.99"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99")`)
}})
```

`get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` 应该返回 `"$358.40"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rental_cost(rented, returned, tier):

    return rented
```

# --solutions--

```py
from datetime import datetime, timedelta
import math

def get_rental_cost(rented, returned, tier):
    pricing = {
        1: {"base": 4.99, "late": 3.99},
        3: {"base": 3.99, "late": 2.99},
        7: {"base": 2.99, "late": 0.99},
    }
    rented_dt = datetime.fromisoformat(rented.replace("Z", "+00:00"))
    returned_dt = datetime.fromisoformat(returned.replace("Z", "+00:00"))
    due = rented_dt + timedelta(days=tier)
    due = due.replace(hour=12, minute=0, second=0, microsecond=0)
    base = pricing[tier]["base"]
    late = pricing[tier]["late"]
    if returned_dt <= due:
        return f"${base:.2f}"
    days_late = math.ceil((returned_dt - due).total_seconds() / 86400)
    return f"${base + days_late * late:.2f}"
```
