---
id: 6a151d32d772271dd2448c2c
title: "挑战 310：英式英语转美式英语"
challengeType: 29
dashedName: challenge-310
---

# --description--

给定一个句子，使用下表将任何英式英语拼写转换为美式英语对应拼写，并返回更新后的句子：

| 英式 | 美式 |
| - | - |
| `"colour"` | `"color"` |
| `"flavour"` | `"flavor"` |
| `"honour"` | `"honor"` |
| `"neighbour"` | `"neighbor"` |
| `"labour"` | `"labor"` |
| `"humour"` | `"humor"` |
| `"centre"` | `"center"` |
| `"fibre"` | `"fiber"` |
| `"defence"` | `"defense"` |
| `"offence"` | `"offense"` |
| `"organise"` | `"organize"` |
| `"recognise"` | `"recognize"` |
| `"analyse"` | `"analyze"` |

- 替换应不区分大小写。例如，`"Colour"` 应该变成 `"Color"`。
- 输入中可能包含基于表中根词拼写构建的单词，这些单词也需要被更改。例如，`"colouring"` 应该变成 `"coloring"`，`"disorganised"` 应该变成 `"disorganized"`。

# --hints--

`british_to_american("I love the colour blue.")` 应该返回 `"I love the color blue."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("I love the colour blue."), "I love the color blue.")`)
}})
```

`british_to_american("The fibre optic cable is new.")` 应该返回 `"The fiber optic cable is new."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("The fibre optic cable is new."), "The fiber optic cable is new.")`)
}})
```

`british_to_american("It's an honour to meet someone with such humour.")` 应该返回 `"It's an honor to meet someone with such humor."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("It's an honour to meet someone with such humour."), "It's an honor to meet someone with such humor.")`)
}})
```

`british_to_american("The unrecognised artist analysed his colour palette at the centre.")` 应该返回 `"The unrecognized artist analyzed his color palette at the center."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("The unrecognised artist analysed his colour palette at the centre."), "The unrecognized artist analyzed his color palette at the center.")`)
}})
```

`british_to_american("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful.")` 应该返回 `The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful."), "The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.")`)
}})
```

# --seed--

## --seed-contents--

```py
def british_to_american(sentence):

    return sentence
```

# --solutions--

```py
import re

def british_to_american(sentence):
    lookup = {
        "colour": "color",
        "flavour": "flavor",
        "honour": "honor",
        "neighbour": "neighbor",
        "labour": "labor",
        "humour": "humor",
        "centre": "center",
        "fibre": "fiber",
        "defence": "defense",
        "offence": "offense",
        "organise": "organize",
        "recognise": "recognize",
        "analyse": "analyze"
    }

    result = sentence
    for british, american in lookup.items():
        def replace(match, american=american):
            return american[0].upper() + american[1:] if match.group()[0].isupper() else american
        result = re.sub(british, replace, result, flags=re.IGNORECASE)

    return result
```
