---
id: 6a15cadf5f240d05a2649559
title: "挑战 315：夏至"
challengeType: 29
dashedName: challenge-315
---

# --description--

今天是夏至，北半球一年中白昼最长的一天，南半球则是最短的一天。给定一个纬度，返回一个表现白昼和夜晚小时数的字串。

- 纬度将在 90（北极）和 -90（南极）之间，包括两端
- 白昼小时数 = 12 + (纬度 / 90) * 12
- 将白昼小时数四舍五入到最接近的偶数

返回一个 24 字符的字串，使用 `"☀️"` 表示白昼小时，使用 `"🌑"` 表示夜晚小时，其中：

- 每个字符代表一小时，从午夜（第 0 小时）开始
- 日出和日落对称分布在正午两侧

例如，纬度为 `0`（赤道）时白昼为 12 小时，所以日出时间是上午 6:00，日落时间是下午 6:00。返回：`"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`。

# --hints--

`get_daytime_hours(0)` 应该返回 `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(0), "🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(90)` 应该返回 `"☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(90), "☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️")`)
}})
```

`get_daytime_hours(-90)` 应该返回 `"🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-90), "🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(-33)` 应该返回 `"🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-33), "🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(66.5)` 应该返回 `"🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(66.5), "🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑")`)
}})
```

`get_daytime_hours(40)` 应该返回 `"🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(40), "🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑")`)
}})
```

`get_daytime_hours(-50)` 应该返回 `"🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-50), "🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_daytime_hours(latitude):

    return latitude
```

# --solutions--

```py
def get_daytime_hours(latitude):
    daylight_hours = round((12 + (latitude / 90) * 12) / 2) * 2
    sunrise = 12 - daylight_hours // 2
    sunset = 12 + daylight_hours // 2
    result = ""
    for hour in range(24):
        result += "☀️" if sunrise <= hour < sunset else "🌑"
    return result
```
