---
id: 6a15cadf5f240d05a264955d
title: "挑战 319：Frontmatter 解析器"
challengeType: 29
dashedName: challenge-319
---

# --description--

给定一个表示 frontmatter 块的字串，解析它并返回一个包含密钥和值的对象（JavaScript）或字典（Python）。

Frontmatter 用 `---` 分隔符包裹，里面包含每行一个的 `key: value` 点对。例如：

```md
---
title: My Post
draft: false
views: 100
---
```

应返回：

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- 数字、布尔和字串都应以它们各自的类型返回。
- 给定的字串中的换行符用换行字符（`"\n"`）分隔。上面的例子会给出为：`"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`。

# --hints--

`parse_frontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` 应该返回 `{ title: "My Post", draft: False, views: 100 }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\ntitle: My Post\\ndraft: false\\nviews: 100\\n---"), {"title": "My Post", "draft": False, "views": 100})`)
}})
```

`parse_frontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` 应该返回 `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: False }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nid: 6a174db57256a112f932195c\\ntitle: My Book\\nlocale: en\\nwordCount: 10000\\npublished: false\\n---"), {"id": "6a174db57256a112f932195c", "title": "My Book", "locale": "en", "wordCount": 10000, "published": False})`)
}})
```

`parse_frontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` 应该返回 `{ version: "1.0.0", url: "https://example.com", private: True }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nversion: 1.0.0\\nurl: https://example.com\\nprivate: true\\n---"), {"version": "1.0.0", "url": "https://example.com", "private": True})`)
}})
```

`parse_frontmatter("---\nrating: 4.5\nprice: 9.99\n---")` 应该返回 `{ rating: 4.5, price: 9.99 }`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nrating: 4.5\\nprice: 9.99\\n---"), {"rating": 4.5, "price": 9.99})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_frontmatter(s):

    return s
```

# --solutions--

```py
def parse_frontmatter(s):
    lines = [line for line in s.split("\n") if line != "---" and line != ""]
    result = {}
    for line in lines:
        key, *rest = line.split(":")
        value = ":".join(rest).strip()
        if value == "true":
            result[key.strip()] = True
        elif value == "false":
            result[key.strip()] = False
        else:
            try:
                n = float(value)
                result[key.strip()] = int(n) if n == int(n) else n
            except ValueError:
                result[key.strip()] = value
    return result
```
