---
id: 6a15cadf5f240d05a264955c
title: "挑战 318：DNA 突变"
challengeType: 29
dashedName: challenge-318
---

# --description--

给定两个长度相等的 DNA 链，返回一个数组，包含链中不同（突变）的位置索引。

- DNA 链是由字符 `"A"`、`"T"`、`"C"` 和 `"G"` 组成的字串
- 返回的索引按升序排列
- 如果没有突变，返回空数组

# --hints--

`detect_mutations("ATCG", "ATGG")` 应该返回 `[2]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ATCG", "ATGG"), [2])`)
}})
```

`detect_mutations("ATGCGTACGTTAGC", "ATGCATACGATTGC")` 应该返回 `[4, 9, 11]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ATGCGTACGTTAGC", "ATGCATACGATTGC"), [4, 9, 11])`)
}})
```

`detect_mutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG")` 应该返回 `[]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG"), [])`)
}})
```

`detect_mutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG")` 应该返回 `[15, 16, 17, 18]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG"), [15, 16, 17, 18])`)
}})
```

`detect_mutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT")` 应该返回 `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28])`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_mutations(strand1, strand2):

    return strand1
```

# --solutions--

```py
def detect_mutations(strand1, strand2):
    return [i for i, (a, b) in enumerate(zip(strand1, strand2)) if a != b]
```
