---
id: 6a19b1062d1b153d8ac76d73
title: "挑战 325：幸运数字"
challengeType: 29
dashedName: challenge-325
---

# --description--

给定一个包含某人名字和姓氏的字串，使用以下规则计算他们的幸运数字：

- 名字和姓氏之间用空格分隔
- 查找每个名字中的元音和构造数量
- 将较小的元音和构造数量相乘，再乘以较短名字的长度
- 对较大的两个数量和较长的名字做同样的操作
- 用较大的值减去较小的值，得到他们的幸运数字

如果最终值是零（`0`），则返回 `13`。

# --hints--

`get_lucky_number("John Doe")` 应该返回 `21`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("John Doe"), 21)`)
}})
```

`get_lucky_number("Olivia Lewis")` 应该返回 `52`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Olivia Lewis"), 52)`)
}})
```

`get_lucky_number("James Wilson")` 应该返回 `18`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("James Wilson"), 18)`)
}})
```

`get_lucky_number("Elizabeth Hernandez")` 应该返回 `81`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Elizabeth Hernandez"), 81)`)
}})
```

`get_lucky_number("Mike Walker")` 应该返回 `32`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Mike Walker"), 32)`)
}})
```

`get_lucky_number("Chloe Perez")` 应该返回 `13`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Chloe Perez"), 13)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_lucky_number(name):

    return name
```

# --solutions--

```py
def get_lucky_number(name):
    vowels = "aeiouAEIOU"

    def name_stats(s):
        vowel_count = sum(1 for c in s if c in vowels)
        consonant_count = sum(1 for c in s if c.isalpha() and c not in vowels)
        return vowel_count, consonant_count, len(s)

    first, last = name.split(" ")
    first_vowels, first_consonants, first_len = name_stats(first)
    last_vowels, last_consonants, last_len = name_stats(last)

    larger_vowels = max(first_vowels, last_vowels)
    smaller_vowels = min(first_vowels, last_vowels)
    larger_consonants = max(first_consonants, last_consonants)
    smaller_consonants = min(first_consonants, last_consonants)
    larger_length = max(first_len, last_len)
    smaller_length = min(first_len, last_len)

    larger_value = larger_vowels * larger_consonants * larger_length
    smaller_value = smaller_vowels * smaller_consonants * smaller_length
    result = larger_value - smaller_value

    return 13 if result == 0 else result
```
