---
id: 6a19b1062d1b153d8ac76d72
title: "挑战 324：副本字符计数"
challengeType: 29
dashedName: challenge-324
---

# --description--

给定两个字串，返回第二个字串中可以在第一个字串中找到的字符数量。

- 第二个字串中的副本字符分别计数。

# --hints--

`duplicate_character_count("aloha", "hei")` 应该返回 `1`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("aloha", "hei"), 1)`)
}})
```

`duplicate_character_count("jambo", "bonjour")` 应该返回 `4`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("jambo", "bonjour"), 4)`)
}})
```

`duplicate_character_count("hello", "hola")` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello", "hola"), 3)`)
}})
```

`duplicate_character_count("ola", "hej")` 应该返回 `0`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ola", "hej"), 0)`)
}})
```

`duplicate_character_count("ciao", "konnichiwa")` 应该返回 `5`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ciao", "konnichiwa"), 5)`)
}})
```

`duplicate_character_count("merhaba", "xin chao")` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("merhaba", "xin chao"), 2)`)
}})
```

`duplicate_character_count("hello world", "hello to everyone around the world")` 应该返回 `26`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello world", "hello to everyone around the world"), 26)`)
}})
```

# --seed--

## --seed-contents--

```py
def duplicate_character_count(str1, str2):

    return str1
```

# --solutions--

```py
def duplicate_character_count(str1, str2):
    return sum(1 for char in str2 if char in str1)
```
