---
id: 6a1d9f98e819ed70a0e994db
title: "挑战 329：桶填充"
challengeType: 29
dashedName: challenge-329
---

# --description--

给定一个二维网格、一个起始位置（`[row, col]`）和一个新值，将起始位置及所有相同值且相连的单元格的值替换为新值。

- 如果单元格在水平方向或垂直方向相邻（不包括对角线方向），则它们是相连的。

返回更新后的网格。

# --hints--

`bucket_fill([["R", "G"], ["R", "G"]], [0, 1], "B")` 应该返回 `[["R", "B"], ["R", "B"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["R", "G"], ["R", "G"]], [0, 1], "B"), [["R", "B"], ["R", "B"]])`)
}})
```

`bucket_fill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B")` 应该返回 `[["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B"), [["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]])`)
}})
```

`bucket_fill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R")` 应该返回 `[["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R"), [["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]])`)
}})
```

`bucket_fill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y")` 应该返回 `[["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y"), [["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]])`)
}})
```

`bucket_fill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G")` 应该返回 `[["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G"), [["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def bucket_fill(grid, pos, new_value):

    return grid
```

# --solutions--

```py
def bucket_fill(grid, pos, new_value):
    row, col = pos
    target = grid[row][col]
    if target == new_value:
        return grid

    def fill(r, c):
        if r < 0 or r >= len(grid):
            return
        if c < 0 or c >= len(grid[0]):
            return
        if grid[r][c] != target:
            return
        grid[r][c] = new_value
        fill(r - 1, c)
        fill(r + 1, c)
        fill(r, c - 1)
        fill(r, c + 1)

    fill(row, col)
    return grid
```
