---
id: 6a1d9f98e819ed70a0e994e2
title: "挑战 336：星座匹配"
challengeType: 29
dashedName: challenge-336
---

# --description--

给定两个星座字串，返回它们的兼容性百分比。

星座按以下顺序排列在一个 12 个位置的轮盘上：`"Aries"`、`"Taurus"`、`"Gemini"`、`"Cancer"`、`"Leo"`、`"Virgo"`、`"Libra"`、`"Scorpio"`、`"Sagittarius"`、`"Capricorn"`、`"Aquarius"`、`"Pisces"`，在`"Pisces"`之后回绕到`"Aries"`。查找两个星座之间的最短距离并返回兼容性：

| 距离 | 兼容性 |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscope_match("Libra", "Sagittarius")` 应该返回 `"80%"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Libra", "Sagittarius"), "80%")`)
}})
```

`horoscope_match("Gemini", "Scorpio")` 应该返回 `"20%"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Gemini", "Scorpio"), "20%")`)
}})
```

`horoscope_match("Pisces", "Aries")` 应该返回 `"40%"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Pisces", "Aries"), "40%")`)
}})
```

`horoscope_match("Capricorn", "Cancer")` 应该返回 `"50%"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Capricorn", "Cancer"), "50%")`)
}})
```

`horoscope_match("Aquarius", "Aquarius")` 应该返回 `"100%"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Aquarius", "Aquarius"), "100%")`)
}})
```

`horoscope_match("Virgo", "Taurus")` 应该返回 `"90%"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Virgo", "Taurus"), "90%")`)
}})
```

`horoscope_match("Leo", "Scorpio")` 应该返回 `"30%"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Leo", "Scorpio"), "30%")`)
}})
```

# --seed--

## --seed-contents--

```py
def horoscope_match(sign1, sign2):

    return sign1
```

# --solutions--

```py
def horoscope_match(sign1, sign2):
    signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"]
    a = signs.index(sign1)
    b = signs.index(sign2)
    diff = abs(a - b)
    distance = min(diff, 12 - diff)
    compatibility = {0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%"}
    return compatibility[distance]
```
