---
id: 6a1d9f98e819ed70a0e994dd
title: "挑战 331：最接近的倍数"
challengeType: 29
dashedName: challenge-331
---

# --description--

给定两个整数，将第一个整数四舍五入到最接近的第二个整数的倍数。

# --hints--

`round_to_nearest_multiple(5, 3)` 应该返回 `6`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(5, 3), 6)`)
}})
```

`round_to_nearest_multiple(17, 4)` 应该返回 `16`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(17, 4), 16)`)
}})
```

`round_to_nearest_multiple(43, 5)` 应该返回 `45`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(43, 5), 45)`)
}})
```

`round_to_nearest_multiple(38, 11)` 应该返回 `33`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(38, 11), 33)`)
}})
```

`round_to_nearest_multiple(93, 12)` 应该返回 `96`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(93, 12), 96)`)
}})
```

# --seed--

## --seed-contents--

```py
def round_to_nearest_multiple(num, multiple):

    return num
```

# --solutions--

```py
def round_to_nearest_multiple(num, multiple):
    return round(num / multiple) * multiple
```
