---
id: 6a1d9f98e819ed70a0e994df
title: "挑战 333：问题分类 2"
challengeType: 29
dashedName: challenge-333
---

# --description--

给定一个问题标题和当前标签的数组，根据以下规则返回更新后的标签数组：

如果该问题没有任何标签，添加：

- 如果标题包含 `"error"` 或 `"bug"`，添加 `"bug"` 和 `"needs triage"`
- 如果标题包含 `"feature"` 或 `"add"`，添加 `"enhancement"` 和 `"discussing"`

否则，如果给定的标签包含：

- 如果包含 `"needs triage"` 且标题包含 `"simple"` 或 `"easy"`，移除 `"needs triage"` 并添加 `"good first issue"`
- 如果包含 `"discussing"` 且标题包含 `"planned"` 或 `"next"`，移除 `"discussing"` 并添加 `"on the roadmap"`
- 否则，如果存在 `"needs triage"` 或 `"discussing"`，移除它并添加 `"help wanted"`

如果标题包含：

- `"security"`，添加一个 `"critical"` 标签

# --hints--

`triage_issue("app crashes with error", [])` 应该返回 `["bug", "needs triage"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", [])), sorted(["bug", "needs triage"]))`)
}})
```

`triage_issue("app crashes with error", ["bug", "needs triage"])` 应该返回 `["bug", "help wanted"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", ["bug", "needs triage"])), sorted(["bug", "help wanted"]))`)
}})
```

`triage_issue("add dark mode", [])` 应该返回 `["enhancement", "discussing"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", [])), sorted(["enhancement", "discussing"]))`)
}})
```

`triage_issue("add dark mode", ["enhancement", "discussing"])` 应该返回 `["enhancement", "help wanted"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted"]))`)
}})
```

`triage_issue("xss security bug", [])` 应该返回 `["bug", "needs triage", "critical"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("xss security bug", [])), sorted(["bug", "needs triage", "critical"]))`)
}})
```

`triage_issue("security vulnerability in auth", [])` 应该返回 `["critical"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("security vulnerability in auth", [])), sorted(["critical"]))`)
}})
```

`triage_issue("easy a11y fix", ["bug", "needs triage"])` 应该返回 `["bug", "good first issue"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("easy a11y fix", ["bug", "needs triage"])), sorted(["bug", "good first issue"]))`)
}})
```

`triage_issue("planned api migration", ["enhancement", "discussing"])` 应该返回 `["enhancement", "on the roadmap"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("planned api migration", ["enhancement", "discussing"])), sorted(["enhancement", "on the roadmap"]))`)
}})
```

`triage_issue("improve security", ["enhancement", "discussing"])` 应该返回 `["enhancement", "help wanted", "critical"]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("improve security", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted", "critical"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(title, labels):

    return title
```

# --solutions--

```py
def triage_issue(title, labels):
    result = list(labels)
    t = title.lower()

    def has(label):
        return label in result

    def add(label):
        if not has(label):
            result.append(label)

    def remove(label):
        if label in result:
            result.remove(label)

    if len(result) == 0:
        if "error" in t or "bug" in t:
            add("bug")
            add("needs triage")
        elif "feature" in t or "add" in t:
            add("enhancement")
            add("discussing")
    else:
        if has("needs triage") and ("simple" in t or "easy" in t):
            remove("needs triage")
            add("good first issue")
        elif has("discussing") and ("planned" in t or "next" in t):
            remove("discussing")
            add("on the roadmap")
        elif has("needs triage") or has("discussing"):
            remove("needs triage")
            remove("discussing")
            add("help wanted")

    if "security" in t:
        add("critical")

    return result
```
