---
id: 6a1d9f98e819ed70a0e994de
title: "挑战 332：问题分支处理"
challengeType: 29
dashedName: challenge-332
---

# --description--

给定自上次在某个问题上发布消息以来的毫秒数，以及该问题上最后发布的消息，根据以下规则确定你应该如何处理该问题：

- 如果最后一条消息是在不到 7 天前，返回 `"leave it"`
- 如果最后一条消息是在 7 天或更久之前，且其内容包含 `"bump"`（不区分大小写），返回 `"close it"`
- 否则，返回 `"bump it"`

# --hints--

`triage_issue(86400000, "Lets fix it")` 应该返回 `"leave it"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(86400000, "Lets fix it"), "leave it")`)
}})
```

`triage_issue(1209600000, "still waiting")` 应该返回 `"bump it"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(1209600000, "still waiting"), "bump it")`)
}})
```

`triage_issue(864000000, "bump")` 应该返回 `"close it"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(864000000, "bump"), "close it")`)
}})
```

`triage_issue(604800000, "Do we still want this?")` 应该返回 `"bump it"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Do we still want this?"), "bump it")`)
}})
```

`triage_issue(604800000, "Bumping this")` 应该返回 `"close it"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Bumping this"), "close it")`)
}})
```

`triage_issue(345600000, "I'll make a PR")` 应该返回 `"leave it"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(345600000, "I'll make a PR"), "leave it")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(ms, message):

    return ms
```

# --solutions--

```py
def triage_issue(ms, message):
    seven_days = 7 * 24 * 60 * 60 * 1000
    if ms < seven_days:
        return "leave it"
    if "bump" in message.lower():
        return "close it"
    return "bump it"
```
