---
id: 6a2037a68a0bc2aef0006003
title: "挑战 341：生日倒计时"
challengeType: 29
dashedName: challenge-341
---

# --description--

给定今天的日期和一个生日，返回距离此人下一个生日还有多少天。

- 今天的日期以带前导零的 `"YYYY-MM-DD"` 字串形式给出，例如：`"2026-07-16"`。
- 生日以不带前导零的 `"M/D"` 字串形式给出，例如：`"9/7"`。
- 如果今天是他们的生日，返回距离下一个生日的天数（不是 `0`）。
- 需要考虑闰年。

# --hints--

`days_until_birthday("2026-07-16", "9/7")` 应该返回 `53`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "9/7"), 53)`)
}})
```

`days_until_birthday("2026-07-16", "3/22")` 应该返回 `249`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "3/22"), 249)`)
}})
```

`days_until_birthday("2026-07-16", "7/16")` 应该返回 `365`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "7/16"), 365)`)
}})
```

`days_until_birthday("2024-02-28", "3/1")` 应该返回 `2`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-02-28", "3/1"), 2)`)
}})
```

`days_until_birthday("2023-04-24", "12/30")` 应该返回 `250`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2023-04-24", "12/30"), 250)`)
}})
```

`days_until_birthday("2024-03-01", "2/29")` 应该返回 `1460`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-03-01", "2/29"), 1460)`)
}})
```

`days_until_birthday("2096-03-01", "2/29")` 应该返回 `2920`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2096-03-01", "2/29"), 2920)`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_birthday(today, birthday):

    return today
```

# --solutions--

```py
def days_until_birthday(today, birthday):
    from datetime import date

    year, month, day = [int(x) for x in today.split("-")]
    b_month, b_day = [int(x) for x in birthday.split("/")]

    today_date = date(year, month, day)

    def is_leap_year(y):
        return (y % 4 == 0 and y % 100 != 0) or y % 400 == 0

    def next_valid_birthday(from_year):
        for y in range(from_year, from_year + 9):
            if b_month == 2 and b_day == 29 and not is_leap_year(y):
                continue
            d = date(y, b_month, b_day)
            if d > today_date:
                return d

    return (next_valid_birthday(year) - today_date).days
```
