---
id: 6a2037a68a0bc2aef0006007
title: "挑战 345：单词混合器"
challengeType: 29
dashedName: challenge-345
---

# --description--

给定两个单词，通过将第一个单词的前半部分与第二个单词的后半部分组合，返回一个新单词。

- 对于长度为奇数的单词，前半部分较短。

# --hints--

`blend_words("turtle", "toucan")` 应该返回 `"turcan"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("turtle", "toucan"), "turcan")`)
}})
```

`blend_words("chipmunk", "flamingo")` 应该返回 `"chipingo"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("chipmunk", "flamingo"), "chipingo")`)
}})
```

`blend_words("falcon", "pelican")` 应该返回 `"falican"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("falcon", "pelican"), "falican")`)
}})
```

`blend_words("hyena", "iguana")` 应该返回 `"hyana"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("hyena", "iguana"), "hyana")`)
}})
```

`blend_words("scorpion", "gorilla")` 应该返回 `"scorilla"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("scorpion", "gorilla"), "scorilla")`)
}})
```

`blend_words("platypus", "wolverine")` 应该返回 `"platerine"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("platypus", "wolverine"), "platerine")`)
}})
```

# --seed--

## --seed-contents--

```py
def blend_words(word1, word2):

    return word1
```

# --solutions--

```py
def blend_words(word1, word2):
    first_half = word1[:len(word1) // 2]
    second_half = word2[len(word2) // 2:]
    return first_half + second_half
```
