---
id: 6a2037a68a0bc2aef0006002
title: "挑战 340：猪拉丁文转换器"
challengeType: 29
dashedName: challenge-340
---

# --description--

给定一个字串，使用以下规则将其转换为猪拉丁文：

- 如果一个单词以元音（`"a"`、`"e"`、`"i"`、`"o"` 或 `"u"`）开头，添加 `"way"` 到末尾。例如，`"universe"` 转换为 `"universeway"`。
- 如果一个单词以一个或多个辅音开头，将它们移到末尾并添加 `"ay"`。例如，`"hello"` 转换为 `"ellohay"`。
- 保持首字母的大小写。例如，`"Hello"` 转换为 `"Ellohay"`。

# --hints--

`pig_latin("universe")` 应该返回 `"universeway"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("universe"), "universeway")`)
}})
```

`pig_latin("hello")` 应该返回 `"ellohay"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello"), "ellohay")`)
}})
```

`pig_latin("hello universe")` 应该返回 `"ellohay universeway"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello universe"), "ellohay universeway")`)
}})
```

`pig_latin("Hello universe")` 应该返回 `"Ellohay universeway"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Hello universe"), "Ellohay universeway")`)
}})
```

`pig_latin("Pig Latin is fun")` 应该返回 `"Igpay Atinlay isway unfay"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Pig Latin is fun"), "Igpay Atinlay isway unfay")`)
}})
```

`pig_latin("The quick brown fox jumped over the lazy dog")` 应该返回 `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday")`)
}})
```

# --seed--

## --seed-contents--

```py
def pig_latin(s):

    return s
```

# --solutions--

```py
def pig_latin(s):
    vowels = "aeiou"

    def convert_word(word):
        lower_word = word.lower()
        if lower_word[0] in vowels:
            return word + "way"
        i = 0
        while i < len(word) and lower_word[i] not in vowels:
            i += 1
        consonants = word[:i]
        rest = word[i:]
        result = rest + consonants + "ay"
        if word[0].isupper():
            return result[0].upper() + result[1:].lower()
        return result

    return " ".join(convert_word(word) for word in s.split(" "))
```
