---
id: 6a2037a68a0bc2aef0006004
title: "挑战 342：骰子概率"
challengeType: 29
dashedName: challenge-342
---

# --description--

给定要掷的六面骰数量和一个目标和，返回以字符串形式表示掷出该和的概率，格式为 `"1 in X"`。

- 骰子数量在 1 到 6 之间。
- 给定数量的骰子总能达到目标和。
- 将 `"X"` 四舍五入到最接近的整数。

# --hints--

`get_odds(1, 5)` 应该返回 `"1 in 6"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(1, 5), "1 in 6")`)
}})
```

`get_odds(2, 4)` 应该返回 `"1 in 12"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(2, 4), "1 in 12")`)
}})
```

`get_odds(3, 10)` 应该返回 `"1 in 8"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(3, 10), "1 in 8")`)
}})
```

`get_odds(4, 7)` 应该返回 `"1 in 65"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(4, 7), "1 in 65")`)
}})
```

`get_odds(5, 26)` 应该返回 `"1 in 111"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(5, 26), "1 in 111")`)
}})
```

`get_odds(6, 35)` 应该返回 `"1 in 7776"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(6, 35), "1 in 7776")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_odds(dice, target):

    return dice
```

# --solutions--

```py
def get_odds(dice, target):
    valid_combos = 0
    total_combos = 6 ** dice

    def roll(dice_left, total):
        nonlocal valid_combos
        if dice_left == 0:
            if total == target:
                valid_combos += 1
            return
        for face in range(1, 7):
            roll(dice_left - 1, total + face)

    roll(dice, 0)
    return f"1 in {round(total_combos / valid_combos)}"
```
