---
id: 6a2037a68a0bc2aef0006001
title: "挑战 339：数组块"
challengeType: 29
dashedName: challenge-339
---

# --description--

给定一个数组和一个块大小，返回按该大小拆分成子数组的数组。

- 如果数组不能被均匀分割，最后一个块可能会更小。

# --hints--

`chunk_array([1, 2, 3, 4, 5, 6], 3)` 应该返回 `[[1, 2, 3], [4, 5, 6]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]])`)
}})
```

`chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` 应该返回 `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]])`)
}})
```

`chunk_array([1, 2, 3, 4, 5], 3)` 应该返回 `[[1, 2, 3], [4, 5]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]])`)
}})
```

`chunk_array(["a", "b", "c", "d", "e"], 1)` 应该返回 `[["a"], ["b"], ["c"], ["d"], ["e"]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]])`)
}})
```

`chunk_array([1, 2, 3], 5)` 应该返回 `[[1, 2, 3]]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3], 5), [[1, 2, 3]])`)
}})
```

# --seed--

## --seed-contents--

```py
def chunk_array(arr, size):

    return arr
```

# --solutions--

```py
def chunk_array(arr, size):
    return [arr[i:i + size] for i in range(0, len(arr), size)]
```
