---
id: 6a2037a68a0bc2aef0006006
title: "挑战 344：黄金比例"
challengeType: 29
dashedName: challenge-344
---

# --description--

给定两个数字，判断它们的比值是否接近黄金比例。

- 使用黄金比例 `1.618`
- 允许误差为 `0.01`

# --hints--

`is_golden_ratio(21, 34)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(21, 34), True)`)
}})
```

`is_golden_ratio(15, 20)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(15, 20), False)`)
}})
```

`is_golden_ratio(8, 13)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(8, 13), True)`)
}})
```

`is_golden_ratio(10, 16)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(10, 16), False)`)
}})
```

`is_golden_ratio(1618, 1000)` 应该返回 `True`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(1618, 1000), True)`)
}})
```

`is_golden_ratio(88, 55)` 应该返回 `False`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(88, 55), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_golden_ratio(a, b):

    return a
```

# --solutions--

```py
def is_golden_ratio(a, b):
    golden_ratio = 1.618
    ratio = max(a, b) / min(a, b)
    return abs(ratio - golden_ratio) <= 0.01
```
