---
id: 6a22d77ddf034bc4e35b1d54
title: "挑战 347：博弈论"
challengeType: 29
dashedName: challenge-347
---

# --description--

给定两个等长的字串，分别表现两个玩家在一个游戏中的策略，返回得分数组 `[player1, player2]`。

- 给定的字串只包含两个字母之一：`"C"`（合作）或 `"D"`（背叛）。
- 每个字符代表一轮，得分规则如下：
  - 如果双方都合作，每人得 3 分。
  - 如果双方都背叛，每人得 1 分。
  - 如果一方背叛另一方合作，背叛者得 5 分，合作方得 0 分。

# --hints--

`play_game("CCCC", "CCCC")` 应该返回 `[12, 12]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` 应该返回 `[4, 4]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` 应该返回 `[5, 10]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` 应该返回 `[24, 34]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` 应该返回 `[66, 21]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
