---
id: 6a22d77ddf034bc4e35b1d59
title: "挑战 352：对比度评级 1"
challengeType: 29
dashedName: challenge-352
---

# --description--

给定一个对比度比率和一个布尔值表示文本是否为大号，使用下表返回 WCAG 评级：

| 评级 | 正常文本 | 大号文本 |
|--------|-------------|------------|
| `"AAA"` | 7.0 及以上 | 4.5 及以上 |
| `"AA"` | 4.5 及以上 | 3.0 及以上 |
| `"Fail"` | 低于 4.5 | 低于 3.0 |

# --hints--

`get_contrast_rating("7.5", False)` 应该返回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("7.5", False), "AAA")`)
}})
```

`get_contrast_rating("4.8", False)` 应该返回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.8", False), "AA")`)
}})
```

`get_contrast_rating("4.2", False)` 应该返回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.2", False), "Fail")`)
}})
```

`get_contrast_rating("4.5", True)` 应该返回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.5", True), "AAA")`)
}})
```

`get_contrast_rating("3.0", True)` 应该返回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("3.0", True), "AA")`)
}})
```

`get_contrast_rating("2.7", False)` 应该返回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("2.7", False), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(ratio, is_large_text):

    return ratio
```

# --solutions--

```py
def get_contrast_rating(ratio, is_large_text):
    ratio = float(ratio)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
