---
id: 6a22d77ddf034bc4e35b1d5a
title: "挑战 353：对比度评级 2"
challengeType: 29
dashedName: challenge-353
---

# --description--

给定两个相对亮度值和一个布尔值表示文本是否较大，使用以下方法返回 WCAG 对比度评级：

通过给每个亮度值添加 0.05，然后用较亮的值除以较暗的值来计算对比度比率。较亮的值始终是第一个参数。

根据对比度比率使用下表返回评级：

| 评级 | 正常文本 | 大号文本 |
|--------|-------------|------------|
| `"AAA"` | 7.0 及以上 | 4.5 及以上 |
| `"AA"` | 4.5 及以上 | 3.0 及以上 |
| `"Fail"` | 低于 4.5 | 低于 3.0 |

# --hints--

`get_contrast_rating(1.0, 0.0, False)` 应该返回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(1.0, 0.0, False), "AAA")`)
}})
```

`get_contrast_rating(0.9015, 0.1364, False)` 应该返回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.9015, 0.1364, False), "AA")`)
}})
```

`get_contrast_rating(0.8965, 0.1628, False)` 应该返回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.8965, 0.1628, False), "Fail")`)
}})
```

`get_contrast_rating(0.7469, 0.0957, True)` 应该返回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7469, 0.0957, True), "AAA")`)
}})
```

`get_contrast_rating(0.7489, 0.2018, True)` 应该返回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7489, 0.2018, True), "AA")`)
}})
```

`get_contrast_rating(0.6571, 0.1974, True)` 应该返回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.6571, 0.1974, True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(l1, l2, is_large_text):

    return l1
```

# --solutions--

```py
def get_contrast_rating(l1, l2, is_large_text):
    ratio = (l1 + 0.05) / (l2 + 0.05)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
