---
id: 6a22d77ddf034bc4e35b1d5b
title: "挑战 354：对比度评级 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

给定两个表示 RGB 值的数组和一个布尔值表示文本是否为大号，使用以下方法返回 WCAG 对比度评级：

首先，将每个 RGB 值转换为相对亮度：

- 将每个通道 `[R, G, B]` 除以 255，得到介于 `0` 和 `1` 之间的值
- 对每个通道应用伽马校正公式：
  - 如果通道值小于或等于 `0.04045`：`channel / 12.92`
  - 否则：`((channel + 0.055) / 1.055) ^ 2.4`
- 计算亮度：`0.2126 * R + 0.7152 * G + 0.0722 * B`

然后，通过给每个亮度值加 0.05，再将较亮的值除以较暗的值来计算对比度比率。较亮的值总是第一个参数。

根据对比度比率使用下表返回评级：

| 评级 | 正常文本 | 大号文本 |
|--------|-------------|------------|
| `"AAA"` | 7.0 及以上 | 4.5 及以上 |
| `"AA"` | 4.5 及以上 | 3.0 及以上 |
| `"Fail"` | 低于 4.5 | 低于 3.0 |

# --hints--

`get_contrast_rating([255, 255, 255], [0, 0, 0], False)` 应该返回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

`get_contrast_rating([215, 188, 188], [55, 55, 55], False)` 应该返回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

`get_contrast_rating([143, 144, 210], [46, 47, 61], False)` 应该返回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

`get_contrast_rating([167, 167, 210], [53, 10, 53], True)` 应该返回 `"AAA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

`get_contrast_rating([135, 147, 155], [60, 70, 90], True)` 应该返回 `"AA"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

`get_contrast_rating([125, 210, 195], [105, 130, 90], True)` 应该返回 `"Fail"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
