---
id: 6a22d77ddf034bc4e35b1d57
title: "挑战 350：字母距离"
challengeType: 29
dashedName: challenge-350
---

# --description--

给定两个长度相等的字串，返回每对字符之间最短距离的总和。

- 输入只包含小写字母
- 字母表被视为一个圆环，因此 `a` 和 `z` 之间的距离是 1。

# --hints--

`letter_distance("abc", "bcd")` 应该返回 `3`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "bcd"), 3)`)
}})
```

`letter_distance("abc", "xyz")` 应该返回 `9`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "xyz"), 9)`)
}})
```

`letter_distance("encrypt", "decrypt")` 应该返回 `10`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("encrypt", "decrypt"), 10)`)
}})
```

`letter_distance("algorithm", "codeblock")` 应该返回 `43`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("algorithm", "codeblock"), 43)`)
}})
```

`letter_distance("lobster", "penguin")` 应该返回 `47`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("lobster", "penguin"), 47)`)
}})
```

`letter_distance("alligator", "crocodile")` 应该返回 `55`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("alligator", "crocodile"), 55)`)
}})
```

# --seed--

## --seed-contents--

```py
def letter_distance(str1, str2):

    return str1
```

# --solutions--

```py
def letter_distance(str1, str2):
    total = 0
    for a, b in zip(str1, str2):
        diff = abs(ord(a) - ord(b))
        total += min(diff, 26 - diff)
    return total
```
