---
id: 6a22d77ddf034bc4e35b1d56
title: "挑战 349：手机信号"
challengeType: 29
dashedName: challenge-349
---

# --description--

给定一个包含三个手机信号塔读数的网格，确定手机的位置。

- 网格中的每个单元格要么是 `0`（无信号塔），要么是一个正整数，表示从该单元格到手机的单元格数，测量方式为直线：水平、垂直或对角线。
- 返回距离所有三个信号塔单元格数都正确的单元格的 `[row, col]`。
- 总是恰好有一个方案。

# --hints--

`find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` 应该返回 `[1, 2]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2])`)
}})
```

`find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` 应该返回 `[2, 1]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1])`)
}})
```

`find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` 应该返回 `[2, 2]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2])`)
}})
```

`find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` 应该返回 `[3, 4]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4])`)
}})
```

`find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` 应该返回 `[3, 3]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_signal(grid):

    return grid
```

# --solutions--

```py
def find_signal(grid):
    rows = len(grid)
    cols = len(grid[0])

    towers = []
    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                towers.append((r, c, grid[r][c]))

    def is_valid(r, c, tr, tc, dist):
        dr = abs(r - tr)
        dc = abs(c - tc)
        return (dr == dist or dc == dist) and (dr == 0 or dc == 0 or dr == dc)

    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                continue
            if all(is_valid(r, c, tr, tc, dist) for tr, tc, dist in towers):
                return [r, c]
```
