---
id: 6a26df3e2988bcdded204896
title: "挑战 360：口语时长"
challengeType: 29
dashedName: challenge-360
---

# --description--

给定一个秒数，返回用口语表达的时长。

- 将时长拆分为小时、分钟和秒。
- 跳过任何为零的值。
- 根据情况使用单数或复数（`"1 hour"`，`"2 hours"`）。
- 如果有多个单位，最后两个单位之间用`"and"`连接，倒数第二和倒数第三个单位之间用逗号连接（`"1 hour, 2 minutes and 3 seconds"`）。

# --hints--

`get_spoken_duration(3723)` 应该返回 `"1 hour, 2 minutes and 3 seconds"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(3723), "1 hour, 2 minutes and 3 seconds")`)
}})
```

`get_spoken_duration(7295)` 应该返回 `"2 hours, 1 minute and 35 seconds"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(7295), "2 hours, 1 minute and 35 seconds")`)
}})
```

`get_spoken_duration(8521)` 应该返回 `"2 hours, 22 minutes and 1 second"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(8521), "2 hours, 22 minutes and 1 second")`)
}})
```

`get_spoken_duration(435)` 应该返回 `"7 minutes and 15 seconds"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(435), "7 minutes and 15 seconds")`)
}})
```

`get_spoken_duration(14455)` 应该返回 `"4 hours and 55 seconds"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(14455), "4 hours and 55 seconds")`)
}})
```

`get_spoken_duration(72000)` 应该返回 `"20 hours"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(72000), "20 hours")`)
}})
```

`get_spoken_duration(1)` 应该返回 `"1 second"`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(1), "1 second")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_duration(seconds):

    return seconds
```

# --solutions--

```py
def get_spoken_duration(seconds):
    hours = seconds // 3600
    minutes = (seconds % 3600) // 60
    secs = seconds % 60

    parts = []
    if hours:
        parts.append(f"{hours} {'hour' if hours == 1 else 'hours'}")
    if minutes:
        parts.append(f"{minutes} {'minute' if minutes == 1 else 'minutes'}")
    if secs:
        parts.append(f"{secs} {'second' if secs == 1 else 'seconds'}")

    if len(parts) == 1:
        return parts[0]
    if len(parts) == 2:
        return f"{parts[0]} and {parts[1]}"
    return f"{parts[0]}, {parts[1]} and {parts[2]}"
```
