---
id: 6a26df95efa55a2524399745
title: "挑战 364：两个桶之间"
challengeType: 29
dashedName: challenge-364
---

# --description--

给定两个油漆桶，每个都有一个 RGB 颜色和一个满度等级，返回混合后的 RGB 颜色，作为一个包含三个整数的数组。

- 每个桶是一个对象（JavaScript）或字典（Python），具有一个 `color` 属性（一个包含三个整数 `[r, g, b]` 的数组）和一个 `fullness` 属性（0–100）。
- 混合颜色是基于满度等级对两个颜色中每个通道的加权平均，每个通道四舍五入到最接近的整数。

# --hints--

`mix_paint({"color": [250, 250, 250], "fullness": 50}, {"color": [0, 0, 0], "fullness": 50})` 应该返回 `[125, 125, 125]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [250, 250, 250], "fullness": 50}, {"color": [0, 0, 0], "fullness": 50}), [125, 125, 125])`)
}})
```

`mix_paint({"color": [250, 250, 250], "fullness": 80}, {"color": [0, 0, 0], "fullness": 20})` 应该返回 `[200, 200, 200]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [250, 250, 250], "fullness": 80}, {"color": [0, 0, 0], "fullness": 20}), [200, 200, 200])`)
}})
```

`mix_paint({"color": [100, 150, 200], "fullness": 30}, {"color": [100, 150, 200], "fullness": 70})` 应该返回 `[100, 150, 200]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [100, 150, 200], "fullness": 30}, {"color": [100, 150, 200], "fullness": 70}), [100, 150, 200])`)
}})
```

`mix_paint({"color": [143, 143, 101], "fullness": 45}, {"color": [100, 204, 204], "fullness": 90})` 应该返回 `[114, 184, 170]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [143, 143, 101], "fullness": 45}, {"color": [100, 204, 204], "fullness": 90}), [114, 184, 170])`)
}})
```

`mix_paint({"color": [15, 134, 249], "fullness": 29}, {"color": [97, 178, 55], "fullness": 54})` 应该返回 `[68, 163, 123]`。

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [15, 134, 249], "fullness": 29}, {"color": [97, 178, 55], "fullness": 54}), [68, 163, 123])`)
}})
```

# --seed--

## --seed-contents--

```py
def mix_paint(bucket1, bucket2):

    return bucket1
```

# --solutions--

```py
def mix_paint(bucket1, bucket2):
    total = bucket1['fullness'] + bucket2['fullness']
    return [
        round((c1 * bucket1['fullness'] + c2 * bucket2['fullness']) / total)
        for c1, c2 in zip(bucket1['color'], bucket2['color'])
    ]
```
