---
id: 6721089b71655f803ce3bef2
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

次に、ドロップダウンに4つのオプション、`All`、`Woodwinds`、`Brass`、`Percussion`をこの順番で追加してください。各オプションには、小文字の option を使って `id` 属性と `value` 属性を指定してください。

# --hints--

`option` 要素の中に `select` 要素が4つあるはずです。

```js
assert.lengthOf(document.querySelectorAll("select>option"), 4);
```

最初の `option` 要素にはテキスト `All` があるはずです。

```js
assert.equal(document.querySelectorAll("select>option")[0].innerText,"All");
```

`All` のオプションには `id` 属性と、値が `value` の `all` 属性があります。

```js
assert.equal(document.querySelectorAll("select>option")[0].getAttribute("id"),"all");
assert.equal(document.querySelectorAll("select>option")[0].getAttribute("value"),"all");
```

2番目の `option` 要素にはテキスト `Woodwinds` があるはずです。

```js
assert.equal(document.querySelectorAll("select>option")[1].innerText,"Woodwinds");
```

`Woodwinds` のオプションには `id` 属性と、値が `value` の `woodwinds` 属性があります。

```js
assert.equal(document.querySelectorAll("select>option")[1].getAttribute("id"),"woodwinds");
assert.equal(document.querySelectorAll("select>option")[1].getAttribute("value"),"woodwinds");
```

3番目の `option` 要素にはテキスト `Brass` があるはずです。

```js
assert.equal(document.querySelectorAll("select>option")[2].innerText,"Brass");
```

`Brass` のオプションには `id` 属性と、値が `value` の `brass` 属性があります。

```js
assert.equal(document.querySelectorAll("select>option")[2].getAttribute("id"),"brass");
assert.equal(document.querySelectorAll("select>option")[2].getAttribute("value"),"brass");
```

4番目の `option` 要素にはテキスト `Percussion` があるはずです。

```js
assert.equal(document.querySelectorAll("select>option")[3].innerText,"Percussion");
```

`Percussion` のオプションには `id` 属性と、値が `value` の `percussion` 属性があります。

```js
assert.equal(document.querySelectorAll("select>option")[3].getAttribute("id"),"percussion");
assert.equal(document.querySelectorAll("select>option")[3].getAttribute("value"),"percussion");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music Instruments product page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Student Instruments</h1>

    <main>
--fcc-editable-region--
      <select class="select-container">

      </select>
--fcc-editable-region--
      <div class="products-container">
        <div class="card">
          <h2>Flute</h2>
          <p>$500</p>
        </div>
        <div class="card">
          <h2>Clarinet</h2>
          <p>$200</p>
        </div>
        <div class="card">
          <h2>Oboe</h2>
          <p>$4000</p>
        </div>

        <div class="card">
          <h2>Trumpet</h2>
          <p>$200</p>
        </div>
        <div class="card">
          <h2>Trombone</h2>
          <p>$300</p>
        </div>
        <div class="card">
          <h2>French Horn</h2>
          <p>$4300</p>
        </div>

        <div class="card">
          <h2>Drum Set</h2>
          <p>$500</p>
        </div>
        <div class="card">
          <h2>Xylophone</h2>
          <p>$3000</p>
        </div>
        <div class="card">
          <h2>Cymbals</h2>
          <p>$200</p>
        </div>
        <div class="card">
          <h2>Marimba</h2>
          <p>$3000</p>
        </div>
      </div>
    </main>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body {
  background-color: var(--dark-grey);
  color: var(--white);
}

h1 {
  text-align: center;
  margin-top: 20px;
}

.select-container {
  display: block;
  margin: 25px auto;
  padding: 8px;
  border: 4px solid var(--white);
  border-radius: 4px;
  width: 200px;
}

.products-container {
  display: flex;
  flex-direction: column;
  flex-wrap: wrap;
  align-items: center;
  justify-content: center;
  gap: 20px;
}

@media (min-width: 760px) {
  .products-container {
    flex-direction: row;
  }
}

.card {
  background-color: var(--white);
  color: var(--dark-grey);
  border: 4px solid var(--yellow);
  border-radius: 5px;
  padding: 10px;
  width: 200px;
}
```
