---
id: 6690e10ebe2181212abc9652
title: 设计一张名片
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个带有值为 `business-card` 的 `class` 属性的 `div`，它将包含所有其他元素。
1. 在 `.business-card` 元素内，应有一个带有 `class` 属性且值为 `profile-image` 的 `img` 元素。你可以将图像源设置为 `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`（如果你愿意）。应有一个带有有意义描述的 `alt` 属性。
1. 在 `.business-card` 元素内，`img` 元素之后，你应该有三个带有 `class` 属性的 `p` 元素，属性值分别为 `full-name`、`designation` 和 `company`。
1. 第一个 `p` 元素应包含你的名字。
1. 第二个 `p` 元素应包含你的称谓。
1. 第三个 `p` 元素应包含你的公司名称。
1. 第三个 `p` 元素下方应有一个 `hr` 元素。
1. 在 `hr` 元素之后，应有两个 `p` 元素。在第一个 `p` 元素中，文本应为一个电子邮件地址。在第二个 `p` 元素中，文本应为一个电话号码。
1. 在两个 `p` 元素之后，应有一个类为 `portfolio-link` 的 `a` 元素，文本为 `Portfolio`，并且它应链接到一个有效的 URL。
1. 在包含作品集链接的 `a` 元素后应有一个 `hr` 元素。
1. 你应该有另一个带有 `class` 属性且值为 `social-media` 的 `div` 元素。在此元素内，应有一个 `h2` 元素，其文本为 `Connect with me`。
1. 在 `.social-media` 元素内，应有三个 `a` 元素。`a` 元素应分别包含文本 `Twitter`、`LinkedIn` 和 `GitHub`，并链接到有效的网站。
1. 你应该在 HTML 的 `head` 部分使用 `link` 标签链接样式表 `styles.css`。
1. 你的页面背景应为 `rosybrown`。整体字体应为 `Arial`，备选字体为 `sans-serif`。
1. `.business-card` 选择器应具有将元素宽度设置为 `300px` 和你选择的背景颜色的属性。你还应设置四周的 `padding` 为 `20px`，以及顶部 `margin` 为 `100px`。文本应居中对齐，字体大小应为 `16px`。左右边距应设置为 `auto`。
1. `.profile-image` 选择器应具有值为 `100%` 的 `max-width` 属性。
1. 所有的 `p` 元素应具有上下 `margin` 为 `5px`。
1. 页面上的所有链接都不应带下划线。

# --hints--

你应该有一个主 `div`，其 `class` 属性值为 `business-card`。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

在 `.business-card` 元素内，应有一个带有 `class` 属性且值为 `profile-image` 的 `img` 元素。

```js
assert.exists(document.querySelector('img.profile-image'));
```

图像源应设置为有效的图像。

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

图像的 `alt` 属性应设置为有意义的描述。

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

在 `.business-card` 元素内，`img` 元素之后，应有一个带有值为 `full-name` 的 `class` 属性的 `p` 元素。

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

在 `.full-name` 元素之后，应有一个带有值为 `designation` 的 `class` 属性的 `p` 元素。

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

在 `.designation` 元素之后，应有一个带有 `class` 属性且值为 `company` 的 `p` 元素。

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

第一个 `p` 元素，即具有 `full-name` 类的那个，应该包含你的名字。

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

第二个 `p` 元素，即具有 `class` 属性且值为 `designation` 的那个，应包含你的职位。

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

第三个 `p` 元素，即 `class` 属性值为 `company` 的那个，应该包含你的公司名称。

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

在第三个 `p` 元素之后，即具有值为 `company` 的 `class` 属性的那个，应该有一个 `hr` 元素。

```js
assert.exists(document.querySelector('p.company + hr'));
```

在第一个 `hr` 元素之后，应有一个 `p` 元素，其文本为电子邮件地址。

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

在电子邮件的 `p` 元素之后，应有另一个 `p` 元素，其文本为你的电话号码。

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

在电话号码 `p` 元素之后，应有一个带有类 `portfolio-link` 的 `a` 元素。 

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

`.portfolio-link` 元素的文本应为 `Portfolio`。

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

你的 `.portfolio-link` 元素应具有有效的 `href`。

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

在包含你作品集的 `a` 元素之后，应有另一个 `hr` 元素作为分隔符。

```js
assert.exists(document.querySelector('a + hr'));
```

在第二个 `hr` 元素之后，应有一个带有 `class` 属性且值为 `social-media` 的 `div` 元素。

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

在 `.social-media` 元素内，应有一个带有文本 `Connect with me` 的 `h2` 元素。

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

在 `.social-media` 元素内应有三个 `a` 元素。

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

第一个 `a` 元素应具有 `Twitter` 文本。

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

第一个 `a` 元素的 `href` 应指向一个有效的链接。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

第二个 `a` 元素的文本应为 `LinkedIn`。

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

第二个 `a` 元素的 `href` 应指向有效的链接。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

第三个 `a` 元素应具有 `GitHub` 文本。

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

第三个 `a` 元素的 `href` 应指向有效的链接。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

你的 `link` 元素应具有一个 `rel` 属性，其值为 `stylesheet`。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

你的 `link` 元素应具有一个 `href` 属性，其值为 `styles.css`。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

你应该在 `body` 选择器中设置页面背景颜色为 `rosybrown`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

你应该在 `body` 元素中将页面字体设置为 `Arial`，并以 `sans-serif` 作为后备字体。

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

你的 `.business-card` 选择器应具有值为 `300px` 的 `width` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

你的 `.business-card` 选择器应具有带有效颜色值的 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

你的 `.business-card` 选择器应具有值为 `20px` 的 `padding` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

你的 `.business-card` 选择器应有 `100px` 的上边距。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

`.business-card` 元素内的文本应居中对齐。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

`.business-card` 元素内文本的字体大小应为 `16px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

`.business-card` 元素的左右边距应设置为 `auto`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

你的 `.profile-image` 选择器应具有值为 `100%` 的 `max-width` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

你的 `p` 元素应具有上下 `margin` 为 `5px`。

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

页面的 `links` 不应有下划线。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
