---
id: 6720b1f2680ec706ed534acf
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, practicar치s trabajar con diferentes pseudo-clases y pseudo-elementos dise침ando una tarjeta de felicitaci칩n. Se ha proporcionado la plantilla b치sica de HTML para ti. 

Inicia el taller enlazando el archivo `styles.css`.

# --hints--

Su elemento `link` debe estar dentro de su elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Su elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    
  </head>
  <body>
  </body>
</html>

--fcc-editable-region--
```

```css

```
