---
id: 6720b3191bbafd269bae0ae8
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Crea un elemento `div` que tenga un `id` de `greeting-card` y una `class` de `card`.

Dentro del elemento `div`, agrega un `h1` con el texto `Happy Birthday!`. Luego, añade un elemento de párrafo con una `class` llamada `message` y el texto `Wishing you all the happiness and joy on your special day!`.

# --hints--

Debe tener un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

El elemento `div` debe tener un `id` de `greeting-card`.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

El elemento `div` debe tener una `class` de `card`.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

Debes tener un elemento `h1` dentro del elemento `div`.

```js
assert.exists(document.querySelector("div > h1"));
```

El elemento `h1` debería contener el texto `Happy Birthday!`.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

Debes tener un elemento `p` dentro del `div`.

```js
assert.exists(document.querySelector("div > p"));
```

El elemento `p` debería tener una `class` de `message`.

```js
assert.exists(document.querySelector("div > p.message"));
```

El elemento `p` debería contener el texto `Wishing you all the happiness and joy on your special day!`.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
