---
id: 672b4434dd6d6a1edc405419
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Agrega dos elementos `a` dentro del elemento `.card-links`.

El primero debe tener un texto de `Send Card`, una `class` de `send-link` y un `href` de `#send`.

El segundo debe tener un texto de `Share on Social Media`, una `class` de `share-link` y un `href` de `#share`.

# --hints--

El primer elemento `a` dentro de `.card-links` debe tener un texto de `Send Card`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

El primer elemento `a` debe tener una `class` con un valor de `send-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

El primer elemento `a` debe tener un `href` con un valor de `#send`.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

El segundo elemento `a` dentro de `.card-links` debe tener un texto de `Share on Social Media`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

El segundo elemento `a` debe tener una `class` con un valor de `share-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

El segundo elemento `a` debe tener un `href` con un valor de `#share`.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>

--fcc-editable-region--

      <div class="card-links"></div>

--fcc-editable-region--

  	</div>
  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
