---
id: 67333082ab35e414eb4ea70d
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Cuando se pasa el mouse sobre los elementos `a`, el color de fondo hace una transición a un color diferente. Puedes regular cómo ocurre esa transición con la propiedad `transition`:

```css
a {
  transition: color 1s linear;
}
```

Los valores que acepta la propiedad `transition` son, en orden, la propiedad a la que se aplica la transición, la duración de la transición y el momento de la misma.

Si hay múltiples propiedades que tienen una transición, puedes escribir los valores para cada una separados por una coma:

```css
p {
  transition: property1 0.1s, property2 0.6s linear;
}
```

Si se omite un valor, como el timing para la primera propiedad, se aplica un valor predeterminado.

Agrega al selector `.card` `transition: transform 0.3s, background-color 0.3s ease`.

Pruébalo, la transición al pasar el ratón está completa.

# --hints--

El selector `.card` debe tener una propiedad `transition` establecida en `transform 0.3s, background-color 0.3s ease`.

```js
assert.oneOf(
  new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("transition"),
  [
    "transform 0.3s ease 0s, background-color 0.3s ease 0s", // this is what comes out of `getPropertyValue`
    "transform 0.3s, background-color 0.3s" // different environment got this instead
  ]
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
--fcc-editable-region--
  
--fcc-editable-region--
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

```
