---
id: 587d8247367417b2b2512c38
title: Mitigate el riesgo de hacer clic con el casco.frameguard()
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

Como recordatorio, este proyecto se está construyendo sobre la siguiente estructura inicial clonada desde <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Tu página podría ser puesta en un `<frame>` o `<iframe>` sin tu consentimiento. Esto puede resultar en ataques de clickjacking, entre otras cosas. El clickjacking es una técnica para engañar a un usuario y hacer que interactúe con una página diferente a la que cree. Esto se puede lograr ejecutando tu página en un contexto malicioso, mediante iframing. En ese contexto, un hacker puede poner una capa oculta sobre tu página. Botones ocultos pueden usarse para ejecutar scripts maliciosos. Este middleware establece el encabezado X-Frame-Options. Restringe quién puede poner tu sitio en un frame. Tiene tres modos: DENY, SAMEORIGIN y ALLOW-FROM.

No necesitamos que nuestra aplicación se enmarque.

# --instructions--

Utilice `helmet.frameguard()` pasando con el objeto de configuración `{action: 'deny'}`.

# --hints--

helmet.frameguard() middleware debería montarse correctamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

helmet.frameguard() 'action' debe establecerse en 'DENY'

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```

