---
id: 587d8247367417b2b2512c37
title: Ocultar Información potencialmente Peligrosa usando helmet.hidePoweredBy()
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

Como recordatorio, este proyecto se está construyendo sobre la siguiente estructura inicial clonada desde <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Los hackers pueden explotar vulnerabilidades conocidas en Express/Node si ven que tu sitio está impulsado por Express. `X-Powered-By: Express` se envía en cada solicitud que proviene de Express por defecto. Usa el middleware `helmet.hidePoweredBy()` para eliminar el encabezado X-Powered-By.

# --hints--

helmet.hidePoweredBy() middleware debería montarse correctamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```

