---
id: 587d8247367417b2b2512c39
title: Mitigar el riesgo de ataques de cross site scripting (XSS) con helmet.xssFilter()
  Mitigate the Risk of Cross Site Scripting (XSS) Attacks with
  helmet.xssFilter()
challengeType: 2
forumTopicId: 301583
dashedName: mitigate-the-risk-of-cross-site-scripting-xss-attacks-with-helmet-xssfilter
---

# --description--

Como recordatorio, este proyecto se está construyendo sobre la siguiente estructura inicial clonada desde <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Secuencias de Comando de Sitio-Cruzado (XSS) es un ataque de tipo frecuente donde instrucciones maliciosas son insertadas dentro de páginas vulnerables, con el propósito de robar información sensible como información almacenada de sesión (cookies), o contraseñas.

La regla básica para reducir el riesgo de un ataque XSS es simple: "Nunca confíes en la entrada del usuario". Como desarrollador siempre debes sanitizar toda la entrada que viene de afuera. Esto incluye datos que vienen de formularios, consultas GET en URLs e incluso de cuerpos POST. Sanitizar significa que debes encontrar y codificar los caracteres que pueden ser peligrosos, por ejemplo &lt;, >.

Navegadores actuales pueden ayudar a mitigar el riesgo mediante la adopción de mejores estrategias de software. A menudo esto es configurable via cabeceras http.

La cabecera X-XSS-Protection HTTP es una protección básica. El navegador detecta una pontencial sentencia maliciosa insertada usando un filtro heurístico. Si la cabecera esta habilitada, el navegador cambia el código de la sentencia maliciosa, neutralizándola. Todavía tiene un soporte limitado.

# --instructions--

Usa `helmet.xssFilter()` para sanear la entrada envíada a su servidor.

# --hints--

helmet.xssFilter() middleware debe ser montado correctamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'xXssProtection');
assert.property(data.headers, 'x-xss-protection');
```

