---
id: 587d8248367417b2b2512c3d
title: Desactivar Prefetching DNS con helmet.dnsPrefetchControl()
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

Como recordatorio, este proyecto se está construyendo sobre la siguiente estructura inicial clonada desde <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Para mejorar el rendimiento, la mayoría de los navegadores prefetch DNS records para los enlaces en una página. De esa manera, la ip de destino ya se conoce cuando el usuario hace clic en un enlace. Esto puede llevar a un uso excesivo del servicio DNS (si tienes un sitio web grande, visitado por millones de personas…), problemas de privacidad (un espía podría inferir que estás en una página determinada) o alteración de las estadísticas de la página (algunos enlaces pueden aparecer como visitados aunque no lo estén). Si tienes necesidades de alta seguridad, puedes desactivar el DNS prefetching, a costa de una penalización en el rendimiento.

# --instructions--

Usa el método `helmet.dnsPrefetchControl()` en tu servidor.

# --hints--

helmet.dnsPrefetchControl() middleware debería montarse correctamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```

