---
id: 587d8248367417b2b2512c3a
title: Evita inferir en la Respuesta MIME Tipo con helmet.noSniff()
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

Como recordatorio, este proyecto se está construyendo sobre la siguiente estructura inicial clonada desde <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Los navegadores pueden usar la detección de contenido o MIME para anular el encabezado `Content-Type` de una respuesta y así adivinar y procesar los datos usando un tipo de contenido implícito. Aunque esto puede ser conveniente en algunos escenarios, también puede conducir a ataques peligrosos. Este middleware establece el encabezado `X-Content-Type-Options` en `nosniff`, indicando al navegador que no omita el `Content-Type` proporcionado.

# --instructions--

Usa el método `helmet.noSniff()` en tu servidor.

# --hints--

helmet.noSniff() middleware debe ser montado correctamente

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```

