---
id: 68599633407374c0d574587b
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Hatua inayofuata ni kuongeza `button` zitakazohusika na kubadili kati ya sayari tofauti na taarifa zao.

Ndani ya kipengele cha `tablist`, ongeza vipengele vitatu vya `button` kila kikiwa na sifa ya `role` iliyowekwa kuwa `"tab"`. Kipengele cha kwanza cha `button` kinapaswa kuwa na maandishi ya `🌍 Earth`. Kipengele cha pili cha `button` kinapaswa kuwa na maandishi ya `🪐 Saturn`. Kipengele cha tatu cha `button` kinapaswa kuwa na maandishi ya `🔴 Mars`.  

# --hints--

Unapaswa kuwa na vipengele vitatu vya `button` ndani ya kipengele cha `tablist`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

Kila kipengele chako cha `button` kinapaswa kuwa na sifa ya `role` iliyowekwa kuwa `"tab"`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

Kipengele chako cha kwanza cha `button` kinapaswa kuwa na maandishi ya `🌍 Earth`.

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

Kipengele chako cha pili cha `button` kinapaswa kuwa na maandishi ya `🪐 Saturn`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

Kipengele chako cha tatu cha `button` kinapaswa kuwa na maandishi ya `🔴 Mars`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./style.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
