---
id: 6859ae3ef795c4c9d4badf9c
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Sasa ni wakati wa kuanza kuongeza utendakazi kwenye orodha ya tabo. Hatua ya kwanza ni kupata upatikanaji wa vipengele vyote vya tabo na paneli katika orodha ya tabo.

Tengeneza kigezo kinachoitwa `tabs` na weka thamani yake kuwa matokeo ya hoja ya utafutaji ya kupata vipengele vyote vyenye sifa ya kipengele `role` ikiwa imewekwa kuwa `tab`.

Kisha, tengeneza kigezo kinachoitwa `panels` na weka thamani yake kuwa matokeo ya hoja ya utafutaji ya kupata vipengele vyote vyenye sifa ya kipengele `role` ikiwa imewekwa kuwa `tabpanel`.

Kumbuka kwamba unaweza kutumia njia ya `querySelectorAll` kuchagua vipengele vingi kwa wakati mmoja.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `tabs`.

```js
assert.isDefined(tabs);
```

Unapaswa kuweka `document.querySelectorAll('[role="tab"]')` kwenye kigezo cha `tabs`.

```js
assert.deepEqual(
  Array.from(tabs),
  Array.from(document.querySelectorAll('[role="tab"]'))
);
```

Unapaswa kuwa na kigezo kinachoitwa `panels`.

```js
assert.isDefined(panels);
```

Unapaswa kuweka `document.querySelectorAll('[role="tabpanel"]')` kwenye kigezo cha `panels`.

```js
assert.deepEqual(
  Array.from(panels),
  Array.from(document.querySelectorAll('[role="tabpanel"]'))
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div id="panel-earth" role="tabpanel" aria-labelledby="tab-earth">
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.
        </p>
      </div>
      <div id="panel-saturn" role="tabpanel" aria-labelledby="tab-saturn" hidden>
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div id="panel-mars" role="tabpanel" aria-labelledby="tab-mars" hidden>
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.
        </p>
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
