---
id: 6859a1270ac255c571352189
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Sasa ni wakati wa kuunda alama za uandishi kwa kila moja ya paneli zitakazoonyesha ukweli kuhusu kila sayari.

Anza kwa kuongeza vipengele vitatu vya `div` kila kikiwa na sifa ya kipengele ya `role` iliyowekwa kuwa `"tabpanel"`. Ndani ya kila kipengele cha `div`, panapaswa kuwepo kipengele cha `p`.

# --hints--

Unapaswa kuwa na vipengele vitatu vya `div` vyenye sifa ya kipengele ya `role` iliyowekwa kuwa `"tabpanel"`.

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"]'), 3);
```

Kila kipengele cha `tabpanel` kinapaswa kuwa na kipengele cha `p` ndani yake.

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"] p'), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
