---
id: 685976a92ffab8b2dc744b72
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Ndani ya kipengele chako cha `.tabs`, ongeza kipengele cha `div` chenye sifa ya kipengele ya `role` iliyowekwa kuwa `"tablist"` na sifa ya kipengele ya `aria-labelledby` iliyowekwa kuwa `"tabs-title"`.  

Sifa ya kipengele ya `aria-labelledby` inahitajika hapa kuhusisha orodha ya tabo na lebo inayoonekana, kusaidia teknolojia za msaada kama kisomaji cha skrini kutangaza madhumuni ya orodha ya tabo kwa watumizi.

# --hints--

Unapaswa kuwa na kipengele cha `div` ndani ya kipengele chako cha `.tabs`.

```js
assert.exists(document.querySelector(".tabs div"));
```

Kipengele chako cha `div` kinapaswa kuwa na sifa ya kipengele ya `role` iliyowekwa kuwa `"tablist"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

Kipengele chako cha `div` kinapaswa kuwa na sifa ya kipengele ya `aria-labelledby` iliyowekwa kuwa `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
