---
id: 68599af28ead38c22590a89d
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Inapaswa kuwepo na njia ya kuunganisha vifungo na maudhui yanayowakilisha. Hii hufanyika kwa kutumia sifa ya kipengele ya `aria-controls` kwenye kila kifungo. Baadaye, utaongeza maudhui na sifa za kipengele za `id` zinazolingana kwenye vipengele vya maudhui.

Anza kwa kuongeza sifa ya kipengele ya `aria-controls` kwenye kila kipengele cha `button`. Thamani ya sifa ya kipengele ya `aria-controls` ya kifungo cha kwanza inapaswa kuwekwa kuwa `"panel-earth"`. Sifa ya kipengele ya `aria-controls` ya kifungo cha pili inapaswa kuwekwa kuwa `"panel-saturn"`. Sifa ya kipengele ya `aria-controls` ya kifungo cha tatu inapaswa kuwekwa kuwa `"panel-mars"`.

# --hints--

Kila kipengele chako cha `button` kinapaswa kuwa na sifa ya kipengele ya `aria-controls`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

Kipengele cha kwanza cha `button` kinapaswa kuwa na sifa ya kipengele ya `aria-controls` iliyowekwa kuwa `"panel-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

Kipengele cha pili cha `button` kinapaswa kuwa na sifa ya kipengele ya `aria-controls` iliyowekwa kuwa `"panel-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

Kipengele cha tatu cha `button` kinapaswa kuwa na sifa ya kipengele ya `aria-controls` iliyowekwa kuwa `"panel-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
