---
id: 685db0bd1280154ae85bcf9f
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utajenga orodha ya tabo za sayari zinazopatikana kwa kutumia HTML, CSS na JavaScript. Muundo wa msingi wa HTML na CSS vimetolewa tayari kwako.

Katika hatua hizi za mwanzo, utajenga sehemu zilizobaki za HTML kwa ajili ya orodha ya tabo.

Anza kwa kuunda kipengele cha `div` chenye darasa la `"tabs"`. Ndani ya kipengele hicho cha `div`, ongeza kipengele cha `h2` chenye sifa ya kipengele ya `id` iliyowekwa kuwa `"tabs-title"` na maandishi ya `Planets`.

# --hints--

Unapaswa kuwa na kipengele cha `div`.

```js
assert.exists(document.querySelector("div"));
```

Kipengele chako cha `div` kinapaswa kuwa na darasa la `"tabs"`.

```js
assert.exists(document.querySelector("div.tabs"));
```

Kipengele chako cha `div` kinapaswa kuwa na kipengele cha `h2` ndani yake.

```js
assert.exists(document.querySelector(".tabs h2"));
```

Kipengele chako cha `h2` kinapaswa kuwa na sifa ya kipengele ya `id` iliyowekwa kuwa `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

Kipengele chako cha `h2` kinapaswa kuwa na maandishi ya `Planets`.

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
