---
id: 699d1f8e51fcdc8230f779fb
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將建置一個文化遺產函式庫目錄系統。你會解析未經處理的書籍資料，搜尋並群組條目，驗證資料品質，並將結果匯出成不同格式。

每張目錄卡片都以單一字串儲存，格式為 `"Title | Author | Year | Location"`，欄位以管線字元（`|`）分隔。

你將在整個工作坊中使用這些字串的陣列來建置和測試你的目錄函式。創建一個名為 `rawCatalogCards` 的陣列，並將以下兩個字串加入其中：

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

你應該宣告一個名為 `rawCatalogCards` 的變數。

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` 應該是一個陣列。

```js
assert.isArray(rawCatalogCards);
```

`rawCatalogCards` 的第一個元素應該是 `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`。

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

`rawCatalogCards` 的第二個元素應該是 `"The Shining | King, Stephen | 1977 | Shelf K1"`。

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
