---
id: 699d20623e57e8ffee885459
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

`parseCard` 已經過測試並且運作正常。`cardResult` 變數已不再需要。

首先刪除 `const cardResult` 行和 `console.log(cardResult)` 行。接著創建一個名為 `parseCatalog` 的函式，該函式接受 `rawCards` 作為參數。暫時在函式內宣告一個空的陣列 `catalog` 並傳回它。

在函式下方，呼叫 `parseCatalog(rawCatalogCards)`，將結果指定給名為 `catalog` 的變數，並將 `catalog.length` 輸出到主控台。

# --hints--

你應該有一個名為 `parseCatalog` 的函式（程式），並帶有一個 `rawCards` 參數。

```js
const explorer = await __helpers.Explorer(code);
const { parseCatalog: parseCatalogFn } = explorer.allFunctions;
assert.exists(parseCatalogFn);
assert.equal(parseCatalogFn.parameters[0]?.toString(), 'rawCards');
```

你應該宣告一個名為 `catalog` 的變數，將它初始化為空陣列，然後傳回它。

```js
const explorer = await __helpers.Explorer(code);
const { parseCatalog: parseCatalogFn } = explorer.allFunctions;
assert.exists(parseCatalogFn);
const { catalog } = parseCatalogFn.variables;
assert.exists(catalog);
const catalogValue = eval(catalog.value?.toString());
assert.isArray(catalogValue);
assert.isEmpty(catalogValue);
assert.isTrue(parseCatalogFn.hasReturn('catalog'));
```

`catalog` 應該被定義並設為 `parseCatalog(rawCatalogCards)`。

```js
const explorer = await __helpers.Explorer(code);
const { catalog } = explorer.variables;
assert.exists(catalog);
assert.isTrue(catalog.value.matches('parseCatalog(rawCatalogCards)'));
```

你應該刪除 `cardResult` 變數和相關的 `console.log()` 呼叫。

```js
assert.notMatch(__helpers.removeJSComments(code), /\bcardResult\b/);
```

你應該將 `catalog.length` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*catalog\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

--fcc-editable-region--
const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
--fcc-editable-region--
```
