---
id: 699d20623e57e8ffee885451
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

`parseCard` 函式會接收一個未經處理的目錄字串並傳回一個結構化的物件。在撰寫你的函式之前，先建立一個方法來查看它的輸出。

使用 `parseCard` 作為引數呼叫 `rawCatalogCards[2]`，將結果指定給名為 `cardResult` 的變數，並將 `cardResult` 輸出到主控台。

目前 `parseCard` 沒有函式主體，因此會傳回 `undefined`。當你在接下來的步驟中為函式新增程式碼時，主控台輸出會準確顯示函式在每個階段傳回的內容。

# --hints--

你應該使用 `parseCard` 作為引數呼叫 `rawCatalogCards[2]`。

```js
assert.match(__helpers.removeJSComments(code), /parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)/);
```

你應該將呼叫 `parseCard` 並帶入 `rawCatalogCards[2]` 的結果指定給名為 `cardResult` 的變數。

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+cardResult\s*=\s*parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)\s*;?/);
```

你應該將 `cardResult` 輸出到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*cardResult\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {}

--fcc-editable-region--

--fcc-editable-region--
```
