---
id: 699d20623e57e8ffee885457
title: 第 8 步驟
challengeType: 1
dashedName: step-8
---

# --description--

有些目錄卡的欄位缺失。例如，條目 16 的作者欄是空白的，所以 `parseCard` 目前對該欄位傳回空字串。

`||` 運算子提供了備援：當 `a || b` 為真值時，`a` 求值為 `a`，否則為 `b`。由於空字串是偽值，`"" || "Unknown"` 求值為 `"Unknown"`。

更新回傳述語中的 `title`、`author` 和 `location`，使用 `|| "Unknown"` 作為備援。`year` 留待下一步使用。

# --hints--

當呼叫時字串的作者欄是空白，回傳物件的 `author` 屬性應該是 `"Unknown"`。

```js
assert.strictEqual(parseCard("X |  | 2000 | Y").author, "Unknown");
```

當呼叫時字串的標題欄是空白，回傳物件的 `title` 屬性應該是 `"Unknown"`。

```js
assert.strictEqual(parseCard(" | A | 2000 | Y").title, "Unknown");
```

當呼叫時字串的地點欄是空白，回傳物件的 `location` 屬性應該是 `"Unknown"`。

```js
assert.strictEqual(parseCard("X | A | 2000 | ").location, "Unknown");
```

你還不應該為 `year` 添加備援。

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "");
```

當所有欄位都有值時，`parseCard` 應該傳回它們的實際值，而非 `"Unknown"`。

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title,
    author: author,
    year: year,
    location: location
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
