---
id: 699d20623e57e8ffee885455
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

從主控台的輸出可以看到，分割後，每個部分仍然有多餘的前置和／或後置空白。要清理每個元素，你需要逐一處理它們。

這正是使用 `for` 迴圈的完美時機，該迴圈會對可迭代物件（例如陣列）中的每個元素重複執行一段程式碼：

```js
for (let i = 0; i < array.length; i++) {
  // code to run on each iteration
}
```

`i` 從 `0`（第一個索引）開始，當 `i` 小於陣列長度時持續執行迴圈，並且 `i++` 在每次迭代後將 `i` 增加一。

在 `parseCard` 裡，宣告一個名為 `trimmedParts` 的空陣列。接著加入一個 `for` 迴圈來迭代 `parts`。目前先將迴圈主體留空，下一步你會填入內容。

# --hints--

你應該宣告一個名為 `trimmedParts` 的變數，並將它設為空陣列。

```js
assert.match(parseCard.toString(), /(const|let|var)\s+trimmedParts\s*=\s*\[\s*\]/);
```

你應該使用 `for` 迴圈來迭代 `parts`。

```js
assert.match(parseCard.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*parts\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--

--fcc-editable-region--
  return parts;
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
