---
id: 699d20623e57e8ffee88545d
title: 第 19 步驟
challengeType: 1
dashedName: step-19
---

# --description--

在 `"Unknown"` 檢查之後，為具有數字年份的書籍添加以十年為群組的邏輯。宣告 `const decade = Math.floor(book.year / 10) * 10`。例如，`Math.floor(1977 / 10)` 會得到 `197`，乘以 `10` 會得到 `1970`。

接著宣告 `const decadeKey =` `` `${decade}s` `` to produce a string like `"1970s"`。

使用括號表示法來群組書籍：檢查 `grouped[decadeKey]` 是否尚未存在，若是則初始化為空陣列，然後將 `book` 推入其中。這裡必須使用括號表示法，因為鍵是存放在變數中。

# --hints--

你應該宣告一個名為 `decade` 的變數，並將其設為 `Math.floor(book.year / 10) * 10`。

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+decade\s*=\s*Math\.floor\s*\(\s*book\.year\s*\/\s*10\s*\)\s*\*\s*10/);
```

你應該宣告一個名為 `decadeKey` 的變數，該變數包含 `decade` 和字母 `"s"`。

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+decadeKey\s*=.*decade/s);
assert.match(groupByDecade.toString(), /decadeKey.*s/s);
```

你應該檢查 `grouped[decadeKey]` 是否尚未存在，並將其初始化為空陣列。

```js
assert.match(groupByDecade.toString(), /!\s*grouped\s*\[\s*decadeKey\s*\]/);
assert.match(groupByDecade.toString(), /grouped\s*\[\s*decadeKey\s*\]\s*=\s*\[\s*\]/);
```

你應該將 `book` 推入 `grouped[decadeKey]`。

```js
assert.match(groupByDecade.toString(), /grouped\s*\[\s*decadeKey\s*\]\s*\.push\s*\(\s*book\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

function groupByDecade(catalog) {
  const grouped = {};
  for (let i = 0; i < catalog.length; i++) {
    const book = catalog[i];
    if (book.year === "Unknown") {
      if (!grouped["Unknown"]) {
        grouped["Unknown"] = [];
      }
      grouped["Unknown"].push(book);
      continue;
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
