---
id: 69b83e35f19ba26ba1fa517a
title: Eine Ampel-Folgensteuerung bauen
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

In diesem Labor simulieren Sie konfigurierbare Ampelzyklen und protokollieren Anomalien.

Sie arbeiten mit Konfigurationsobjekten, die die Phasen einer Ampel beschreiben. Jedes Konfigurationsobjekt hat folgende Eigenschaften:

- `fault`: ein Boolean-Flag, das bei `true` eine vorzeitige Beendigung auslöst.
- `phases`: ein Array von Phasenobjekten.

Jedes Phasenobjekt innerhalb von `phases` hat folgende Eigenschaften:

- `color`: ein String, der die Lichtfarbe darstellt (`"green"`, `"yellow"` oder `"red"`).
- `duration`: eine positive ganze Zahl, die angibt, wie lange die Phase in Sekunden dauert.

Sie können die bereitgestellten Objekte `config1`, `config2`, `config3` und `config4` als Beispiele möglicher Konfigurationsobjekte verwenden.

**Hinweis**: Fügen Sie keine zusätzlichen `console.log()`-Anweisungen hinzu, da diese die Tests zum Scheitern bringen können.

**Ziel**: Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

# User Stories

1. Sie sollten eine Funktion namens `runSequence` mit zwei Parametern haben: `config` und `cycles`, wobei `config` ein Konfigurationsobjekt darstellt und `cycles` die maximale Anzahl der Durchläufe der Folge vor dem Stoppen angibt.

1. Sie sollten `runSequence(config, cycles)` mit einer `for`- oder `while`-Schleife implementieren, um jede Phase über die angegebene Anzahl von Zyklen zu durchlaufen.

1. Die Funktion `runSequence` sollte:
    - `No phases found` protokollieren und sofort zurückkehren, wenn `config.phases` leer ist.
    - `Faulted phase!` protokollieren und die Simulation vorzeitig stoppen, wenn `config.fault` auf `true` gesetzt ist.
    - `Invalid phase detected` protokollieren, wenn `duration <= 0` ist.
    - `Switching to [color] for [duration] s` für jede gültige Phase protokollieren. Ersetzen Sie `[color]` und `[duration]` durch die entsprechenden Eigenschaften des Phasenobjekts.
    - Zum Beispiel sollte `runSequence(config1, 1)` folgendes protokollieren:

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. Sie sollten eine Funktion namens `generateTimeline` mit zwei Parametern haben: `config` und `cycles`.

1. Die Funktion `generateTimeline` sollte:
    - Die kumulierte verstrichene Zeit nach jeder Phase über die Zyklen in einem Array aufzeichnen, indem Sie die `duration` jeder Phase zum laufenden Gesamtwert addieren, während Sie iterieren.
    - Alle fehlerhaften und ungültigen Phasen ohne Validierung verarbeiten, selbst wenn `config.fault` `true` ist oder `duration <= 0`.
    - Das Array der kumulierten Zeitstempel zurückgeben.
    - Zum Beispiel sollte `generateTimeline(config1, 1)` das Array `[5, 7, 11]` zurückgeben.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

Sie sollten keine Änderungen an den Objekten `config1`, `config2`, `config3` und `config4` vornehmen.

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

Sie sollten eine Funktion namens `runSequence` haben.

```js
assert.isFunction(runSequence);
```

Ihre Funktion `runSequence` sollte zwei Parameter akzeptieren, `config` und `cycles`.

```js
const regex = __helpers.functionRegex('runSequence', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regex);
```

`runSequence(config1, 1)` sollte in der Reihenfolge `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` protokollieren.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` sollte in der Reihenfolge `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`, `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` protokollieren.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` sollte in der Reihenfolge `Switching to red for 3 s`, `Invalid phase detected`, `Switching to green for 6 s` protokollieren.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` sollte `Faulted phase!` protokollieren.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` sollte `No phases found` protokollieren.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

Sie sollten eine Funktion namens `generateTimeline` haben.

```js
assert.isFunction(generateTimeline);
```

Ihre Funktion `generateTimeline` sollte zwei Parameter akzeptieren, `config` und `cycles`.

```js
const regEx = __helpers.functionRegex('generateTimeline', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regEx);
```

`generateTimeline(config1, 1)` sollte `[5, 7, 11]` zurückgeben.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` sollte `[5, 7, 11, 16, 18, 22]` zurückgeben.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` sollte `[3, 1, 7, 10, 8, 14]` zurückgeben.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` sollte `[5, 7, 13]` zurückgeben.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` sollte `[]` zurückgeben.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
