---
id: 56533eb9ac21ba0edf2244b3
title: Badilisha Celsius kuwa Fahrenheit
challengeType: 1
forumTopicId: 16806
dashedName: convert-celsius-to-fahrenheit
---

# --description--

Fomula ya kubadilisha kutoka Celsius kwenda Fahrenheit ni joto katika Celsius mara `9/5`, kisha ongeza `32`.

Umepatiwa kigezo `celsius` kinachowakilisha joto katika Celsius. Tumia kigezo `fahrenheit` kilichowekwa tayari na weka thamani yake kuwa joto sawa la Fahrenheit kwa joto lililotolewa la Celsius. Tumia fomula iliyotajwa hapo juu kusaidia kubadilisha joto la Celsius kuwa Fahrenheit.

# --hints--

`convertCtoF(0)` inapaswa kurudisha nambari

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` inapaswa kurudisha thamani ya `-22`

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` inapaswa kurudisha thamani ya `14`

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` inapaswa kurudisha thamani ya `32`

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` inapaswa kurudisha thamani ya `68`

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` inapaswa kurudisha thamani ya `86`

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js
function convertCtoF(celsius) {
  let fahrenheit;
  return fahrenheit;
}

convertCtoF(30);
```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
