---
id: a202eed8fc186c8434cb6d61
title: Geuza mfuatano wa herufi
challengeType: 1
forumTopicId: 16043
dashedName: reverse-a-string
---

# --description--

Geuza mfuatano wa herufi uliotolewa na urudishe mfuatano wa herufi uliogeuzwa.

Kwa mfano, `"hello"` inapaswa kuwa `"olleh"`.

# --hints--

`reverseString("hello")` inapaswa kurudisha mfuatano wa herufi.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` inapaswa kurudisha mfuatano wa herufi `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` inapaswa kurudisha mfuatano wa herufi `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` inapaswa kurudisha mfuatano wa herufi `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js
function reverseString(str) {
  return str;
}

reverseString('hello');
```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
