---
id: a24c1a4622e3c05097f71d67
title: Nina Mahali Pa Nini
challengeType: 1
forumTopicId: 16094
dashedName: where-do-i-belong
---

# --description--

Rudisha kielezo cha chini kabisa ambacho thamani (hoja ya pili) inapaswa kuingizwa katika safu ya data (hoja ya kwanza) mara tu itakapopangwa. Thamani inayorudishwa inapaswa kuwa nambari.

Kwa mfano, `getIndexToIns([1,2,3,4], 1.5)` inapaswa kurudisha `1` kwa sababu ni kubwa kuliko `1` (kielezo 0), lakini ni ndogo kuliko `2` (kielezo 1).

Vivyo hivyo, `getIndexToIns([20,3,5], 19)` inapaswa kurudisha `2` kwa sababu mara safu ya data itakapopangwa itaonekana kama `[3,5,20]` na `19` ni ndogo kuliko `20` (kielezo 2) na kubwa kuliko `5` (kielezo 1).

# --hints--

`getIndexToIns([10, 20, 30, 40, 50], 35)` inapaswa kurudisha nambari.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` inapaswa kurudisha `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` inapaswa kurudisha nambari.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 30));
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` inapaswa kurudisha `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` inapaswa kurudisha nambari.

```js
assert.isNumber(getIndexToIns([40, 60], 50));
```

`getIndexToIns([40, 60], 50)` inapaswa kurudisha `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` inapaswa kurudisha nambari.

```js
assert.isNumber(getIndexToIns([3, 10, 5], 3));
```

`getIndexToIns([3, 10, 5], 3)` inapaswa kurudisha `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` inapaswa kurudisha nambari.

```js
assert.isNumber(getIndexToIns([5, 3, 20, 3], 5));
```

`getIndexToIns([5, 3, 20, 3], 5)` inapaswa kurudisha `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` inapaswa kurudisha `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([2, 20, 10], 19)` inapaswa kurudisha nambari.

```js
assert.isNumber(getIndexToIns([2, 20, 10], 19));
```

`getIndexToIns([2, 5, 10], 15)` inapaswa kurudisha `3`.

```js
assert.strictEqual(getIndexToIns([2, 5, 10], 15), 3);
```

`getIndexToIns([2, 5, 10], 15)` inapaswa kurudisha nambari.

```js
assert.isNumber(getIndexToIns([2, 5, 10], 15));
```

`getIndexToIns([], 1)` inapaswa kurudisha nambari.

```js
assert.isNumber(getIndexToIns([], 1));
```

`getIndexToIns([], 1)` inapaswa kurudisha `0`.

```js
assert.strictEqual(getIndexToIns([], 1), 0);
```

# --seed--

## --seed-contents--

```js
function getIndexToIns(arr, num) {
  return num;
}

getIndexToIns([40, 60], 50);
```

# --solutions--

```js
function getIndexToIns(arr, num) {
  arr = arr.sort((a, b) => a - b);

  for (let i = 0; i < arr.length; i++) {
    if (arr[i] >= num) {
      return i;
    }
  }

  return arr.length;
}

getIndexToIns([40, 60], 50);
```
