---
id: a302f7aae1aa3152a5b413bc
title: Fanya faktoria ya nambari
challengeType: 1
forumTopicId: 16013
dashedName: factorialize-a-number
---

# --description--

Rudisha faktoria ya nambari kamili iliyotolewa.

Ikiwa nambari kamili inawakilishwa na herufi `n`, faktoria ni bidhaa ya nambari kamili chanya zote zilizo chini au sawa na `n`.

Faktoria mara nyingi huwakilishwa kwa alama fupi `n!`

Kwa mfano: `5! = 1 * 2 * 3 * 4 * 5 = 120`

Nambari kamili zitakazotolewa kwa kitendakazi zitakuwa kubwa au sawa na sifuri tu.

# --hints--

`factorialize(5)` inapaswa kurudisha nambari.

```js
assert.isNumber(factorialize(5));
```

`factorialize(5)` inapaswa kurudisha `120`.

```js
assert.strictEqual(factorialize(5), 120);
```

`factorialize(10)` inapaswa kurudisha `3628800`.

```js
assert.strictEqual(factorialize(10), 3628800);
```

`factorialize(20)` inapaswa kurudisha `2432902008176640000`.

```js
assert.strictEqual(factorialize(20), 2432902008176640000);
```

`factorialize(0)` inapaswa kurudisha `1`.

```js
assert.strictEqual(factorialize(0), 1);
```

# --seed--

## --seed-contents--

```js
function factorialize(num) {
  return num;
}

factorialize(5);
```

# --solutions--

```js
function factorialize(num) {
  return num < 1 ? 1 : num * factorialize(num - 1);
}

factorialize(5);
```
