---
id: a6e40f1041b06c996f7b2406
title: Wapata Wanaweka
challengeType: 1
forumTopicId: 16016
dashedName: finders-keepers
---

# --description--

Tengeneza kitendakazi kinachotafuta kupitia safu ya data `arr` na kurudisha kipengele cha kwanza ndani yake kinachopita 'jaribio la kweli'. Hii inamaanisha kwamba kwa kipengele `x`, 'jaribio la kweli' linapita ikiwa `func(x)` ni `true`. Ikiwa hakuna kipengele kinachopita jaribio, rudisha `undefined`.

# --hints--

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` inapaswa kurudisha `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` inapaswa kurudisha `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js
function findElement(arr, func) {
  let num = 0;
  return num;
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
