---
id: a77dbc43c33f39daa4429b4f
title: Boo nani
challengeType: 1
forumTopicId: 16000
dashedName: boo-who
---

# --description--

Kagua kama thamani imeainishwa kama thamani ya kweli au siyo kweli ya msingi. Rudisha `true` au `false`.

Thamani za msingi za thamani ya kweli au siyo kweli ni `true` na `false`.

# --hints--

`booWho(true)` inapaswa kurudisha `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` inapaswa kurudisha `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` inapaswa kurudisha `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js
function booWho(bool) {
  return bool;
}

booWho(null);
```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

booWho(null);
```
