---
id: ab6137d4e35944e21037b769
title: Andika Sentensi kwa Herufi Kuu Kwanza ya Kila Neno
challengeType: 1
forumTopicId: 16088
dashedName: title-case-a-sentence
---

# --description--

Rudisha mfuatano wa herufi uliotolewa ukiwa na herufi kubwa ya kwanza ya kila neno. Hakikisha sehemu nyingine ya neno iko kwa herufi ndogo.

Kwa madhumuni ya zoezi hili, unapaswa pia kuweka herufi kubwa kwa maneno ya kuunganisha kama `the` na `of`.

# --hints--

`titleCase("I'm a little tea pot")` inapaswa kurudisha mfuatano wa herufi.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` inapaswa kurudisha mfuatano wa herufi `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` inapaswa kurudisha mfuatano wa herufi `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` inapaswa kurudisha mfuatano wa herufi `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js
function titleCase(str) {
  return str;
}

titleCase("I'm a little tea pot");
```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}

titleCase("I'm a little tea pot");
```
