---
id: acda2fb1324d9b0fa741e6b5
title: Thibitisha Mwisho
challengeType: 1
forumTopicId: 16006
dashedName: confirm-the-ending
---

# --description--

Kagua kama mfuatano wa herufi (hoja ya kwanza, `str`) unamalizika na mfuatano wa herufi uliopewa kama lengo (hoja ya pili, `target`).

Changamoto hii _inaweza_ kutatuliwa kwa kutumia njia ya `.endsWith()`, ambayo ilianzishwa katika ES2015. Lakini kwa madhumuni ya changamoto hii, tungependa utumie moja ya njia za JavaScript za substring badala yake.

# --hints--

`confirmEnding("Bastian", "n")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` inapaswa kurudisha `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` inapaswa kurudisha `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` inapaswa kurudisha `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` inapaswa kurudisha `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing", "mountain")` inapaswa kurudisha `false`.

```js
assert.isFalse(
  confirmEnding(
    'If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing',
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` inapaswa kurudisha `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Msimbo wako haupaswi kutumia njia iliyo ndani ya `.endsWith()` kutatua changamoto hii.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js
function confirmEnding(str, target) {
  return str;
}

confirmEnding('Bastian', 'n');
```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
