---
id: afcc8d540bea9ea2669306b6
title: Rudia mfuatano wa herufi rudia mfuatano wa herufi
challengeType: 1
forumTopicId: 16041
dashedName: repeat-a-string-repeat-a-string
---

# --description--

Rudia mfuatano wa herufi uliotolewa `str` (hoja ya kwanza) kwa mara `num` (hoja ya pili). Rudisha mfuatano wa herufi tupu ikiwa `num` si nambari chanya. Kwa madhumuni ya changamoto hii, usitumie njia iliyo jengwa ndani ya `.repeat()`.

# --hints--

`repeatStringNumTimes("*", 3)` inapaswa kurudisha mfuatano wa herufi `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` inapaswa kurudisha mfuatano wa herufi `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` inapaswa kurudisha mfuatano wa herufi `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` inapaswa kurudisha mfuatano wa herufi `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` inapaswa kurudisha mfuatano wa herufi `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` inapaswa kurudisha mfuatano wa herufi tupu (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

Njia iliyo jengwa ndani ya `repeat()` haipaswi kutumika.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

`repeatStringNumTimes("abc", 0)` inapaswa kurudisha `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

# --seed--

## --seed-contents--

```js
function repeatStringNumTimes(str, num) {
  return str;
}

repeatStringNumTimes('abc', 3);
```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
