---
id: 587d7dab367417b2b2512b6f
title: "Die some-Methode verwenden, um zu prüfen, ob ein Element in einem Array ein Kriterium erfüllt"
challengeType: 1
forumTopicId: 301314
dashedName: use-the-some-method-to-check-that-any-elements-in-an-array-meet-a-criteria
---

# --description--

Die `some`-Methode arbeitet mit Arrays, um zu prüfen, ob *irgendein* Element einen bestimmten Test besteht. Sie gibt einen Boolean-Wert zurück – `true`, wenn eines der Werte das Kriterium erfüllt, `false`, wenn nicht.

Zum Beispiel würde der folgende Code prüfen, ob ein Element im `numbers`-Array kleiner als 10 ist:

```js
const numbers = [10, 50, 8, 220, 110, 11];

numbers.some(function(currentValue) {
  return currentValue < 10;
});
```

Die `some`-Methode würde `true` zurückgeben.

# --instructions--

Verwenden Sie die `some`-Methode innerhalb der `checkPositive`-Funktion, um zu prüfen, ob ein Element in `arr` positiv ist. Die Funktion soll einen Boolean-Wert zurückgeben.

# --hints--

Ihr Code sollte die `some`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.some/g));
```

`checkPositive([1, 2, 3, -4, 5])` sollte `true` zurückgeben.

```js
assert(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` sollte `true` zurückgeben.

```js
assert(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([-1, -2, -3, -4, -5])` sollte `false` zurückgeben.

```js
assert(!checkPositive([-1, -2, -3, -4, -5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.some(elem => elem > 0);
}
```
