---
id: 587d7b88367417b2b2512b45
title: 'Verwenden Sie höherwertige Funktionen map, filter oder reduce, um ein komplexes Problem zu lösen'
challengeType: 1
forumTopicId: 301311
dashedName: use-higher-order-functions-map-filter-or-reduce-to-solve-a-complex-problem
---

# --description--

Nachdem Sie einige Aufgaben mit höherwertigen Funktionen wie `map()`, `filter()` und `reduce()` bearbeitet haben, können Sie diese nun anwenden, um eine komplexere Aufgabe zu lösen.

# --instructions--

Vervollständigen Sie den Code für die `squareList`-Funktion mit einer beliebigen Kombination aus `map()`, `filter()` und `reduce()`. Die Funktion soll ein neues Array zurückgeben, das nur die Quadrate der positiven ganzen Zahlen (Dezimalzahlen sind keine ganzen Zahlen) enthält, wenn ihr ein Array mit reellen Zahlen übergeben wird. Ein Beispiel für ein Array mit reellen Zahlen ist `[-3, 4.8, 5, 3, -3.2]`.

**Hinweis:** Ihre Funktion darf keine Art von `for`- oder `while`-Schleifen oder die `forEach()`-Funktion verwenden.

# --hints--

`squareList` sollte eine `function` sein.

```js
assert.typeOf(squareList, 'function'),
  '<code>squareList</code> should be a <code>function</code>';
```

`for`, `while` und `forEach` sollten nicht verwendet werden.

```js
assert(!__helpers.removeJSComments(code).match(/for|while|forEach/g));
```

`map`, `filter` oder `reduce` sollten verwendet werden.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/\.(map|filter|reduce)\(/g)
);
```

Die Funktion sollte ein `array` zurückgeben.

```js
assert(Array.isArray(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])));
```

`squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])` sollte `[16, 1764, 36]` zurückgeben.

```js
assert.deepStrictEqual(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2]), [
  16,
  1764,
  36
]);
```

`squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3])` sollte `[9, 100, 49]` zurückgeben.

```js
assert.deepStrictEqual(squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3]), [
  9,
  100,
  49
]);
```

# --seed--

## --seed-contents--

```js
const squareList = arr => {
  // Only change code below this line
  return arr;
  // Only change code above this line
};

const squaredIntegers = squareList([-3, 4.8, 5, 3, -3.2]);
console.log(squaredIntegers);
```

# --solutions--

```js
const squareList = arr => {
  const positiveIntegers = arr.filter(num => {
    return num >= 0 && Number.isInteger(num);
  });
  const squaredIntegers = positiveIntegers.map(num => {
    return num ** 2;
  });
  return squaredIntegers;
};
```
