---
id: 587d7daa367417b2b2512b6b
title: Einen String mit der Methode split in ein Array aufteilen
challengeType: 1
forumTopicId: 18305
dashedName: split-a-string-into-an-array-using-the-split-method
---

# --description--

Die `split`-Methode teilt einen String in ein Array von Strings auf. Sie nimmt ein Argument für das Trennzeichen entgegen, das ein Zeichen sein kann, mit dem der String aufgeteilt wird, oder ein regulärer Ausdruck. Wenn das Trennzeichen beispielsweise ein Leerzeichen ist, erhalten Sie ein Array von Wörtern, und wenn das Trennzeichen ein leerer String ist, erhalten Sie ein Array mit jedem einzelnen Zeichen des Strings.

Hier sind zwei Beispiele, die einen String anhand von Leerzeichen und einen anderen anhand von Ziffern mit einem regulären Ausdruck aufteilen:

```js
const str = "Hello World";
const bySpace = str.split(" ");

const otherString = "How9are7you2today";
const byDigits = otherString.split(/\d/);
```

`bySpace` hätte den Wert `["Hello", "World"]` und `byDigits` hätte den Wert `["How", "are", "you", "today"]`.

Da Strings unveränderlich sind, erleichtert die `split`-Methode die Arbeit mit ihnen.

# --instructions--

Verwenden Sie die `split`-Methode innerhalb der `splitify`-Funktion, um `str` in ein Array von Wörtern aufzuteilen. Die Funktion soll das Array zurückgeben. Beachten Sie, dass die Wörter nicht immer durch Leerzeichen getrennt sind und das Array keine Satzzeichen enthalten sollte.

# --hints--

Ihr Code sollte die `split`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.split/g));
```

`splitify("Hello World,I-am code")` sollte `["Hello", "World", "I", "am", "code"]` zurückgeben.

```js
assert(
  JSON.stringify(splitify('Hello World,I-am code')) ===
    JSON.stringify(['Hello', 'World', 'I', 'am', 'code'])
);
```

`splitify("Earth-is-our home")` sollte `["Earth", "is", "our", "home"]` zurückgeben.

```js
assert(
  JSON.stringify(splitify('Earth-is-our home')) ===
    JSON.stringify(['Earth', 'is', 'our', 'home'])
);
```

`splitify("This.is.a-sentence")` sollte `["This", "is", "a", "sentence"]` zurückgeben.

```js
assert(
  JSON.stringify(splitify('This.is.a-sentence')) ===
    JSON.stringify(['This', 'is', 'a', 'sentence'])
);
```

# --seed--

## --seed-contents--

```js
function splitify(str) {
  // Only change code below this line


  // Only change code above this line
}

splitify("Hello World,I-am code");
```

# --solutions--

```js
function splitify(str) {
  return str.split(/\W/);
}
```
