---
id: 587d7b90367417b2b2512b65
title: Teil eines Arrays mit der Methode slice zurückgeben
challengeType: 1
forumTopicId: 301239
dashedName: return-part-of-an-array-using-the-slice-method
---

# --description--

Die `slice`-Methode gibt eine Kopie bestimmter Elemente eines Arrays zurück. Sie kann zwei Argumente annehmen: Das erste gibt den Index an, an dem der Ausschnitt beginnen soll, das zweite ist der Index, an dem der Ausschnitt enden soll (und dieser ist nicht inklusive). Wenn keine Argumente angegeben werden, ist der Standardwert, am Anfang des Arrays bis zum Ende zu starten, was eine einfache Möglichkeit ist, eine Kopie des gesamten Arrays zu erstellen. Die `slice`-Methode verändert das ursprüngliche Array nicht, sondern gibt ein neues zurück.

Hier ist ein Beispiel:

```js
const arr = ["Cat", "Dog", "Tiger", "Zebra"];
const newArray = arr.slice(1, 3);
```

`newArray` hätte den Wert `["Dog", "Tiger"]`.

# --instructions--

Verwenden Sie die `slice`-Methode in der `sliceArray`-Funktion, um einen Teil des `anim`-Arrays zurückzugeben, der durch die angegebenen `beginSlice`- und `endSlice`-Indizes bestimmt wird. Die Funktion soll ein Array zurückgeben.

# --hints--

Ihr Code sollte die `slice`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.slice/g));
```

Die Variable `inputAnim` darf sich nicht ändern.

```js
assert(
  JSON.stringify(inputAnim) ===
    JSON.stringify(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 3)` sollte `["Dog", "Tiger"]` zurückgeben.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 3)) ===
    JSON.stringify(['Dog', 'Tiger'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 0, 1)` sollte `["Cat"]` zurückgeben.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 0, 1)) ===
    JSON.stringify(['Cat'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 4)` sollte `["Dog", "Tiger", "Zebra"]` zurückgeben.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 4)) ===
    JSON.stringify(['Dog', 'Tiger', 'Zebra'])
);
```

# --seed--

## --seed-contents--

```js
function sliceArray(anim, beginSlice, endSlice) {
  // Only change code below this line


  // Only change code above this line
}

const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
sliceArray(inputAnim, 1, 3);
```

# --solutions--

```js
function sliceArray(anim, beginSlice, endSlice) {
  return anim.slice(beginSlice, endSlice);
}
const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
```
