---
id: 587d7da9367417b2b2512b69
title: Ein Array alphabetisch mit der Methode sort sortieren
challengeType: 1
forumTopicId: 18303
dashedName: sort-an-array-alphabetically-using-the-sort-method
---

# --description--

Die `sort`-Methode sortiert die Elemente eines Arrays entsprechend der Callback-Funktion.

Zum Beispiel:

```js
function ascendingOrder(arr) {
  return arr.sort(function(a, b) {
    return a - b;
  });
}

ascendingOrder([1, 5, 2, 3, 4]);
```

Dies würde den Wert `[1, 2, 3, 4, 5]` zurückgeben.

```js
function reverseAlpha(arr) {
  return arr.sort(function(a, b) {
    return a === b ? 0 : a < b ? 1 : -1;
  });
}

reverseAlpha(['l', 'h', 'z', 'b', 's']);
```

Dies würde den Wert `['z', 's', 'l', 'h', 'b']` zurückgeben.

Die Standard-Sortiermethode von JavaScript erfolgt nach dem Unicode-Punktwert von Strings, was unerwartete Ergebnisse liefern kann. Deshalb wird empfohlen, eine Callback-Funktion anzugeben, um zu bestimmen, wie die Array-Elemente sortiert werden sollen. Wenn eine solche Callback-Funktion, normalerweise `compareFunction` genannt, übergeben wird, werden die Array-Elemente entsprechend dem Rückgabewert der `compareFunction` sortiert: Wenn `compareFunction(a,b)` für zwei Elemente `a` und `b` einen Wert kleiner als 0 zurückgibt, dann wird `a` vor `b` stehen. Wenn `compareFunction(a,b)` für zwei Elemente `a` und `b` einen Wert größer als 0 zurückgibt, dann wird `b` vor `a` stehen. Wenn `compareFunction(a,b)` für zwei Elemente `a` und `b` den Wert 0 zurückgibt, bleiben `a` und `b` unverändert.

# --instructions--

Verwenden Sie die `sort`-Methode in der `alphabeticalOrder`-Funktion, um die Elemente von `arr` alphabetisch zu sortieren. Die Funktion soll das sortierte Array zurückgeben.

# --hints--

Ihr Code sollte die `sort`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.sort/g));
```

`alphabeticalOrder(["a", "d", "c", "a", "z", "g"])` sollte `["a", "a", "c", "d", "g", "z"]` zurückgeben.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'd', 'c', 'a', 'z', 'g'])) ===
    JSON.stringify(['a', 'a', 'c', 'd', 'g', 'z'])
);
```

`alphabeticalOrder(["x", "h", "a", "m", "n", "m"])` sollte `["a", "h", "m", "m", "n", "x"]` zurückgeben.

```js
assert(
  JSON.stringify(alphabeticalOrder(['x', 'h', 'a', 'm', 'n', 'm'])) ===
    JSON.stringify(['a', 'h', 'm', 'm', 'n', 'x'])
);
```

`alphabeticalOrder(["a", "a", "a", "a", "x", "t"])` sollte `["a", "a", "a", "a", "t", "x"]` zurückgeben.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'a', 'a', 'a', 'x', 't'])) ===
    JSON.stringify(['a', 'a', 'a', 'a', 't', 'x'])
);
```

# --seed--

## --seed-contents--

```js
function alphabeticalOrder(arr) {
  // Only change code below this line

  return arr
  // Only change code above this line
}

alphabeticalOrder(["a", "d", "c", "a", "z", "g"]);
```

# --solutions--

```js
function alphabeticalOrder(arr) {
  return arr.sort();
}
```
