---
id: 587d7b8e367417b2b2512b5d
title: Die Gefahren der Verwendung von imperativem Code verstehen
challengeType: 1
forumTopicId: 301241
dashedName: understand-the-hazards-of-using-imperative-code
---

# --description--

Funktionale Programmierung ist eine gute Gewohnheit. Sie hält Ihren Code leicht handhabbar und bewahrt Sie vor heimtückischen Bugs. Aber bevor wir dazu kommen, schauen wir uns einen imperativen Programmieransatz an, um hervorzuheben, wo Probleme auftreten können.

Im Englischen (und vielen anderen Sprachen) wird der Imperativ verwendet, um Befehle zu geben. Ähnlich ist ein imperativer Stil in der Programmierung einer, der dem Computer eine Reihe von Anweisungen gibt, um eine Aufgabe auszuführen.

Oft ändern diese Anweisungen den Zustand des Programms, zum Beispiel durch Aktualisieren globaler Variablen. Ein klassisches Beispiel ist das Schreiben einer `for`-Schleife, die genaue Anweisungen gibt, um die Indizes eines Arrays durchzulaufen.

Im Gegensatz dazu ist funktionale Programmierung eine Form der deklarativen Programmierung. Sie sagen dem Computer, was erledigt werden soll, indem Sie eine Methode oder Funktion aufrufen.

JavaScript bietet viele vordefinierte Methoden, die häufige Aufgaben übernehmen, sodass Sie nicht schreiben müssen, wie der Computer sie ausführen soll. Zum Beispiel könnten Sie statt der oben erwähnten `for`-Schleife die `map`-Methode aufrufen, die die Details des Durchlaufens eines Arrays übernimmt. Das hilft, semantische Fehler zu vermeiden, wie die „Off By One Errors“, die im Debugging-Abschnitt behandelt wurden.

Betrachten Sie folgendes Szenario: Sie surfen im Web in Ihrem Browser und möchten die Tabs verfolgen, die Sie geöffnet haben. Versuchen wir, dies mit einfachem objektorientiertem Code zu modellieren.

Ein Window-Objekt besteht aus Tabs, und normalerweise haben Sie mehr als ein Window geöffnet. Die Titel jeder geöffneten Seite in jedem Window-Objekt werden in einem Array gespeichert. Nachdem Sie im Browser gearbeitet haben (neue Tabs öffnen, Fenster zusammenführen und Tabs schließen), möchten Sie die Tabs ausgeben, die noch geöffnet sind. Geschlossene Tabs werden aus dem Array entfernt, und neue Tabs (zur Vereinfachung) werden am Ende hinzugefügt.

Der Code-Editor zeigt eine Implementierung dieser Funktionalität mit Funktionen für `tabOpen()`, `tabClose()` und `join()`. Das Array `tabs` ist Teil des Window-Objekts und speichert die Namen der geöffneten Seiten.

# --instructions--

Untersuchen Sie den Code im Editor. Er verwendet eine Methode, die Seiteneffekte im Programm verursacht und zu falschem Verhalten führt. Die endgültige Liste der geöffneten Tabs, gespeichert in `finalTabs.tabs`, sollte `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']` sein, aber die vom Code erzeugte Liste ist leicht anders.

Ändern Sie `Window.prototype.tabClose`, sodass der richtige Tab entfernt wird.

# --hints--

`finalTabs.tabs` sollte `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']` sein

```js
assert.deepEqual(finalTabs.tabs, [
  'FB',
  'Gitter',
  'Reddit',
  'Twitter',
  'Medium',
  'new tab',
  'Netflix',
  'YouTube',
  'Vine',
  'GMail',
  'Work mail',
  'Docs',
  'freeCodeCamp',
  'new tab'
]);
```

# --seed--

## --seed-contents--

```js
// tabs is an array of titles of each site open within the window
const Window = function(tabs) {
  this.tabs = tabs; // We keep a record of the array inside the object
};

// When you join two windows into one window
Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

// When you open a new tab at the end
Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab'); // Let's open a new tab for now
  return this;
};

// When you close a tab
Window.prototype.tabClose = function(index) {

  // Only change code below this line

  const tabsBeforeIndex = this.tabs.splice(0, index); // Get the tabs before the tab
  const tabsAfterIndex = this.tabs.splice(index + 1); // Get the tabs after the tab

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex); // Join them together

  // Only change code above this line

  return this;
 };

// Let's create three browser windows
const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']); // Your mailbox, drive, and other work sites
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']); // Social sites
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']); // Entertainment sites

// Now perform the tab opening, closing, and other operations
const finalTabs = socialWindow
  .tabOpen() // Open a new tab for cat memes
  .join(videoWindow.tabClose(2)) // Close third tab in video window, and join
  .join(workWindow.tabClose(1).tabOpen());
console.log(finalTabs.tabs);
```

# --solutions--

```js
const Window = function(tabs) {
  this.tabs = tabs;
};

Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab');
  return this;
};

Window.prototype.tabClose = function(index) {
  const tabsBeforeIndex = this.tabs.slice(0, index);
  const tabsAfterIndex = this.tabs.slice(index + 1);

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex);
  return this;
 };

const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']);
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']);
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']);

const finalTabs = socialWindow
  .tabOpen()
  .join(videoWindow.tabClose(2))
  .join(workWindow.tabClose(1).tabOpen());
```
