---
id: 587d7da9367417b2b2512b66
title: Zwei Arrays mit der Methode concat verbinden
challengeType: 1
forumTopicId: 301229
dashedName: combine-two-arrays-using-the-concat-method
---

# --description--

<dfn>Verkettung</dfn> bedeutet, Elemente aneinandergereiht zu verbinden. JavaScript bietet die `concat`-Methode sowohl für Strings als auch für Arrays an, die auf dieselbe Weise funktioniert. Bei Arrays wird die Methode auf einem Array aufgerufen, dann wird ein weiteres Array als Argument an `concat` übergeben, das an das Ende des ersten Arrays angehängt wird. Es wird ein neues Array zurückgegeben, und keines der ursprünglichen Arrays wird verändert. Hier ein Beispiel:

```js
[1, 2, 3].concat([4, 5, 6]);
```

Das zurückgegebene Array wäre `[1, 2, 3, 4, 5, 6]`.

# --instructions--

Verwenden Sie die `concat`-Methode in der `nonMutatingConcat`-Funktion, um `attach` an das Ende von `original` anzuhängen. Die Funktion soll das verkettete Array zurückgeben.

# --hints--

Ihr Code sollte die `concat`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

Das `first`-Array darf sich nicht ändern.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

Das `second`-Array darf sich nicht ändern.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingConcat([1, 2, 3], [4, 5])` soll `[1, 2, 3, 4, 5]` zurückgeben.

```js
assert(
  JSON.stringify(nonMutatingConcat([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingConcat(original, attach) {
  // Only change code below this line


  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingConcat(first, second);
```

# --solutions--

```js
function nonMutatingConcat(original, attach) {
  return original.concat(attach);
}
const first = [1, 2, 3];
const second = [4, 5];
```
