---
id: 587d7daa367417b2b2512b6c
title: Ein Array mit der Methode join zu einem String verbinden
challengeType: 1
forumTopicId: 18221
dashedName: combine-an-array-into-a-string-using-the-join-method
---

# --description--

Die `join`-Methode wird verwendet, um die Elemente eines Arrays zu einem String zusammenzufügen. Sie nimmt ein Argument für das Trennzeichen, das verwendet wird, um die Array-Elemente im String zu trennen.

Hier ist ein Beispiel:

```js
const arr = ["Hello", "World"];
const str = arr.join(" ");
```

`str` hätte den Wert des Strings `Hello World`.
# --instructions--

Verwenden Sie die `join`-Methode (unter anderem) innerhalb der `sentensify`-Funktion, um aus den Wörtern im String `str` einen Satz zu bilden. Die Funktion soll einen String zurückgeben. Zum Beispiel würde `I-like-Star-Wars` in `I like Star Wars` umgewandelt werden. Für diese Aufgabe verwenden Sie bitte nicht die `replace`-Methode.

# --hints--

Ihr Code sollte die `join`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.join/g));
```

Ihr Code sollte die `replace`-Methode nicht verwenden.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?replace/g));
```

`sentensify("May-the-force-be-with-you")` sollte einen String zurückgeben.

```js
assert(typeof sentensify('May-the-force-be-with-you') === 'string');
```

`sentensify("May-the-force-be-with-you")` sollte den String `May the force be with you` zurückgeben.

```js
assert(sentensify('May-the-force-be-with-you') === 'May the force be with you');
```

`sentensify("The.force.is.strong.with.this.one")` sollte den String `The force is strong with this one` zurückgeben.

```js
assert(
  sentensify('The.force.is.strong.with.this.one') ===
    'The force is strong with this one'
);
```

`sentensify("There,has,been,an,awakening")` sollte den String `There has been an awakening` zurückgeben.

```js
assert(
  sentensify('There,has,been,an,awakening') === 'There has been an awakening'
);
```

# --seed--

## --seed-contents--

```js
function sentensify(str) {
  // Only change code below this line


  // Only change code above this line
}

sentensify("May-the-force-be-with-you");
```

# --solutions--

```js
function sentensify(str) {
  return str.split(/\W/).join(' ');
}
```
