---
id: 587d7b8f367417b2b2512b60
title: Globale Variablen aus Funktionen heraus refaktorisieren
challengeType: 1
forumTopicId: 301235
dashedName: refactor-global-variables-out-of-functions
---

# --description--

Bisher haben wir zwei unterschiedliche Grundsätze der funktionalen Programmierung kennengelernt:

1) Ändern Sie keine Variable oder kein Objekt – erstellen Sie neue Variablen und Objekte und geben Sie diese bei Bedarf aus einer Funktion zurück. Hinweis: Wenn Sie etwas wie `const newArr = arrVar` verwenden, wobei `arrVar` ein Array ist, wird lediglich eine Referenz auf die bestehende Variable erstellt und keine Kopie. Das Ändern eines Werts in `newArr` würde also den Wert in `arrVar` ändern.

2) Deklarieren Sie Funktionsparameter – jede Berechnung innerhalb einer Funktion hängt nur von den an die Funktion übergebenen Argumenten ab und nicht von einem globalen Objekt oder einer globalen Variable.

Eine Zahl um eins zu erhöhen ist nicht sehr spannend, aber wir können diese Grundsätze anwenden, wenn wir mit Arrays oder komplexeren Objekten arbeiten.

# --instructions--

Schreiben Sie den Code so um, dass das globale Array `bookList` innerhalb keiner der beiden Funktionen verändert wird. Die `add`-Funktion soll das gegebene `bookName` an das Ende des übergebenen Arrays anhängen und ein neues Array (Liste) zurückgeben. Die `remove`-Funktion soll das gegebene `bookName` aus dem übergebenen Array entfernen.

**Hinweis:** Beide Funktionen sollen ein Array zurückgeben, und alle neuen Parameter sollen vor dem `bookName`-Parameter hinzugefügt werden.

# --hints--

`bookList` sollte sich nicht ändern und weiterhin `["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"]` entsprechen.

```js
add(bookList, "Test");
remove(bookList, "The Hound of the Baskervilles");
assert(
  JSON.stringify(bookList) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'On The Electrodynamics of Moving Bodies',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae'
    ])
);
```

`add(bookList, "A Brief History of Time")` sollte `["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae", "A Brief History of Time"]` zurückgeben.

```js
assert(
  JSON.stringify(add(bookList, "A Brief History of Time")) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'On The Electrodynamics of Moving Bodies',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae',
      'A Brief History of Time'
    ])
);
```

`remove(bookList, "On The Electrodynamics of Moving Bodies")` sollte `["The Hound of the Baskervilles", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"]` zurückgeben.

```js
assert(
  JSON.stringify(remove(bookList, 'On The Electrodynamics of Moving Bodies')) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae'
    ])
);
```

`remove(add(bookList, "A Brief History of Time"), "On The Electrodynamics of Moving Bodies");` sollte `["The Hound of the Baskervilles", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae", "A Brief History of Time"]` entsprechen.

```js
assert(
  JSON.stringify(remove(add(bookList, 'A Brief History of Time'), 'On The Electrodynamics of Moving Bodies')) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae',
      'A Brief History of Time'
    ])
);
```

# --seed--

## --seed-contents--

```js
// The global variable
const bookList = ["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"];

// Change code below this line
function add(bookName) {

  bookList.push(bookName);
  return bookList;
  
  // Change code above this line
}

// Change code below this line
function remove(bookName) {
  const book_index = bookList.indexOf(bookName);
  if (book_index >= 0) {

    bookList.splice(book_index, 1);
    return bookList;

    // Change code above this line
    }
}
```

# --solutions--

```js
// The global variable
const bookList = ["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"];

function add(bookList, bookName) {
  return [...bookList, bookName];
}

function remove(bookList, bookName) {
  const bookListCopy = [...bookList];
  const bookNameIndex = bookList.indexOf(bookName);
  if (bookNameIndex >= 0) {
    bookListCopy.splice(bookNameIndex, 1);
  }
  return bookListCopy;
}
```
