---
id: 587d7dab367417b2b2512b6e
title: "Die Methode every verwenden, um zu prüfen, ob jedes Element in einem Array ein Kriterium erfüllt"
challengeType: 1
forumTopicId: 301312
dashedName: use-the-every-method-to-check-that-every-element-in-an-array-meets-a-criteria
---

# --description--

Die `every`-Methode arbeitet mit Arrays, um zu prüfen, ob *jedes* Element einen bestimmten Test besteht. Sie gibt einen Boolean-Wert zurück – `true`, wenn alle Werte das Kriterium erfüllen, `false`, wenn nicht.

Zum Beispiel würde der folgende Code prüfen, ob jedes Element im `numbers`-Array kleiner als 10 ist:

```js
const numbers = [1, 5, 8, 0, 10, 11];

numbers.every(function(currentValue) {
  return currentValue < 10;
});
```

Die `every`-Methode würde hier `false` zurückgeben.

# --instructions--

Verwenden Sie die `every`-Methode innerhalb der `checkPositive`-Funktion, um zu prüfen, ob jedes Element in `arr` positiv ist. Die Funktion soll einen Boolean-Wert zurückgeben.

# --hints--

Ihr Code sollte die `every`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.every/g));
```

`checkPositive([1, 2, 3, -4, 5])` sollte `false` zurückgeben.

```js
assert.isFalse(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` sollte `true` zurückgeben.

```js
assert.isTrue(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([1, -2, 3, -4, 5])` sollte `false` zurückgeben.

```js
assert.isFalse(checkPositive([1, -2, 3, -4, 5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.every(num => num > 0);
}
```
