---
id: 587d7da9367417b2b2512b67
title: Elemente am Ende eines Arrays mit concat statt push hinzufügen
challengeType: 1
forumTopicId: 301226
dashedName: add-elements-to-the-end-of-an-array-using-concat-instead-of-push
---

# --description--

Funktionale Programmierung dreht sich darum, nicht-mutierende Funktionen zu erstellen und zu verwenden.

Die letzte Aufgabe stellte die `concat`-Methode als eine Möglichkeit vor, Arrays zu einem neuen Array zusammenzuführen, ohne die ursprünglichen Arrays zu verändern. Vergleichen Sie `concat` mit der `push`-Methode. `push` fügt Elemente am Ende desselben Arrays hinzu, auf dem sie aufgerufen wird, wodurch dieses Array verändert wird. Hier ist ein Beispiel:

```js
const arr = [1, 2, 3];
arr.push(4, 5, 6);
```

`arr` hätte dann den veränderten Wert `[1, 2, 3, 4, 5, 6]`, was nicht dem Prinzip der funktionalen Programmierung entspricht.

`concat` bietet eine Möglichkeit, neue Elemente am Ende eines Arrays zusammenzuführen, ohne dabei mutierende Nebeneffekte zu verursachen.

# --instructions--

Ändern Sie die `nonMutatingPush`-Funktion so, dass sie `concat` verwendet, um `newItem` am Ende von `original` zusammenzuführen, ohne die Arrays `original` oder `newItem` zu verändern. Die Funktion soll ein Array zurückgeben.

# --hints--

Ihr Code sollte die `concat`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

Ihr Code sollte die `push`-Methode nicht verwenden.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?push/g));
```

Das `first`-Array darf sich nicht ändern.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

Das `second`-Array darf sich nicht ändern.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingPush([1, 2, 3], [4, 5])` sollte `[1, 2, 3, 4, 5]` zurückgeben.

```js
assert(
  JSON.stringify(nonMutatingPush([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingPush(original, newItem) {
  // Only change code below this line
  return original.push(newItem);

  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingPush(first, second);
```

# --solutions--

```js
function nonMutatingPush(original, newItem) {
  return original.concat(newItem);
}
const first = [1, 2, 3];
const second = [4, 5];
```
