---
id: 587d7b8f367417b2b2512b62
title: Map auf einem Prototype implementieren
challengeType: 1
forumTopicId: 301230
dashedName: implement-map-on-a-prototype
---

# --description--

Wie Sie bereits beim Anwenden von `Array.prototype.map()` oder einfach `map()` gesehen haben, gibt die `map`-Methode ein Array mit der gleichen Länge zurück wie das, auf dem sie aufgerufen wurde. Außerdem verändert sie das ursprüngliche Array nicht, solange ihre Callback-Funktion das nicht tut.

Mit anderen Worten, `map` ist eine reine Funktion, und ihr Ergebnis hängt ausschließlich von ihren Eingaben ab. Außerdem nimmt sie eine weitere Funktion als Argument.

Sie können viel über die `map`-Methode lernen, wenn Sie Ihre eigene Version davon implementieren. Es wird empfohlen, eine `for`-Schleife oder `Array.prototype.forEach()` zu verwenden.

# --instructions--

Schreiben Sie Ihre eigene `Array.prototype.myMap()`, die sich genau wie `Array.prototype.map()` verhalten soll. Sie sollten die eingebaute `map`-Methode nicht verwenden. Die `Array`-Instanz kann in der `myMap`-Methode über `this` aufgerufen werden.

# --hints--

`[23, 65, 98, 5, 13].myMap(item => item * 2)` sollte `[46, 130, 196, 10, 26]` entsprechen.

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item * 2;
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`["naomi", "quincy", "camperbot"].myMap(element => element.toUpperCase())` sollte `["NAOMI", "QUINCY", "CAMPERBOT"]` zurückgeben.

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element.toUpperCase();
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`[1, 1, 2, 5, 2].myMap((element, index, array) => array[index + 1] || array[0])` sollte `[1, 2, 5, 2, 1]` zurückgeben.

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array[index + 1] || array[0];
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

Ihr Code sollte die `map`-Methode nicht verwenden.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?map/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    newArray.push(callback(this[i], i, this));
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const doubled_s = s.myMap(item => item * 2);
```
