---
id: 587d7b8d367417b2b2512b5b
title: Über funktionale Programmierung lernen
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

Funktionale Programmierung ist ein Programmierstil, bei dem Lösungen einfache, isolierte Funktionen sind, ohne Nebeneffekte außerhalb des Gültigkeitsbereichs der Funktion: `INPUT -> PROCESS -> OUTPUT`

Funktionale Programmierung bedeutet:

1) Isolierte Funktionen – es gibt keine Abhängigkeit vom Zustand des Programms, einschließlich globaler Variablen, die sich ändern können

2) Reine Funktionen – dieselbe Eingabe liefert immer dieselbe Ausgabe

3) Funktionen mit begrenzten Nebeneffekten – alle Änderungen oder Mutationen am Zustand des Programms außerhalb der Funktion werden sorgfältig kontrolliert

# --instructions--

Die Mitglieder von freeCodeCamp lieben zufällig Tee.

Im Code-Editor sind die Funktionen `prepareTea` und `getTea` bereits für Sie definiert. Rufen Sie die Funktion `getTea` auf, um 40 Tassen Tee für das Team zu bekommen, und speichern Sie sie in der Variablen `tea4TeamFCC`.

# --hints--

Die Variable `tea4TeamFCC` sollte 40 Tassen Tee für das Team speichern.

```js
assert(tea4TeamFCC.length === 40);
```

Die Variable `tea4TeamFCC` sollte Tassen grünen Tee speichern.

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40); 
```
