---
id: 64221007887f38213fa57827
title: 步骤18
challengeType: 0
dashedName: step-18
---

# --description--

关于这个表达式还有最后一件事。 你实际上并不需要捕获组中的匹配值，因此你可以将其转换为<dfn>非捕获组</dfn>。 这将允许你将字符组合在一起而不保留结果。

要在正则表达式中创建非捕获组，可以在组的左括号后添加 `?:`。 例如，`(?:a|b)` 将匹配 `a` 或 `b`，但不会捕获结果。

更新你的正则表达式以使用非捕获组。

# --hints--

你的 `dollarRegex` 应该使用 `?:`。

```js
assert.match(dollarRegex.source, /\?:/);
```

你的 `dollarRegex` 应该使用非捕获组。

```js
assert.match(dollarRegex.source, /\(\?:.*\)/);
```

你的 `(hundred|thousand|million|billion)` 应该是一个非捕获组。

```js
assert.match(dollarRegex.source, /\(\?:hundred\|thousand\|million\|billion\)/);
```

你的 `dollarRegex` 应该仍然与 `10hundred dollars` 匹配。

```js
let dollarStr = '10hundred dollars'
assert.deepEqual(dollarStr.match(dollarRegex), dollarStr.match(/[0-9]+\s*(?:hundred|thousand|million|billion)?\s+dollars/i));
```

你的 `dollarRegex` 应该仍然与 `10 hundred  dollars` 匹配。

```js
let dollarStr = '10 hundred  dollars'
assert.deepEqual(dollarStr.match(dollarRegex), dollarStr.match(/[0-9]+\s*(?:hundred|thousand|million|billion)?\s+dollars/i));
```

你的 `dollarRegex` 应该仍然与 `100 dollars` 匹配。

```js
let dollarStr = '100 dollars'
assert.deepEqual(dollarStr.match(dollarRegex), dollarStr.match(/[0-9]+\s*(?:hundred|thousand|million|billion)?\s+dollars/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
--fcc-editable-region--
const dollarRegex = /[0-9]+\s*(hundred|thousand|million|billion)?\s+dollars/i;
--fcc-editable-region--

const denyList = [helpRegex, dollarRegex];

const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
