---
id: 56533eb9ac21ba0edf2244b3
title: تحويل درجة مئوية إلى فهرنهايت
challengeType: 1
forumTopicId: 16806
dashedName: convert-celsius-to-fahrenheit
---

# --description--

الصيغة لتحويل من درجة مئوية إلى فهرنهايت هي درجة الحرارة بالمئوية مضروبة في `9/5`، زائد `32`.

أنت تملك متغيرًا اسمه `celsius` يمثل درجة حرارة بالمئوية. استخدم المتغير `fahrenheit` المعرف مسبقًا وقم بتعيين درجة الحرارة المكافئة بالفهرنهايت لدرجة الحرارة المعطاة بالمئوية. استخدم الصيغة المذكورة أعلاه لمساعدتك في تحويل درجة الحرارة من مئوية إلى فهرنهايت.

# --hints--

يجب أن تُرجع الدالة `convertCtoF(0)` رقمًا

```js
assert.isNumber(convertCtoF(0));
```

يجب أن تُرجع الدالة `convertCtoF(-30)` القيمة `-22`

```js
assert.strictEqual(convertCtoF(-30), -22);
```

يجب أن تُرجع الدالة `convertCtoF(-10)` القيمة `14`

```js
assert.strictEqual(convertCtoF(-10), 14);
```

يجب أن تُرجع الدالة `convertCtoF(0)` القيمة `32`

```js
assert.strictEqual(convertCtoF(0), 32);
```

يجب أن تُرجع الدالة `convertCtoF(20)` القيمة `68`

```js
assert.strictEqual(convertCtoF(20), 68);
```

يجب أن تُرجع الدالة `convertCtoF(30)` القيمة `86`

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js
function convertCtoF(celsius) {
  let fahrenheit;
  return fahrenheit;
}

convertCtoF(30);
```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
