---
id: 579e2a2c335b9d72dd32e05c
title: القص والدمج
challengeType: 1
forumTopicId: 301148
dashedName: slice-and-splice
---

# --description--

تُعطى مصفوفتان وفهرس.

انسخ كل عنصر من المصفوفة الأولى إلى المصفوفة الثانية، بالترتيب.

ابدأ بإدخال العناصر عند الفهرس `n` في المصفوفة الثانية.

أرجع المصفوفة الناتجة. يجب أن تبقى المصفوفتان المدخلتان كما هما بعد تنفيذ الدالة.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` يجب أن تُرجع `[4, 1, 2, 3, 5]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` يجب أن تُرجع `["a", 1, 2, "b"]`.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` يجب أن تُرجع `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

يجب إضافة كل العناصر من المصفوفة الأولى إلى المصفوفة الثانية بترتيبها الأصلي. `frankenSplice([1, 2, 3, 4], [], 0)` يجب أن تُرجع `[1, 2, 3, 4]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

يجب أن تبقى المصفوفة الأولى كما هي بعد تنفيذ الدالة.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

يجب أن تبقى المصفوفة الثانية كما هي بعد تنفيذ الدالة.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js
function frankenSplice(arr1, arr2, n) {
  return arr2;
}

frankenSplice([1, 2, 3], [4, 5, 6], 1);
```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  // It's alive. It's alive!
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}

frankenSplice([1, 2, 3], [4, 5], 1);
```
