---
id: a202eed8fc186c8434cb6d61
title: عكس سلسلة نصية
challengeType: 1
forumTopicId: 16043
dashedName: reverse-a-string
---

# --description--

اعكس السلسلة النصية المقدمة وأعد السلسلة المعكوسة.

على سبيل المثال، يجب أن تتحول `"hello"` إلى `"olleh"`.

# --hints--

يجب أن تُعيد الدالة `reverseString("hello")` سلسلة نصية.

```js
assert.isString(reverseString('hello'));
```

يجب أن تُعيد الدالة `reverseString("hello")` السلسلة النصية `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

يجب أن تُعيد الدالة `reverseString("Howdy")` السلسلة النصية `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

يجب أن تُعيد الدالة `reverseString("Greetings from Earth")` السلسلة النصية `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js
function reverseString(str) {
  return str;
}

reverseString('hello');
```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
