---
id: a24c1a4622e3c05097f71d67
title: أين أنتمي
challengeType: 1
forumTopicId: 16094
dashedName: where-do-i-belong
---

# --description--

أعد أدنى فهرس يجب إدخال قيمة (المعلمة الثانية) فيه داخل مصفوفة (المعلمة الأولى) بعد فرزها. يجب أن تكون القيمة المعادة رقمًا.

على سبيل المثال، يجب أن تُعيد الدالة `getIndexToIns([1,2,3,4], 1.5)` القيمة `1` لأنها أكبر من `1` (الفهرس 0)، لكنها أقل من `2` (الفهرس 1).

بنفس الطريقة، يجب أن تُعيد الدالة `getIndexToIns([20,3,5], 19)` القيمة `2` لأنه بعد فرز المصفوفة ستصبح `[3,5,20]` و`19` أقل من `20` (الفهرس 2) وأكبر من `5` (الفهرس 1).

# --hints--

يجب أن تُعيد الدالة `getIndexToIns([10, 20, 30, 40, 50], 35)` رقمًا.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

يجب أن تُعيد الدالة `getIndexToIns([10, 20, 30, 40, 50], 35)` القيمة `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

يجب أن تُعيد الدالة `getIndexToIns([10, 20, 30, 40, 50], 30)` رقمًا.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 30));
```

يجب أن تُعيد الدالة `getIndexToIns([10, 20, 30, 40, 50], 30)` القيمة `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

يجب أن تُعيد الدالة `getIndexToIns([40, 60], 50)` رقمًا.

```js
assert.isNumber(getIndexToIns([40, 60], 50));
```

يجب أن تُعيد الدالة `getIndexToIns([40, 60], 50)` القيمة `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

يجب أن تُعيد الدالة `getIndexToIns([3, 10, 5], 3)` رقمًا.

```js
assert.isNumber(getIndexToIns([3, 10, 5], 3));
```

يجب أن تُعيد الدالة `getIndexToIns([3, 10, 5], 3)` القيمة `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

يجب أن تُعيد الدالة `getIndexToIns([5, 3, 20, 3], 5)` رقمًا.

```js
assert.isNumber(getIndexToIns([5, 3, 20, 3], 5));
```

يجب أن تُعيد الدالة `getIndexToIns([5, 3, 20, 3], 5)` القيمة `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

يجب أن تُعيد الدالة `getIndexToIns([2, 20, 10], 19)` القيمة `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

يجب أن تُعيد الدالة `getIndexToIns([2, 20, 10], 19)` رقمًا.

```js
assert.isNumber(getIndexToIns([2, 20, 10], 19));
```

يجب أن تُعيد الدالة `getIndexToIns([2, 5, 10], 15)` القيمة `3`.

```js
assert.strictEqual(getIndexToIns([2, 5, 10], 15), 3);
```

يجب أن تُعيد الدالة `getIndexToIns([2, 5, 10], 15)` رقمًا.

```js
assert.isNumber(getIndexToIns([2, 5, 10], 15));
```

يجب أن تُعيد الدالة `getIndexToIns([], 1)` رقمًا.

```js
assert.isNumber(getIndexToIns([], 1));
```

يجب أن تُعيد الدالة `getIndexToIns([], 1)` القيمة `0`.

```js
assert.strictEqual(getIndexToIns([], 1), 0);
```

# --seed--

## --seed-contents--

```js
function getIndexToIns(arr, num) {
  return num;
}

getIndexToIns([40, 60], 50);
```

# --solutions--

```js
function getIndexToIns(arr, num) {
  arr = arr.sort((a, b) => a - b);

  for (let i = 0; i < arr.length; i++) {
    if (arr[i] >= num) {
      return i;
    }
  }

  return arr.length;
}

getIndexToIns([40, 60], 50);
```
