---
id: a26cbbe9ad8655a977e1ceb5
title: إيجاد أطول كلمة في سلسلة نصية
challengeType: 1
forumTopicId: 16015
dashedName: find-the-longest-word-in-a-string
---

# --description--

أرجع طول أطول كلمة في الجملة المقدمة.

يجب أن تكون استجابتك رقمًا.

# --hints--

يجب أن تُرجع `findLongestWordLength("The quick brown fox jumped over the lazy dog")` رقمًا.

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

يجب أن تُرجع `findLongestWordLength("The quick brown fox jumped over the lazy dog")` القيمة `6`.

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

يجب أن تُرجع `findLongestWordLength("May the force be with you")` القيمة `5`.

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

يجب أن تُرجع `findLongestWordLength("Google do a barrel roll")` القيمة `6`.

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

يجب أن تُرجع `findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` القيمة `8`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

يجب أن تُرجع `findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` القيمة `19`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js
function findLongestWordLength(str) {
  return str.length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
