---
id: a302f7aae1aa3152a5b413bc
title: حساب مضروب عدد
challengeType: 1
forumTopicId: 16013
dashedName: factorialize-a-number
---

# --description--

أرجع مضروب العدد الصحيح المقدم.

إذا تم تمثيل العدد الصحيح بالحرف `n`، فإن المضروب هو حاصل ضرب كل الأعداد الصحيحة الموجبة التي تقل أو تساوي `n`.

غالبًا ما يُمثّل المضروب بالاختصار `n!`

على سبيل المثال: `5! = 1 * 2 * 3 * 4 * 5 = 120`

سيتم تزويد الدالة فقط بأعداد صحيحة أكبر من أو تساوي الصفر.

# --hints--

يجب أن تُرجع `factorialize(5)` رقمًا.

```js
assert.isNumber(factorialize(5));
```

يجب أن تُرجع `factorialize(5)` القيمة `120`.

```js
assert.strictEqual(factorialize(5), 120);
```

يجب أن تُرجع `factorialize(10)` القيمة `3628800`.

```js
assert.strictEqual(factorialize(10), 3628800);
```

يجب أن تُرجع `factorialize(20)` القيمة `2432902008176640000`.

```js
assert.strictEqual(factorialize(20), 2432902008176640000);
```

يجب أن تُرجع `factorialize(0)` القيمة `1`.

```js
assert.strictEqual(factorialize(0), 1);
```

# --seed--

## --seed-contents--

```js
function factorialize(num) {
  return num;
}

factorialize(5);
```

# --solutions--

```js
function factorialize(num) {
  return num < 1 ? 1 : num * factorialize(num - 1);
}

factorialize(5);
```
