---
id: a6e40f1041b06c996f7b2406
title: الباحثون يحتفظون
challengeType: 1
forumTopicId: 16016
dashedName: finders-keepers
---

# --description--

أنشئ دالة تبحث في مصفوفة `arr` وتُرجع أول عنصر فيها ينجح في "اختبار الحقيقة". هذا يعني أنه إذا كان العنصر `x`، فإن "اختبار الحقيقة" ينجح إذا كانت نتيجة `func(x)` هي `true`. إذا لم ينجح أي عنصر في الاختبار، أرجع `undefined`.

# --hints--

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` يجب أن تُرجع `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` يجب أن تُرجع `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js
function findElement(arr, func) {
  let num = 0;
  return num;
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
