---
id: a77dbc43c33f39daa4429b4f
title: من هو بو
challengeType: 1
forumTopicId: 16000
dashedName: boo-who
---

# --description--

تحقق مما إذا كانت القيمة مصنفة كقيمة منطقية أولية. أرجع `true` أو `false`.

القيم المنطقية الأولية هي `true` و `false`.

# --hints--

يجب أن تُرجع الدالة `booWho(true)` القيمة `true`.

```js
assert.isTrue(booWho(true));
```

يجب أن تُرجع الدالة `booWho(false)` القيمة `true`.

```js
assert.isTrue(booWho(false));
```

يجب أن تُرجع الدالة `booWho([1, 2, 3])` القيمة `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

يجب أن تُرجع الدالة `booWho([].slice)` القيمة `false`.

```js
assert.isFalse(booWho([].slice));
```

يجب أن تُرجع الدالة `booWho({ "a": 1 })` القيمة `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

يجب أن تُرجع الدالة `booWho(1)` القيمة `false`.

```js
assert.isFalse(booWho(1));
```

يجب أن تُرجع الدالة `booWho(NaN)` القيمة `false`.

```js
assert.isFalse(booWho(NaN));
```

يجب أن تُرجع الدالة `booWho("a")` القيمة `false`.

```js
assert.isFalse(booWho('a'));
```

يجب أن تُرجع الدالة `booWho("true")` القيمة `false`.

```js
assert.isFalse(booWho('true'));
```

يجب أن تُرجع الدالة `booWho("false")` القيمة `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js
function booWho(bool) {
  return bool;
}

booWho(null);
```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

booWho(null);
```
