---
id: ab6137d4e35944e21037b769
title: كتابة الجملة بحروف كبيرة في بداية كل كلمة
challengeType: 1
forumTopicId: 16088
dashedName: title-case-a-sentence
---

# --description--

أعد السلسلة النصية المقدمة مع جعل الحرف الأول من كل كلمة بحرف كبير. تأكد من أن بقية الكلمة بحروف صغيرة.

لهذا التمرين، يجب أيضًا كتابة الكلمات الوصل مثل `the` و `of` بحروف كبيرة.

# --hints--

`titleCase("I'm a little tea pot")` يجب أن تُعيد سلسلة نصية.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` يجب أن تُعيد السلسلة النصية `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` يجب أن تُعيد السلسلة النصية `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` يجب أن تُعيد السلسلة النصية `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js
function titleCase(str) {
  return str;
}

titleCase("I'm a little tea pot");
```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}

titleCase("I'm a little tea pot");
```
