---
id: acda2fb1324d9b0fa741e6b5
title: تأكيد النهاية
challengeType: 1
forumTopicId: 16006
dashedName: confirm-the-ending
---

# --description--

تحقق مما إذا كانت سلسلة نصية (المعلمة الأولى، `str`) تنتهي بالسلسلة النصية الهدف المعطاة (المعلمة الثانية، `target`).

يمكن حل هذا التحدي باستخدام الطريقة `.endsWith()`، التي أُدخلت في ES2015. لكن لغرض هذا التحدي، نود منك استخدام إحدى طرق السلاسل النصية الفرعية في JavaScript بدلاً من ذلك.

# --hints--

يجب أن تُرجع `confirmEnding("Bastian", "n")` القيمة `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

يجب أن تُرجع `confirmEnding("Congratulation", "on")` القيمة `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

يجب أن تُرجع `confirmEnding("Connor", "n")` القيمة `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

يجب أن تُرجع `confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` القيمة `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

يجب أن تُرجع `confirmEnding("He has to give me a new name", "name")` القيمة `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

يجب أن تُرجع `confirmEnding("Open sesame", "same")` القيمة `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

يجب أن تُرجع `confirmEnding("Open sesame", "sage")` القيمة `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

يجب أن تُرجع `confirmEnding("Open sesame", "game")` القيمة `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing", "mountain")` يجب أن تُرجع `false`.

```js
assert.isFalse(
  confirmEnding(
    'If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing',
    'mountain'
  )
);
```

يجب أن تُرجع `confirmEnding("Abstraction", "action")` القيمة `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

يجب ألا يستخدم كودك الطريقة المدمجة `.endsWith()` لحل التحدي.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js
function confirmEnding(str, target) {
  return str;
}

confirmEnding('Bastian', 'n');
```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
