---
id: afcc8d540bea9ea2669306b6
title: كرر سلسلة نصية كرر سلسلة نصية
challengeType: 1
forumTopicId: 16041
dashedName: repeat-a-string-repeat-a-string
---

# --description--

كرر السلسلة النصية المعطاة `str` (المعلمة الأولى) عدد `num` مرات (المعلمة الثانية). أرجع سلسلة نصية فارغة إذا لم يكن `num` عددًا موجبًا. لغرض هذا التحدٍّ، لا تستخدم الطريقة المدمجة `.repeat()`.

# --hints--

`repeatStringNumTimes("*", 3)` يجب أن تُرجع السلسلة النصية `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` يجب أن تُرجع السلسلة النصية `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` يجب أن تُرجع السلسلة النصية `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` يجب أن تُرجع السلسلة النصية `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` يجب أن تُرجع السلسلة النصية `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` يجب أن تُرجع سلسلة نصية فارغة (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

لا يجب استخدام الطريقة المدمجة `repeat()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

`repeatStringNumTimes("abc", 0)` يجب أن تُرجع `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

# --seed--

## --seed-contents--

```js
function repeatStringNumTimes(str, num) {
  return str;
}

repeatStringNumTimes('abc', 3);
```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
