---
id: 5f331e55dfab7a896e53c3a1
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

The `title` is one of several elements that provide extra information not visible on the web page, but it is useful for search engines or how the page gets displayed.

Verschachtele innerhalb des `head` Elements ein `meta-` -Element mit einem Attribut mit dem Namen `charset`, das auf den Wert `utf-8` gesetzt ist, um dem Browser mitzuteilen, wie Zeichen für die Seite entschlüsselt werden sollen.

# --hints--

Du solltest ein `meta`-Tag haben.

```js
assert.match(code, /<meta.*?\/?>/is);
```

Das `meta` Element ist ein leeres Element, es sollte kein schließendes Tag `</meta>`haben.

```js
assert.notMatch(code, /<\/meta>/i);
```

Your `meta` tag should have a `charset` attribute.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Your `charset` attribute should have a value of `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
