---
id: 5f33071498eb2472b87ddee4
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS is the language used to style an HTML document. It describes how HTML elements should be displayed on the screen.

Wie du in den letzten Schritten der Katzenfoto-App bereits gelernt hast, ist eine grundlegende Struktur erforderlich, um deine Webseite zu erstellen. Jedes HTML-Dokument sollte eine `DOCTYPE` Deklaration und ein `HTML-` Element enthalten. Der `DOCTYPE` teilt dem Browser mit, in welcher HTML-Version das Dokument verfasst ist. Und das `HTML-` Element stellt das Wurzelelement dar, das alle anderen Elemente enthält.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Add the `<!DOCTYPE html>` tag, and an `html` element with a `lang` attribute of `en`.

# --hints--

Du solltest eine `<!DOCTYPE html>`-Deklaration haben.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Du solltest ein öffnendes `<html lang="en">` Element haben, das die Sprache auf Englisch stellt.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Du solltest ein abschließendes `</html>`-Tag haben.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
