---
id: 5f331e55dfab7a896e53c3a1
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Der `title` ist eines von mehreren Elementen, die zus채tzliche Informationen enthalten, die auf der Webseite nicht sichtbar aber n체tzlich f체r Suchmaschinen sind oder die Art und Weise 채ndern, wie die Seite dargestellt wird.

Inside the `head` element, nest a `meta` element with an attribute named `charset` set to the value `utf-8` to tell the browser how to encode characters for the page.

# --hints--

Du solltest ein `meta`-Tag haben.

```js
assert.match(code, /<meta.*?\/?>/is);
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Dein `meta`-Tag sollte ein `charset`-Attribut haben.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Dein `charset` sollte einen Wert von `utf-8` haben.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
