---
id: 64dc97005bc7943e2978df0a
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

现在你可以删除你的 `print()` 调用。 然后，在 `move()` 函数中，删除 `pass` 关键字并打印 `rods` 字典的内容。

# --hints--

你的代码中不应该有 `print(number_of_moves)`。

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

你应该删除 `pass` 关键字并打印 `rods` 字典。

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
