---
id: 64dc922df2919e3c38cead61
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

在代码顶部，声明一个名为 `NUMBER_OF_DISKS` 的变量来存储磁盘数量，并赋予其 `3` 的值。 然后，用新变量替换传递给 `range()` 函数的第一个参数。

# --hints--

你应该声明一个名为 `NUMBER_OF_DISKS` 的变量。

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

`NUMBER_OF_DISKS` 变量的值应为 `3`。

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

你应该将 `range()` 函数的第一个参数的值更改为 `NUMBER_OF_DISKS` 变量。

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
