---
id: 64dcdee07a14f0299ec30559
title: 步骤 22
challengeType: 20
dashedName: step-22
---

# --description--

你编写了代码来查找杆之间允许的移动，但实际移动可以是双向的。 在 `print()` 调用之后，声明一个名为 `forward` 的变量并将其设置为 `False`。 你将使用该变量来检查需要在杆之间移动圆盘的方向。

# --hints--

你应该声明一个名为 `forward` 的变量。

```js
({ test: () => assert.match(code, /forward\s*=\s*./) })
```

你的 `forward` 变量应设置为 `False`。

```js
({ test: () => assert.match(code, /forward\s*=\s*False/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
