---
id: 64dc90949ff85e3b37be40d0
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

`range()` 函数返回一个不可变的数字序列。 如你所见，`rods['A']` 的数据类型是 `range`，但你希望它是一个列表。

将你的 `range()` 调用传递给 `list()` 函数来执行此操作。

# --hints--

你应该将 `range()` 调用传递给 `list()` 函数。

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
