---
id: 64dc963d866fff3dd0329993
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

打印你在上一步中声明的变量，并随意更改圆盘的数量，以查看所需的最小移动次数增加的速度。

# --hints--

你应该打印 `number_of_moves` 变量。

```js
({ test: () => assert.match(code, /print\s*\(\s*number_of_moves\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move():
    pass
```
