---
id: 64dc916138dbe03bdfe6cec6
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

汉诺塔的目标是将所有圆盘移至最后一根杆子。 要做到这一点，你必须遵循三个简单的规则：

1. 你只能移动最顶层的圆盘
2. 一次只能移动一个圆盘
3. 不能将较大的圆盘放在较小的圆盘上

在现有代码下方，声明一个名为 `move` 的空函数。 稍后，你将使用该功能在杆之间移动圆盘。 现在，为了避免错误，请在函数体中使用 `pass` 关键字。

# --hints--

你应该声明一个名为 `move` 的空函数。 请记住在函数体内使用 `pass` 关键字并正确缩进。

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
