---
id: 64dc92a9718fb53ca3d1bad7
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

汉诺塔谜题只需 2<sup>n</sup> - 1 步即可解决，其中 n 是圆盘的数量。 声明一个名为 `number_of_moves` 的变量，并将总移动次数分配给该变量。

Python 中的幂运算符是 `**`。

# --hints--

你应该声明一个名为 `number_of_moves` 的变量。

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

`number_of_moves` 的值应该是计算移动次数的表达式。

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
