---
id: 64dc8ea01436383a88256d7f
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

谜题从第一根杆上堆积的圆盘开始，按大小依次减小，最小的圆盘放在顶部，最大的圆盘放在底部。 你需要在第一个列表中填充代表各种圆盘大小的数字。

无需手动将项目添加到第一个列表，而是使用 `range()` 函数生成从 `3` 到 `1` 倒数的数字序列，并将其分配给 `rods['A']`。 这里，`3` 表示堆底部最大的圆盘，而 `1` 表示堆顶部最小的圆盘。

语法为 `range(x, y, h)`，其中 `x` 是起始整数（含），`y` 是最后一个整数（不含），`h` 是序列中一个数字与下一个数字之间的差值。

# --hints--

你应该使用 `range()` 函数将一系列数字分配给 `rods['A']`。 调用 `range()` 函数的语法是 `range(x, y, h)`。

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

你应该使用 `range()` 函数来分配从 `3` 到 `1` 到 `rods['A']` 的数字序列。

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
