---
id: 64dc985e6720a23edac63b51
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

杆之间允许的圆盘运动呈现出每三次移动发生一次的重复模式。 例如，在第一、第四和第七步中，允许杆 `A` 和杆 `C` 之间移动，其中手数与 3 之差的余数等于 1。

在先前创建的 `for` 循环内，将现有的 `print()` 调用替换为在 `(i + 1) % 3 == 1` 时触发的 `if` 语句。 在此 `if` 语句中，使用 f 字符串打印 `f'Move {i + 1} allowed between {source} and {target}'`。 请注意，`i + 1` 是移动次数，因为在第一次迭代期间 `i` 为零。

# --hints--

你应该有一个 `if` 语句，该语句应在 `(i + 1) % 3 == 1` 时执行。 

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

你应该使用 f 字符串打印字符串 `f'Move {i + 1} allowed between {source} and {target}'`。

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
