---
id: 64dcd9bbc2268127e7898d77
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

最后，当除以3的移动数值没有余数时，允许在 `'B'` 和 `'C'` 之间移动。

为此添加 `elif` 语句。 然后，如果满足条件，则打印适当的字符串。 

# --hints--

你应该有一个 `elif` 语句在 `remainder == 0` 时执行。 

```js
({test: () => assert.match(code, /elif\s+remainder\s*==\s*0\s*:/)})
```

你应该打印字符串 `f'Move {i + 1} allowed between {auxiliary} and {target}'`。 

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*auxiliary\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
