---
id: 64dc9004f658183afa85f4f1
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

现在通过将 `'A'` 键传递给 `type()` 函数来检查其数据类型，并在终端上打印结果。

# --hints--

你应该将 `'A'` 键传递给 `type()` 函数。

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

你应该使用 `print()` 函数打印 `rods['A']` 的类型。

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
