---
id: 64dc8b1faf9c60304ca8a9b6
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在这个项目中，你将解决被称为汉诺塔的数学难题。 该问题由三根杆和多个不同直径的圆盘组成。

这个谜题的目标是将圆盘从第一根杆移动到第三根杆，并遵循特定的规则，限制将较大的圆盘放在较小的圆盘上。

首先创建一个名为 `rods` 的空字典来表示杆。 

# --hints--

你应该有一个名为 `rods` 的变量。

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

你的 `rods` 变量应该是一个空字典。

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
