---
id: 64dc976bf864693e668d67e8
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

现在调用你的函数并查看终端上的输出。

# --hints--

你应该调用你的 `move()` 函数。

```js
({test: () => assert.match(code, /^move\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    print(rods)
--fcc-editable-region--

--fcc-editable-region--
```
