---
id: 64dc90e837670a3ba26a1e61
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

现在类型已经符合要求，`list`，你可以删除 `print()` 调用。

# --hints--

你应该从代码中删除 `print(type(rods['A']))`。

```js
({ test: () => assert.match(code, /\}(?!\s+print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")A\1\s*\]\s*\)\s*\))/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
