---
id: 64ddd280d69fa818ca782d50
title: 步骤 29
challengeType: 20
dashedName: step-29
---

# --description--

可以看到，圆盘 1 每移动三次就会来回移动一次。 发生这种情况是因为你仍然需要注意其他杆之间的移动。 不要重复你在前几个步骤中编写的相同代码并更改杆，而是最好将该代码移到函数中以在每个条件语句中调用。 声明一个名为 `make_allowed_move()` 的空函数，并且不要忘记 `pass` 关键字。

# --hints--

你应该声明一个名为 `make_allowed_move` 的空函数。 请记住在函数体内使用 `pass` 关键字并正确缩进。

```js
({ test: () => assert(runPython(`
      import inspect
      inspect.isfunction(make_allowed_move)
  `)) 
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if not rods[target]:
                forward = True
            elif rods[source] and rods[source][-1] < rods[target][-1]:
                forward = True              
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
            else:
                print(f'Moving disk {rods[target][-1]} from {target} to {source}')
                rods[source].append(rods[target].pop())
            
            # display our progress
            print(rods)
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
