---
id: 64de1469fccce61940f498c4
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

`rods` 字典将表示三根杆及其圆盘。 将字符串 `'A'`、`'B'` 和 `'C'` 作为键，并将它们分别设置为空列表。

# --hints--

你的 `rods` 字典应该有一个 `'A'` 键。

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` 应该是一个空列表。

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

你的 `rods` 字典应该有一个 `'B'` 键。

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` 应该是一个空列表。

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

你的 `rods` 字典应该有一个 `'C'` 键。

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` 应该是一个空列表。

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
