---
id: 64df496c6a8ddf1b38db1ed6
title: 步骤 50
challengeType: 20
dashedName: step-50
---

# --description--

在前面的步骤中，你编写了代码将子问题的最大圆盘移动到目标杆。

现在，你需要做的就是添加另一个递归调用来移动你已经移位的 `n - 1` 个圆盘。 复制第一个递归调用并将其粘贴到 `if` 块的末尾。

请注意，函数参数的顺序不正确。 尝试找出正确的顺序。

# --hints--

你应该在 `move` 函数末尾调用 `move(n - 1,  auxiliary, source, target)`。

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
