---
id: 64df47b32b92301a815d5ef8
title: 步骤 49
challengeType: 20
dashedName: step-49
---

# --description--

首先，刚刚添加的递归调用处理将 `n - 1` 个圆盘移动到第二根杆的子问题。

因此，`target` 参数对应于第二根杆，而 `auxiliary` 参数对应于第三根杆。 请记住，随着递归的进行，它们将继续交换。

修复递归调用中交换 `target` 和 `auxiliary` 的参数顺序。

# --hints--

你应该修改 `move(n - 1, source, auxiliary, target)` 调用中的参数顺序。

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
