---
id: 64e34146860065146733883b
title: 步骤 53
challengeType: 20
dashedName: step-53
---

# --description--

尽管递归有时不太容易理解，但它使你能够创建更简洁的代码。 在这种情况下，你甚至不需要区分圆盘的偶数和奇数。 

将 `NUMBER_OF_DISKS` 设置为 `5` 并检查输出。

# --hints--

你应该将 `NUMBER_OF_DISKS` 变量设置为 5。

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 5) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
A = list(range(NUMBER_OF_DISKS, 0, -1))
B = []
C = []

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        target.append(source.pop())
        
        # display our progress
        print(A, B, C, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, A, B, C)
--fcc-editable-region--
```
