---
id: 650c6082e5586f9e3acfcd3b
title: 第 52 步
challengeType: 20
dashedName: step-52
---

# --description--

现在，删除 `rods` 字典并将其键分别转换为变量 `A`、`B` 和 `C`，保留它们的值。 重构你的代码以反映这些变化。 如果你在终端上看到输出，则表示你正确执行了此步骤。 

# --hints--

你应该将 `rods['A']` 转换为名为 `A` 的变量并保留其值。

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

你应该将 `rods['B']` 转换为名为 `B` 的变量并保留其值。

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

你应该将 `rods['C']` 转换为名为 `C` 的变量并保留其值。

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

你应该修改 `print()` 调用以打印 `A`、`B`、`C`，而不是 `rods` 对象。 在 `print()` 调用中保留换行符。

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

你的代码中应该有 `target.append(source.pop())`。

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```

