---
id: 65782342850feb3b8d62f936
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

由于你将多次使用表达式 `(i + 1) % 3`，因此将其存储在变量中会很方便。

在 `if` 语句上方，声明一个 `remainder` 变量并将值 `(i + 1) % 3` 分配给该变量。

# --hints--

你应该在 `if` 块上方声明一个 `remainder` 变量。

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=/m) })
```

你的 `remainder` 变量的值应该是表达式 `(i + 1) % 3`。

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=\s*\(\s*i\s*\+\s*1\s*\)\s*%\s*3/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
