---
id: 66ec4c8e9878d8441956516f
title: Crea una tabla de catálogo de libros
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

Cumple las historias de usuario a continuación y haz que todas las pruebas pasen para completar el laboratorio.

**Instrucciones:**

1. Debe crear un elemento `table` que liste información de libros.
1. Su tabla debe tener un elemento de encabezado de tabla con una fila.
1. La fila en el elemento de encabezado de su tabla debe tener cuatro elementos de encabezado de tabla, con el texto de `Title`, `Author`, `Genre`, y `Publication Year`, en ese orden.
1. Su tabla debe tener un elemento de cuerpo de tabla con al menos cinco filas.
1. Cada fila en el cuerpo de su tabla debe tener cuatro elementos de datos de tabla que muestren el Título, Autor, Género y Año de Publicación del libro.
1. Su tabla debe tener un elemento de pie de tabla con una fila.
1. La fila en el elemento de pie de su tabla debe tener un elemento de datos de tabla que abarque cuatro columnas y tenga el texto `Total Books: N` donde `N` debe ser reemplazado por el número de libros en su tabla.

# --hints--

Debe tener un elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Debe tener un elemento `thead` dentro de su elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

Debe tener un elemento `tr` dentro de su elemento `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

Debe tener cuatro elementos `th` dentro de la fila de su elemento `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

Sus cuatro elementos `th` deben tener el texto `Title`, `Author`, `Genre`, y `Publication Year`, en ese orden.

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent, 'Title');
assert.equal(ths[1]?.textContent, 'Author');
assert.equal(ths[2]?.textContent, 'Genre');
assert.equal(ths[3]?.textContent, 'Publication Year');
```

Debe tener un elemento `tbody` dentro de su elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

Su elemento `tbody` debe tener al menos cinco filas.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

Cada fila en su elemento `tbody` debe tener exactamente cuatro elementos `td` como hijos.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

Cada elemento `td` en su cuerpo de tabla debe tener texto con información del libro.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.length, 1);
});
```

Debe tener un elemento `tfoot` dentro de su elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

Debe tener exactamente un elemento `tr` en su elemento `tfoot`.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

El elemento `td` en la fila de su elemento `tfoot` debe tener su atributo `colspan` configurado a `4`.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

El elemento `td` en la fila de su elemento `tfoot` debe tener el texto `Total Books: [número de libros en su tabla]`.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent, `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
