---
id: 619665c9abd72906f3ad30f9
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Costruirai un felice Flappy Penguin, esercitandoti ulteriormente con le trasformazioni e le animazioni CSS.

Inizia collegando il tuo foglio di stile alla pagina.

# --hints--

Il tuo codice dovrebbe avere un elemento `link`.

```js
assert.exists(document.querySelector('link'));
```

Il tuo elemento `link` dovrebbe essere all'interno del tuo elemento `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con valore `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con valore `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
