---
id: 587d7fab367417b2b2512bd9
title: 산점도 원에 레이블 추가
challengeType: 6
forumTopicId: 301477
dashedName: add-labels-to-scatter-plot-circles
---

# --description--

산점도의 점에 레이블을 만들기 위해 텍스트를 추가할 수 있습니다.

목표는 `x`의 각 항목에 대해 첫 번째(`y`)와 두 번째(`dataset`) 필드의 쉼표로 구분된 값을 표시하는 것입니다.

`text` 노드는 SVG에서 위치를 지정하기 위해 `x`와 `y` 속성이 필요합니다. 이 도전 과제에서는 높이를 결정하는 `y` 값이 `circle`가 `cy` 속성에 사용하는 값과 같을 수 있습니다. `x` 값은 `cx`의 `circle` 값보다 약간 커서 레이블이 보이게 됩니다. 이렇게 하면 레이블이 표시된 점의 오른쪽으로 밀립니다.

# --instructions--

`text` 요소를 사용해 산점도의 각 점에 레이블을 붙이세요. 레이블의 텍스트는 쉼표와 공백으로 구분된 두 값이어야 합니다. 예를 들어, 첫 번째 점의 레이블은 `34, 78`입니다. `x` 속성은 `5`의 `cx` 속성에 사용한 값보다 `circle` 단위 더 크게 설정하세요. `y` 속성은 `cy`의 `circle` 값과 같은 방식으로 설정하세요.

# --hints--

코드에는 10개의 `text` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

첫 번째 레이블은 텍스트가 `34, 78`, `x` 값이 `39`, `y` 값이 `422`이어야 합니다.

```js
const labelOne = document.querySelector('text');
assert.strictEqual(labelOne?.textContent, '34, 78');
assert.strictEqual(labelOne?.getAttribute('x'), '39');
assert.strictEqual(labelOne?.getAttribute('y'), '422');
```

두 번째 레이블은 텍스트가 `109, 280`, `x` 값이 `114`, `y` 값이 `220`이어야 합니다.

```js
const labelTwo = document.querySelectorAll('text')[1];
assert.strictEqual(labelTwo?.textContent, '109, 280');
assert.strictEqual(labelTwo?.getAttribute('x'), '114');
assert.strictEqual(labelTwo?.getAttribute('y'), '220');
```

세 번째 레이블은 텍스트가 `310, 120`, `x` 값이 `315`, `y` 값이 `380`이어야 합니다.

```js
const labelThree = document.querySelectorAll('text')[2];
assert.strictEqual(labelThree?.textContent, '310, 120');
assert.strictEqual(labelThree?.getAttribute('x'), '315');
assert.strictEqual(labelThree?.getAttribute('y'), '380');
```

네 번째 레이블은 텍스트가 `79, 411`, `x` 값이 `84`, `y` 값이 `89`이어야 합니다.

```js
const labelFour = document.querySelectorAll('text')[3];
assert.strictEqual(labelFour?.textContent, '79, 411');
assert.strictEqual(labelFour?.getAttribute('x'), '84');
assert.strictEqual(labelFour?.getAttribute('y'), '89');
```

다섯 번째 레이블은 텍스트가 `420, 220`, `x` 값이 `425`, `y` 값이 `280`이어야 합니다.

```js
const labelFive = document.querySelectorAll('text')[4];
assert.strictEqual(labelFive?.textContent, '420, 220');
assert.strictEqual(labelFive?.getAttribute('x'), '425');
assert.strictEqual(labelFive?.getAttribute('y'), '280');
```

여섯 번째 레이블은 텍스트가 `233, 145`, `x` 값이 `238`, `y` 값이 `355`이어야 합니다.

```js
const labelSix = document.querySelectorAll('text')[5];
assert.strictEqual(labelSix?.textContent, '233, 145');
assert.strictEqual(labelSix?.getAttribute('x'), '238');
assert.strictEqual(labelSix?.getAttribute('y'), '355');
```

일곱 번째 레이블은 텍스트가 `333, 96`, `x` 값이 `338`, `y` 값이 `404`이어야 합니다.

```js
const labelSeven = document.querySelectorAll('text')[6];
assert.strictEqual(labelSeven?.textContent, '333, 96');
assert.strictEqual(labelSeven?.getAttribute('x'), '338');
assert.strictEqual(labelSeven?.getAttribute('y'), '404');
```

여덟 번째 레이블은 텍스트가 `222, 333`, `x` 값이 `227`, `y` 값이 `167`이어야 합니다.

```js
const labelEight = document.querySelectorAll('text')[7];
assert.strictEqual(labelEight?.textContent, '222, 333');
assert.strictEqual(labelEight?.getAttribute('x'), '227');
assert.strictEqual(labelEight?.getAttribute('y'), '167');
```

아홉 번째 레이블은 텍스트가 `78, 320`, `x` 값이 `83`, `y` 값이 `180`이어야 합니다.

```js
const labelNine = document.querySelectorAll('text')[8];
assert.strictEqual(labelNine?.textContent, '78, 320');
assert.strictEqual(labelNine?.getAttribute('x'), '83');
assert.strictEqual(labelNine?.getAttribute('y'), '180');
```

열 번째 레이블은 텍스트가 `21, 123`, `x` 값이 `26`, `y` 값이 `377`이어야 합니다.

```js
const labelTen = document.querySelectorAll('text')[9];
assert.strictEqual(labelTen?.textContent, '21, 123');
assert.strictEqual(labelTen?.getAttribute('x'), '26');
assert.strictEqual(labelTen?.getAttribute('y'), '377');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg.selectAll('text').data(dataset).enter().append('text');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', d => d[0] + 5)
      .attr('y', d => h - d[1])
      .text(d => d[0] + ', ' + d[1]);
  </script>
</body>
```
