---
id: 587d7fac367417b2b2512bde
title: 미리 정의된 스케일을 사용해 요소 배치하기
challengeType: 6
forumTopicId: 301494
dashedName: use-a-pre-defined-scale-to-place-elements
---

# --description--

스케일을 설정했으니 이제 산점도를 다시 매핑할 시간입니다. 스케일은 `x`와 `y` 원시 데이터를 SVG에 맞게 올바르게 렌더링되는 값으로 변환하는 처리 함수와 같습니다. 이들은 데이터를 화면의 플로팅 영역 안에 유지합니다.

SVG 도형의 좌표 속성 값을 스케일링 함수로 설정합니다. 여기에는 `x` 또는 `y` 요소의 `rect`와 `text` 속성, 또는 `cx`의 `cy`와 `circles`가 포함됩니다. 예시는 다음과 같습니다:

```js
shape.attr('x', d => xScale(d[0]));
```

스케일은 도형 좌표 속성을 설정해 데이터 포인트를 SVG에 배치합니다. 툴팁이나 레이블의 `text()` 메서드처럼 실제 데이터 값을 표시할 때는 스케일을 적용할 필요가 없습니다.

# --instructions--

`xScale`와 `yScale`을 사용해 `circle`와 `text` 도형을 SVG에 배치하세요. `circles`에는 스케일을 적용해 `cx`와 `cy` 속성을 설정하세요. 반지름은 `5` 단위로 지정하세요.

`text` 요소에는 스케일을 적용해 `x`와 `y` 속성을 설정하세요. 레이블은 점 오른쪽에 오프셋되어야 합니다. 이를 위해 `10` 데이터 값에 `x` 단위를 더한 후 `xScale`에 전달하세요.

# --hints--

코드에는 10개의 `circle` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

첫 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `91`이고 `cy` 값이 대략 `368`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const firstCircle = document.querySelectorAll('circle')[0];
assert.strictEqual(Math.round(firstCircle?.getAttribute('cx')), 91);
assert.strictEqual(Math.round(firstCircle?.getAttribute('cy')), 368);
assert.strictEqual(Math.round(firstCircle?.getAttribute('r')), 5);
```

두 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `159`이고 `cy` 값이 대략 `181`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const secondCircle = document.querySelectorAll('circle')[1];
assert.strictEqual(Math.round(secondCircle?.getAttribute('cx')), 159);
assert.strictEqual(Math.round(secondCircle?.getAttribute('cy')), 181);
assert.strictEqual(Math.round(secondCircle?.getAttribute('r')), 5);
```

세 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `340`이고 `cy` 값이 대략 `329`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const thirdCircle = document.querySelectorAll('circle')[2];
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cx')), 340);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cy')), 329);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('r')), 5);
```

네 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `131`이고 `cy` 값이 대략 `60`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const fourthCircle = document.querySelectorAll('circle')[3];
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cy')), 60);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('r')), 5);
```

다섯 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `440`이고 `cy` 값이 대략 `237`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const fifthCircle = document.querySelectorAll('circle')[4];
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cx')), 440);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cy')), 237);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('r')), 5);
```

여섯 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `271`이고 `cy` 값이 대략 `306`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const sixthCircle = document.querySelectorAll('circle')[5];
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cx')), 271);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cy')), 306);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('r')), 5);
```

일곱 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `361`이고 `cy` 값이 대략 `351`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const seventhCircle = document.querySelectorAll('circle')[6];
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cx')), 361);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cy')), 351);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('r')), 5);
```

여덟 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `261`이고 `cy` 값이 대략 `132`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const eighthCircle = document.querySelectorAll('circle')[7];
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cx')), 261);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cy')), 132);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('r')), 5);
```

아홉 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `131`이고 `cy` 값이 대략 `144`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const ninthCircle = document.querySelectorAll('circle')[8];
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cy')), 144);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('r')), 5);
```

열 번째 `circle` 요소는 스케일을 적용한 후 `cx` 값이 대략 `79`이고 `cy` 값이 대략 `326`여야 합니다. 또한 `r` 값은 `5`이어야 합니다.

```js
const tenthCircle = document.querySelectorAll('circle')[9];
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cx')), 79);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cy')), 326);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('r')), 5);
```

코드에는 10개의 `text` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

첫 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `100`이고 `y` 값이 대략 `368`여야 합니다.

```js
const firstLabel = document.querySelectorAll('text')[0];
assert.strictEqual(Math.round(firstLabel?.getAttribute('x')), 100);
assert.strictEqual(Math.round(firstLabel?.getAttribute('y')), 368);
```

두 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `168`이고 `y` 값이 대략 `181`여야 합니다.

```js
const secondLabel = document.querySelectorAll('text')[1];
assert.strictEqual(Math.round(secondLabel?.getAttribute('x')), 168);
assert.strictEqual(Math.round(secondLabel?.getAttribute('y')), 181);
```

세 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `350`이고 `y` 값이 대략 `329`여야 합니다.

```js
const thirdLabel = document.querySelectorAll('text')[2];
assert.strictEqual(Math.round(thirdLabel?.getAttribute('x')), 350);
assert.strictEqual(Math.round(thirdLabel?.getAttribute('y')), 329);
```

네 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `141`이고 `y` 값이 대략 `60`여야 합니다.

```js
const fourthLabel = document.querySelectorAll('text')[3];
assert.strictEqual(Math.round(fourthLabel?.getAttribute('x')), 141);
assert.strictEqual(Math.round(fourthLabel?.getAttribute('y')), 60);
```

다섯 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `449`이고 `y` 값이 대략 `237`여야 합니다.

```js
const fifthLabel = document.querySelectorAll('text')[4];
assert.strictEqual(Math.round(fifthLabel?.getAttribute('x')), 449);
assert.strictEqual(Math.round(fifthLabel?.getAttribute('y')), 237);
```

여섯 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `280`이고 `y` 값이 대략 `306`여야 합니다.

```js
const sixthLabel = document.querySelectorAll('text')[5];
assert.strictEqual(Math.round(sixthLabel?.getAttribute('x')), 280);
assert.strictEqual(Math.round(sixthLabel?.getAttribute('y')), 306);
```

일곱 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `370`이고 `y` 값이 대략 `351`여야 합니다.

```js
const seventhLabel = document.querySelectorAll('text')[6];
assert.strictEqual(Math.round(seventhLabel?.getAttribute('x')), 370);
assert.strictEqual(Math.round(seventhLabel?.getAttribute('y')), 351);
```

여덟 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `270`이고 `y` 값이 대략 `132`여야 합니다.

```js
const eighthLabel = document.querySelectorAll('text')[7];
assert.strictEqual(Math.round(eighthLabel?.getAttribute('x')), 270);
assert.strictEqual(Math.round(eighthLabel?.getAttribute('y')), 132);
```

아홉 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `140`이고 `y` 값이 대략 `144`여야 합니다.

```js
const ninthLabel = document.querySelectorAll('text')[8];
assert.strictEqual(Math.round(ninthLabel?.getAttribute('x')), 140);
assert.strictEqual(Math.round(ninthLabel?.getAttribute('y')), 144);
```

열 번째 레이블은 스케일을 적용한 후 `x` 값이 대략 `88`이고 `y` 값이 대략 `326`여야 합니다.

```js
const tenthLabel = document.querySelectorAll('text')[9];
assert.strictEqual(Math.round(tenthLabel?.getAttribute('x')), 88);
assert.strictEqual(Math.round(tenthLabel?.getAttribute('y')), 326);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1]);
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));
  </script>
</body>
```
