---
id: 587d7fa8367417b2b2512bcd
title: 데이터 세트의 각 데이터 포인트에 막대 만들기
challengeType: 6
forumTopicId: 301482
dashedName: create-a-bar-for-each-data-point-in-the-set
---

# --description--

이전 도전 과제에서는 막대를 나타내기 위해 `svg` 요소에 사각형 하나만 추가했습니다. 여기서는 지금까지 배운 `data()`, `enter()`, 그리고 SVG 도형에 관한 내용을 결합하여 `dataset`의 각 데이터 포인트마다 사각형을 생성하고 추가할 것입니다.

이전 도전 과제에서는 `div`의 각 항목에 대해 `dataset`를 생성하고 추가하는 형식을 보여주었습니다:

```js
d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
```

`rect` 요소 대신 `div` 요소를 다룰 때 몇 가지 차이점이 있습니다. `rect` 요소는 `svg`에 직접 추가하는 것이 아니라 `body` 요소에 추가해야 합니다. 또한, 각 `rect`을 `svg` 영역 내 어디에 배치할지 D3에 알려야 합니다. 막대 배치는 다음 도전 과제에서 다룰 예정입니다.

# --instructions--

`data()`, `enter()`, 그리고 `append()` 메서드를 사용하여 `rect`의 각 항목에 대해 `dataset`를 생성하고 추가하세요. 막대들은 모두 겹쳐서 표시될 것입니다; 이는 다음 도전 과제에서 수정할 예정입니다.

# --hints--

문서에는 9개의 `rect` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('rect'), 9);
```

코드는 `data()` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.data/g);
```

코드는 `enter()` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.enter/g);
```

코드는 `append()` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.append/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      // Add your code below this line

      // Add your code above this line
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
