---
id: 587d7fa8367417b2b2512bc9
title: 요소의 높이를 동적으로 업데이트하기
challengeType: 6
forumTopicId: 301493
dashedName: update-the-height-of-an-element-dynamically
---

# --description--

이전 도전 과제에서는 배열에서 데이터를 표시하는 방법과 CSS 클래스를 추가하는 방법을 다뤘습니다. 이 수업을 결합해 간단한 막대 그래프를 만들 수 있습니다. 단계는 두 가지입니다:

1. 배열의 각 데이터 포인트마다 `div`를 만드세요

2. `div` 메서드에서 콜백 함수를 사용해 각 `style()`에 동적으로 높이를 지정하세요. 높이는 데이터 값과 같게 설정합니다

콜백 함수를 사용해 스타일을 설정하는 형식을 기억하세요:

```js
selection.style('cssProperty', d => d);
```

# --instructions--

편집기 내 코드에 `style()` 메서드를 추가해 각 요소의 `height` 속성을 설정하세요. 콜백 함수를 사용해 데이터 포인트 값에 문자열 `px`를 덧붙여 반환하세요.

# --hints--

첫 번째 `div`는 `height`가 `12` 픽셀이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('div')[0]?.style.height, '12px');
```

두 번째 `div`는 `height`가 `31` 픽셀이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('div')[1]?.style.height, '31px');
```

세 번째 `div`는 `height`가 `22` 픽셀이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('div')[2]?.style.height, '22px');
```

네 번째 `div`는 `height`가 `17` 픽셀이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('div')[3]?.style.height, '17px');
```

다섯 번째 `div`는 `height`가 `25` 픽셀이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('div')[4]?.style.height, '25px');
```

여섯 번째 `div`는 `height`가 `18` 픽셀이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('div')[5]?.style.height, '18px');
```

일곱 번째 `div`는 `height`가 `29` 픽셀이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('div')[6]?.style.height, '29px');
```

여덟 번째 `div`는 `height`가 `14` 픽셀이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('div')[7]?.style.height, '14px');
```

아홉 번째 `div`는 `height`가 `9` 픽셀이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('div')[8]?.style.height, '9px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => `${d}px`);
  </script>
</body>
```
