---
id: 587d7fa7367417b2b2512bc4
title: D3에서 데이터 다루기
challengeType: 6
forumTopicId: 301497
dashedName: work-with-data-in-d3
---

# --description--

D3 라이브러리는 데이터 기반 접근법에 포커스를 맞춥니다. 데이터 집합이 있으면 D3 메서드를 적용해 페이지에 표시할 수 있습니다. 데이터는 다양한 형식으로 제공되지만, 이 도전 과제에서는 간단한 숫자 배열을 사용합니다.

첫 번째 단계는 D3가 데이터를 인식하도록 하는 것입니다. `data()` 메서드는 DOM 요소 선택에 사용되어 해당 요소에 데이터를 연결합니다. 데이터 집합은 메서드에 인수로 전달됩니다.

일반적인 작업 흐름 패턴은 데이터 집합의 각 항목마다 문서에 새 요소를 만드는 것입니다. D3에는 이를 위한 `enter()` 메서드가 있습니다.

`enter()`가 `data()` 메서드와 결합되면, 페이지에서 선택된 요소와 데이터 집합의 항목 수를 비교합니다. 요소가 데이터 항목보다 적으면, 부족한 요소를 생성합니다.

다음은 `ul` 요소를 선택하고 배열의 항목 수에 따라 새 목록 항목을 만드는 예시입니다:

```html
<body>
  <ul></ul>
  <script>
    const dataset = ['a', 'b', 'c'];
    d3.select('ul')
      .selectAll('li')
      .data(dataset)
      .enter()
      .append('li')
      .text('New item');
  </script>
</body>
```

아직 존재하지 않는 요소를 선택하는 것이 혼란스러울 수 있습니다. 이 코드는 먼저 페이지에서 `ul`를 선택하라고 D3에 지시합니다. 다음으로 모든 목록 항목을 선택하는데, 이는 빈 선택을 반환합니다. 그다음 `data()` 메서드가 데이터 집합을 검토하고 배열의 각 항목마다 세 번 다음 코드를 실행합니다. `enter()` 메서드는 페이지에 `li` 요소가 없음을 확인하지만, `dataset`의 각 데이터 항목마다 하나씩 총 3개가 필요합니다. 새 `li` 요소가 `ul`에 추가되고 텍스트는 `New item`입니다.

# --instructions--

`body` 노드를 선택한 다음 모든 `h2` 요소를 선택하세요. D3가 `h2` 배열의 각 항목마다 `dataset` 태그를 생성하고 추가하게 하세요. `h2`의 텍스트는 `New Title`이어야 합니다. 코드에는 `data()`와 `enter()` 메서드를 사용하세요.

# --hints--

문서에는 `h2` 요소가 9개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h2'), 9);
```

`h2` 요소의 텍스트는 `New Title`이어야 합니다. 대소문자와 공백이 정확히 일치해야 합니다.

```js
const h2Elements = document.querySelectorAll('h2');
for (let i = 0; i < h2Elements.length; i++) {
  assert.match(h2Elements[i]?.textContent, /New Title/g);
}
```

코드에 `data()` 메서드를 사용하세요.

```js
assert.match(code, /\.data/g);
```

코드에 `enter()` 메서드를 사용하세요.

```js
assert.match(code, /\.enter/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text('New Title');
  </script>
</body>
```
