---
id: 587d7fa6367417b2b2512bc2
title: D3로 문서 요소 추가하기
challengeType: 6
forumTopicId: 301474
dashedName: add-document-elements-with-d3
---

# --description--

D3에는 문서에 요소를 추가하고 변경할 수 있는 여러 메서드가 있습니다.

`select()` 메서드는 문서에서 하나의 요소를 선택합니다. 원하는 요소 이름을 인수로 받아 문서에서 이름과 일치하는 첫 번째 요소의 HTML 노드를 반환합니다. 예시는 다음과 같습니다.

```js
const anchor = d3.select('a');
```

위 예시는 페이지에서 첫 번째 앵커 태그를 찾아 HTML 노드를 `anchor` 변수에 저장합니다. 선택한 요소는 다른 메서드와 함께 사용할 수 있습니다. 예시의 `d3` 부분은 D3 객체를 참조하는 것으로, D3 메서드에 접근하는 방법입니다.

다른 유용한 메서드로는 `append()`와 `text()`가 있습니다.

`append()` 메서드는 문서에 추가할 요소 이름을 인수로 받습니다. 선택한 항목에 HTML 노드를 추가하고, 그 노드에 대한 핸들을 반환합니다.

`text()` 메서드는 선택한 노드의 텍스트를 설정하거나 현재 텍스트를 가져옵니다. 값을 설정하려면 소괄호 안에 문자열을 인수로 전달하세요.

다음은 순서 없는 목록을 선택하고, 목록 항목을 추가한 뒤 텍스트를 넣는 예시입니다.

```js
d3.select('ul')?.append('li').text('Very important item');
```

D3는 여러 메서드를 점으로 연결해 연속적으로 여러 작업을 수행할 수 있게 합니다.

# --instructions--

`select` 메서드를 사용해 문서에서 `body` 태그를 선택하세요. 그런 다음 `append` 메서드로 `h1` 태그를 추가하고, `Learning D3` 요소에 `h1` 텍스트를 넣으세요.

# --hints--

`body`에는 `h1` 요소가 하나 있어야 합니다.

```js
const body = document.querySelector('body');
const headers = body?.querySelectorAll('h1');
assert.lengthOf(headers, 1);
```

`h1` 요소에는 `Learning D3` 텍스트가 있어야 합니다.

```js
assert.strictEqual(document.querySelector('h1')?.textContent, 'Learning D3');
```

코드는 `d3` 객체에 접근해야 합니다.

```js
assert.match(code, /d3/g);
```

코드는 `select` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.select/g);
```

코드는 `append` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.append/g);
```

코드는 `text` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.text/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    d3.select('body').append('h1').text('Learning D3');
  </script>
</body>
```
