---
id: 587d7fac367417b2b2512bdb
title: 스케일에 도메인과 레인지를 설정하기
challengeType: 6
forumTopicId: 301491
dashedName: set-a-domain-and-a-range-on-a-scale
---

# --description--

기본적으로 스케일은 항등 관계를 사용합니다. 즉, 입력 값이 출력 값에 매핑됩니다. 하지만 스케일은 훨씬 더 유연하고 흥미로울 수 있습니다.

예를 들어, 데이터셋에 50에서 480까지의 값이 있다고 가정해 보세요. 이것이 스케일의 입력 정보, 즉 <dfn>도메인</dfn>입니다.

SVG의 `x` 축을 따라 10 단위에서 500 단위 사이에 이 점들을 매핑하고 싶습니다. 이것이 출력 정보, 즉 <dfn>레인지</dfn>입니다.

`domain()`과 `range()` 메서드가 스케일에 이 값을 설정합니다. 두 메서드 모두 최소 두 개 요소가 있는 배열을 인수로 받습니다. 예시는 다음과 같습니다:

```js
scale.domain([50, 480]);
scale.range([10, 500]);
scale(50);
scale(480);
scale(325);
scale(750);
d3.scaleLinear();
```

순서대로 콘솔에 표시될 값은 `10`, `500`, `323.37`, 그리고 `807.67`입니다.

스케일이 도메인과 레인지 값 사이의 선형 관계를 사용해 주어진 숫자에 대해 출력이 무엇이어야 하는지 계산하는 점에 주목하세요. 도메인의 최솟값(50)은 레인지의 최솟값(10)에 매핑됩니다.

# --instructions--

스케일을 생성하고 도메인을 `[250, 500]`로, 레인지를 `[10, 150]`로 설정하세요.

**참고:** `domain()` 변수에 `range()`와 `scale` 메서드를 체인으로 연결할 수 있습니다.

# --hints--

코드는 `domain()` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.domain/g);
```

`domain()`의 `scale`은 `[250, 500]`로 설정해야 합니다.

```js
assert.deepEqual(scale.domain(), [250, 500]);
```

코드는 `range()` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.range/g);
```

`range()`의 `scale`은 `[10, 150]`로 설정해야 합니다.

```js
assert.deepEqual(scale.range(), [10, 150]);
```

`h2`의 텍스트는 `-102`이어야 합니다.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '-102');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line
    const scale = d3.scaleLinear();



    // Add your code above this line
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    scale.domain([250, 500]);
    scale.range([10, 150]);
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
