---
id: 587d7fa8367417b2b2512bca
title: 막대 그래프의 표시 방식 변경
challengeType: 6
forumTopicId: 301481
dashedName: change-the-presentation-of-a-bar-chart
---

# --description--

마지막 도전 과제에서는 막대 그래프를 만들었지만, 다음과 같은 몇 가지 형식 변경으로 그래프를 더 개선할 수 있습니다:

1. 각 막대 사이에 여백을 추가해 시각적으로 구분되도록 하는데, 이는 `bar` 클래스의 CSS에 여백을 추가하여 구현합니다.

2. 막대의 높이를 키워 값 차이를 더 잘 보여주는데, 이는 값을 숫자로 곱해 높이를 조정하는 방식으로 합니다.

# --instructions--

먼저, `margin` 태그의 `2px` 클래스에 `bar`의 `style`을 추가하세요. 다음으로, `style()` 메서드의 콜백 함수에서 원래 데이터 값에 `10`을 곱한 값(여기에 `px`을 더한 값)을 반환하도록 변경하세요.

**참고:** 각 데이터 포인트에 _같은_ 상수를 곱하는 것은 단지 스케일만 변경합니다. 확대하는 것과 같으며, 기본 데이터의 의미는 바뀌지 않습니다.

# --hints--

첫 번째 `div`에는 `height` 픽셀의 `120`와 `margin` 픽셀의 `2`이 있어야 합니다.

```js
const firstDiv = document.querySelector('div');
assert.exists(firstDiv); 
const height = window.getComputedStyle(firstDiv)['height'];
const margin = window.getComputedStyle(firstDiv)['margin-right'];
assert.strictEqual(height, '120px');
assert.strictEqual(margin, '2px');
```

두 번째 `div`에는 `height` 픽셀의 `310`와 `margin` 픽셀의 `2`이 있어야 합니다.

```js
const secondDiv = document.querySelectorAll('div')[1];
assert.exists(secondDiv); 
const height = window.getComputedStyle(secondDiv)['height'];
const margin = window.getComputedStyle(secondDiv)['margin-right'];
assert.strictEqual(height, '310px');
assert.strictEqual(margin, '2px');
```

세 번째 `div`에는 `height` 픽셀의 `220`와 `margin` 픽셀의 `2`이 있어야 합니다.

```js
const thirdDiv = document.querySelectorAll('div')[2];
assert.exists(thirdDiv); 
const height = window.getComputedStyle(thirdDiv)['height'];
const margin = window.getComputedStyle(thirdDiv)['margin-right'];
assert.strictEqual(height, '220px');
assert.strictEqual(margin, '2px');
```

네 번째 `div`에는 `height` 픽셀의 `170`와 `margin` 픽셀의 `2`이 있어야 합니다.

```js
const fourthDiv = document.querySelectorAll('div')[3];
assert.exists(fourthDiv); 
const height = window.getComputedStyle(fourthDiv)['height'];
const margin = window.getComputedStyle(fourthDiv)['margin-right'];
assert.strictEqual(height, '170px');
assert.strictEqual(margin, '2px');
```

다섯 번째 `div`에는 `height` 픽셀의 `250`와 `margin` 픽셀의 `2`이 있어야 합니다.

```js
const fifthDiv = document.querySelectorAll('div')[4];
assert.exists(fifthDiv); 
const height = window.getComputedStyle(fifthDiv)['height'];
const margin = window.getComputedStyle(fifthDiv)['margin-right'];
assert.strictEqual(height, '250px');
assert.strictEqual(margin, '2px');
```

여섯 번째 `div`에는 `height` 픽셀의 `180`와 `margin` 픽셀의 `2`이 있어야 합니다.

```js
const sixthDiv = document.querySelectorAll('div')[5];
assert.exists(sixthDiv); 
const height = window.getComputedStyle(sixthDiv)['height'];
const margin = window.getComputedStyle(sixthDiv)['margin-right'];
assert.strictEqual(height, '180px');
assert.strictEqual(margin, '2px');
```

일곱 번째 `div`에는 `height` 픽셀의 `290`와 `margin` 픽셀의 `2`이 있어야 합니다.

```js
const seventhDiv = document.querySelectorAll('div')[6];
assert.exists(seventhDiv); 
const height = window.getComputedStyle(seventhDiv)['height'];
const margin = window.getComputedStyle(seventhDiv)['margin-right'];
assert.strictEqual(height, '290px');
assert.strictEqual(margin, '2px');
```

여덟 번째 `div`에는 `height` 픽셀의 `140`와 `margin` 픽셀의 `2`이 있어야 합니다.

```js
const eighthDiv = document.querySelectorAll('div')[7];
assert.exists(eighthDiv); 
const height = window.getComputedStyle(eighthDiv)['height'];
const margin = window.getComputedStyle(eighthDiv)['margin-right'];
assert.strictEqual(height, '140px');
assert.strictEqual(margin, '2px');
```

아홉 번째 `div`에는 `height` 픽셀의 `90`와 `margin` 픽셀의 `2`이 있어야 합니다.

```js
const ninthDiv = document.querySelectorAll('div')[8];
assert.exists(ninthDiv); 
const height = window.getComputedStyle(ninthDiv)['height'];
const margin = window.getComputedStyle(ninthDiv)['margin-right'];
assert.strictEqual(height, '90px');
assert.strictEqual(margin, '2px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    /* Add your code below this line */

    
    /* Add your code above this line */
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d + 'px'); // Change this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    margin: 2px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d * 10 + 'px');
  </script>
</body>
```
