---
id: 587d7fa7367417b2b2512bc6
title: 요소에 인라인 스타일 추가하기
challengeType: 6
forumTopicId: 301475
dashedName: add-inline-styling-to-elements
---

# --description--

D3는 `style()` 메서드를 사용해 동적 요소에 인라인 CSS 스타일을 추가할 수 있습니다.

`style()` 메서드는 쉼표로 구분된 키-값 쌍을 인수로 받습니다. 선택한 요소의 텍스트 색상을 파란색으로 설정하는 예시는 다음과 같습니다.

```js
selection.style('color', 'blue');
```

# --instructions--

편집기 내 코드를 수정해 모든 표시된 텍스트에 `style()`가 `font-family`가 되도록 `verdana` 메서드를 추가하세요.

# --hints--

`h2` 요소들은 `font-family`인 `verdana`를 가져야 합니다.

```js
const headingTwo = document.querySelector('h2');
assert.exists(headingTwo); 
const fontFamily = window.getComputedStyle(headingTwo)['font-family'];
assert.strictEqual(fontFamily, 'verdana');
```

코드는 `style()` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.style/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('font-family', 'verdana');
  </script>
</body>
```
