---
id: 587d7fa6367417b2b2512bc3
title: D3로 요소 그룹 선택하기
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--

D3에는 요소 그룹을 선택하는 `selectAll()` 메서드도 있습니다. 이 메서드는 입력 문자열과 일치하는 문서 내 모든 항목에 대해 HTML 노드 배열을 반환합니다. 문서 내 모든 앵커 태그를 선택하는 예시는 다음과 같습니다.

```js
const anchors = d3.selectAll('a');
```

`select()` 메서드처럼 `selectAll()`도 메서드 체이닝을 지원하며, 다른 메서드와 함께 사용할 수 있습니다.

# --instructions--

문서 내 모든 `li` 태그를 선택하고, `list item` 메서드를 체이닝하여 텍스트를 문자열 `.text()`로 변경하세요.

# --hints--

페이지에 `li` 요소가 3개 있어야 하며, 각각의 텍스트는 `list item`이어야 합니다. 대소문자와 공백이 정확히 일치해야 합니다.

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

코드는 `d3` 객체에 접근해야 합니다.

```js
assert.match(code, /d3/g);
```

코드는 `selectAll` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```
