---
id: 587d7fa9367417b2b2512bd1
title: SVG 요소의 색상 변경
challengeType: 6
forumTopicId: 301480
dashedName: change-the-color-of-an-svg-element
---

# --description--

막대는 올바른 위치에 있지만 모두 같은 검은색입니다. SVG에서는 막대의 색상을 변경하는 방법이 있습니다.

SVG에서 `rect` 도형은 `fill` 속성으로 색상을 지정합니다. 이 속성은 16진수 코드, 색상 이름, rgb 값뿐 아니라 그라디언트나 투명도 같은 더 복잡한 옵션도 지원합니다.

# --instructions--

모든 막대의 `attr()`을 네이비 색상으로 설정하는 `fill` 메서드를 추가하세요.

# --hints--

모든 막대는 `fill` 색상이 네이비여야 합니다.

```js
const rect = document.querySelector('rect');
assert.exists(rect); 
const fill = window.getComputedStyle(rect)['fill'];
assert.strictEqual(fill, 'rgb(0, 0, 128)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');
  </script>
</body>
```
