---
id: 587d7fab367417b2b2512bda
title: D3로 선형 스케일 만들기
challengeType: 6
forumTopicId: 301483
dashedName: create-a-linear-scale-with-d3
---

# --description--

막대 차트와 산점도 차트는 모두 데이터를 SVG에 직접 그렸습니다. 하지만 막대의 높이나 데이터 점 중 하나가 SVG 높이나 너비 값보다 크면 SVG 영역 밖으로 나가게 됩니다.

D3에는 데이터를 그리는 데 도움이 되는 스케일이 있습니다. `scales`는 원시 데이터 집합을 SVG 픽셀에 어떻게 매핑할지 프로그램에 알려주는 함수입니다.

예를 들어, 100x500 크기의 SVG가 있고 여러 국가의 국내총생산(GDP)을 그리려 한다고 가정해 보세요. 숫자 집합은 수십억 또는 수조 달러 범위일 것입니다. D3에 스케일 유형을 제공하여 큰 GDP 값을 100x500 크기 영역에 배치하는 방법을 알려줍니다.

원시 데이터를 그대로 그리는 경우는 거의 없습니다. 그리기 전에 전체 데이터 집합에 대한 스케일을 설정하여 `x`과 `y` 값이 SVG 너비와 높이에 맞도록 합니다.

D3에는 여러 스케일 유형이 있습니다. 선형 스케일(보통 정량적 데이터에 사용)에는 D3 메서드 `scaleLinear()`가 있습니다:

```js
const scale = d3.scaleLinear();
```

기본적으로 스케일은 항등 관계를 사용합니다. 입력 값과 출력 값이 같습니다. 이 부분을 변경하는 방법은 별도의 도전 과제에서 다룹니다.

# --instructions--

`scale` 변수를 변경하여 선형 스케일을 만드세요. 그런 다음 `output` 변수를 `50`를 인수로 호출한 스케일로 설정하세요.

# --hints--

`h2`의 텍스트는 `50`이어야 합니다.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '50');
```

코드는 `scaleLinear()` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.scaleLinear/g);
```

`output` 변수는 `scale`을 `50` 인수로 호출해야 합니다.

```js
assert.strictEqual(output, 50);
assert.match(code, /scale\(\s*50\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line

    const scale = undefined; // Create the scale here
    const output = scale(); // Call scale with an argument here

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    const output = scale(50);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
