---
id: 587d7fa7367417b2b2512bc7
title: 데이터에 따라 스타일 변경
challengeType: 6
forumTopicId: 301479
dashedName: change-styles-based-on-data
---

# --description--

D3는 데이터 시각화와 표현에 관한 것입니다. 데이터에 따라 요소의 스타일을 변경하고 싶을 때가 많습니다.
예를 들어, 값이 20보다 작으면 데이터 포인트를 파란색으로, 그렇지 않으면 빨간색으로 색칠하고 싶을 수 있습니다. `style()` 메서드에 콜백 함수를 사용하고 조건문을 포함할 수 있습니다. 콜백 함수는 `d` 매개변수를 사용해 데이터 포인트를 나타냅니다:

```js
selection.style('color', d => {});
```

`style()` 메서드는 `color` 설정에만 국한되지 않고 다른 CSS 속성에도 사용할 수 있습니다.

# --instructions--

편집기의 코드에 `style()` 메서드를 추가해 `color` 요소의 `h2`를 조건부로 설정하세요. 데이터 값이 20보다 작으면 빨간색을 반환하고, 그렇지 않으면 초록색을 반환하는 콜백 함수를 작성하세요.

**참고:** if-else 문이나 삼항 연산자를 사용할 수 있습니다.

# --hints--

첫 번째 `h2`은 빨간색 `color`을 가져야 합니다.

```js
const headingTwoFirst = document.querySelector('h2');
assert.exists(headingTwoFirst); 
const color = window.getComputedStyle(headingTwoFirst)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

두 번째 `h2`은 초록색 `color`을 가져야 합니다.

```js
const headingTwoSecond = document.querySelectorAll('h2')[1];
assert.exists(headingTwoSecond); 
const color = window.getComputedStyle(headingTwoSecond)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

세 번째 `h2`은 초록색 `color`을 가져야 합니다.

```js
const headingTwoThird = document.querySelectorAll('h2')[2];
assert.exists(headingTwoThird); 
const color = window.getComputedStyle(headingTwoThird)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

네 번째 `h2`은 빨간색 `color`을 가져야 합니다.

```js
const headingTwoFourth = document.querySelectorAll('h2')[3];
assert.exists(headingTwoFourth); 
const color = window.getComputedStyle(headingTwoFourth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

다섯 번째 `h2`은 초록색 `color`을 가져야 합니다.

```js
const headingTwoFifth = document.querySelectorAll('h2')[4];
assert.exists(headingTwoFifth); 
const color = window.getComputedStyle(headingTwoFifth)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

여섯 번째 `h2`은 빨간색 `color`을 가져야 합니다.

```js
const headingTwoSixth = document.querySelectorAll('h2')[5];
assert.exists(headingTwoSixth); 
const color = window.getComputedStyle(headingTwoSixth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

일곱 번째 `h2`은 초록색 `color`을 가져야 합니다.

```js
const headingTwoSeventh = document.querySelectorAll('h2')[6];
assert.exists(headingTwoSeventh); 
const color = window.getComputedStyle(headingTwoSeventh)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

여덟 번째 `h2`은 빨간색 `color`을 가져야 합니다.

```js
const headingTwoEighth = document.querySelectorAll('h2')[7];
assert.exists(headingTwoEighth); 
const color = window.getComputedStyle(headingTwoEighth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

아홉 번째 `h2`은 빨간색 `color`을 가져야 합니다.

```js
const headingTwoNinth = document.querySelectorAll('h2')[8];
assert.exists(headingTwoNinth); 
const color = window.getComputedStyle(headingTwoNinth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('color', d => (d < 20 ? 'red' : 'green'));
  </script>
</body>
```
