---
id: 587d7faa367417b2b2512bd4
title: D3 요소에 호버 효과 추가
challengeType: 6
forumTopicId: 301469
dashedName: add-a-hover-effect-to-a-d3-element
---

# --description--

사용자가 마우스를 올렸을 때 막대를 강조하는 효과를 추가할 수 있습니다. 지금까지는 내장된 D3와 SVG 메서드로 사각형 스타일을 적용했지만, CSS도 사용할 수 있습니다.

SVG 요소에 CSS 클래스를 `attr()` 메서드로 설정하세요. 그런 다음 새 클래스에 대한 `:hover` 가상 클래스가 호버 효과에 대한 스타일 규칙을 담고 있습니다.

# --instructions--

`attr()` 메서드를 사용해 모든 `bar` 요소에 `rect` 클래스를 추가하세요. 이렇게 하면 마우스를 올렸을 때 막대의 `fill` 색상이 갈색으로 바뀝니다.

# --hints--

`rect` 요소에 `bar` 클래스가 있어야 합니다.

```js
assert.isTrue(document.querySelector('rect')?.classList.contains('bar'));
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy')
    // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```
