---
id: 587d7fab367417b2b2512bd8
title: 원 요소에 속성 추가하기
challengeType: 6
forumTopicId: 301471
dashedName: add-attributes-to-the-circle-elements
---

# --description--

마지막 도전 과제에서는 `circle` 요소를 `dataset`의 각 점마다 생성하고 SVG에 추가했습니다. 하지만 D3가 각 `circle`의 위치와 크기에 관한 더 많은 정보가 있어야 올바르게 표시할 수 있습니다.

SVG의 `circle`에는 세 가지 주요 속성이 있습니다. `cx`와 `cy` 속성은 좌표로, 도형의 _중심_을 SVG 내 어디에 위치시킬지 D3에 알려줍니다. 반지름(`r` 속성)은 `circle`의 크기를 나타냅니다.

`rect` `y` 좌표처럼 `cy` 속성도 `circle`의 경우 SVG의 아래가 아니라 위에서부터 측정합니다.

세 속성 모두 콜백 함수를 사용해 값을 동적으로 설정할 수 있습니다. `data(dataset)` 뒤에 연결된 모든 메서드는 `dataset`의 각 항목마다 한 번씩 실행된다는 점을 기억하세요. 콜백 함수의 `d` 매개변수는 `dataset`의 현재 항목을 가리키며, 이는 각 점마다 배열입니다. `d[0]`처럼 대괄호 표기법을 사용해 배열 내 값을 접근합니다.

# --instructions--

`cx` 요소에 `cy`, `r`, `circle` 속성을 추가하세요. `cx` 값은 `dataset`의 각 항목 배열에서 첫 번째 숫자여야 합니다. `cy` 값은 배열의 두 번째 숫자를 기반으로 해야 하지만, 차트가 뒤집히지 않고 올바른 방향으로 표시되도록 하세요. `r` 값은 모든 원에 대해 `5`여야 합니다.

# --hints--

코드에는 10개의 `circle` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

첫 번째 `circle` 요소는 `cx` 값이 `34`, `cy` 값이 `422`, `r` 값이 `5`이어야 합니다.

```js
const circle1 = document.querySelector('circle');
assert.strictEqual(circle1?.getAttribute('cx'), '34');
assert.strictEqual(circle1?.getAttribute('cy'), '422');
assert.strictEqual(circle1?.getAttribute('r'), '5');
```

두 번째 `circle` 요소는 `cx` 값이 `109`, `cy` 값이 `220`, `r` 값이 `5`이어야 합니다.

```js
const circle2 = document.querySelectorAll('circle')[1];
assert.strictEqual(circle2?.getAttribute('cx'), '109');
assert.strictEqual(circle2?.getAttribute('cy'), '220');
assert.strictEqual(circle2?.getAttribute('r'), '5');
```

세 번째 `circle` 요소는 `cx` 값이 `310`, `cy` 값이 `380`, `r` 값이 `5`이어야 합니다.

```js
const circle3 = document.querySelectorAll('circle')[2];
assert.strictEqual(circle3?.getAttribute('cx'), '310');
assert.strictEqual(circle3?.getAttribute('cy'), '380');
assert.strictEqual(circle3?.getAttribute('r'), '5');
```

네 번째 `circle` 요소는 `cx` 값이 `79`, `cy` 값이 `89`, `r` 값이 `5`이어야 합니다.

```js
const circle4 = document.querySelectorAll('circle')[3];
assert.strictEqual(circle4?.getAttribute('cx'), '79');
assert.strictEqual(circle4?.getAttribute('cy'), '89');
assert.strictEqual(circle4?.getAttribute('r'), '5');
```

다섯 번째 `circle` 요소는 `cx` 값이 `420`, `cy` 값이 `280`, `r` 값이 `5`이어야 합니다.

```js
const circle5 = document.querySelectorAll('circle')[4];
assert.strictEqual(circle5?.getAttribute('cx'), '420');
assert.strictEqual(circle5?.getAttribute('cy'), '280');
assert.strictEqual(circle5?.getAttribute('r'), '5');
```

여섯 번째 `circle` 요소는 `cx` 값이 `233`, `cy` 값이 `355`, `r` 값이 `5`이어야 합니다.

```js
const circle6 = document.querySelectorAll('circle')[5];
assert.strictEqual(circle6?.getAttribute('cx'), '233');
assert.strictEqual(circle6?.getAttribute('cy'), '355');
assert.strictEqual(circle6?.getAttribute('r'), '5');
```

일곱 번째 `circle` 요소는 `cx` 값이 `333`, `cy` 값이 `404`, `r` 값이 `5`이어야 합니다.

```js
const circle7 = document.querySelectorAll('circle')[6];
assert.strictEqual(circle7?.getAttribute('cx'), '333');
assert.strictEqual(circle7?.getAttribute('cy'), '404');
assert.strictEqual(circle7?.getAttribute('r'), '5');
```

여덟 번째 `circle` 요소는 `cx` 값이 `222`, `cy` 값이 `167`, `r` 값이 `5`이어야 합니다.

```js
const circle8 = document.querySelectorAll('circle')[7];
assert.strictEqual(circle8?.getAttribute('cx'), '222');
assert.strictEqual(circle8?.getAttribute('cy'), '167');
assert.strictEqual(circle8?.getAttribute('r'), '5');
```

아홉 번째 `circle` 요소는 `cx` 값이 `78`, `cy` 값이 `180`, `r` 값이 `5`이어야 합니다.

```js
const circle9 = document.querySelectorAll('circle')[8];
assert.strictEqual(circle9?.getAttribute('cx'), '78');
assert.strictEqual(circle9?.getAttribute('cy'), '180');
assert.strictEqual(circle9?.getAttribute('r'), '5');
```

열 번째 `circle` 요소는 `cx` 값이 `21`, `cy` 값이 `377`, `r` 값이 `5`이어야 합니다.

```js
const circle10 = document.querySelectorAll('circle')[9];
assert.strictEqual(circle10?.getAttribute('cx'), '21');
assert.strictEqual(circle10?.getAttribute('cy'), '377');
assert.strictEqual(circle10?.getAttribute('r'), '5');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => d[0])
      .attr('cy', d => h - d[1])
      .attr('r', 5);
  </script>
</body>
```
