---
id: 587d7fa9367417b2b2512bd0
title: SVG 요소 뒤집기
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--

막대 그래프가 거꾸로, 즉 뒤집혀 보였을 수 있습니다. 이는 SVG가 (x, y) 좌표를 사용하는 방식 때문입니다.

SVG에서 좌표의 원점은 왼쪽 위 모서리에 있습니다. `x` 좌표가 0이면 도형이 SVG 영역의 왼쪽 가장자리에 위치합니다. `y` 좌표가 0이면 도형이 SVG 영역의 위쪽 가장자리에 위치합니다. `x` 값이 커질수록 사각형이 오른쪽으로 이동합니다. `y` 값이 커질수록 사각형이 아래로 이동합니다.

막대를 올바른 방향으로 세우려면 `y` 좌표를 계산하는 방식을 바꿔야 합니다. 막대의 높이와 SVG 영역 전체 높이 둘 다를 고려해야 합니다.

SVG 영역의 높이는 100입니다. 데이터 집합에 0이라는 값이 있으면 막대가 SVG 영역의 아래쪽에서 시작하도록 해야 합니다(위쪽이 아니라). 이를 위해 `y` 좌표는 100의 값을 가져야 합니다. 데이터 값이 1이라면 막대를 아래쪽에 놓기 위해 `y` 좌표를 100으로 시작합니다. 그리고 막대 높이 1을 고려해야 하므로 최종 `y` 좌표는 99가 됩니다.

`y` 좌표가 `y = heightOfSVG - heightOfBar` 이면 막대가 올바른 방향으로 세워집니다.

# --instructions--

`y` 속성의 콜백 함수를 변경하여 막대를 올바른 방향으로 세우세요. 막대의 `height` 은 데이터 값 `d` 의 3배라는 점을 기억하세요.

**참고:** 일반적으로 관계식은 `y = h - m * d` 이며, 여기서 `m` 는 데이터 값을 스케일하는 상수입니다.

# --hints--

첫 번째 `rect` 은 `y` 값이 `64` 이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

두 번째 `rect` 은 `y` 값이 `7` 이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

세 번째 `rect` 은 `y` 값이 `34` 이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

네 번째 `rect` 은 `y` 값이 `49` 이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

다섯 번째 `rect` 은 `y` 값이 `25` 이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

여섯 번째 `rect` 은 `y` 값이 `46` 이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

일곱 번째 `rect` 은 `y` 값이 `13` 이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

여덟 번째 `rect` 은 `y` 값이 `58` 이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

아홉 번째 `rect` 은 `y` 값이 `73` 이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```
