---
id: 587d7faa367417b2b2512bd6
title: D3 요소에 툴팁 추가하기
challengeType: 6
forumTopicId: 301470
dashedName: add-a-tooltip-to-a-d3-element
---

# --description--

툴팁은 사용자가 페이지의 항목 위에 마우스를 올렸을 때 해당 항목에 대한 추가 정보를 보여줍니다. 시각화에 툴팁을 추가하는 방법은 여러 가지가 있습니다. 이 도전 과제에서는 SVG `title` 요소를 사용합니다.

`title`은 `text()` 메서드와 함께 막대에 데이터를 동적으로 추가하는 데 사용됩니다.

# --instructions--

각 `title` 노드 아래에 `rect` 요소를 추가하세요. 그런 다음 콜백 함수를 사용해 `text()` 메서드를 호출하여 텍스트가 데이터 값을 표시하도록 하세요.

# --hints--

코드에는 9개의 `title` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('title'), 9);
```

첫 번째 `title` 요소는 `12`의 툴팁 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('title')[0]?.textContent, '12');
```

두 번째 `title` 요소는 `31`의 툴팁 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('title')[1]?.textContent, '31');
```

세 번째 `title` 요소는 `22`의 툴팁 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('title')[2]?.textContent, '22');
```

네 번째 `title` 요소는 `17`의 툴팁 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('title')[3]?.textContent, '17');
```

다섯 번째 `title` 요소는 `25`의 툴팁 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('title')[4]?.textContent, '25');
```

여섯 번째 `title` 요소는 `18`의 툴팁 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('title')[5]?.textContent, '18');
```

일곱 번째 `title` 요소는 `29`의 툴팁 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('title')[6]?.textContent, '29');
```

여덟 번째 `title` 요소는 `14`의 툴팁 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('title')[7]?.textContent, '14');
```

아홉 번째 `title` 요소는 `9`의 툴팁 텍스트를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('title')[8]?.textContent, '9');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar')
      .append('title')
      .text(d => d);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```
