---
id: 587d7fa9367417b2b2512bce
title: 각 막대의 좌표를 동적으로 설정하기
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--

마지막 도전 과제에서는 `svg`의 각 점마다 막대를 나타내기 위해 `dataset` 요소에 사각형을 생성하고 추가했습니다. 하지만 모두 겹쳐서 쌓여 있었습니다.

사각형의 위치는 `x`와 `y` 속성으로 처리합니다. 이 속성들은 D3에게 `svg` 영역에서 도형을 어디서부터 그리기 시작할지 알려줍니다. 지난 도전 과제에서는 이 값을 모두 0으로 설정해서 모든 막대가 좌상단 모서리에 놓였습니다.

막대 그래프에서는 모든 막대가 같은 수직 위치에 있어야 하므로 `y` 값은 모든 막대에 대해 동일하게 (0으로) 유지됩니다. 하지만 `x` 값은 새 막대를 추가할 때마다 달라져야 합니다. `x` 값이 클수록 항목이 오른쪽으로 더 멀리 밀린다는 점을 기억하세요. `dataset` 배열 요소를 순회하면서 `x` 값은 증가해야 합니다.

D3의 `attr()` 메서드는 해당 속성을 동적으로 설정할 콜백 함수를 받습니다. 콜백 함수는 두 개의 인수를 받는데, 하나는 데이터 점 자체(보통 `d`)이고 다른 하나는 배열 내 데이터 점의 인덱스입니다. 인덱스에 대한 두 번째 인수는 선택 사항입니다. 형식은 다음과 같습니다:

```js
selection.attr('property', (d, i) => {});
```

중요한 점은 데이터 집합의 항목을 반복하기 위해 `for` 루프를 작성하거나 `forEach()`를 사용할 필요가 없다는 것입니다. `data()` 메서드가 데이터 집합을 파싱하며, `data()` 이후에 체인된 모든 메서드는 데이터 집합의 각 항목에 대해 한 번씩 실행됩니다.

# --instructions--

`x` 속성 콜백 함수를 인덱스에 30을 곱한 값을 반환하도록 변경하세요.

**참고:** 각 막대의 너비는 25이므로 `x` 값을 30씩 증가시키면 막대 사이에 약간의 간격이 생깁니다. 이 예제에서는 25보다 큰 값이면 모두 작동합니다.

# --hints--

첫 번째 `rect`은 `x` 값이 `0`이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

두 번째 `rect`은 `x` 값이 `30`이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

세 번째 `rect`은 `x` 값이 `60`이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

네 번째 `rect`은 `x` 값이 `90`이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

다섯 번째 `rect`은 `x` 값이 `120`이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

여섯 번째 `rect`은 `x` 값이 `150`이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

일곱 번째 `rect`은 `x` 값이 `180`이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

여덟 번째 `rect`은 `x` 값이 `210`이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

아홉 번째 `rect`은 `x` 값이 `240`이어야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
