---
id: 587d7fa8367417b2b2512bcb
title: D3에서 SVG 배우기
challengeType: 6
forumTopicId: 301489
dashedName: learn-about-svg-in-d3
---

# --description--

<dfn>SVG</dfn>는 <dfn>Scalable Vector Graphics</dfn>의 약자입니다.

여기서 "scalable"은 객체를 확대하거나 축소해도 픽셀이 깨져 보이지 않는다는 뜻입니다. 작은 모바일 화면이든 큰 TV 모니터든 디스플레이 시스템에 맞게 크기가 조절됩니다.

SVG는 일반적인 기하학적 도형을 만드는 데 사용됩니다. D3는 데이터를 시각적으로 나타내기 때문에 시각화를 위한 도형을 만들 때 SVG를 사용합니다. 웹 페이지의 SVG 도형은 HTML `svg` 태그 안에 있어야 합니다.

CSS는 스타일에 상대 단위(`vh`, `vw`, 또는 백분율)를 사용하면 확장 가능하지만, 데이터 시각화를 만들 때는 SVG가 더 유연합니다.

# --instructions--

`svg`에 `body`를 사용해 `append()` 노드를 추가하세요. `width` 또는 `w` 메서드를 사용해 제공된 `height` 상수로 `h` 속성을, 제공된 `attr()` 상수로 `style()` 속성을 각각 설정하세요. `background-color` 태그에 핑크색 `style`가 적용되어 있어 출력에서 확인할 수 있습니다.

**참고:** `attr()` 너비와 높이 속성은 단위를 사용하지 않습니다. 이것이 크기 조절의 기본 단위로, 확대 수준에 관계없이 요소는 항상 너비와 높이가 5:1 비율을 유지합니다.

# --hints--

문서에 `svg` 요소가 1개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('svg'), 1);
```

`svg` 요소는 `width`로 설정된 `500` 속성을 가지거나 너비가 `500px`로 스타일이 적용되어야 합니다.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const width = window.getComputedStyle(svg)['width'];
assert.isTrue(svg?.getAttribute('width') === '500' || width === '500px');
```

`svg` 요소는 `height`로 설정된 `100` 속성을 가지거나 높이가 `100px`로 스타일이 적용되어야 합니다.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const height = window.getComputedStyle(svg)['height'];
assert.isTrue(svg?.getAttribute('height') === '100' || height === '100px');
```

# --seed--

## --seed-contents--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3.select('body');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
  </script>
</body>
```
