---
id: 587d7faa367417b2b2512bd2
title: D3 요소에 레이블 추가하기
challengeType: 6
forumTopicId: 301476
dashedName: add-labels-to-d3-elements
---

# --description--

D3는 막대 그래프 같은 그래프 요소에 SVG `text` 요소를 사용해 레이블을 붙일 수 있게 해줍니다.

`rect` 요소처럼 `text` 요소도 SVG에 위치시키기 위해 `x`와 `y` 속성이 필요합니다. 또한 표시할 값을 데이터에서 접근해야 합니다.

D3는 막대에 레이블을 붙이는 방식을 세밀하게 제어할 수 있게 해줍니다.

# --instructions--

편집기 내 코드는 이미 각 새 `text` 요소에 데이터를 바인딩하고 있습니다. 먼저 `text`에 `svg` 노드를 추가하세요. 다음으로 `x`과 `y` 좌표에 대한 속성을 추가하세요. 이 값들은 `rect` 좌표를 계산하는 방식과 같아야 하지만, `y`의 `text` 값은 레이블이 막대보다 3 단위 위에 위치하도록 해야 합니다. 마지막으로 D3 `text()` 메서드를 사용해 레이블을 데이터 포인트 값과 같게 설정하세요.

**참고:** 레이블이 막대보다 위에 위치하려면 `y`의 `text` 값이 막대의 `y` 값보다 3 더 크거나 3 더 작아야 하는지 결정하세요.

# --hints--

첫 번째 `text` 요소는 `12` 레이블과 `y` 값이 `61`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('text')[0]?.textContent, '12');
assert.strictEqual(
  document.querySelectorAll('text')[0]?.getAttribute('y'),
  '61'
);
```

두 번째 `text` 요소는 `31` 레이블과 `y` 값이 `4`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('text')[1]?.textContent, '31');
assert.strictEqual(
  document.querySelectorAll('text')[1]?.getAttribute('y'),
  '4'
);
```

세 번째 `text` 요소는 `22` 레이블과 `y` 값이 `31`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('text')[2]?.textContent, '22');
assert.strictEqual(
  document.querySelectorAll('text')[2]?.getAttribute('y'),
  '31'
);
```

네 번째 `text` 요소는 `17` 레이블과 `y` 값이 `46`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('text')[3]?.textContent, '17');
assert.strictEqual(
  document.querySelectorAll('text')[3]?.getAttribute('y'),
  '46'
);
```

다섯 번째 `text` 요소는 `25` 레이블과 `y` 값이 `22`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('text')[4]?.textContent, '25');
assert.strictEqual(
  document.querySelectorAll('text')[4]?.getAttribute('y'),
  '22'
);
```

여섯 번째 `text` 요소는 `18` 레이블과 `y` 값이 `43`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('text')[5]?.textContent, '18');
assert.strictEqual(
  document.querySelectorAll('text')[5]?.getAttribute('y'),
  '43'
);
```

일곱 번째 `text` 요소는 `29` 레이블과 `y` 값이 `10`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('text')[6]?.textContent, '29');
assert.strictEqual(
  document.querySelectorAll('text')[6]?.getAttribute('y'),
  '10'
);
```

여덟 번째 `text` 요소는 `14` 레이블과 `y` 값이 `55`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('text')[7]?.textContent, '14');
assert.strictEqual(
  document.querySelectorAll('text')[7]?.getAttribute('y'),
  '55'
);
```

아홉 번째 `text` 요소는 `9` 레이블과 `y` 값이 `70`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('text')[8]?.textContent, '9');
assert.strictEqual(
  document.querySelectorAll('text')[8]?.getAttribute('y'),
  '70'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg.selectAll('text').data(dataset).enter();
    // Add your code below this line




    // Add your code above this line
  </script>
  <body></body>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3)
      .text(d => d);
  </script>
  <body></body>
</body>
```
