---
id: 587d7fab367417b2b2512bd7
title: SVG 원으로 산점도 만들기
challengeType: 6
forumTopicId: 301484
dashedName: create-a-scatterplot-with-svg-circles
---

# --description--

산점도는 또 다른 유형의 시각화입니다. 보통 두 개의 값을 가진 데이터 포인트를 원으로 표시합니다. 이 값들은 `x` 축과 `y` 축에 연결되어 시각화에서 원의 위치를 정하는 데 사용됩니다.

SVG에는 원 모양을 만들기 위한 `circle` 태그가 있습니다. 막대 그래프에 사용한 `rect` 요소와 매우 비슷하게 작동합니다.

# --instructions--

`data()`, `enter()`, `append()` 메서드를 사용해 `dataset`을 SVG에 새로 추가된 `circle` 요소에 바인딩하세요.

**참고:** 아직 속성을 설정하지 않아 원이 보이지 않습니다. 다음 도전 과제에서 설정할 예정입니다.

# --hints--

코드에는 10개의 `circle` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
  </script>
</body>
```
