---
id: 587d7fa9367417b2b2512bcf
title: 각 막대의 높이를 동적으로 변경하기
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--

각 막대의 높이는 배열에 있는 데이터 포인트 값으로 설정할 수 있습니다. 이는 `x` 값이 동적으로 설정된 방식과 유사합니다.

```js
selection.attr('property', (d, i) => {});
```

여기서 `d`은 데이터 포인트 값이고, `i`는 배열에서 데이터 포인트의 인덱스입니다.

# --instructions--

`height` 속성에 대한 콜백 함수를 데이터 값에 3을 곱한 값을 반환하도록 변경하세요.

**참고:** 모든 데이터 포인트에 같은 상수를 곱하면 데이터가 확대되는 효과가 있습니다(확대하는 것과 비슷합니다). 이 예제에서는 막대 값 간 차이를 더 잘 볼 수 있게 도와줍니다.

# --hints--

첫 번째 `rect`는 `height`의 `36`를 가져야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

두 번째 `rect`는 `height`의 `93`를 가져야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

세 번째 `rect`는 `height`의 `66`를 가져야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

네 번째 `rect`는 `height`의 `51`를 가져야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

다섯 번째 `rect`는 `height`의 `75`를 가져야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

여섯 번째 `rect`는 `height`의 `54`를 가져야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

일곱 번째 `rect`는 `height`의 `87`를 가져야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

여덟 번째 `rect`는 `height`의 `42`를 가져야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

아홉 번째 `rect`는 `height`의 `27`를 가져야 합니다.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body>
```
