---
id: 587d7fa7367417b2b2512bc5
title: D3에서 동적 데이터 다루기
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--

마지막 두 도전 과제는 `data()`와 `enter()` 메서드를 사용해 D3로 데이터를 동적으로 표시하는 기본을 다룹니다. 이 메서드들은 데이터 집합을 받아 `append()` 메서드와 함께 데이터 집합의 각 항목에 대해 새로운 DOM 요소를 만듭니다.

이전 도전 과제에서는 `h2` 배열의 각 항목에 대해 새로운 `dataset` 요소를 만들었지만, 모두 같은 텍스트 `New Title`를 포함했습니다. 이는 각 `h2` 요소에 바인딩된 데이터를 활용하지 않았기 때문입니다.

D3 `text()` 메서드는 문자열이나 콜백 함수를 인수로 받을 수 있습니다:

```js
selection.text(d => d);
```

위 예제에서 매개변수 `d`는 선택 항목이 바인딩된 데이터셋의 단일 항목을 가리킵니다.

현재 예제를 기준으로 첫 번째 `h2` 요소는 12에 바인딩되고, 두 번째 `h2` 요소는 31에, 세 번째 `h2` 요소는 22에 바인딩되는 식입니다.

# --instructions--

`text()` 메서드를 변경하여 각 `h2` 요소가 `dataset` 배열의 해당 값을 한 칸 띄우고 문자열 `USD`를 표시하게 하세요. 예를 들어, 첫 번째 제목은 `12 USD`이어야 합니다.

# --hints--

첫 번째 `h2`는 텍스트 `12 USD`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

두 번째 `h2`는 텍스트 `31 USD`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

세 번째 `h2`는 텍스트 `22 USD`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

네 번째 `h2`는 텍스트 `17 USD`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

다섯 번째 `h2`는 텍스트 `25 USD`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

여섯 번째 `h2`는 텍스트 `18 USD`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

일곱 번째 `h2`는 텍스트 `29 USD`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

여덟 번째 `h2`는 텍스트 `14 USD`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

아홉 번째 `h2`는 텍스트 `9 USD`를 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```
