---
id: 587d7fac367417b2b2512bdc
title: >-
  d3.max와 d3.min 함수를 사용해 데이터셋에서 최솟값과 최댓값 찾기
challengeType: 6
forumTopicId: 301496
dashedName: >-
  use-the-d3-max-and-d3-min-functions-to-find-minimum-and-maximum-values-in-a-dataset
---

# --description--

D3 메서드 `domain()`와 `range()`는 데이터에 기반해 스케일의 정보를 설정합니다. 이를 더 쉽게 만드는 몇 가지 메서드가 있습니다.

도메인을 설정할 때 보통 데이터셋 내 최솟값과 최댓값을 사용하고 싶을 것입니다. 특히 큰 데이터셋에서 이 값을 수동으로 찾으려 하면 오류가 발생할 수 있습니다.

D3에는 이 정보를 반환하는 두 가지 메서드 `min()`와 `max()`가 있습니다. 예시는 다음과 같습니다:

```js
const exampleData = [34, 234, 73, 90, 6, 52];
d3.min(exampleData);
d3.max(exampleData);
```

데이터셋에 산점도 예제에 있던 `[x, y]` 좌표 쌍처럼 중첩된 배열이 있을 수 있습니다. 이 경우 D3에 최댓값과 최솟값을 계산하는 방법을 알려줘야 합니다. 다행히 `min()`와 `max()` 메서드 모두 콜백 함수를 받습니다. 이 예제에서 콜백 함수의 인수 `d`는 현재 내부 배열입니다. 콜백은 최댓값이나 최솟값을 계산할 내부 배열의 요소(`x` 또는 `y` 값)를 반환해야 합니다. 배열의 배열에서 최솟값과 최댓값을 찾는 예시는 다음과 같습니다:

```js
const locationData = [
  [1, 7],
  [6, 3],
  [8, 3]
];
const minX = d3.min(locationData, d => d[0]);
```

`minX`는 `1` 값을 갖게 됩니다.

# --instructions--

`positionData` 배열은 x, y, z 좌표의 하위 배열을 담고 있습니다. D3 메서드를 사용해 배열에서 z 좌표(세 번째 값)의 최댓값을 찾아 `output` 변수에 저장하세요.

# --hints--

`h2`의 텍스트는 `8`이어야 합니다.

```js
assert.strictEqual(output, 8);
assert.strictEqual(document.querySelector('h2')?.textContent, '8');
```

코드는 `max()` 메서드를 사용해야 합니다.

```js
assert.match(
  code,
  /\.max/g,
  'Your code should use the <code>max()</code> method.'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];
    // Add your code below this line

    const output = undefined; // Change this line

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];

    const output = d3.max(positionData, d => d[2]);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
