---
id: 587d7fa7367417b2b2512bc8
title: D3로 클래스 추가하기
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--

HTML 요소에 인라인 스타일을 많이 사용하는 것은 작은 앱에서도 관리하기 어려워집니다. 요소에 클래스를 추가하고 CSS 규칙으로 그 클래스를 한 번만 스타일링하는 것이 더 쉽습니다. D3에는 클래스 이름을 포함해 어떤 HTML 속성이든 요소에 추가할 수 있는 `attr()` 메서드가 있습니다.

`attr()` 메서드는 `style()`와 동일하게 작동합니다. 쉼표로 구분된 값을 받고, 콜백 함수를 사용할 수도 있습니다. 다음은 선택 항목에 `container` 클래스를 추가하는 예입니다:

```js
selection.attr('class', 'container');
```

`class` 매개변수는 클래스를 추가할 때마다 동일하게 유지되고, `container` 매개변수만 변경된다는 점을 참고하세요.

# --instructions--

편집기 코드에 `attr()` 메서드를 추가하고 `bar` 요소에 `div` 클래스를 넣으세요.

# --hints--

`div` 요소에 `bar` 클래스가 있어야 합니다.

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

코드는 `attr()` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body>
```
