---
id: 587d7fa8367417b2b2512bcc
title: SVG로 도형 표시하기
challengeType: 6
forumTopicId: 301485
dashedName: display-shapes-with-svg
---

# --description--

마지막 도전 과제에서는 주어진 너비와 높이를 가진 `svg` 요소를 만들었고, `background-color` 태그에 `style`가 적용되어 보여졌습니다. 코드는 주어진 너비와 높이만큼 공간을 만들었습니다.

다음 단계는 `svg` 영역에 넣을 도형을 만드는 것입니다. SVG에는 사각형과 원 같은 여러 지원되는 도형이 있습니다. 이들은 데이터를 표시하는 데 사용됩니다. 예를 들어, 사각형(`<rect>`) SVG 도형은 막대 그래프에서 막대를 만들 수 있습니다.

도형을 `svg` 영역에 배치할 때 `x`와 `y` 좌표로 위치를 지정할 수 있습니다. 원점 (0, 0)은 왼쪽 위 모서리에 있습니다. `x`의 양수 값은 도형을 오른쪽으로 밀고, `y`의 양수 값은 원점에서 아래로 도형을 밀어냅니다.

지난 도전 과제에서 500(너비) x 100(높이) `svg`의 가운데에 도형을 배치하려면 `x` 좌표는 250, `y` 좌표는 50이 됩니다.

SVG `rect`에는 네 개의 속성이 있습니다. `x` 영역에 배치되는 위치를 지정하는 `y`와 `svg` 좌표가 있습니다. 또한 크기를 지정하는 `height`와 `width`도 있습니다.

# --instructions--

`rect`에 `svg` 도형을 `append()`를 사용해 추가하고, `width` 속성에 `25` 값을, `height` 속성에 `100` 값을 지정하세요. 또한 `rect`에 `x`와 `y` 속성을 각각 `0` 값으로 지정하세요.

# --hints--

문서에는 `rect` 요소가 1개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('rect'), 1);
```

`rect` 요소는 `width` 속성이 `25`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('width'), '25');
```

`rect` 요소는 `height` 속성이 `100`로 설정되어야 합니다.

```js
assert.strictEqual(
  document.querySelector('rect')?.getAttribute('height'),
  '100'
);
```

`rect` 요소는 `x` 속성이 `0`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('x'), '0');
```

`rect` 요소는 `y` 속성이 `0`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('y'), '0');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h)
      .append('rect')
      .attr('width', 25)
      .attr('height', 100)
      .attr('x', 0)
      .attr('y', 0);
  </script>
</body>
```
