---
id: 587d7fad367417b2b2512bdf
title: 시각화에 축 추가하기
challengeType: 6
forumTopicId: 301472
dashedName: add-axes-to-a-visualization
---

# --description--

산점도를 개선하는 또 다른 방법은 x축과 y축을 추가하는 것입니다.

D3에는 각각 y축과 x축을 렌더링하는 `axisLeft()`와 `axisBottom()`라는 두 가지 메서드가 있습니다. 이전 도전 과제의 `xScale`를 기반으로 x축을 만드는 예시는 다음과 같습니다.

```js
const xAxis = d3.axisBottom(xScale);
```

다음 단계는 SVG에 축을 렌더링하는 것입니다. 이를 위해 일반 SVG 컴포넌트인 `g` 요소를 사용할 수 있습니다. `g`는 그룹을 의미합니다. `rect`, `circle`, `text`와 달리 축은 렌더링될 때 단순한 직선입니다. 단순한 도형이기 때문에 `g`를 사용하는 것이 적합합니다. 마지막 단계는 `transform` 속성을 적용해 축을 SVG 내 적절한 위치에 배치하는 것입니다. 그렇지 않으면 선이 SVG 테두리를 따라 렌더링되어 보이지 않을 것입니다. SVG는 다양한 유형의 `transforms`를 지원하지만, 축을 배치하려면 `translate`가 필요합니다. `g` 요소에 적용하면 주어진 수치만큼 그룹 전체를 오른쪽과 아래로 이동시킵니다. 예시는 다음과 같습니다.

```js
const xAxis = d3.axisBottom(xScale);

svg
  .append('g')
  .attr('transform', 'translate(0, ' + (h - padding) + ')')
  .call(xAxis);
```

위 코드는 x축을 SVG 하단에 배치합니다. 그런 다음 `call()` 메서드에 인수로 전달됩니다. y축도 같은 방식으로 작동하지만 `translate` 인수는 `(x, 0)` 형태입니다. 위 `translate` 메서드에서 `attr()`가 문자열이므로 인수에 변수 값을 포함하려면 연결을 사용할 수 있습니다.

# --instructions--

산점도에 이제 x축이 있습니다. `yAxis` 메서드를 사용해 `axisLeft()`라는 변수에 y축을 만드세요. 그런 다음 `g` 요소를 사용해 축을 렌더링하세요. 축을 오른쪽으로 패딩 단위만큼, 아래로 `transform` 단위만큼 이동시키기 위해 `0` 속성을 꼭 사용하세요. 축을 `call()`하는 것도 잊지 마세요.

# --hints--

코드는 `axisLeft()` 메서드를 `yScale` 인수와 함께 사용해야 합니다.

```js
assert.match(code, /\.axisLeft\(yScale\)/g);
```

y축 `g` 요소에는 `transform`만큼 축을 이동시키는 `(60, 0)` 속성이 있어야 합니다.

```js
const yAxis = document.querySelectorAll('g')[10];
assert.match(yAxis?.getAttribute('transform'), /translate\(60\s*,\s*0\)/g);
```

코드는 `yAxis`를 호출해야 합니다.

```js
assert.match(code, /\.call\(\s*yAxis\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);
    // Add your code below this line
    const yAxis = undefined;
    // Add your code above this line

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);

    const yAxis = d3.axisLeft(yScale);

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    svg
      .append('g')
      .attr('transform', 'translate(' + padding + ',0)')
      .call(yAxis);
  </script>
</body>
```
