---
id: 587d7faa367417b2b2512bd3
title: D3 레이블 스타일 지정
challengeType: 6
forumTopicId: 301492
dashedName: style-d3-labels
---

# --description--

D3 메서드는 막대 레이블에 스타일을 추가할 수 있습니다. `fill` 속성은 `text` 노드의 텍스트 색상을 설정합니다. `style()` 메서드는 `font-family` 또는 `font-size` 같은 다른 스타일에 대한 CSS 규칙을 설정합니다.

# --instructions--

`font-size` 요소의 `text`를 `25px`로 설정하고, 텍스트 색상을 빨간색으로 지정하세요.

# --hints--

레이블은 모두 빨간색 `fill`를 가져야 합니다.

```js
const text = document.querySelector('text');
assert.exists(text); 
const fill = window.getComputedStyle(text)['fill'];
assert.strictEqual(fill, 'rgb(255, 0, 0)');
```

레이블은 모두 `font-size` 픽셀의 `25`를 가져야 합니다.

```js
const text = document.querySelector('text');
assert.exists(text); 
const fontSize = window.getComputedStyle(text)['font-size'];
assert.strictEqual(fontSize, '25px');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy');

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy');

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3)
      .style('font-size', 25)
      .attr('fill', 'red');
  </script>
</body>
```
