---
id: 587d7fac367417b2b2512bdd
title: 동적 스케일 사용하기
challengeType: 6
forumTopicId: 301495
dashedName: use-dynamic-scales
---

# --description--

D3의 `min()`와 `max()` 메서드는 스케일 설정에 유용합니다.

복잡한 데이터 세트가 주어졌을 때, 우선순위 중 하나는 시각화가 SVG 컨테이너의 너비와 높이에 맞도록 스케일을 설정하는 것입니다. 모든 데이터가 SVG 안에 플로팅되어 웹 페이지에서 보이도록 해야 합니다.

아래 예제는 산점도 데이터의 x축 스케일을 설정합니다. `domain()` 메서드는 플롯의 원시 데이터 값에 대한 정보를 스케일에 전달합니다. `range()` 메서드는 시각화를 위한 웹 페이지 상의 실제 공간에 대한 정보를 제공합니다.

예제에서 도메인은 0부터 집합 내 최대값까지입니다. 배열의 x 값에 기반한 콜백 함수를 사용하는 `max()` 메서드를 사용합니다. 범위는 SVG의 너비(`w`)를 사용하지만, 패딩도 포함합니다. 이는 산점도 점과 SVG 가장자리 사이에 공간을 만듭니다.

```js
const dataset = [
  [34, 78],
  [109, 280],
  [310, 120],
  [79, 411],
  [420, 220],
  [233, 145],
  [333, 96],
  [222, 333],
  [78, 320],
  [21, 123]
];
const w = 500;
const h = 500;

const padding = 30;
const xScale = d3
  .scaleLinear()
  .domain([0, d3.max(dataset, d => d[0])])
  .range([padding, w - padding]);
```

패딩이 처음에는 혼란스러울 수 있습니다. x축을 0에서 500까지의 수평선으로 생각하세요(SVG의 너비 값). `range()` 메서드에 패딩을 포함하면 플롯이 그 선을 따라 0이 아닌 30에서 시작하고, 500이 아닌 470에서 끝나도록 강제합니다.

# --instructions--

`yScale` 변수를 사용해 선형 y축 스케일을 만드세요. 도메인은 0부터 집합 내 최대 `y` 값까지여야 합니다. 범위는 SVG 높이(`h`)를 사용하고 패딩을 포함하세요.

**참고:** 플롯이 올바른 방향으로 표시되도록 하세요. y 좌표의 범위를 설정할 때, 더 큰 값(높이에서 패딩을 뺀 값)이 첫 번째 인수이고, 더 작은 값이 두 번째 인수입니다.

# --hints--

`h2` 안의 텍스트는 `30` 이어야 합니다.

```js
assert.strictEqual(output, 30);
assert.strictEqual(document.querySelector('h2')?.textContent, '30');
```

yScale의 `domain()` 는 `[0, 411]` 와 같아야 합니다.

```js
assert.deepEqual(yScale.domain(), [0, 411]);
```

yScale의 `range()` 는 `[470, 30]` 와 같아야 합니다.

```js
assert.deepEqual(yScale.range(), [470, 30]);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    // Padding between the SVG boundary and the plot
    const padding = 30;

    // Create an x and y scale

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    // Add your code below this line

    const yScale = undefined;


    // Add your code above this line

    const output = yScale(411); // Returns 30
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const padding = 30;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const output = yScale(411);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
