---
id: 5d822fd413a79914d39e98ca
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Añade una etiqueta de apertura y cierre `html` abajo de `DOCTYPE`para que tengas espacio para comenzar a poner código. Asegúrate de establecer el idioma al español.

# --hints--

Su declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Tu elemento `html` debe tener una etiqueta de apertura, y un atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Tu etiqueta `html` debe estar en el orden correcto.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Solamente debes de tener una etiqueta `html`.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
