---
id: 6735a7370e0ae93a4577c689
title: 新聞レイアウトのデザイン
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

このラボでは、CSS Gridを使って新聞のレイアウトをデザインします。グリッドの行や列の概念も含まれます。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. ページにはクラス`main`を持つ`newspaper-layout`要素が含まれている必要があります。
1. `.newspaper-layout`には、新聞の名前を表示する`header`を含むクラス`title`の`h1`が含まれている必要があります。
1. `.newspaper-layout`には、主要なニュース記事用のクラス`article`の`feature-article`が含まれている必要があります。
1. `.feature-article`には、記事タイトル用の`h2`要素と記事内容用の`p`要素が含まれている必要があります。
1. 3つの小さな記事用に、クラス`article`、`small-article1`、`small-article2`を持つ`small-article3`要素をさらに追加してください。
1. 各小さな記事には、記事タイトル用の`h3`要素と記事内容用の`p`要素が含まれている必要があります。
1. `.newspaper-layout`には、追加ニュースセクション用のクラス`article`を持つ`secondary-article`要素が含まれている必要があります。
1. `.secondary-article`には、記事タイトル用の`h2`要素と記事内容用の`p`要素が含まれている必要があります。
1. `.newspaper-layout`には、新聞の内容を表す画像を表示するクラス`figure`を持つ`cover-image`が含まれている必要があります。
1. クラス`title`、`feature-article`、`secondary-article`、`cover-image`、`small-article1`、`small-article2`、`small-article3`を持つ要素には、同じクラス名の`grid-area`プロパティが設定されている必要があります。
1. `.newspaper-layout`はCSS Gridを使用し、セクションの配置を定義する`grid-template-areas`プロパティを持つ必要があります:
   - `.title`は最上段の行にまたがる必要があります。
   - `.feature-article`は2行目で2列にまたがり、`.cover-image`は3列目に配置されます。
   - `.secondary-article`は3行目で2列にまたがり、`.cover-image`は3列目に配置されます。
   - 3つの小さな記事は4行目に表示されます。
1. `.newspaper-layout`は3つの等しい列に分割する`grid-template-columns`プロパティを持つ必要があります。
1. `.newspaper-layout`の`grid-template-rows`プロパティは`auto`に設定し、最初の行を指定し、残りのスペースは他の行で等分割してください。
1. グリッドアイテム間にギャップを追加してください。
1. `.cover-image`内の画像が指定されたスペースに収まるように、`max-width`を`100%`に設定してください。

**注意:** HTMLにスタイルシートをリンクし、CSSを適用してください。

# --hints--

`main`要素が必要です。

```js
assert.exists(document.querySelector('main'));
```

`main`要素には`class`が`newspaper-layout`である必要があります。

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

`header`要素内に`class`の`title`を持つ`.newspaper-layout`要素が必要です。

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

`header`要素には`h1`要素が含まれている必要があります。

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

`.title`要素にはテキストが含まれている必要があります。

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

各記事にはクラス`article`、`feature-article`、`small-article1`、`small-article2`、`small-article3`を持つ`secondary-article`要素が必要です。

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

`.feature-article`には`h2`要素と`p`要素が続けて含まれている必要があります。

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`、`.small-article2`、`.small-article3`には`h3`要素と`p`要素が続けて含まれている必要があります。

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

`.secondary-article`には`h2`要素と`p`要素が続けて含まれている必要があります。

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`figure`内にクラス`cover-image`を持つ`.newspaper-layout`要素が必要です。

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

`.cover-image`には`img`要素が含まれている必要があります。

```js
assert.exists(document.querySelector(".cover-image img"));
```

`display`要素の`.newspaper-layout`プロパティは`grid`に設定してください。

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

`.newspaper-layout`要素には、最初の行の3列すべてにまたがる`grid-template-areas`を持つ`title`プロパティが必要です。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

グリッドテンプレートの2行目には、2列にまたがる`feature-article`と`cover-image`が含まれます。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

グリッドテンプレートの3行目には、2列にまたがる`secondary-article`と`cover-image`が含まれます。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

グリッドテンプレートの4行目には`small-article1`、`small-article2`、`small-article3`が含まれます。

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

`.newspaper-layout`セレクターは、スペースを3つの等しい部分に分割する`grid-template-columns`プロパティを使用してください。`repeat`関数を使うか、各分数単位を手動で定義できます。

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

`.newspaper-layout`セレクターには`grid-template-rows`プロパティが`auto 1fr 1fr 1fr`に設定されている必要があります。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

グリッドアイテム間にギャップを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title`には`grid-area`が`title`に設定されています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article`には`grid-area`が`feature-article`に設定されています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image`には`grid-area`が`cover-image`に設定されています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article`には`grid-area`が`secondary-article`に設定されています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1`には`grid-area`が`small-article1`に設定されています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2`には`grid-area`が`small-article2`に設定されています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3`には`grid-area`が`small-article3`に設定されています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

`img`内の`.cover-image`には`max-width`が`100%`に設定されています。

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
