---
id: 615f2d4150fe0d4cbd0f2628
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `h1` 元素下方，添加一個帶有文字 `8 servings per container` 的 `p` 元素。

# --hints--

你應該添加一個新的 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

你的 `p` 元素應該位於你的 `body` 元素內。

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

你的 `p` 元素應該位於你的 `h1` 元素之後。

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

你的 `p` 元素應該有文字 `8 servings per container`。

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
