---
id: 615f3949f58e12577dcefb00
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

善用空白可以讓你的頁面重要元素更為突出，並幫助引導使用者的目光瀏覽你的文字。

為你的 `.label` 選擇器添加一個設定為 `270px` 的 `width` 屬性。

# --hints--

你的 `.label` 選擇器應該設定 `width` 屬性為 `270px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.width === '270px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--
.label {
  border: 2px solid black;
}
--fcc-editable-region--
```
