---
id: 615f41c979787462e76dab90
title: 第 20 步
challengeType: 0
dashedName: step-20
---

# --description--

你的新類別尚未有任何樣式。創建一個 `.bold` 選擇器並為它設定 `font-weight` 屬性為 `800`，以使文字變粗。

接著也從你的 `h1` 選擇器中移除 `font-weight` 屬性。

# --hints--

你應該有一個 `.bold` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold'));
```

你的 `.bold` 選擇器應該有一個 `font-weight` 屬性設定為 `800`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold')?.fontWeight === '800');
```

你的 `h1` 選擇器不應該有 `font-weight` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === "");
```

你不應該移除你的 `h1` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p class="bold">Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--
h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

--fcc-editable-region--
```
