---
id: 615f38cabc64e3556f98cc1a
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

邊框可用來將內容群組並設定優先順序。

創建 `.label` 選擇器並為它設定 `border` 為 `2px solid black`。

# --hints--

你應該有一個 `.label` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.label'));
```

你的 `.label` 選擇器應該有一個設定為 `2px solid black` 的 `border` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.border === '2px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--

--fcc-editable-region--
```
