---
id: 615f38279e5c3d53692ea441
title: 步驟7
challengeType: 0
dashedName: step-7
---

# --description--

將你的 `h1` 和 `p` 元素包裝在 `div` 元素中。 將該 `div` 的 `class` 屬性設置爲 `label`。

# --hints--

你應該創建一個新的 `div` 元素。

```js
assert(document.querySelector('div'));
```

你的新的 `div` 元素應將 `class` 屬性設置爲 `label`。

```js
assert(document.querySelector('div')?.classList?.contains('label'));
```

你的 `h1` 和 `p` 元素應該在新的 `.label` 元素中。

```js
const children = [...document.querySelectorAll('h1'), ...document.querySelectorAll('p')];
assert(children?.every(child => child?.parentElement?.classList?.contains('label')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

--fcc-editable-region--
<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
--fcc-editable-region--
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}
```
