---
id: 615f357957e370510f21ea16
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

創建一個 `body` 選擇器，並將其 `font-family` 設定為 `Open Sans`，備用字型為 `sans-serif`。

請記得名稱中有空格的字型必須用引號包起來以供 CSS 使用。

# --hints--

你應該有一個 `body` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

你的 `body` 選擇器應該設定 `font-family` 屬性為 `"Open Sans", sans-serif`。

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
