---
id: 615f34ecc1091b4fd5a8a484
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在你的 `head` 元素中，為你的文件添加一個 `link` 元素，其 `rel` 屬性設為 `stylesheet`，`href` 屬性設為 `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`。

這將匯入 `Open Sans` 字型家族，字型粗細值為 `400`、`700` 和 `800`。

同時為你的 `styles.css` 檔案添加一個 `link` 元素。

# --hints--

你的程式碼應該有兩個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

你的 `link` 元素應該是空元素，不應該有結束標籤 `</link>`。

```js
assert.notMatch(code, /<\/link>/);
```

你的兩個 `link` 元素應該在 `head` 元素內。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

你的兩個 `link` 元素都應該將 `rel` 屬性設為 `stylesheet`。

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

你的其中一個 `link` 元素應該有一個 `href` 屬性設定為 `styles.css`。

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

你的其中一個 `link` 元素應該有一個 `href` 屬性設定為 `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
