---
id: 642df9df4b5216350de7b0d2
title: 步驟6
challengeType: 0
dashedName: step-6
---

# --description--

你將需要一個函數來生成一系列數字。

聲明一個空的 `range` 函數，該函數接受 `start` 和 `end` 參數。 使用 `Array()` 構造函數並隱式返回一個空數組。

# --hints--

你應該聲明一個 `range` 變量。

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

你應該使用 `const` 來聲明你的 `range` 變量。

```js
assert.match(code, /const\s+range/);
```

你的 `range` 變量應該是一個函數。

```js
assert.isFunction(range);
```

你的 `range` 函數應該使用箭頭語法。

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

你的 `range` 函數應該首先接受 `start` 參數。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

你的 `range` 函數應該第二接受 `end` 參數。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

你的 `range` 函數應該使用隱式返回。 請記住，這意味着你將不會使用花括號。

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

你的 `range` 函數應該使用 `Array()` 構造函數。 原始構造函數不需要 `new` 關鍵字。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

你不應將任何內容傳遞給 `Array()` 構造函數。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
