---
id: 642db8c409d9991d0b3b2f0d
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

你的項目從一個基本的 HTML 容器和一些相應的 CSS 開始。 你的首要任務是以編程方式生成電子表格的單元格。

全局 `window` 對象代表瀏覽器窗口（或選項卡）。 它有一個 `onload` 屬性，允許你定義窗口加載整個頁面（包括樣式表和腳本）時的行爲。

首先將 `window` 的 `onload` 屬性設置爲沒有參數的箭頭函數。 在函數中聲明一個 `container` 變量，併爲其賦值通過 `id` 爲 `"container"` 獲取元素的值。

# --hints--

你應該訪問 `window` 對象的 `onload` 屬性。

```js
assert.match(code, /window\.onload/);
```

你應該將 `onload` 屬性設置爲一個函數。

```js
assert.isFunction(window.onload);
```

你應該使用箭頭語法。

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

你的 `onload` 函數不應接受任何參數。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

你應該在 `onload` 函數中聲明一個 `container` 變量。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

你的 `container` 變量應使用 `const` 聲明。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

你應該使用 `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

你應該有一個 `id` 爲 `"container"` 的元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

你應該將 `#container` 元素分配給 `container`。

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
