---
id: 642dfb07e7fa6736251541c8
title: 步驟7
challengeType: 0
dashedName: step-7
---

# --description--

你的數組需要是該範圍的大小。 你可以通過查找 `end` 和 `start` 之間的差值，然後將 `1` 添加到結果中來計算這一點。

將此計算作爲 `Array()` 構造函數的參數傳遞。

# --hints--

你應該從 `end` 中減去 `start`。

```js
assert.match(code, /end\s*-\s*start/);
```

你應該將 `1` 添加到 `end - start` 計算中。

```js
assert.match(code, /end\s*-\s*start\s*\+\s*1/);
```

你應該將計算傳遞給 `Array()` 構造函數。

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
const range = (start, end) => Array();
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
