---
id: 642e02be7845f13b014cd2b0
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

現在你有了 `range` 函數，你也可以使用它來創建字母範圍。

使用 `const` 和箭頭語法聲明 `charRange` 函數。 它應該採用 `start` 和 `end` 參數。 該函數應隱式返回以 `start` 和 `end` 作爲參數調用 `range()` 的結果。

# --hints--

你應該聲明一個 `charRange` 變量。

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

你的 `charRange` 變量應使用 `const` 聲明。

```js
assert.match(code, /const\s+charRange/);
```

你的 `charRange` 變量應該是一個函數。

```js
assert.isFunction(charRange);
```

你的 `charRange` 函數應該使用箭頭語法。

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

你的 `charRange` 函數應該將 `start` 作爲第一個參數。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

你的 `charRange` 函數應該將 `end` 作爲第二個參數。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

你的 `charRange` 函數應該使用隱式返回。

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

你的 `charRange` 函數應該調用你的 `range` 函數。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

你應該將 `start` 和 `end` 作爲參數傳遞給 `range` 調用。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
