---
id: 643498755d54c6279ba09078
title: 步驟 20
challengeType: 0
dashedName: step-20
---

# --description--

大多數電子表格程序都包含用於計算的內置函數。

聲明一個 `sum` 函數，該函數接受一個 `nums` 參數，該參數是一個數字數組。 它應該返回對數組調用 `reduce` 對所有數字求和的結果。

# --hints--

你應該聲明一個 `sum` 變量。

```js
assert.match(code, /(?:let|const|var)\s+sum/);
```

你應該使用 `const` 來聲明你的 `sum` 變量。

```js
assert.match(code, /const\s+sum/);
```

你的 `sum` 變量應該是一個函數。

```js
assert.isFunction(sum);
```

你的 `sum` 函數應該使用箭頭語法。

```js
assert.match(code, /const\s+sum\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

你的 `sum` 函數應該有一個 `nums` 參數。

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>/);
```

你的 `sum` 函數應該使用隱式返回。

```js
assert.notMatch(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

你的 `sum` 函數應該返回在 `nums` 上調用 `.reduce()` 的結果。

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*nums\.reduce\(/);
```

你的 `sum` 函數應該返回 `nums` 中所有數字的總和。

```js
const numbers = [1, 2, 3, 4, 5];
assert.equal(sum(numbers), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
