---
id: 6434552bcc0a951a0a99df3b
title: 步驟11
challengeType: 0
dashedName: step-11
---

# --description--

你的 `range` 函數需要數字，但你的 `start` 和 `end` 值將是字符串（具體來說，它們將是單個字符，例如 `A`）。

通過使用 `.charCodeAt()` 方法將 `range()` 調用中的 `start` 和 `end` 值轉換爲數字，並將數字 `0` 作爲參數傳遞給該方法。

# --hints--

你應該使用 `.charCodeAt()` 方法。

```js
assert.match(code, /\.charCodeAt\(/);
```

你應該在 `start` 上調用 `.charCodeAt()` 方法。

```js
assert.match(code, /start\.charCodeAt\(/);
```

你應該將 `0` 傳遞給 `start` 的 `.charCodeAt()` 方法。

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

你應該在 `end` 上調用 `.charCodeAt()` 方法。

```js
assert.match(code, /end\.charCodeAt\(/);
```

你應該將 `0` 傳遞給 `end` 的 `.charCodeAt()` 方法。

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

你應該在 `range` 調用中直接使用 `.charCodeAt()` 方法。

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
