---
id: 64345c560591891f64976f7a
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

現在你的輔助函數已完成，回到你的 `onload` 事件處理程序中，你應該聲明一個 `letters` 變量。 將以字母 `"A"` 和 `"J"` 作爲參數調用 `charRange()` 的結果賦值給它。

# --hints--

你應該聲明一個 `letters` 變量。

```js
assert.match(code, /(?:let|const|var)\s+letters/);
```

你應該使用 `const` 來聲明你的 `letters` 變量。

```js
assert.match(code, /const\s+letters/);
```

你應該爲你的 `letters` 變量分配一個 `charRange()` 調用。

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(/);
```

你應該將 `"A"` 作爲第一個參數傳遞給 `charRange()` 調用。

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1/);
```

你應該將 `"J"` 作爲第二個參數傳遞給 `charRange()` 調用。

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1\s*,\s*('|"|`)J\2\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
