---
id: 643715013330824ecaa70442
title: 步驟 23
challengeType: 0
dashedName: step-23
---

# --description--

你的下一個函數將計算數字數組的中值。 首先聲明一個接受 `nums` 參數的 `median` 箭頭函數。

在函數中，聲明一個 `sorted` 變量，併爲其分配對 `nums` 數組的副本進行排序的值。

你應該使用 `slice()` 方法來創建數組的淺拷貝。

# --hints--

你應該聲明一個 `median` 變量。

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

你應該使用 `const` 來聲明你的 `median` 變量。

```js
assert.match(code, /const\s+median/);
```

你的 `median` 變量應該是一個函數。

```js
assert.isFunction(median);
```

你的 `median` 函數應該使用箭頭語法。

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

你的 `median` 函數應該有一個 `nums` 參數。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

你的 `median` 函數不應使用隱式返回。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

你的 `median` 函數應該有一個 `sorted` 變量。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

你應該使用 `const` 來聲明你的 `sorted` 變量。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

你應該使用 `.slice()` 將 `nums` 數組的副本分配給 `sorted`。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

你應該將 `.sort()` 方法鏈接到 `.slice()` 方法。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

你應該將回調函數傳遞給 `sort` 方法，以便準確地按升序對數字進行排序。 爲了清楚起見，使用隱式返回。

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
