---
id: 68af273960e009597dd861e2
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

تريد عرضًا مقروءًا لتصحيح الأخطاء والطباعة. أنشئ طريقة `__str__` تُرجع سلسلة نصية بالتنسيق `title (year) - duration min, director` (حيث يجب استبدال `title` و`year` و`duration` و`director` بقيم الخصائص المقابلة) بحيث تعطي طباعة الفيلم ملخصًا جميلًا.

# --hints--

يجب أن يحتوي صنف `Movie` على طريقة `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

يجب أن تُرجع طريقة `__str__` سلسلة نصية بالتنسيق `title (year) - duration min, director`، حيث يجب استبدال `title` و`year` و`duration` و`director` بقيم الخصائص المقابلة.

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
--fcc-editable-region--
    
--fcc-editable-region--
```
