---
id: 68af274560e009597dd861e3
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

الآن، أنشئ نسخة من `Movie` مع تمرير المعلمات التي تختارها وقم بتعيينها إلى متغير باسم `movie1`. ثم اطبع `movie1` في وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير باسم `movie1`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

يجب تعيين نسخة من `Movie` إلى المتغير `movie1`.

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

يجب طباعة `movie1` في وحدة التحكم.

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

--fcc-editable-region--

--fcc-editable-region--
```
