---
id: 68af275860e009597dd861e5
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

تم إنتاج أول فيلم في عام 1895. أنشئ جملة `if` أخرى تتحقق مما إذا كانت قيمة `year` أقل من `1895` وارفع استثناء `ValueError` مع الرسالة `Year must be 1895 or later`.

# --hints--

يجب أن ترفع استثناء `ValueError` مع الرسالة `Year must be 1895 or later` عندما يكون `year` أقل من `1895`.

```js
({
  test: () => runPython(`
  _msg = 'Year must be 1895 or later'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie('The Matrix', 1894, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('The Matrix', 0, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

لا يجب رفع استثناء `ValueError` عندما تكون قيمة `year` رقمًا صحيحًا.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1895, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid year'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
--fcc-editable-region--
        
--fcc-editable-region--
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
    
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
