---
id: 68af272f60e009597dd861e1
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستنشئ كتالوجًا للوسائط. ابدأ بتعريف صنف باسم `Movie`.

أضف إلى الصنف طريقة `__init__` بالمعلمات التالية: `self`، `title`، `year`، `director`، و`duration`. داخل طريقة `__init__`، عيّن معلمات `title`، `year`، `director`، و`duration` إلى خصائص الكائن بنفس الأسماء.

# --hints--

يجب أن تنشئ صنفًا باسم `Movie`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

يجب أن يحتوي صنف `Movie` على طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

يجب أن تحتوي طريقة `__init__` على المعلمات التالية: `self`، `title`، `year`، `director`، و`duration`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

يجب أن تعيّن `title` إلى `self.title` داخل طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

يجب أن تعيّن `year` إلى `self.year` داخل طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

يجب أن تعيّن `director` إلى `self.director` داخل طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

يجب أن تعيّن `duration` إلى `self.duration` داخل طريقة `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
