---
id: 68af274e60e009597dd861e4
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

في الخطوات السابقة، أنشأت كائن فيلم أساسي. الآن، أضف بعض التحقق للتأكد من عدم إمكانية إنشاء أفلام بخصائص غير صالحة.

في بداية طريقة `__init__`، أنشئ جملة `if` تتحقق مما إذا كان `title` إما فارغًا أو يحتوي فقط على مسافات بيضاء (يمكنك استخدام طريقة `strip` لبناء الشرط).

إذا كان الأمر كذلك، ارفع استثناء `ValueError` مع الرسالة `Title cannot be empty`.

# --hints--

يجب أن ترفع استثناء `ValueError` مع الرسالة `Title cannot be empty` عندما يكون `title` إما فارغًا أو يحتوي فقط على مسافات بيضاء.

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

لا يجب أن ترفع استثناء `ValueError` عندما يكون `title` سلسلة نصية صالحة.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
--fcc-editable-region--
        
--fcc-editable-region--
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
