---
id: 69757cc0faae0152c1418aad
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將練習使用數字和數學操作來建置一個帳單分裂視窗。這個工具會在加入餐費和小費後計算每個人應付的金額。

要開始，你需要一種方法來追蹤隨著成本增加的總金額。在 Python 中，你可以使用 `variable` 來儲存會隨時間改變的整數（整數）。

例如，你可能會寫：

```py
my_number = 2
```

創建一個名為 `running_total` 的變數並指定其值為 `0`。

# --hints--

你應該有一個名為 `running_total` 的變數。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

你應該將整數 `0` 指定給你的 `running_total` 變數。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
