---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

每個課程都有費用。你需要將這些金額儲存在變數中以便稍後使用。由於這些金額包含分，因此你將使用用來表述小數的 `float` 型別。以下是一個具有浮點值的變數範例：

```py
change = 2.35
```

創建四個變數：`appetizers` 設為 `37.89`、`main_courses` 設為 `57.34`、`desserts` 設為 `39.39`，以及 `drinks` 設為 `64.21`。 

# --hints--

你應該定義一個名為 `appetizers` 的變數並指定其值為 `37.89`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

你應該定義一個名為 `main_courses` 的變數並將其指定為值 `57.34`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

你應該定義一個名為 `desserts` 的變數並將其指定為值 `39.39`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

你應該定義一個名為 `drinks` 的變數並指定它的值為 `64.21`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
