---
id: 6976cc0b0c135686a4620b94
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

接下來，你需要考慮分攤帳單的人數。將此數值儲存在一個變數中，就像你在前一步驟所做的那樣。

創建一個名為 `num_of_friends` 的變數並指定其值為 `4`。這將在工作坊稍後用於計算最終分配。

# --hints--

你應該有一個名為 `num_of_friends` 的變數。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

你應該將整數 `4` 指定給你的 `num_of_friends` 變數。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
