---
id: 6977a4306bb3e856690a4890
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

現在你已經計算出小費，你需要將它添加到你的 `running_total` 以找出最終帳單金額。

在 Python 中，你可以使用增強型指定運算子 `+=` 將值加到變數並同時更新該變數。舉例來說，`total += 5` 是寫成 `total = total + 5` 的簡寫方式。

使用 `+=` 運算子將 `tip` 的值添加到你的 `running_total`。最後，使用 `print()` 列印字串 `Total with tip:`，後面接一個空格和 `running_total` 的值。

# --hints--

你應該使用 `+=` 運算子將 `tip` 的值添加到你的 `running_total`。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

你應該列印字串 `Total with tip:`，後面接一個空格和 `running_total` 變數。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total with tip:', running_total)") or _Node(_code).has_call("print(f'Total with tip: {running_total}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total = appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
