---
id: 69779dcb44016eccc05c15a4
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

服務非常優秀，所以群組決定留下 25％ 的小費。要在 Python 中計算百分比，你可以將總額乘以百分比的十進位等價。

例如，要找出一個值的 10％，你會使用 `*` 運算子將它乘以 `0.10`：

```py
tax = total * 0.10
```

創建一個名為 `tip` 的變數，並將 `running_total` 乘以 `0.25` 的結果指定給它。

最後，使用 `print()` 列印字串 `Tip amount:`，後面接一個空格和你的 `tip` 變數的值。

# --hints--

你應該有一個名為 `tip` 的變數。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

你的 `tip` 變數應該是 `running_total * 0.25` 的結果。

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

你應該列印字串 `Tip amount:`，後面接一個空格和 `tip` 變數。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total = appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
