---
id: 6977a8bbcbdd557886bd0718
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

現在包含小費後，你有整個群組的最終金額。你必須將帳單總額除以朋友人數，以決定每個人應付多少。 

在 Python 中，你使用斜線 `/` 來執行除法。例如：

```py
half = 10 / 2
```

創建一個名為 `final_bill` 的變數，並將 `running_total` 除以 `num_of_friends` 的結果指定給它。 

最後，使用 `print()` 函式列印字串 `Bill per person:`，後面接一個空格和 `final_bill` 的值。

# --hints--

你應該定義一個名為 `final_bill` 的變數。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

你應該使用 `/` 運算子將 `running_total` 除以 `num_of_friends`，並將結果指定到你的 `final_bill` 變數。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

你應該列印字串 `Bill per person:`，後面接一個空格和 `final_bill` 變數。

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
