---
id: 6982684f3a25f379e195a5fc
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

現在你已經儲存了各別的成本，你可以計算總和。在 Python 中，你使用加法運算子 `+` 將數值相加。 

將 `appetizers`、`main_courses`、`desserts` 和 `drinks` 相加，更新 `running_total` 變數。 

最後，使用 `print()` 列印字串 `Total bill so far:`，後面接一個空格和 `running_total` 的值。

**注意：**你可能會注意到輸出有比預期更多的小數位數。正如你在之前的課程中學到的，這是因為數字是以二進位儲存，且許多十進位值無法在此格式中被精確表述，導致四捨五入誤差。

# --hints--

你應該使用 `+` 運算子更新 `running_total`，以加總所有四個課程的費用。 

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total = {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

你應該列印字串 `Total bill so far:`，後面接一個空格和 `running_total` 的值。

```js
({
    test: () => runPython(`
    sol1 = "print('Total bill so far:', running_total)"
    sol2 = "print(f'Total bill so far: {running_total}')"
    assert _Node(_code).has_call(sol1) or _Node(_code).has_call(sol2)`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
