---
id: 67457ff2dc531fbed928a0f0
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Agora é hora de construir o componente `Navbar` usando JSX.

Você aprendeu em lições anteriores que JSX é uma extensão de sintaxe para JavaScript que permite escrever código semelhante a HTML dentro do JavaScript, facilitando a criação e o gerenciamento de componentes React.

Aqui está um lembrete de como o JSX se parece:

```jsx
const App = () => {
  return (
    <div>
      <h1>freeCodeCamp React Curriculum</h1>
      <p>I love learning with freeCodeCamp!</p>
    </div>
  );
}
```


Dentro do componente `Navbar`, retorne um elemento `nav` que tenha um `className` de `navbar`.

A razão pela qual `className` é usado aqui em vez de `class` é porque `class` é uma palavra-chave reservada em JavaScript. Então, o React usa `className` como uma forma de aplicar classes aos elementos.

# --hints--

Você deve retornar um elemento `nav`.

```js
assert.exists(document.querySelector('nav'));
```

Seu elemento `nav` deve ter um atributo `className` com o valor `navbar`.

```js
assert.equal(document.querySelector('nav')?.getAttribute('class'), 'navbar');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Navbar</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Navbar } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <Navbar />
      );
    </script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --white: #fff;
  --light-grey: #e1e0e0;
  --dark-purple: #7c0e7c;
  --black: #000;
}

body {
  background-color: var(--light-grey);
}

.navbar {
  background-color: var(--white);
}

.navbar ul {
  display: flex;
  justify-content: space-around;
}

.navbar ul li {
  list-style: none;
  border-radius: 4px;
}

.navbar ul li a {
  text-decoration: none;
  color: var(--black);
  padding: 10px;
  display: inline-block;
  width: 100%;
}

button {
  background: transparent;
  border: none;
  font-family: 'Times New Roman', Times, serif;
  padding: 10px;
  font-size: 1rem;
}

.navbar ul .nav-item a:hover {
  background-color: var(--dark-purple);
  color: var(--white);
}

button:hover {
  background-color: var(--dark-purple);
  color: var(--white);
}

.navbar ul .nav-item .sub-menu {
  visibility: hidden;
  opacity: 0;
  position: absolute;
  right: 5%;
  transition: opacity 0.5s ease;
  display: block;
  background-color: var(--white);
}

@media (min-width: 768px) {
  .navbar ul .nav-item .sub-menu {
    right: 15%;
  }
}

@media (min-width: 1024px) {
  .navbar ul .nav-item .sub-menu {
    right: 13%;
  }
}

.navbar ul .nav-item:hover .sub-menu,
.navbar ul .nav-item:focus-within .sub-menu {
  visibility: visible;
  opacity: 1;
}
```

```jsx
--fcc-editable-region--
export const Navbar = () => {
  
}
--fcc-editable-region--
```
