---
id: 66fe4f33a2cc9b33f4d5cd9b
title: 順列ジェネレーターを作成する
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

このラボでは、文字列を受け取り、その文字列の文字のすべての可能な順列を返す順列ジェネレーターを作成します。例えば、文字列 `cat` の可能な順列は `cat`、`cta`、`act`、`atc`、`tac`、および `tca` です。

文字列の順列を再帰的に作成する方法は、文字列の固定された開始部分（接頭辞）を保持し、残りの部分の順列を作成することで動作します。

例えば、単語 `machine` を考えてみましょう。順列を作成する最初のラウンドでは、文字列の接頭辞として `m` を固定し、その後、残りの文字列 `achine` の順列を作成します。

残りの文字列についても同じ方法で順列が続きます。1文字が接頭辞に追加されます。例えば `c` かもしれません。そうすると接頭辞は `mc` になります。次に、`ahine` の各順列が接頭辞に連結されます。

これが、接頭辞にすべての文字が含まれ、残りの文字列が空になるまで続きます。つまり、1つの順列が作成されたことになります。

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `permuteString` という名前の関数を作成してください。
2. `permuteString` 関数は1つのパラメータ（文字列）を受け取り、さらにデフォルト値を持つ2つのパラメータ（接頭辞の値と結果を格納して返すための空の配列）を受け取る必要があります。接頭辞の値は、順列を形成するために文字を蓄積するために使われる文字列です。この関数は `permuteString("cat")` のように1つの引数で呼び出されます。
3. 関数内で、渡された文字列の長さが `0` かどうかをチェックしてください。そうであれば、現在の接頭辞を結果に追加し、結果を返してください。
4. 入力文字列の各文字を繰り返し処理し、各反復で現在の文字を文字列から削除し、残りの順列を作成するために更新された引数で `permuteString` 関数を再帰的に呼び出してください。
5. 最終的な結果の配列を返してください。
6. 重複を削除して順列が一意になるようにしてください。

# --hints--

`permuteString` という関数を用意してください。

```js
assert.isFunction(permuteString);
```

`permuteString` 関数で再帰を使う必要があります。

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` は `[ "far", "fra", "afr", "arf", "rfa", "raf" ]` を返すべきです。

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` は `[ "fcc", "cfc", "ccf" ]` を返すべきです。

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` は `[ "p" ]` を返すべきです。

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` は `[""]` を返すべきです。

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` は `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]` を返すべきです。
`.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` は正しい結果を返すべきです。

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
