---
id: 6565a536ba1f9f25bd30e88b
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このプロジェクトでは、Luhnアルゴリズムを実装します。このアルゴリズムは、さまざまな識別番号の検証に使われる公式です。

まず、`main`という名前の関数を宣言してください。これはプログラムのエントリーポイントとして機能します。エラーを避けるために`pass`キーワードを使ってください。

# --hints--

コード内に`def main():`があることを確認してください。

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

コード内に`pass`があることを確認してください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
