---
id: 65687946130b0ea10aa19b75
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

スクリプトの最後で `main` 関数を呼び出してください。

# --hints--

`main()` は `main` 関数の外にある必要があります。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        assert.match(transformedCode, /\nmain\(\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)


--fcc-editable-region--
```
