---
id: 656877c202f44995ed952608
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

`pass` の文を `card_number` という名前の変数と `'4111-1111-4555-1142'` という値に置き換えてください。

# --hints--

`card_number = '4111-1111-4555-1142'` 関数の中に `main` があるはずです。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

コードに `pass` があってはいけません。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
