---
id: 65687e0ed12e20c91811a48d
title: ステップ 31
challengeType: 20
dashedName: step-31
---

# --description--

`for` 関数の2つ目の `verify_card_number` ループの下に、`total` という名前の変数を作成し、奇数と偶数の数字の合計の値を代入してください。`total` をコンソールに出力してください。

# --hints--

`total = sum_of_odd_digits + sum_of_even_digits` 関数内に `verify_card_number` があるはずです。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const function_body_indentation = function_body.match(/^\s*/)[0];

        const total_re = new RegExp(`^${function_body_indentation}total\\s*=\\s*(sum_of_odd_digits\\s*\\+\\s*sum_of_even_digits|sum_of_even_digits\\s*\\+\\s*sum_of_odd_digits)`, "m");

        assert.match(function_body, total_re);
    }
})
```

`print(total)` 関数内に `verify_card_number` があるはずです。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*total\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
