---
id: 656877f71bba2b97acedf9af
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

Python には文字列操作に役立つ組み込みクラスが用意されています。そのうちの1つが `str` クラスです。このクラスには `maketrans` というメソッドがあり、翻訳テーブルを作成するのに役立ちます。このテーブルを使うと、文字列内の文字を置き換えることができます。

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

上記は文字列に対して呼び出すと、すべての `t` 文字を `c` に、すべての `l` 文字を `b` に置き換えます。

`card_translation` という変数を作成し、すべての `-` と ` ` 文字を空文字に置き換える翻訳テーブルを代入してください。

# --hints--

`card_translation` 内に `main` 変数を作成してください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

`card_translation` に `str.maketrans({'-': '', ' ': ''})` の値を代入してください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
