---
id: 656879a66338b2a461d5d307
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

Luhnアルゴリズムは次の通りです：

1. 右から左へ、2番目ごとの数字の値を2倍にします。もし積が9より大きければ、その積の各桁の数字を合計します。
2. すべての数字の合計を取ります。
3. すべての数字の合計が10の倍数であれば、その番号は有効です。そうでなければ無効です。

例として、チェックディジットが追加される口座番号「7992739871」があり、形式は7992739871xとなります：

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

`pass`の文を変数`sum_of_odd_digits`と値`0`に置き換えてください。

# --hints--

`sum_of_odd_digits = 0`を`verify_card_number`関数内に含める必要があります。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

`pass`を`verify_card_number`関数内に含めてはいけません。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
