---
id: 65687b48f2201ab32e06c37c
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

現在、文字列を整数に加えようとしているため、スクリプトは`TypeError`をスローしています。組み込みの`digit`関数を使って、`sum_of_odd_digits`変数を`int`に加える前に整数に変換すれば、この問題を修正できます。

```python
my_string = '123'
my_int = int(my_string)
```

`digit`変数を`sum_of_odd_digits`に加える前に整数に変換してください。その後、`print`の呼び出しを`verify_card_number`関数の最後に移動して、`sum_of_odd_digits`の値を出力してください。

# --hints--

`sum_of_odd_digits += int(digit)`ループの中に`for`があるはずです。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

`print(sum_of_odd_digits)`関数の最後に`verify_card_number`があるはずです。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

もう`print(digit)`は使わないでください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
