---
id: 656878585631369a6b2d2191
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

翻訳を定義するだけでは文字列は翻訳されません。翻訳テーブルを引数として、翻訳したい文字列に対して`translate`メソッドを呼び出す必要があります。

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

`translated_card_number`に対して`translate`を引数にして`card_number`メソッドを呼び出した結果を、`card_translation`という変数に代入してください。

# --hints--

`translated_card_number`の中に`main`という変数を作成してください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

`translated_card_number`に`card_number.translate(card_translation)`の値を代入してください。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
