---
id: 656879c6f35a85a59c06b3a7
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

これまでに、インデックス演算子 `[]` を使って文字列の要素（文字）にアクセスしたことがあります。インデックス演算子は `string[start:stop:step]` を使って文字列の範囲をアクセスすることもできます。

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

ここで `start` は開始インデックス（含む）、`stop` は終了インデックス（含まない）、`step` はスキップする文字数です。指定しない場合、`step` はデフォルトで 1 になります。

`card_number_reversed` の最初の4文字を代入した変数 `card_number` を作成してください。

# --hints--

`card_number_reversed = card_number[0:4]` を `verify_card_number` 関数内に含める必要があります。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
