---
id: 65687db6a9fcf3c548a6c1b7
title: ステップ 28
challengeType: 20
dashedName: step-28
---

# --description--

アルゴリズムの一部として、右から数えて2番目ごとに数字を2倍にします。もし2倍にした結果が`10`以上であれば、その2桁の数字を足し合わせます。例えば、数字が`6`の場合、2倍にすると`12`になります。`1`と`2`を足し合わせて`3`を得ます。これは、_整数除算_を使って最初の桁を取得し、モジュラス演算子（`%`）を使って2番目の桁を取得することで実現できます。

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

整数除算は、割り算の商を小数点以下切り捨てで求めます。

`if`文の中で、`number`に`number // 10`（整数除算）の結果と`number`と`10`のモジュラスの和を代入してください。

# --hints--

`number = number // 10 + number % 10`文の中に`if`があるはずです。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
