---
id: 6569b831a07d7154c793301b
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

パラメータ`verify_card_number`を持つ関数`card_number`を定義し、`pass`キーワードを使って関数が何もしないようにしてください。

# --hints--

`verify_card_number`というパラメータを持つ`card_number`という関数があるはずです。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_parameters } = verify_card_number;
        assert.equal(function_parameters.trim(), "card_number");
    }
})
```

`pass`関数の本体で`verify_card_number`キーワードを使う必要があります。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_body } = verify_card_number;
        assert.match(function_body, / +pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

main()
```
