---
id: 6564883669b5af4b69f794cc
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

你的 `all_characters` 變數是由所有小寫和大寫字母、所有 10 個數字以及幾個特殊字元組成的字串。 

就在它之前，添加一個註解說明 `Combine all characters`。

# --hints--

你應該將註解添加在你的 `all_characters` 變數正上方。

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
