---
id: 6564b8c9349bd76dc037967b
title: 第 29 步
challengeType: 20
dashedName: step-29
---

# --description--

`re` 模組中的 `search()` 函式會分析作為引數傳入的字串，尋找正則表達式樣式與字串首次匹配的位置。

宣告一個名為 `quote` 的變數，並將字串 `'Not all those who wander are lost.'` 指定給此變數。然後，列印 `pattern.search(quote)` 的結果。

# --hints--

你應該有一個 `quote` 變數。

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

你應該將提供的字串指定到你的新 `quote` 變數。

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

你應該列印 `pattern.search(quote)`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
