---
id: 65646ab6bced52360f44670f
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

現在宣告一個名為 `all_characters` 的變數，並將串接你現有變數的結果指定給它。

# --hints--

你應該宣告一個名為 `all_characters` 的變數。

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

你應該將 `letters`、`digits` 和 `symbols` 串接起來，並將結果指定給你的 `all_characters` 變數。

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
