---
id: 6564d096a55d707bd77ab67b
title: 第 38 步
challengeType: 20
dashedName: step-38
---

# --description--

插入符號 `^` 放置在字元類別的開頭時，會匹配類別中未指定的所有字元。

在你的字元類別內將 `^` 添加為第一個字元，看看會發生什麼。

# --hints--

你應該在你的字元類別內將 `^` 添加為第一個字元。

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[^a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '[a-z]t'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
