---
id: 6564d75a923d21815caaa445
title: 步驟 43
challengeType: 20
dashedName: step-43
---

# --description--

如果你需要匹配在樣式中具有特殊意義的字元，例如 `.` 或 `+`，你可以在它前面加上反斜線字元 `\` 來跳脫它。 例如，這會匹配字面上的加號：`\+`。

修改 `pattern`，使其符合單一字面上的點。

# --hints--

你應該將你的 `pattern` 變數修改為 `\.`。

```js
({ test: () => assert.match(code, /^pattern\s*=\s*("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '.+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
