---
id: 656474314bf37d3c83a85143
title: 步驟 10
challengeType: 20
dashedName: step-10
---

# --description--

`random` 模塊中的 `choice()` 函數接受一個序列並返回該序列的一個隨機項。

修改你的 `print()` 調用以使用 `choice()` 函數並傳遞 `all_characters` 作爲參數。

# --hints--

You should modify your existing `print()` call to print `random.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
