---
id: 65646ab6bced52360f44670f
title: 步驟5
challengeType: 20
dashedName: step-5
---

# --description--

現在聲明一個名爲 `all_characters` 的變量，並將連接現有變量的結果賦給它。

# --hints--

你應該聲明一個名爲 `all_characters` 的變量。

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

你應該連接 `letters`、`digits` 和 `symbols` 並將結果分配給你的 `all_characters` 變量。

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
