---
id: 65647c71a5d2bd431596f629
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

你的 `generate_password` 函式需要一些參數。首先為你的函式添加一個 `length` 參數。

# --hints--

你的 `generate_password` 函式（程式）應該接受一個 `length` 參數。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == "(length)"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import secrets
import string


def generate_password():    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

--fcc-editable-region--
```
