---
id: 6564d97a5ef99783216a6229
title: 步驟 44
challengeType: 20
dashedName: step-44
---

# --description--

Python 提供了一種特殊類型的字符串，稱爲*原始*字符串。 原始字符串以 `r` 爲前綴。 與常規字符串的主要區別在於它們如何處理反斜槓字符：在原始字符串中，反斜槓被視爲文字字符而不是轉義字符。 編寫正則表達式時，使用原始字符串是一種很好的做法，因爲它們通常可以包含很多 `\` 字符。

通過在 `pattern` 字符串前添加 `r` 作爲前綴，將其轉換爲原始字符串。

# --hints--

你應該將 `pattern` 變量修改爲 `r'\.'`。

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
