---
id: 6564efa70114b591b74d5679
title: 步驟 50
challengeType: 20
dashedName: step-50
---

# --description--

空白字元和最後的句號被匹配，因為它們是字串中唯一的非英數字元。

現在將你的 `quote` 字串轉換成單一底線字元。

# --hints--

你的 `quote` 變數應該是 `'_'`。

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "_") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\W'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
