---
id: 656494269ccce754411a2112
title: 步驟 27
challengeType: 20
dashedName: step-27
---

# --description--

`re` 模組允許你在你的程式碼中使用 *regular expressions*。你很快就會學到更多關於 regular expressions 的知識。

目前，為你的程式碼頂部添加一個 `import` 述語以匯入 `re` 模組。

# --hints--

你應該匯入 `re` 模組。

```js
({ test: () => assert.match(code, /^import\s+re/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

import secrets
import string
--fcc-editable-region--

def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
