---
id: 6564eebf0d2d6390b9377197
title: 步驟 49
challengeType: 20
dashedName: step-49
---

# --description--

現在，將 `pattern` 轉換為非字母數字字元的簡寫類別。

# --hints--

你的 `pattern` 變數應該是 `'\W'`。

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\W\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
