---
id: 6564c67db26c417561ab510d
title: 步驟 31
challengeType: 20
dashedName: step-31
---

# --description--

如你從輸出中所見，現在你的 regex 符合字串中第一個 `l`。

在你的樣式中，你可以在字元後面添加量詞，以指定該字元應該重複多少次。例如，`+` 量詞表示它應該重複一次或多次。

為你的 `pattern` 添加 `+` 量詞。

# --hints--

你應該將你的 `pattern` 變數修改為 `re.compile('l+')`。

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
