---
id: 656475bbf1c2573de1d2c69c
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

每次執行程式碼時，你都應該看到來自 `all_characters` 字串的隨機字元。這正是你想要達成的，以創建隨機密碼。

然而，`random` 所依賴的演算法使得產生的偽隨機數可被預測。因此，雖然 `random` 模組適用於大多數常見應用，但由於其決定性特性，無法用於密碼學用途。 

不要匯入 `random`，改為匯入 `secrets` 模組。然後將 `print()` 呼叫改為使用 `secrets.choice(all_characters)`。

# --hints--

你應該匯入 `secrets` 模組，而不是 `random` 模組。

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

你應該修改你現有的 `print()` 呼叫來列印 `secrets.choice(all_characters)`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
