---
id: 65649a631440b757f9f7b428
title: 步驟 28
challengeType: 20
dashedName: step-28
---

# --description--

正規表達式，或稱 regex，是用來匹配字串中某特定字元組合的樣式。當你需要在字串中匹配或尋找樣式以進行驗證、字元取代等用途時，它是 `if`/`else` 條件述語的有效替代方案。

`re` 模組中的 `compile()` 函式會將作為引數傳入的字串編譯成可被其他 `re` 方法使用的正規表達式物件。

宣告一個新的 `pattern` 變數，並將 `re.compile('i')` 的值指定給此變數。

# --hints--

你應該宣告一個 `pattern` 變數。

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

你應該將 `re.compile('i')` 指定到你的 `pattern` 變數。

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
