---
id: 656479aa5f298441c190bf8f
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

 宣告一個 `generate_password` 函式，並將除了 `import` 行以外的所有程式碼寫在函式主體內。

# --hints--

你應該宣告一個名為 `generate_password` 的函式（程式）。

```js
({
  test: () => assert(runPython(`
  _Node(_code).has_function('generate_password')
  `))
})
```

`import` 述語仍應該在函式外部。

```js
({
  test: () => assert(runPython(`
  imports_list = _Node(_code).find_imports()
  any([imp.is_equivalent('import secrets') for imp in imports_list]) and any([imp.is_equivalent('import string') for imp in imports_list])
  `))
})
```

這四個變數宣告應該移動到函式內部。

```js
({
  test: () => assert(runPython(`
  func = _Node(_code).find_function('generate_password')
  vars = ['letters', 'digits', 'symbols', 'all_characters']
  func.find_variable('letters').is_equivalent('letters = string.ascii_letters') and \\
  func.find_variable('digits').is_equivalent('digits = string.digits') and \\
  func.find_variable('symbols').is_equivalent('symbols = string.punctuation') and \\
  func.find_variable('all_characters').is_equivalent('all_characters = letters + digits + symbols') and \\
  all(not _Node(_code).has_variable(var) for var in vars)
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

--fcc-editable-region--
```
