---
id: 656469e62377ab34a5bcb8d1
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

宣告兩個名為 `digits` 和 `symbols` 的新變數，並分別將它們指定為 `string.digits` 和 `string.punctuation`。

# --hints--

你應該宣告一個名為 `digits` 的變數。

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

你應該將 `string.digits` 指定到你的 `digits` 變數。

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

你應該宣告一個名為 `symbols` 的變數。

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

你應該將 `string.punctuation` 指定到你的 `symbols` 變數。

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
