---
id: 656474314bf37d3c83a85143
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

`random` 模組中的 `choice()` 函式會接受一個序列並傳回該序列中的隨機項目。

修改你的 `print()` 呼叫以使用 `choice()` 函式（程式），並將 `all_characters` 作為引數傳遞。

# --hints--

你應該修改你現有的 `print()` 呼叫來列印 `random.choice(all_characters)`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
