---
id: 65648f4b2281ba50051ae39c
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

接著，寫一個以 `i` 為迴圈變數的 `for` 迴圈。使用 `range()` 函式來迭代到 `length` 的值。

在迴圈中，使用加法指定運算子將來自 `all_characters` 的隨機字元加到 `password` 的目前值。為此，使用 `secrets` 模組中的 `choice()` 函式（程式）。

# --hints--

你應該撰寫一個 `for` 迴圈，對 `range(length)` 進行迭代。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

你應該使用 `+=` 運算子，將 `all_characters` 中的隨機字元添加到 `password` 的當前值。

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
