---
id: 6564f0279e23ce924eedd1b2
title: 步驟 51
challengeType: 20
dashedName: step-51
---

# --description--

由於底線字元是變數名稱的有效字元，因此它包含在 `\w` 字元類別中（等價於 `[a-zA-Z0-9_]`），可方便地用來匹配變數名稱。

因此，` \W` 字元類別等同於 `[^a-zA-Z0-9_]`，其中底線字元不會被匹配。基於此原因，你無法使用它來匹配你所有的特殊字元。

刪除你程式碼中的最後三行。

# --hints--

你應該刪除你程式碼中的最後 3 行。

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  assert.notMatch(commentless_code, /pattern\s*=\s*r("|')\\W\1/);
  assert.notMatch(commentless_code, /quote\s*=\s*("|')_\1/);
  assert.notMatch(commentless_code, /print\(\s*re\.findall\(\s*pattern\s*,\s*quote\s*\)\s*\)/);
}})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\W'
quote = '_'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
