---
id: 6564f32b18480893cf7799fd
title: 第 52 步
challengeType: 20
dashedName: step-52
---

# --description--

現在，將原始字符串與 f 字符串結合起來，並在字符類中插入 `symbols` 變量。 請記住，你可以使用花括號 `{ }` 在 f 字符串中插入變量。

# --hints--

第四個約束元組中的第二項應該是字符串 `fr'[{symbols}]'`。

```js
({ test: () => assert.match(code, /\(\s*special_chars\s*,\s*(fr|rf)("|')\[\{\s*symbols\s*\}\]\2/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
--fcc-editable-region--            
            (special_chars, r'\W')
--fcc-editable-region--            
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
