---
id: 6564a767a7241362f7d8d664
title: 步驟18
challengeType: 20
dashedName: step-18
---

# --description--

獨立的單下劃線用於表示你不關心的值或不會在代碼中使用的值。 你的迭代變量實際上並未被使用。

將你的 `i` 變量修改爲單個下劃線。

# --hints--

你應該將 `i` 變量修改爲 `_`。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/for\s+_\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--
    password = ''
    # Generate password
    for i in range(length):
        password += secrets.choice(all_characters)

--fcc-editable-region--
```
