---
id: 65649122c7f77f519aaf0975
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

最後，宣告一個變數 `new_password` 並將呼叫 `generate_password` 的結果指定給它。將 `8` 作為引數傳遞給你的 `generate_password` 呼叫。

# --hints--

你應該呼叫 `generate_password`，並傳入 `8` 作為引數。

```js
({ test: () => assert.match(code, /generate_password\s*\(\s*8\s*\)/) })
```

你應該將 `generate_password(8)` 指定給變數 `new_password`。

```js
({
  test: () => {
    assert(__userGlobals.has("new_password"));
    assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*\)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--

--fcc-editable-region--
```
