---
id: 6564cbb6311a40783b3f5de6
title: 步驟 36
challengeType: 20
dashedName: step-36
---

# --description--

字符類還允許你指示要匹配的字符範圍。 你需要指定以連字符分隔的起始字符和結束字符，`-`。 字符遵循 Unicode 順序。

修改你的 `pattern` 變量以匹配 `quote` 變量中以小寫字母開頭的任何字母 `t`。 爲此，請使用從 `a` 到 `z` 的字符範圍。

# --hints--

你應該使用 `[a-z]` 類修改 `pattern` 變量，使其匹配 `quote` 中以小寫字母開頭的任何字母 `t`。

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, '[0123456789]')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
