---
id: 656477845006313fbfea0ad1
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

雖然效果看起來可能與 `random.choice()` 相等，但 `secrets` 確保你獲得作業系統能提供的最安全的隨機來源。

現在，刪除你的兩個 `print()` 呼叫。

# --hints--

你應該刪除你的兩個 `print()` 呼叫。

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
