---
id: 6564683821b2ee3174e7250e
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

你可以透過點記法存取匯入模組內定義的工具。點記法的運作方式是將一個點以及工具名稱附加到模組名稱後面。例如，以下是如何存取 `ascii_lowercase` 常數：

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

在這個專案中，你將使用來自 `string` 模組的不同常數。 

宣告一個名為 `letters` 的新變數，並將 `string.ascii_letters` 指定給此變數。

# --hints--

你應該宣告一個名為 `letters` 的變數。

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

你應該將 `string.ascii_letters` 指定到你的 `letters` 變數。

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
