---
id: 65649a631440b757f9f7b428
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

正則表達式是用於匹配字符串中特定字符組合的模式。 當你需要在字符串中匹配或查找模式以進行驗證、字符替換等操作時，它是 `if`/`else` 條件語句的有效替代方案。

`re` 模塊中的 `compile()` 函數將作爲參數傳遞的字符串編譯成正則表達式對象，該對象可被其他 `re` 方法使用。

聲明一個新的 `pattern` 變量，並將 `re.compile('i')` 的值賦給該變量。

# --hints--

你應該聲明一個 `pattern` 變量。

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

你應該將 `re.compile('i')` 賦值給你的 `pattern` 變量。

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
