---
id: 6564d97a5ef99783216a6229
title: 步驟 44
challengeType: 20
dashedName: step-44
---

# --description--

Python 提供一種特定型別的字串，稱為 *未經處理的* 字串。未經處理的字串以 `r` 作為前置式。與一般字串的主要區別在於它們如何處理反斜線字元：在未經處理的字串中，反斜線被視為字面字元，而非控制代碼。在撰寫正規表達式時，使用未經處理的字串是一種良好做法，因為它們通常可以包含大量的 `\` 字元。

將你的 `pattern` 字串透過在前面加上 `r` 轉換成未經處理的字串。

# --hints--

你應該將你的 `pattern` 變數修改為 `r'\.'`。

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
