---
id: 65646ffeaed2d238c562a014
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

將 `import` 述語放在你的程式碼頂部是常見的慣例。此外，如果有多個 `import` 述語，請按字母順序排序以提升可讀性。

在你的程式碼頂部匯入 `random` 模組。

# --hints--

你應該匯入 `random` 模組。

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
