---
id: 65648e4c5b316c4ec5e4fddc
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

在你的函式（程式）底部，宣告一個 `password` 變數並將空字串指定給此變數。

# --hints--

你應該在你的函式中有一個 `password` 變數，其值為空字串。請注意縮排。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/password\s*=\s*("|')\1/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--

--fcc-editable-region--
```
