---
id: 6564d27c8e4c197d367c3b1c
title: 步驟 39
challengeType: 20
dashedName: step-39
---

# --description--

將一個新的 tuple 添加到 `constraints` 列表。使用 `lowercase` 作為第一個項目，並使用符合單一小寫字母的正則表達式樣式作為第二個項目。

# --hints--

你應該為你的 `constraints` 列表添加第二個元組。使用 `lowercase` 參數作為第一個條款，並使用 `[a-z]` 作為第二個條款。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\[0-9\]\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*("|')\[a-z\]\2\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, '[0-9]')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = '[^a-z]t'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
