---
id: 6564f32b18480893cf7799fd
title: 步驟 52
challengeType: 20
dashedName: step-52
---

# --description--

現在，將你的未經處理的字串與 f-string 結合，並在字元類別中插入你的 `symbols` 變數。請記住，你可以使用大括號 `{ }` 在 f-string 中插入變數。

# --hints--

你第四個條件元組中的第二個項目應該是字串 `fr'[{symbols}]'`。

```js
({ test: () => assert.match(code, /\(\s*special_chars\s*,\s*(fr|rf)("|')\[\{\s*symbols\s*\}\]\2/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
--fcc-editable-region--            
            (special_chars, r'\W')
--fcc-editable-region--            
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
