---
id: 6564f0279e23ce924eedd1b2
title: 步驟51
challengeType: 20
dashedName: step-51
---

# --description--

由於下劃線字符是變量名的有效字符，因此它被包含在 `\w` 字符類中（相當於 `[a-zA-Z0-9_]`），可以方便地用於匹配變量名。

因此，`\W` 字符類等同於帶有不匹配下劃線字符的 `[^a-zA-Z0-9_]`。 所以，你不能使用它來匹配所有特殊字符。

刪除代碼中的最後三行。

# --hints--

你應該刪除代碼中的最後三行。

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  assert.notMatch(commentless_code, /pattern\s*=\s*r("|')\\W\1/);
  assert.notMatch(commentless_code, /quote\s*=\s*("|')_\1/);
  assert.notMatch(commentless_code, /print\(\s*re\.findall\(\s*pattern\s*,\s*quote\s*\)\s*\)/);
}})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\W'
quote = '_'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
