---
id: 656485a8a3496d4a36b1496a
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

這三個變數構成了在產生密碼時可供選擇的可能字元。

就在它們之前，添加一個註解說明 `Define the possible characters for the password`。

# --hints--

你應該在 `letters` 變數上方添加註解。

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
