---
id: 656470d517833a39bb8b5608
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

`random` 模組包含一個偽隨機數產生器。它的大部分功能依賴於 `random()` 函式，該函式傳回一個介於 `0.0`（含）和 `1.0`（不含）之間的浮點數。

從 `random` 模組呼叫 `random()` 函式並列印結果。

# --hints--

你應該列印呼叫 `random.random()` 的結果。

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
