---
id: 6564d68c34027a8072a704f4
title: 步驟 42
challengeType: 20
dashedName: step-42
---

# --description--

點字元是一個萬用符號，預設會匹配字串中的任何字元——除了換行字元。將目前的 `pattern` 替換為 `.` 後接 `+` 量詞，以修改 `pattern` 來匹配整個字串。

# --hints--

你應該修改你的 `pattern` 變數以符合整個 `quote` 字串。

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), ".+") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '[^a-z]t'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
