---
id: 6564642ce403fe29cc23fee7
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

Python 模組是一個包含用於執行特定任務的程式碼的檔案。Python 標準程式庫包含許多你可以匯入並在你的程式碼中使用的模組。你可以使用 `import` 述語，後面接模組名稱來達成此目的。

透過匯入 `string` 模組開始這個專案。

# --hints--

你應該匯入 `string` 模組。

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
