---
id: 6564cbb6311a40783b3f5de6
title: 第 36 步
challengeType: 20
dashedName: step-36
---

# --description--

字元類別也允許你顯示要比對的字元範圍。你需要指定起始字元和結束字元，並以連字號 `-` 分隔。字元依照 Unicode 順序排列。

修改你的 `pattern` 變數，以符合 `quote` 變數中任何前面有小寫字母的字母 `t`。為此，使用從 `a` 到 `z` 的字元範圍。

# --hints--

你應該修改你的 `pattern` 變數，以符合在 `quote` 中由小寫字母接續的任意字母 `t`，使用 `[a-z]` 類別。

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
