---
id: 6565c13fdb798865c161d8f8
title: 步驟 65
challengeType: 20
dashedName: step-65
---

# --description--

現在是測試你的函式（程式）的時候了。取消註解你程式碼中最後兩行，並修改函式（程式）呼叫，傳入 5 個引數。長度使用 `8`，其他四個條件使用 `1`。 

# --hints--

你應該使用提供的引數呼叫 `generate_password`。

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*,\s*1\s*,\s*1\s*,\s*1\s*,\s*1\s*\)/m) })
```

你應該列印你的 `new_password` 變數。

```js
({ test: () => assert.match(code, /^print\s*\(\s*new_password\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
        
        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break

    return password
    
--fcc-editable-region--
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
```
