---
id: 656ef89b4b486b7a16077864
title: 步驟 26
challengeType: 20
dashedName: step-26
---

# --description--

元組是 Python 中另一種內建的資料結構。元組非常類似列表，但它們是用小括弧 `()` 定義，而不是方括號。此外，元組是不可變的，與列表不同。

```py
my_tuple = ('larch', 1, True)
```

你的 `constraints` 列表將會儲存元組。每個元組的第一個項目將會是條件參數。

修改 `constraints` 列表指派，為你的列表添加一個元組。使用 `nums` 作為第一個項目，空字串作為第二個項目。

# --hints--

你應該將一個 tuple 添加到 `constraints` 列表中。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

你應該將 `(nums, '')` 添加到 `constraints` 列表。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
