---
id: 656ef5bc5c0cd464be1df675
title: 步驟-23
challengeType: 20
dashedName: step-23
---

# --description--

接下來，你將爲你的函數提供更多參數，這些參數將作爲生成密碼的約束。

通過在現有的 `length` 參數後按順序添加 `nums`、`special_chars`、`uppercase` 和 `lowercase` 來修改函數聲明。

# --hints--

你的函數應該以 `length`、`nums`、`special_chars`、`uppercase` 和 `lowercase` 作爲參數。 The order matters.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)

    return password

# new_password = generate_password(8)
# print(new_password)
```
