---
id: 656ef54070c72161e6feb90f
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

看起來一切正常，但如果有一種方法可以檢查產生的密碼是否符合特定特性會更好。例如，特殊字元、數字或大寫／小寫字母的最小數量。你很快就會處理這個問題。

目前，將你程式碼的最後兩行註解起來。

# --hints--

你應該將你的程式碼最後兩行變成註解。

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
