---
id: 656ef5bc5c0cd464be1df675
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

接下來，你將為你的函式添加更多參數，這些參數將作為產生密碼的條件。

修改你的函式（程式）宣告式，在現有的 `length` 參數之後，依序添加 `nums`、`special_chars`、`uppercase` 和 `lowercase`。

# --hints--

你的函式（程式）應該接受 `length`、`nums`、`special_chars`、`uppercase` 和 `lowercase` 作為參數。順序很重要。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password

# new_password = generate_password(8)
# print(new_password)
```
