---
id: 656f0bc259c1f6b57486ed68
title: 步驟 33
challengeType: 20
dashedName: step-33
---

# --description--

要檢查生成的密碼是否滿足所需的特性，你將使用 `re` 模塊中的 `findall()` 函數。 它類似於 `search`，但它返回一個包含所有匹配模式出現的列表。

將 `search()` 調用替換爲 `findall()` 並檢查輸出。

# --hints--

你應該修改現有的 `print()` 調用，將 `search()` 替換爲 `findall()`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.search(pattern, quote))
--fcc-editable-region--
```
