---
id: 656f04b20397d9a574cc7eb2
title: 步驟 32
challengeType: 20
dashedName: step-32
---

# --description--

你可以在不使用 `compile()` 函式（程式）的情況下取得相同結果。將你的 `pattern` 變數修改為字面字串 `'l+'`。然後，將 `print()` 呼叫改為列印 `re.search(pattern, quote)`。

# --hints--

你應該將你的 `pattern` 變數修改為字面字串 `'l+'`。

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "l+") })
```

你應該列印 `re.search(pattern, quote)`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.search\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l+')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
