---
id: 656f0d6bc60b58b9777a36c9
title: 步驟 34
challengeType: 20
dashedName: step-34
---

# --description--

字元類別由方括號顯示，並且會匹配括號中指定的其中一個字元。例如，`ma[dnt]` 會匹配 `mad`、`man` 或 `mat`。

修改你的樣式以符合一個 `w`，後面接著 `h` 或 `a`。

# --hints--

你的樣式應該使用字元類別來匹配一個 `w`，後面跟著 `h` 或 `a`。

```js
({ test: () => assert(runPython(`
    pattern == "w[ha]" or pattern == "w[ah]"    
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
