---
id: 656f10458134d4c4e283a2f1
title: 步驟 35
challengeType: 20
dashedName: step-35
---

# --description--

現在，將條件元組中的空字串轉換成用於匹配單一數字的正規表達式樣式。使用指定從 `0` 到 `9` 每個字元的字元類別。

# --hints--

你應該在你的 `constraints` 列表中有 `(nums, '[0123456789]')` 元組。

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\s*\(\s*nums\s*,\s*("|')\[\d{10}\]\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, '')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
