---
id: 657732654845d8e2fb1217e6
title: 步驟 62
challengeType: 20
dashedName: step-62
---

# --description--

現在，`all()` 正在將一個空列表作為引數。使用推導式語法填充該空列表，使該列表儲存對 `constraints` 列表中每個 `constraint`-`pattern` 元組評估表達式 `constraint <= len(re.findall(pattern, password))` 的結果。

這樣，只有在所有需求都達成後，你才會跳出 `while` 迴圈。

# --hints--

你應該將 `[constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints]` 傳遞給 `all()` 函式（程式）。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("all([constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints])")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if all([]):
--fcc-editable-region--        
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
