---
id: 6577319039f4f7de9251b822
title: 步驟 61
challengeType: 20
dashedName: step-61
---

# --description--

除了使用循環和計數器變量之外，你還可以使用不同的方法實現相同的結果，你將在接下來的幾個步驟中實現該方法。

`all()` 是一個內置 Python 函數，如果給定可迭代對象內的所有元素的計算結果都爲 `True`，則返回 `True`。 否則，返回 `False`。

用單個 `if` 語句替換現有的 `for` 循環和兩個 `if` 語句。 對於 `if` 條件，使用對 `all()` 函數的調用並將空列表作爲函數調用的參數傳遞。

# --hints--

你應該用單個 `if` 語句替換現有的 `for` 循環和兩個 `if` 語句。

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

你的新 `if` 條件應該是 `all([])`。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

你的新 `if` 主體內應該有 `break`。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1

        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
