---
id: 6577333feab1e8e927014f03
title: 步驟 63
challengeType: 20
dashedName: step-63
---

# --description--

擁有 `all([expression for i in iterable])`，表示透過對 `iterable` 中的每個 `i` 評估 `expression`，創建了一個新的列表。當 `all()` 函式（程式）迭代剛創建的列表後，該列表會自動被刪除，因為它不再需要。

可以透過使用產生器表達式來存檔記憶體。產生器表達式遵循列表推導式的語法，但它們使用小括弧而非方括弧。

將你的列表推導式改成產生器表達式，方法是移除方括號。

# --hints--

你應該透過移除方括號，將你的列表推導式轉換成產生器表達式。

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        if all(
            [
                constraint <= len(re.findall(pattern, password))
                for constraint, pattern in constraints
            ]
        ):
--fcc-editable-region--
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
