---
id: 6577319039f4f7de9251b822
title: 步驟 61
challengeType: 20
dashedName: step-61
---

# --description--

你可以用不同的方法達成相同的結果，而不是使用迴圈和計數變數，這個方法你將在接下來的幾個步驟中實作。

`all()` 是 Python 內建的函式（程式），如果給定的可疊代物件中所有元素都計算為 `True`，則傳回 `True`。否則，傳回 `False`。

將你現有的 `for` 迴圈和兩個 `if` 述語取代為單一的 `if` 述語。對於 `if` 條件，使用對 `all()` 函式的呼叫，並將空列表作為引數傳遞給該函式呼叫。

# --hints--

你應該將你現有的 `for` 迴圈和兩個 `if` 述語替換為單一的 `if` 述語。 

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

你的新 `if` 條件應該是 `all([])`。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

你應該在新的 `if` 主體內有 `break`。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
