---
id: 657dadf4d8b93c1704f3a57c
title: 步驟 66
challengeType: 20
dashedName: step-66
---

# --description--

它能運作，但你仍然可以改進幾件事。首先，使用 5 個引數呼叫函式可能會造成混淆，不知道哪個值會被指定到哪個參數。

你可以使用關鍵字引數來呼叫函式，也就是明確地寫出參數名稱，接著使用指定運算子和該值。例如：

```py
def add(x, y):
    return x + y

add(x=1, y=7) # 8
```

修改你的函式呼叫以使用關鍵字引數。

# --hints--

你的 `new_password` 變數應該有 `generate_password(length=8, nums=1, special_chars=1, uppercase=1, lowercase=1)` 的值。

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*,\s*nums\s*=\s*1\s*,\s*special_chars\s*=\s*1\s*,\s*uppercase\s*=\s*1\s*,\s*lowercase\s*=\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
   
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region-- 
new_password = generate_password(8, 1, 1, 1, 1)
print(new_password)
--fcc-editable-region--
```
