---
id: afcc8d540bea9ea2669306b6
title: створити функцію повторення рядка
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

У цій лабораторній роботі ви створите функцію, яка повторює заданий рядок певну кількість разів. Для цієї лабораторної роботи _не_ використовуйте вбудований метод `.repeat()`.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні створити функцію з іменем `repeatStringNumTimes`, яка приймає два параметри: рядок і число.
2. Функція має повертати рядок, повторений вказану кількість разів.
3. Якщо число менше або дорівнює нулю, функція має повертати порожній рядок.

# --hints--

Ви повинні створити функцію з іменем `repeatStringNumTimes`.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` має приймати два параметри.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

Функція `repeatStringNumTimes` завжди має повертати рядок.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` має повертати рядок `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` має повертати рядок `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` має повертати рядок `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` має повертати рядок `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` має повертати рядок `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` має повертати порожній рядок (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` має повертати `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

Вбудований метод `repeat()` не повинен використовуватися.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
