---
id: 6894c1b715b8d18a18ef3203
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся з класами та об’єктами Python, створюючи інвентар музичних інструментів.

Щоб визначити клас у Python, використайте ключове слово `class`, за яким іде назва класу та двокрапка.

```python
class ClassName:
    # class contents go here
```

Пам’ятайте, що в Python назви класів зазвичай пишуться у PascalCase, тобто кожне слово починається з великої літери без пробілів між словами.

Почніть зі створення класу з назвою `MusicalInstrument`. Наразі додайте `pass` як тіло класу.

# --hints--

Ви маєте створити клас з назвою `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
