---
id: 6895e4940e39950288e406c2
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Далі додайте до свого класу ще один метод з назвою `get_fact`, який повертатиме факт про інструмент.

Цей метод має мати лише параметр `self` і повертати f-рядок, який каже `The [instrument name] is part of the [instrument type] family of instruments.`.

На відміну від методу `play`, який виводить повідомлення безпосередньо, цей метод має повертати рядок, який можна використовувати в інших частинах вашого коду.

# --hints--

Ви маєте визначити метод з назвою `get_fact` з параметром `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

Ваш метод `get_fact` має повертати f-рядок, що містить назву та тип інструмента.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
