---
id: 6895e356fd877efd6e3657ee
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--


У класі `MusicalInstrument` визначте ще один метод з назвою `play`, який приймає лише параметр `self`. Параметр `self` дозволяє методам отримувати доступ до атрибутів конкретного екземпляра.

# --hints--

Ви повинні визначити метод з назвою `play`, який приймає параметр `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
