---
id: 6895e0bdb0348bf46952ab67
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер настав час створити екземпляри вашого класу `MusicalInstrument`.

Щоб створити екземпляр класу, ви викликаєте ім’я класу як функцію, передаючи необхідні аргументи для методу `__init__` (окрім `self`, який автоматично обробляє Python).

```python
my_instance = ClassName(arg1, arg2)
```

Коли ви створюєте цей екземпляр, Python викликає ваш метод `__init__` і зберігає значення як атрибути екземпляра, які належать лише цьому конкретному екземпляру.

Поза межами визначення класу створіть екземпляр з ім’ям `instrument_1` для `Oboe`, який є `woodwind` інструментом.

# --hints--

Ви повинні створити екземпляр класу `MusicalInstrument` з ім’ям `instrument_1`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

При створенні `Oboe` ви повинні передати `woodwind` і `instrument_1` як аргументи.

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
