---
id: 6895e3f91c069e0068a42582
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Усередині вашого методу `play` використайте `print()`, щоб показати повідомлення про гру на інструменті.

Використайте f-рядок, щоб включити назву інструмента у повідомлення: `The [instrument name] is fun to play!`. `[instrument name]` має бути замінений на фактичну назву інструмента.

# --hints--

Ваш метод `play` має використовувати f-рядок для відображення повідомлення `The [instrument name] is fun to play!`. Використайте `self.name`, щоб отримати доступ до назви інструмента.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
