---
id: 6895e5337b1914050b0af7bd
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Тепер викличте метод `get_fact()` у екземпляра `instrument_1` і виведіть результат.

Оскільки `get_fact()` повертає рядок замість того, щоб виводити його безпосередньо, потрібно використати `print()` для відображення повернутого рядка.

# --hints--

Ви повинні викликати `get_fact()` у `instrument_1` і вивести результат.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()

--fcc-editable-region--
```
