---
id: 6895e235935106f88281b648
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Ви можете отримати доступ до атрибутів своїх екземплярів, використовуючи точкову нотацію:

```py
instance_name.attribute_name
```

Це дозволяє отримати значення, які були збережені під час ініціалізації.

Спробуйте вивести атрибути `name` і `instrument_type` обох ваших екземплярів. Ви побачите значення, які призначили під час їх створення.

# --hints--

Ви повинні вивести атрибут `name` об’єкта `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

Ви повинні вивести атрибут `instrument_type` об’єкта `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

Ви повинні вивести атрибут `name` об’єкта `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

Ви повинні вивести атрибут `instrument_type` об’єкта `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
