---
id: 6895dffc64e97df0e62e9ef9
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Методи — це функції, визначені всередині класу, які можуть виконувати дії, використовуючи дані класу. Ось як створюють методи:

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

У Python метод `__init__` — це спеціальний метод, відомий як ініціаль, який викликається під час створення нового екземпляра класу.

Всередині вашого класу `MusicalInstrument` додайте метод `__init__` з трьома параметрами: `self`, `name` і `instrument_type`.

Параметр `self` — це стандартний спосіб посилатися на екземпляр класу і він обов’язковий як перший параметр для всіх методів екземпляра в класі.

# --hints--

Ви повинні визначити метод `__init__` всередині вашого класу `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

Ваш метод `__init__` має мати три параметри: `self`, `name` і `instrument_type`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
