---
id: 6895e0452b3370f26f365a46
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

У методі `__init__` потрібно призначити параметри атрибутам екземпляра. Атрибути екземпляра — це змінні, які належать конкретному екземпляру класу.

Усередині вашого методу `__init__` призначте параметр `name` атрибуту `self.name`, а параметр `instrument_type` — атрибуту `self.instrument_type`.

# --hints--

Ви повинні призначити параметр `name` атрибуту `self.name` всередині методу `__init__`.

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

Ви повинні призначити параметр `instrument_type` атрибуту `self.instrument_type` всередині методу `__init__`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        pass
--fcc-editable-region--
```
