---
id: 6895e197180c9af6988e1e19
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Створіть другий екземпляр з іменем `instrument_2` для `Trumpet`, який є `brass` інструментом.

Зверніть увагу, як можна створювати кілька екземплярів з одного шаблону класу. Кожен екземпляр зберігає свої власні значення атрибутів незалежно від інших екземплярів.

# --hints--

Ви повинні створити екземпляр з іменем `instrument_2` класу `MusicalInstrument`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

Під час створення `Trumpet` передайте `brass` та `instrument_2` як аргументи.

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')

--fcc-editable-region--
```
