---
id: a105e963526e7de52b219be9
title: 고유한 정렬된 합집합 구현하기
challengeType: 26
dashedName: implement-a-unique-sorted-union
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `uniteUnique`라는 이름의 함수를 만들어야 합니다.
1. `uniteUnique` 함수는 두 개 이상의 배열을 인수로 받아야 합니다.
1. 함수는 인수 배열에서 처음 발견된 순서대로 고유한 값을 포함하는 새 배열을 반환해야 합니다. 예를 들어, `[1, 2, 4], [2, 3, 5]` 같은 입력은 `[1, 2, 4, 3, 5]` 같은 출력을 가져야 합니다.

# --hints--

`uniteUnique` 함수를 만들어야 합니다.

```js
assert.isFunction(uniteUnique);
```

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])`는 `[1, 3, 2, 5, 4]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])`는 `[1, 2, 3, 5]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])`는 `[1, 2, 3, 5, 4, 6, 7, 8]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])`는 `[1, 3, 2, 5, 4, 6]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])`는 `[1, 3, 2, 5, 4]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const uniteUnique = (...arrays) => {
  return arrays.reduce((a, b) => {
    return [
      ...a,
      ...b.filter((e, currentIndex) => b.indexOf(e) === currentIndex && a.indexOf(e) === -1)
    ];
  }, []);
};
```
