---
id: 587d78b2367417b2b2512b10
title: Rimuovere elementi usando splice()
challengeType: 1
forumTopicId: 301166
dashedName: remove-items-using-splice
---

# --description--

Ok, abbiamo imparato come rimuovere elementi dall'inizio e dalla fine degli array usando `shift()` e `pop()`, ma cosa succede se vogliamo rimuovere un elemento da qualche parte nel mezzo? O rimuovere più di un elemento alla volta? Ecco dove entra in gioco `splice()`. `splice()` ci permette proprio questo: **rimuovere qualsiasi numero di elementi consecutivi** da qualsiasi punto di un array.

`splice()` può prendere fino a 3 parametri, ma per ora ci concentreremo solo sui primi 2. I primi due parametri di `splice()` sono numeri interi che rappresentano gli indici, o posizioni, degli elementi nell'array su cui `splice()` viene chiamato. E ricorda, gli array sono *indicizzati da zero*, quindi per indicare il primo elemento di un array useremmo `0`. Il primo parametro di `splice()` rappresenta l'indice nell'array da cui iniziare a rimuovere gli elementi, mentre il secondo parametro indica il numero di elementi da eliminare. Per esempio:

```js
let array = ['today', 'was', 'not', 'so', 'great'];

array.splice(2, 2);
```

Qui rimuoviamo 2 elementi, a partire dal terzo elemento (all'indice 2). `array` avrebbe il valore `['today', 'was', 'great']`.

`splice()` non solo modifica l'array su cui viene chiamato, ma restituisce anche un nuovo array contenente il valore degli elementi rimossi:

```js
let array = ['I', 'am', 'feeling', 'really', 'happy'];

let newArray = array.splice(3, 2);
```

`newArray` ha il valore `['really', 'happy']`.

# --instructions--

Abbiamo inizializzato un array `arr`. Usa `splice()` per rimuovere elementi da `arr`, in modo che contenga solo elementi la cui somma sia pari al valore di `10`.

# --hints--

Non dovresti modificare la riga originale di `const arr = [2, 4, 5, 1, 7, 5, 2, 1];`.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/constarr=\[2,4,5,1,7,5,2,1\];?/)
);
```

`arr` dovrebbe contenere solo elementi la cui somma è `10`.

```js
assert.strictEqual(
  arr.reduce((a, b) => a + b),
  10
);
```

Il tuo codice dovrebbe utilizzare il metodo `splice()` su `arr`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(/));
```

Lo splice dovrebbe rimuovere solo elementi da `arr` e non aggiungere altri elementi a `arr`.

```js
assert(
  !__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(\d+,\d+,\d+.*\)/g)
);
```

# --seed--

## --seed-contents--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
// Only change code below this line

// Only change code above this line
console.log(arr);
```

# --solutions--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
arr.splice(1, 4);
```
