---
id: 587d7b7d367417b2b2512b1d
title: "Iterare sulle chiavi di un oggetto con un'istruzione for...in"
challengeType: 1
forumTopicId: 301162
dashedName: iterate-through-the-keys-of-an-object-with-a-for---in-statement
---

# --description--

A volte è necessario iterare su tutte le chiavi di un oggetto. Puoi usare un loop <dfn>for...in</dfn> per farlo. Il loop for...in ha questa forma:

```javascript
const refrigerator = {
  'milk': 1,
  'eggs': 12,
};

for (const food in refrigerator) {
  console.log(food, refrigerator[food]);
}
```

Questo codice stampa `milk 1` e `eggs 12`, con ogni coppia chiave-valore su una riga separata.

Abbiamo definito la variabile `food` nell'intestazione del loop e questa variabile è stata impostata su ciascuna chiave dell'oggetto a ogni iterazione, facendo sì che il nome di ogni cibo venga stampato sulla console.

**NOTA:** Gli oggetti non mantengono un ordine delle chiavi archiviate come fanno gli array; quindi la posizione di una chiave in un oggetto, o l'ordine relativo in cui appare, non ha importanza quando si fa riferimento o si consulta quella chiave.

# --instructions--

Abbiamo definito una funzione `countOnline` che accetta un argomento, `allUsers`. Usa un'istruzione <dfn>for...in</dfn> all'interno di questa funzione per iterare sull'oggetto `allUsers` e restituisci il numero di utenti la cui proprietà `online` è impostata su `true`. Un esempio di oggetto che potrebbe essere passato a `countOnline` è mostrato qui sotto. Ogni utente avrà una proprietà `online` impostata su `true` o `false`.

```js
{
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
```

# --hints--

La funzione `countOnline` dovrebbe usare un'istruzione `for in` per iterare sulle chiavi dell'oggetto che le viene passato.

```js
assert.match(
  __helpers.removeJSComments(code),
    /for\s*\(\s*(var|let|const)\s+[a-zA-Z_$]\w*\s+in\s+[a-zA-Z_$]\w*\s*\)/
);
```

La funzione `countOnline` dovrebbe restituire `1` quando le viene passato l'oggetto `{ Alan: { online: false }, Jeff: { online: true }, Sarah: { online: false } }`

```js
const usersObj1 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj1), 1);
```

La funzione `countOnline` dovrebbe restituire `2` quando le viene passato l'oggetto `{ Alan: { online: true }, Jeff: { online: false }, Sarah: { online: true } }`

```js
const usersObj2 = {
  Alan: {
    online: true
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: true
  }
}
assert.equal(countOnline(usersObj2), 2);
```

La funzione `countOnline` dovrebbe restituire `0` quando le viene passato l'oggetto `{ Alan: { online: false }, Jeff: { online: false }, Sarah: { online: false } }`

```js
const usersObj3 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj3), 0);
```

# --seed--

## --seed-contents--

```js
const users = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}

function countOnline(allUsers) {
  // Only change code below this line

  // Only change code above this line
}

console.log(countOnline(users));
```

# --solutions--

```js
function countOnline(allUsers) {
  let numOnline = 0;
  for(const user in allUsers){
    if(allUsers[user].online) {
      numOnline++;
    }
  }
  return numOnline;
}
```
