---
id: 587d7b7c367417b2b2512b1b
title: usa la parola chiave delete per eliminare proprietà degli oggetti
challengeType: 1
forumTopicId: 301168
dashedName: use-the-delete-keyword-to-remove-object-properties
---

# --description--

Ora sai cosa sono gli oggetti e le loro funzionalità e vantaggi di base. In breve, sono archivi di coppie chiave-valore che offrono un modo flessibile e intuitivo per strutturare i dati, ***e*** forniscono tempi di ricerca molto rapidi. Nel resto di queste sfide, descriveremo diverse operazioni comuni che puoi eseguire sugli oggetti per farti prendere confidenza con l'applicazione di queste utili strutture dati nei tuoi programmi.

Nelle sfide precedenti, abbiamo sia aggiunto che modificato le coppie chiave-valore di un oggetto. Qui vedremo come possiamo *eliminare* una coppia chiave-valore da un oggetto.

Rivediamo il nostro esempio di oggetto `foods` un'ultima volta. Se volessimo eliminare la chiave `apples`, possiamo farlo usando la parola chiave `delete` in questo modo:

```js
delete foods.apples;
```

# --instructions--

Usa la parola chiave delete per eliminare le chiavi `oranges`, `plums` e `strawberries` dall'oggetto `foods`.

# --hints--

L'oggetto `foods` dovrebbe avere solo tre chiavi: `apples`, `grapes` e `bananas`.

```js
assert(
  !foods.hasOwnProperty('oranges') &&
    !foods.hasOwnProperty('plums') &&
    !foods.hasOwnProperty('strawberries') &&
    Object.keys(foods).length === 3
);
```

Le chiavi `oranges`, `plums` e `strawberries` dovrebbero essere eliminate usando `delete`.

```js
assert(
  __helpers.removeJSComments(code).search(/oranges:/) !== -1 &&
    __helpers.removeJSComments(code).search(/plums:/) !== -1 &&
    __helpers.removeJSComments(code).search(/strawberries:/) !== -1
);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

delete foods.oranges;
delete foods.plums;
delete foods.strawberries;

console.log(foods);
```
