---
id: 587d7b7c367417b2b2512b18
title: aggiungere coppie chiave-valore agli oggetti JavaScript
challengeType: 1
forumTopicId: 301153
dashedName: add-key-value-pairs-to-javascript-objects
---

# --description--

Alla base, gli oggetti sono semplicemente raccolte di coppie <dfn>chiave-valore</dfn>. In altre parole, sono dati (<dfn>valori</dfn>) associati a identificatori unici chiamati <dfn>proprietà</dfn> (<dfn>chiavi</dfn>). Dai un’occhiata a un esempio:

```js
const tekkenCharacter = {
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true
};
```

Il codice sopra definisce un oggetto personaggio del videogioco Tekken chiamato `tekkenCharacter`. Ha tre proprietà, ognuna delle quali è associata a un valore specifico. Se vuoi aggiungere una proprietà aggiuntiva, come "origin", puoi farlo assegnando `origin` all’oggetto:

```js
tekkenCharacter.origin = 'South Korea';
```

Questo usa la notazione a punto. Se osservi l’oggetto `tekkenCharacter`, ora includerà la proprietà `origin`. Hwoarang aveva anche capelli arancioni distintivi. Puoi aggiungere questa proprietà con la notazione a parentesi facendo:

```js
tekkenCharacter['hair color'] = 'dyed orange';
```

La notazione a parentesi è necessaria se la tua proprietà contiene uno spazio o se vuoi usare una variabile per nominare la proprietà. Nel caso sopra, la proprietà è racchiusa tra virgolette per indicarla come stringa e sarà aggiunta esattamente come mostrato. Senza virgolette, sarà valutata come una variabile e il nome della proprietà sarà il valore di quella variabile. Ecco un esempio con una variabile:

```js
const eyes = 'eye color';

tekkenCharacter[eyes] = 'brown';
```

Dopo aver aggiunto tutti gli esempi, l’oggetto sarà così:

```js
{
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true,
  origin: 'South Korea',
  'hair color': 'dyed orange',
  'eye color': 'brown'
};
```

# --instructions--

È stato creato un oggetto `foods` con tre elementi. Usando la sintassi che preferisci, aggiungi altri tre elementi: `bananas` con valore `13`, `grapes` con valore `35` e `strawberries` con valore `27`.

# --hints--

`foods` dovrebbe essere un oggetto.

```js
assert(typeof foods === 'object');
```

L’oggetto `foods` dovrebbe avere una chiave `bananas` con valore `13`.

```js
assert(foods.bananas === 13);
```

L’oggetto `foods` dovrebbe avere una chiave `grapes` con valore `35`.

```js
assert(foods.grapes === 35);
```

L’oggetto `foods` dovrebbe avere una chiave `strawberries` con valore `27`.

```js
assert(foods.strawberries === 27);
```

La definizione dell’oggetto `foods` non dovrebbe essere modificata.

```js
assert(
  __helpers.removeJSComments(code).search(/let foods/) === -1 &&
  __helpers.removeJSComments(code).search(/const\s+foods\s*=\s*{\s*apples:\s*25,\s*oranges:\s*32,\s*plums:\s*28\s*};/
) !== -1
);
```

# --seed--

## --seed-contents--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

foods['bananas'] = 13;
foods['grapes']  = 35;
foods['strawberries'] = 27;
```
