---
id: 587d7b7e367417b2b2512b20
title: Usa un array per archiviare una raccolta di dati
challengeType: 1
forumTopicId: 301167
dashedName: use-an-array-to-store-a-collection-of-data
---

# --description--

Quanto segue è un esempio della più semplice implementazione di una struttura dati array. Questo è noto come <dfn>array monodimensionale</dfn>, cioè che ha un solo livello, o che non contiene altri array annidati al suo interno. Nota che contiene <dfn>booleani</dfn>, <dfn>stringhe</dfn> e <dfn>numeri</dfn>, tra gli altri tipi di dati validi in JavaScript:

```js
let simpleArray = ['one', 2, 'three', true, false, undefined, null];
console.log(simpleArray.length);
```

La chiamata `console.log` mostra `7`.

Tutti gli array hanno una proprietà lunghezza, che come mostrato sopra, può essere consultata molto facilmente con la sintassi `Array.length`. Di seguito puoi vedere un’implementazione più complessa di un array. Questo è noto come <dfn>array multidimensionale</dfn>, cioè un array che contiene altri array. Nota che questo array contiene anche <dfn>oggetti</dfn> JavaScript, che esamineremo molto da vicino nella nostra prossima sezione, ma per ora, tutto ciò che devi sapere è che gli array possono anche archiviare oggetti complessi.

```js
let complexArray = [
  [
    {
      one: 1,
      two: 2
    },
    {
      three: 3,
      four: 4
    }
  ],
  [
    {
      a: "a",
      b: "b"
    },
    {
      c: "c",
      d: "d"
    }
  ]
];
```

# --instructions--

Abbiamo definito una variabile chiamata `yourArray`. Completa l’istruzione assegnando alla variabile `yourArray` un array di almeno 5 elementi. Il tuo array deve contenere almeno una <dfn>stringa</dfn>, un <dfn>numero</dfn> e un <dfn>booleano</dfn>.

# --hints--

`yourArray` dovrebbe essere un array.

```js
assert.strictEqual(Array.isArray(yourArray), true);
```

`yourArray` dovrebbe avere almeno 5 elementi.

```js
assert.isAtLeast(yourArray.length, 5);
```

`yourArray` dovrebbe contenere almeno un `boolean`.

```js
assert(yourArray.filter((el) => typeof el === 'boolean').length >= 1);
```

`yourArray` dovrebbe contenere almeno un `number`.

```js
assert(yourArray.filter((el) => typeof el === 'number').length >= 1);
```

`yourArray` dovrebbe contenere almeno un `string`.

```js
assert(yourArray.filter((el) => typeof el === 'string').length >= 1);
```

# --seed--

## --seed-contents--

```js
let yourArray; // Change this line
```

# --solutions--

```js
let yourArray = ['a string', 100, true, ['one', 2], 'another string'];
```
