---
id: 587d7b7b367417b2b2512b15
title: Iterare su tutti gli elementi di un array usando i loop for
challengeType: 1
forumTopicId: 301161
dashedName: iterate-through-all-an-arrays-items-using-for-loops
---

# --description--

A volte, quando si lavora con gli array, è molto utile poter iterare su ogni elemento per trovare uno o più elementi di cui potremmo aver bisogno, o per manipolare un array in base a quali elementi dati soddisfano un certo insieme di criteri. JavaScript offre diversi metodi integrati che iterano su array in modi leggermente diversi per ottenere risultati diversi (come `every()`, `forEach()`, `map()`, ecc.), tuttavia la tecnica più flessibile e che ci offre il massimo controllo è un semplice loop `for`.

Considera quanto segue:

```js
function greaterThanTen(arr) {
  let newArr = [];
  for (let i = 0; i < arr.length; i++) {
    if (arr[i] > 10) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}

greaterThanTen([2, 12, 8, 14, 80, 0, 1]);
```

Usando un loop `for`, questa funzione itera e consulta ogni elemento dell'array, sottoponendolo a un semplice test che abbiamo creato. In questo modo, abbiamo determinato facilmente e in modo programmato quali elementi dati sono maggiori di `10`, e restituito un nuovo array, `[12, 14, 80]`, contenente quegli elementi.

# --instructions--

Abbiamo definito una funzione, `filteredArray`, che prende `arr`, un array annidato, e `elem` come argomenti, e restituisce un nuovo array. `elem` rappresenta un elemento che potrebbe essere presente o meno in uno o più degli array annidati all'interno di `arr`. Modifica la funzione, usando un loop `for`, per restituire una versione filtrata dell'array passato in modo che qualsiasi array annidato in `arr` contenente `elem` sia stato rimosso.

# --hints--

`filteredArray([[10, 8, 3], [14, 6, 23], [3, 18, 6]], 18)` dovrebbe restituire `[[10, 8, 3], [14, 6, 23]]`

```js
assert.deepEqual(
  filteredArray(
    [
      [10, 8, 3],
      [14, 6, 23],
      [3, 18, 6]
    ],
    18
  ),
  [
    [10, 8, 3],
    [14, 6, 23]
  ]
);
```

`filteredArray([["trumpets", 2], ["flutes", 4], ["saxophones", 2]], 2)` dovrebbe restituire `[["flutes", 4]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['trumpets', 2],
      ['flutes', 4],
      ['saxophones', 2]
    ],
    2
  ),
  [['flutes', 4]]
);
```

`filteredArray([["amy", "beth", "sam"], ["dave", "sean", "peter"]], "peter")` dovrebbe restituire `[["amy", "beth", "sam"]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['amy', 'beth', 'sam'],
      ['dave', 'sean', 'peter']
    ],
    'peter'
  ),
  [['amy', 'beth', 'sam']]
);
```

`filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3)` dovrebbe restituire `[]`

```js
assert.deepEqual(
  filteredArray(
    [
      [3, 2, 3],
      [1, 6, 3],
      [3, 13, 26],
      [19, 3, 9]
    ],
    3
  ),
  []
);
```

La funzione `filteredArray` dovrebbe utilizzare un loop `for`

```js
assert.notStrictEqual(filteredArray.toString().search(/for/), -1);
```

# --seed--

## --seed-contents--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  // Only change code below this line

  // Only change code above this line
  return newArr;
}

console.log(filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3));
```

# --solutions--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  for (let i = 0; i<arr.length; i++) {
    if (arr[i].indexOf(elem) < 0) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}
```
