---
id: 587d7b7b367417b2b2512b16
title: Crea array multidimensionali complessi
challengeType: 1
forumTopicId: 301159
dashedName: create-complex-multi-dimensional-arrays
---

# --description--

Fantastico! Hai appena imparato tantissimo sugli array! Questa è stata una panoramica piuttosto generale, e c’è ancora molto da scoprire sul lavoro con gli array, gran parte di cui vedrai nelle sezioni successive. Ma prima di passare a guardare i <dfn>oggetti</dfn>, diamo un’ultima occhiata per vedere come gli array possano diventare un po’ più complessi rispetto a quelli visti nelle sfide precedenti.

Una delle funzionalità più potenti quando pensi agli array come strutture dati è che gli array possono contenere, o addirittura essere completamente composti da altri array. Abbiamo già visto array che contengono array nelle sfide precedenti, ma abbastanza semplici. Tuttavia, gli array possono contenere una profondità infinita di array che a loro volta contengono altri array, ognuno con livelli arbitrari di profondità, e così via. In questo modo, un array può diventare molto rapidamente una struttura dati molto complessa, nota come array <dfn>multidimensionale</dfn> o annidato. Considera il seguente esempio:

```js
let nestedArray = [
  ['deep'],
  [
    ['deeper'], ['deeper'] 
  ],
  [
    [
      ['deepest'], ['deepest']
    ],
    [
      [
        ['deepest-est?']
      ]
    ]
  ]
];
```

L’array `deep` è annidato a 2 livelli di profondità. Gli array `deeper` sono a 3 livelli di profondità. Gli array `deepest` sono a 4 livelli, e `deepest-est?` a 5.

Anche se questo esempio può sembrare complicato, questo livello di complessità non è raro, né insolito, quando si lavora con grandi quantità di dati. Tuttavia, possiamo comunque consultare molto facilmente i livelli più profondi di un array così complesso con la notazione a parentesi:

```js
console.log(nestedArray[2][1][0][0][0]);
```

Questo stampa la stringa `deepest-est?`. E ora che sappiamo dove si trova quel dato, possiamo modificarlo se necessario:

```js
nestedArray[2][1][0][0][0] = 'deeper still';

console.log(nestedArray[2][1][0][0][0]);
```

Ora stampa `deeper still`.

# --instructions--

Abbiamo definito una variabile, `myNestedArray`, impostata uguale a un array. Modifica `myNestedArray`, usando qualsiasi combinazione di <dfn>stringhe</dfn>, <dfn>numeri</dfn> e <dfn>booleani</dfn> come elementi dati, in modo che abbia esattamente cinque livelli di profondità (ricorda, l’array più esterno è il livello 1). Da qualche parte al terzo livello, includi la stringa `deep`, al quarto livello includi la stringa `deeper` e al quinto livello la stringa `deepest`.

# --hints--

`myNestedArray` dovrebbe contenere solo numeri, booleani e stringhe come elementi dati

```js
assert.strictEqual(
  (function (arr) {
    let flattened = (function flatten(arr) {
      const flat = [].concat(...arr);
      return flat.some(Array.isArray) ? flatten(flat) : flat;
    })(arr);
    for (let i = 0; i < flattened.length; i++) {
      if (
        typeof flattened[i] !== 'number' &&
        typeof flattened[i] !== 'string' &&
        typeof flattened[i] !== 'boolean'
      ) {
        return false;
      }
    }
    return true;
  })(myNestedArray),
  true
);
```

`myNestedArray` dovrebbe avere esattamente 5 livelli di profondità

```js
assert.strictEqual(
  (function (arr) {
    let depth = 0;
    function arrayDepth(array, i, d) {
      if (Array.isArray(array[i])) {
        arrayDepth(array[i], 0, d + 1);
      } else {
        depth = d > depth ? d : depth;
      }
      if (i < array.length) {
        arrayDepth(array, i + 1, d);
      }
    }
    arrayDepth(arr, 0, 0);
    return depth;
  })(myNestedArray),
  4
);
```

`myNestedArray` dovrebbe contenere esattamente una occorrenza della stringa `deep` in un array annidato a 3 livelli di profondità

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deep').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deep')[0] === 2
);
```

`myNestedArray` dovrebbe contenere esattamente una occorrenza della stringa `deeper` in un array annidato a 4 livelli di profondità

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deeper').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deeper')[0] === 3
);
```

`myNestedArray` dovrebbe contenere esattamente una occorrenza della stringa `deepest` in un array annidato a 5 livelli di profondità

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deepest').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deepest')[0] === 4
);
```

# --seed--

## --seed-contents--

```js
let myNestedArray = [
  // Only change code below this line
  ['unshift', false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
  // Only change code above this line
];
```

# --solutions--

```js
let myNestedArray = [
  ['unshift', ['deep', ['deeper', ['deepest']]],false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
];
```
