---
id: 587d7b7b367417b2b2512b14
title: verifica la presenza di un elemento con indexOf()
challengeType: 1
forumTopicId: 301154
dashedName: check-for-the-presence-of-an-element-with-indexof
---

# --description--

Poiché gli array possono essere modificati, o *mutati*, in qualsiasi momento, non c’è garanzia su dove si trovi un dato elemento in un array, o se quell’elemento esista ancora. Per fortuna, JavaScript ci offre un altro metodo integrato, `indexOf()`, che ci permette di verificare rapidamente e facilmente la presenza di un elemento in un array. `indexOf()` prende un elemento come parametro e, quando viene chiamato, restituisce la posizione, o indice, di quell’elemento, oppure `-1` se l’elemento non esiste nell’array.

Per esempio:

```js
let fruits = ['apples', 'pears', 'oranges', 'peaches', 'pears'];

fruits.indexOf('dates');
fruits.indexOf('oranges');
fruits.indexOf('pears');
```

`indexOf('dates')` restituisce `-1`, `indexOf('oranges')` restituisce `2`, e `indexOf('pears')` restituisce `1` (il primo indice in cui ciascun elemento esiste).

# --instructions--

`indexOf()` può essere incredibilmente utile per verificare rapidamente la presenza di un elemento in un array. Abbiamo definito una funzione, `quickCheck`, che prende un array e un elemento come argomenti. Modifica la funzione usando `indexOf()` in modo che restituisca `true` se l’elemento passato esiste nell’array, e `false` se non esiste.

# --hints--

La funzione `quickCheck` dovrebbe restituire un booleano (`true` o `false`), non una stringa (`"true"` o `"false"`)

```js
assert.isBoolean(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

`quickCheck(["squash", "onions", "shallots"], "mushrooms")` dovrebbe restituire `false`

```js
assert.strictEqual(
  quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'),
  false
);
```

`quickCheck(["onions", "squash", "shallots"], "onions")` dovrebbe restituire `true`

```js
assert.strictEqual(
  quickCheck(['onions', 'squash', 'shallots'], 'onions'),
  true
);
```

`quickCheck([3, 5, 9, 125, 45, 2], 125)` dovrebbe restituire `true`

```js
assert.strictEqual(quickCheck([3, 5, 9, 125, 45, 2], 125), true);
```

`quickCheck([true, false, false], undefined)` dovrebbe restituire `false`

```js
assert.strictEqual(quickCheck([true, false, false], undefined), false);
```

La funzione `quickCheck` dovrebbe utilizzare il metodo `indexOf()`

```js
assert.notStrictEqual(quickCheck.toString().search(/\.indexOf\(/), -1);
```

# --seed--

## --seed-contents--

```js
function quickCheck(arr, elem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

# --solutions--

```js
function quickCheck(arr, elem) {
  return arr.indexOf(elem) >= 0; 
}
```
