---
id: 587d78b2367417b2b2512b0e
title: Aggiungere elementi a un array con push() e unshift()
challengeType: 1
forumTopicId: 301151
dashedName: add-items-to-an-array-with-push-and-unshift
---

# --description--

La lunghezza di un array, come i tipi di dati che può contenere, non è fissa. Gli array possono essere definiti con una lunghezza di qualsiasi numero di elementi, e gli elementi possono essere aggiunti o rimossi nel tempo; in altre parole, gli array sono <dfn>mutabili</dfn>. In questa sfida, vedremo due metodi con cui possiamo modificare programmaticamente un array: `Array.push()` e `Array.unshift()`.

Entrambi i metodi prendono uno o più elementi come parametri e aggiungono quegli elementi all'array su cui viene chiamato il metodo; il metodo `push()` aggiunge elementi alla fine di un array, mentre `unshift()` aggiunge elementi all'inizio. Considera il seguente esempio:

```js
let twentyThree = 'XXIII';
let romanNumerals = ['XXI', 'XXII'];

romanNumerals.unshift('XIX', 'XX');
```

`romanNumerals` avrebbe il valore `['XIX', 'XX', 'XXI', 'XXII']`.

```js
romanNumerals.push(twentyThree);
```

`romanNumerals` avrebbe il valore `['XIX', 'XX', 'XXI', 'XXII', 'XXIII']`. Nota che possiamo anche passare variabili, il che ci permette una flessibilità ancora maggiore nel modificare dinamicamente i dati del nostro array.

# --instructions--

Abbiamo definito una funzione, `mixedNumbers`, a cui stiamo passando un array come argomento. Modifica la funzione usando `push()` e `unshift()` per aggiungere `'I', 2, 'three'` all'inizio dell'array e `7, 'VIII', 9` alla fine, in modo che l'array restituito contenga rappresentazioni dei numeri da 1 a 9 in ordine.

# --hints--

`mixedNumbers(["IV", 5, "six"])` dovrebbe ora restituire `["I", 2, "three", "IV", 5, "six", 7, "VIII", 9]`

```js
assert.deepEqual(mixedNumbers(['IV', 5, 'six']), [
  'I',
  2,
  'three',
  'IV',
  5,
  'six',
  7,
  'VIII',
  9
]);
```

La funzione `mixedNumbers` dovrebbe utilizzare il metodo `push()`

```js
assert(mixedNumbers.toString().match(/\.push/));
```

La funzione `mixedNumbers` dovrebbe utilizzare il metodo `unshift()`

```js
assert(mixedNumbers.toString().match(/\.unshift/));
```

# --seed--

## --seed-contents--

```js
function mixedNumbers(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(mixedNumbers(['IV', 5, 'six']));
```

# --solutions--

```js
function mixedNumbers(arr) {
  arr.push(7,'VIII',9);
  arr.unshift('I',2,'three');
  return arr;
}
```
