---
id: 587d7b7a367417b2b2512b12
title: Copiare elementi di un array usando slice()
challengeType: 1
forumTopicId: 301158
dashedName: copy-array-items-using-slice
---

# --description--

Il prossimo metodo che vedremo è `slice()`. Invece di modificare un array, `slice()` copia o *estrae* un numero dato di elementi in un nuovo array, lasciando intatto l'array su cui viene chiamato. `slice()` prende solo 2 parametri — il primo è l'indice da cui iniziare l'estrazione, e il secondo è l'indice a cui fermarsi (l'estrazione avverrà fino a, ma senza includere, l'elemento a questo indice). Considera questo:

```js
let weatherConditions = ['rain', 'snow', 'sleet', 'hail', 'clear'];

let todaysWeather = weatherConditions.slice(1, 3);
```

`todaysWeather` avrebbe il valore `['snow', 'sleet']`, mentre `weatherConditions` avrebbe ancora `['rain', 'snow', 'sleet', 'hail', 'clear']`.

In pratica, abbiamo creato un nuovo array estraendo elementi da un array esistente.

# --instructions--

Abbiamo definito una funzione, `forecast`, che prende un array come argomento. Modifica la funzione usando `slice()` per estrarre informazioni dall'array passato come argomento e restituisci un nuovo array che contenga gli elementi stringa `warm` e `sunny`.

# --hints--

`forecast` dovrebbe restituire `["warm", "sunny"]`

```js
assert.deepEqual(
  forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']),
  ['warm', 'sunny']
);
```

La funzione `forecast` dovrebbe utilizzare il metodo `slice()`

```js
assert(/\.slice\(/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function forecast(arr) {
  // Only change code below this line

  return arr;
}

// Only change code above this line
console.log(forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']));
```

# --solutions--

```js
function forecast(arr) {
  return arr.slice(2,4);
}
```
