---
id: 587d78b2367417b2b2512b0f
title: Rimuovere elementi da un array con pop() e shift()
challengeType: 1
forumTopicId: 301165
dashedName: remove-items-from-an-array-with-pop-and-shift
---

# --description--

Sia `push()` che `unshift()` hanno metodi corrispondenti che sono quasi opposti funzionalmente: `pop()` e `shift()`. Come avrai già intuito, invece di aggiungere, `pop()` *rimuove* un elemento dalla fine di un array, mentre `shift()` rimuove un elemento dall'inizio. La differenza chiave tra `pop()` e `shift()` e i loro "cugini" `push()` e `unshift()` è che nessuno di questi metodi prende parametri, e ciascuno permette di modificare un array rimuovendo un solo elemento alla volta.

Diamo un'occhiata:

```js
let greetings = ['whats up?', 'hello', 'see ya!'];

greetings.pop();
```

`greetings` avrebbe il valore `['whats up?', 'hello']`.

```js
greetings.shift();
```

`greetings` avrebbe il valore `['hello']`.

Possiamo anche restituire il valore dell'elemento rimosso con entrambi i metodi così:

```js
let popped = greetings.pop();
```

`greetings` avrebbe il valore `[]`, e `popped` avrebbe il valore `hello`.

# --instructions--

Abbiamo definito una funzione, `popShift`, che prende un array come argomento e restituisce un nuovo array. Modifica la funzione, usando `pop()` e `shift()`, per rimuovere il primo e l'ultimo elemento dell'array passato come argomento, e assegna gli elementi rimossi alle rispettive variabili, in modo che l'array restituito contenga i loro valori.

# --hints--

`popShift(["challenge", "is", "not", "complete"])` dovrebbe restituire `["challenge", "complete"]`

```js
assert.deepEqual(popShift(['challenge', 'is', 'not', 'complete']), [
  'challenge',
  'complete'
]);
```

La funzione `popShift` dovrebbe utilizzare il metodo `pop()`

```js
assert.notStrictEqual(popShift.toString().search(/\.pop\(/), -1);
```

La funzione `popShift` dovrebbe utilizzare il metodo `shift()`

```js
assert.notStrictEqual(popShift.toString().search(/\.shift\(/), -1);
```

# --seed--

## --seed-contents--

```js
function popShift(arr) {
  let popped; // Change this line
  let shifted; // Change this line
  return [shifted, popped];
}

console.log(popShift(['challenge', 'is', 'not', 'complete']));
```

# --solutions--

```js
function popShift(arr) {
  let popped = arr.pop(); // Change this line
  let shifted = arr.shift(); // Change this line
  return [shifted, popped];
}
```
