---
id: 587d7b7c367417b2b2512b19
title: "Modificare un oggetto annidato all'interno di un oggetto"
challengeType: 1
forumTopicId: 301164
dashedName: modify-an-object-nested-within-an-object
---

# --description--

Ora diamo un'occhiata a un oggetto un po' piĂą complesso. Le proprietĂ  degli oggetti possono essere annidate a una profonditĂ  arbitraria, e i loro valori possono essere di qualsiasi tipo di dato supportato da JavaScript, inclusi array e persino altri oggetti. Considera il seguente esempio:

```js
let nestedObject = {
  id: 28802695164,
  date: 'December 31, 2016',
  data: {
    totalUsers: 99,
    online: 80,
    onlineStatus: {
      active: 67,
      away: 13,
      busy: 8
    }
  }
};
```

`nestedObject` ha tre proprietĂ : `id` (il valore Ă¨ un numero), `date` (il valore Ă¨ una stringa) e `data` (il valore Ă¨ un oggetto con una struttura annidata). Anche se le strutture possono diventare rapidamente complesse, possiamo comunque usare le stesse notazioni per consultare le informazioni di cui abbiamo bisogno. Per assegnare il valore `10` alla proprietĂ  `busy` dell'oggetto annidato `onlineStatus`, usiamo la notazione a punto per riferirci alla proprietĂ :

```js
nestedObject.data.onlineStatus.busy = 10;
```

# --instructions--

Qui abbiamo definito un oggetto `userActivity`, che include un altro oggetto annidato al suo interno. Imposta il valore della chiave `online` su `45`.

# --hints--

`userActivity` dovrebbe avere le proprietĂ  `id`, `date` e `data`.

```js
assert(
  'id' in userActivity && 'date' in userActivity && 'data' in userActivity
);
```

`userActivity` dovrebbe avere una chiave `data` impostata su un oggetto con le chiavi `totalUsers` e `online`.

```js
assert('totalUsers' in userActivity.data && 'online' in userActivity.data);
```

La proprietĂ  `online` annidata nella chiave `data` di `userActivity` dovrebbe essere impostata su `45`

```js
assert(userActivity.data.online === 45);
```

La proprietĂ  `online` dovrebbe essere impostata usando la notazione a punto o a parentesi.

```js
assert.strictEqual(__helpers.removeJSComments(code).search(/online: 45/), -1);
```

# --seed--

## --seed-contents--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

// Only change code below this line

// Only change code above this line

console.log(userActivity);
```

# --solutions--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

userActivity.data.online = 45;
```
