---
id: 5a661e0f1068aca922b3ef17
title: Consultare il contenuto di un array usando la notazione a parentesi
challengeType: 1
forumTopicId: 301149
dashedName: access-an-arrays-contents-using-bracket-notation
---

# --description--

La funzionalità fondamentale di qualsiasi struttura dati è, naturalmente, la capacità non solo di archiviare dati, ma anche di poterli recuperare su richiesta. Ora che abbiamo imparato a creare un array, iniziamo a pensare a come possiamo consultare le informazioni contenute in quell'array.

Quando definiamo un array semplice come quello mostrato qui sotto, contiene 3 elementi:

```js
let ourArray = ["a", "b", "c"];
```

In un array, ogni elemento ha un <dfn>indice</dfn>. Questo indice indica sia la posizione dell'elemento nell'array sia come fare riferimento a esso. Tuttavia, è importante notare che gli array in JavaScript sono <dfn>indicizzati da zero</dfn>, il che significa che il primo elemento di un array si trova in realtà nella posizione ***zero***, non nella prima. Per recuperare un elemento da un array possiamo racchiudere un indice tra parentesi e aggiungerlo alla fine di un array, o più comunemente, a una variabile che fa riferimento a un oggetto array. Questo è noto come <dfn>notazione a parentesi</dfn>. Per esempio, se vogliamo recuperare `a` da `ourArray` e assegnarlo a una variabile, possiamo farlo con il seguente codice:

```js
let ourVariable = ourArray[0];
```

Ora `ourVariable` ha il valore di `a`.

Oltre a consultare il valore associato a un indice, puoi anche *impostare* un indice a un valore usando la stessa notazione:

```js
ourArray[1] = "not b anymore";
```

Usando la notazione a parentesi, abbiamo ora reimpostato l'elemento all'indice 1 dalla stringa `b` a `not b anymore`. Ora `ourArray` è `["a", "not b anymore", "c"]`.

# --instructions--

Per completare questa sfida, imposta la posizione 2 (indice `1`) di `myArray` a qualsiasi valore tu voglia, tranne la lettera `b`.

# --hints--

`myArray[0]` dovrebbe essere uguale alla lettera `a`

```js
assert.strictEqual(myArray[0], 'a');
```

`myArray[1]` non dovrebbe essere uguale alla lettera `b`

```js
assert.notStrictEqual(myArray[1], 'b');
```

`myArray[2]` dovrebbe essere uguale alla lettera `c`

```js
assert.strictEqual(myArray[2], 'c');
```

`myArray[3]` dovrebbe essere uguale alla lettera `d`

```js
assert.strictEqual(myArray[3], 'd');
```

# --seed--

## --seed-contents--

```js
let myArray = ["a", "b", "c", "d"];
// Only change code below this line

// Only change code above this line
console.log(myArray);
```

# --solutions--

```js
let myArray = ["a", "b", "c", "d"];
myArray[1] = "e";
```
