---
id: 587d78af367417b2b2512b03
title: アンケートフォームを作成する
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `h1` 要素に `id` が `title` のページタイトルがあること。
1. `p` 要素に `id` が `description` の短い説明があること。
1. `form` 要素に `id` が `survey-form` のものがあること。
1. フォーム要素内に、名前を入力するための必須の `input` フィールドがあり、`id` が `name` で `type` が `text` であること。
1. フォーム要素内に、メールアドレスを入力するための必須の `input` フィールドがあり、`id` が `email` であること。
1. 正しくフォーマットされていないメールアドレスを入力すると、HTML5のバリデーションエラーが表示されること。
1. フォーム要素内に、数値を入力するための `input` フィールドがあり、`id` が `number` であること。
1. 数値入力は、非数値を入力できないようにするか、ブラウザによってはHTML5のバリデーションエラーを表示すること。
1. 数値入力の範囲外の値を入力すると、`min` と `max` 属性で定義された範囲外の場合、HTML5のバリデーションエラーが表示されること。
1. 名前、メール、数値の入力フィールドには、それぞれの目的を説明する `label` 要素がフォーム内にあり、IDは `name-label`、`email-label`、`number-label` であること。
1. 名前、メール、数値の入力フィールドには、それぞれのフィールドの説明や指示を示すプレースホルダーテキストがあること。
1. フォーム要素内に、`select` のドロップダウン要素があり、`id` が `dropdown` で、選択肢が少なくとも2つあること。
1. フォーム要素内に、`name` 属性でグループ化された少なくとも2つのラジオボタンのグループから選択できること。
1. フォーム要素内に、`value` 属性を持つ複数のチェックボックスから複数選択できること。
1. フォーム要素内に、追加コメント用の `textarea` があること。
1. フォーム要素内に、すべての入力を提出するための `id` が `submit` のボタンがあること。

# --hints--

`h1` 要素に `id` が `title` のものがあること。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

`#title` は空であってはいけません。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`p` 要素に `id` が `description` のものがあること。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

`#description` は空であってはいけません。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`form` 要素に `id` が `survey-form` のものがあること。

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

`input` 要素に `id` が `name` のものがあること。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#name` は `type` が `text` であること。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

`#name` は入力必須であること。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#name` は `#survey-form` の子孫要素であること。

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

`input` 要素に `id` が `email` のものがあること。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#email` は `type` が `email` であること。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

`#email` は入力必須であること。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#email` は `#survey-form` の子孫要素であること。

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

`input` 要素に `id` が `number` のものがあること。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#number` は `#survey-form` の子孫要素であること。

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

`#number` は `type` が `number` であること。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

`#number` は数値の値を持つ `min` 属性があること。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

`#number` は数値の値を持つ `max` 属性があること。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

`label` 要素に `id` が `name-label` のものがあること。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`label` 要素に `id` が `email-label` のものがあること。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`label` 要素に `id` が `number-label` のものがあること。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`#name-label` は入力を説明するテキストを含むこと。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#email-label` は入力を説明するテキストを含むこと。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#number-label` は入力を説明するテキストを含むこと。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#name-label` は `#survey-form` の子孫要素であること。

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

`#email-label` は `#survey-form` の子孫要素であること。

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

`#number-label` は `#survey-form` の子孫要素であること。

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

`#name` は `placeholder` 属性と値を持つこと。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#email` は `placeholder` 属性と値を持つこと。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#number` は `placeholder` 属性と値を持つこと。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`select` フィールドに `id` が `dropdown` のものがあること。

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

`#dropdown` は少なくとも2つの選択可能（無効化されていない）な `option` 要素を持つこと。

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

`#dropdown` は `#survey-form` の子孫要素であること。

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

少なくとも2つの `input` 要素が `type` が `radio` のもの（ラジオボタン）であること。

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

少なくとも2つのラジオボタンが `#survey-form` の子孫要素であること。

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

すべてのラジオボタンは `value` 属性と値を持つこと。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

すべてのラジオボタンは `name` 属性と値を持つこと。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

すべてのラジオボタングループは少なくとも2つのラジオボタンを持つこと。

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

少なくとも2つの `input` 要素が `type` が `checkbox` のもの（チェックボックス）であり、`#survey-form` の子孫要素であること。

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

`#survey-form` 内のすべてのチェックボックスは `value` 属性と値を持つこと。

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

少なくとも1つの `textarea` 要素が `#survey-form` の子孫要素であること。

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

`input` または `button` 要素に `id` が `submit` のものがあること。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

`#submit` は `type` が `submit` であること。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

`#submit` は `#survey-form` の子孫要素であること。

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
