---
id: 683ec95f2f8781355556ff81
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Necesitas probar las cosas para que entiendas el flujo del algoritmo en esta etapa inicial.

Para hacer eso, primero necesitas salir del bucle. Eso es porque la implementación actual solo permitirá una iteración, así que si la condición no se cumple, habrá un bucle infinito.

Justo después de la declaración `si`, utiliza la palabra clave `break` para salir del bucle `mientras`. Luego, después del bucle `mientras`, devuelve una lista vacía para indicar que el valor no se encontró.

# --hints--

Deberías usar la palabra clave `break` para salir del bucle `mientras` después de la declaración `si`.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

Deberías devolver una lista vacía después de tu bucle `mientras`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```
