---
id: 683ec95f2f8781355556ff7d
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Para comenzar con el bucle, necesitas encontrar el índice medio del espacio de búsqueda actual.

Puedes obtener el índice medio calculando el promedio de `low` y `high`.  Tendrás que usar la división entera para obtener el promedio después de sumar los dos valores, de modo que la respuesta siempre sea un entero redondeado hacia abajo, ya que los índices deben ser enteros.

Remueve la palabra clave `pass` del bucle `while` y crea una variable `mid` establecida en el promedio de `low` y `high`.

# --hints--

Deberías remover la palabra clave `pass` del bucle while y crear una variable `mid` establecida en `(low + high) // 2`.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```
