---
id: 683ec39565dbc7e3375263b1
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En una conferencia anterior, aprendiste sobre la búsqueda binaria y cómo difiere de la búsqueda lineal.

En este taller, lo pondrás en práctica al crear un algoritmo de búsqueda binaria.

Para comenzar, crea una función `binary_search` con parámetros `search_list` y `value`.

`search_list` es la lista de elementos donde la función buscará, y `value` es el objetivo que la función intentará encontrar dentro de la `search_list`.

Agrega la palabra clave `pass` dentro de la función por ahora para que la prueba pueda pasar.

# --hints--

Debes crear una función `binary_search`.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

Tu función `binary_search` debe tener parámetros `search_list` y `value`.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
