---
id: 683ec95f2f8781355556ff7c
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora que has definido los límites de tu búsqueda, es momento de crear el ciclo que ejecutará el algoritmo de búsqueda binaria.

La búsqueda binaria funciona reduciendo repetidamente el espacio de búsqueda. Este proceso continúa mientras haya un rango válido de elementos para verificar.

Puedes expresar esto con un ciclo `while` que continuará mientras tu puntero `low` sea menor o igual a tu puntero `high`.

Si `low` en algún momento se vuelve mayor que `high`, significa que el espacio de búsqueda se ha vaciado, y el `value` no está en la lista.

Dentro de la función `binary_search`, crea un ciclo `while` con una condición que verifique si `low` es menor o igual a `high`. Dentro del ciclo, añade la palabra clave `pass` como un marcador de posición por ahora.

# --hints--

Deberías crear un ciclo `while` que se ejecute mientras `low` sea menor o igual a `high`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--

--fcc-editable-region--
```
