---
id: 683ec95f2f8781355556ff82
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Recuerde que la función toma una lista llamada `search_list` y el `value` que busca.

Ahora, llame a la función con `binary_search([1, 2, 3, 4, 5], 3)` y `binary_search([1, 2, 3, 4, 5, 9], 4)` y imprima las llamadas de inmediato.

# --hints--

Debería llamar a la función con `binary_search([1, 2, 3, 4, 5], 3)` e imprimirla de inmediato.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

Debería llamar a la función con `binary_search([1, 2, 3, 4, 5, 9], 4)` e imprimirla de inmediato.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        break
    return []

--fcc-editable-region--

--fcc-editable-region--
```
