---
id: 683ec95f2f8781355556ff7f
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Ahora, incluye `value_at_middle` en la lista `path_to_target` para rastrear los pasos realizados durante la búsqueda, independientemente de si es el objetivo.

# --hints--

Debes usar el método `append()` para agregar `value_at_middle` a la lista `path_to_target`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid] \\npath_to_target.append(value_at_middle)")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]

--fcc-editable-region--
```
