---
id: 683ec95f2f8781355556ff7e
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Ahora que has calculado el índice `mid`, necesitas realmente recuperar el valor de la `search_list` en ese índice. Este valor es el que compararás contra tu parámetro `value`, que es el objetivo que estás buscando.

Crea una variable `value_at_middle` y asígnale el elemento de la `search_list` que se encuentra en el índice `mid`.

# --hints--

Deberías crear una variable `value_at_middle` establecida en `search_list[mid]`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2


--fcc-editable-region--
```
