---
id: 684169b00a073c9b9f0a703c
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

¿Notaste que la segunda llamada de función ahora tiene los valores `3, 5, 4` en la lista? Eso indica que la búsqueda binaria está funcionando como se esperaba.

Así es como sucedió: El algoritmo primero verificó `3` como el medio de la lista inicial. Dado que `4` es mayor que `3`, la búsqueda se desplazó hacia la mitad derecha. Luego examinó `5` como el nuevo medio. Debido a que `4` es menor que `5`, la búsqueda se movió a la izquierda, identificando finalmente `4` como el medio del rango final.

Para probar la función nuevamente, llámela con `[1, 3, 5, 9, 14, 22], 10` y imprime la llamada de inmediato. Esta es una situación en la que el valor no será encontrado.

# --hints--

Deberías llamar a la función con `binary_search([1, 3, 5, 9, 14, 22], 10)` y imprimirla de inmediato.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
