---
id: 662a6bc12cde72c32fb526f0
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Un <dfn>interface</dfn> è come un progetto per una classe. Un interface contiene un insieme di metodi e prop che una classe dovrebbe implementare.

Inizia questo progetto dichiarando una classe vuota chiamata `Equation`. Userai questa classe per definire un interface, un progetto per un'equazione generica.

# --hints--

Dovresti definire una nuova classe chiamata `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
