---
id: 662bd456896f16d9bd03f1a6
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

All'interno della classe `Equation`, definisci due nuovi metodi di istanza chiamati `solve` e `analyze`.

# --hints--

Dovresti definire un metodo chiamato `solve` all'interno della classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

Il tuo metodo `solve` dovrebbe prendere un parametro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

Dovresti definire un metodo chiamato `analyze` all'interno della classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

Il tuo metodo `analyze` dovrebbe prendere un parametro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
