---
id: 662bdd364bf2cde1487922a9
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

A differenza di altri linguaggi di programmazione, Python non implementa le interfacce nel suo linguaggio di base, ma la libreria standard di Python ti permette di definire interfacce in modo semplice.

Per questo progetto, userai le utility del modulo `abc`. Quindi, importa questo modulo nel tuo codice.

# --hints--

Dovresti importare il modulo `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
