---
id: 662bd552e1c1d2db1b88ba47
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Ora, definisci un'altra classe chiamata `LinearEquation` e falla ereditare da `Equation`. Userai questa classe per rappresentare equazioni lineari.

# --hints--

Dovresti definire una classe chiamata `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

La tua classe `LinearEquation` dovrebbe ereditare dalla classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
