---
id: 662bde88dc84f1e249801b1a
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` indica le *classi base astratte*. La classe `ABC` ti permette di trasformare una classe normale in una classe astratta, cioè una classe che funge da modello per classi concrete.

Modifica la tua istruzione `import` per importare solo la classe `ABC` dal modulo `abc`. Puoi importare un oggetto specifico `x` da un modulo `y` seguendo la sintassi di importazione `from y import x`.

Poi, trasforma la tua classe `Equation` in una classe astratta facendola ereditare da `ABC`.

# --hints--

Dovresti importare `ABC` dal modulo `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

La tua classe `Equation` dovrebbe ereditare da `ABC`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
