---
id: 662bd8260da84bdd5feae419
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Vuoi che la classe `LinearEquation` implementi e non semplicemente erediti tutti i metodi definiti all'interno della classe `Equation`, che dovrebbe agire come un'interfaccia.

Attualmente, la classe `Equation` è semplicemente la classe genitore di `LinearEquation`. Nei prossimi passaggi imparerai come trasformarla in un'interfaccia formale.

Per ora, crea un'istanza di `Equation` e assegnala a una variabile `eq`, e un'istanza di `LinearEquation` e assegnala a una variabile `lin_eq`.

# --hints--

Dovresti dichiarare una variabile `eq` e assegnarle un'istanza di `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

Dovresti dichiarare una variabile `lin_eq` e assegnarle un'istanza di `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
