---
id: 662fc3eba556a6bf800d48c1
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

La funzione integrata `isinstance()` prende due argomenti e restituisce un booleano che indica se l'oggetto passato come primo argomento è un'istanza della classe passata come secondo argomento.

```py
isinstance(7, int) # True
```

Un'altra cosa che vuoi verificare è che ogni argomento sia un numero. Dopo il tuo primo `if`, crea un loop `for` che itera su `args` e controlla se l'argomento all'iterazione corrente non è un'istanza di `int` o `float`. Usa la funzione `isinstance()` e passale una tupla contenente `int` e `float` come secondo argomento.

Se l'argomento non è un numero, genera un `TypeError` con il messaggio `"Coefficients must be of type 'int' or 'float'"`.

# --hints--

Dovresti creare un loop `for` che itera su `args` dopo la tua istruzione `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

Dovresti creare un'istruzione `if` che verifica se il coefficiente corrente non è un'istanza di `int` o `float` all'interno del loop `for`.

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

Dovresti usare la stringa fornita per generare un `TypeError` all'interno dell'istruzione `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
