---
id: 662fbcef5f05e1b84f541a0c
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Ogni oggetto equazione verrà istanziato passando tanti argomenti quanti sono i coefficienti dell’equazione, partendo dal coefficiente di grado n di \\( x \\) fino al coefficiente di grado zero, incluso l’eventuale coefficiente con valore `0`.

Per esempio, `LinearEquation(4, 5)` rappresenterebbe l’equazione \\( 4x + 5 = 0 \\), con `4` che è il coefficiente del primo grado (il più alto qui) e `5` il coefficiente del grado zero.

Devi verificare che venga passato il numero corretto di argomenti per istanziare l’oggetto equazione.

All’interno del metodo `__init__`, crea un’istruzione `if` per controllare se la lunghezza di `args` è diversa dal numero di coefficienti che l’equazione dovrebbe avere (`degree + 1`). Se lo è, solleva un’eccezione `TypeError` e usa la seguente stringa per fornire un messaggio personalizzato: `f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`.

Poi, correggi l’errore passando `2` e `3` per istanziare `lin_eq`.

# --hints--

Dovresti creare un’istruzione `if` che controlli se il numero di coefficienti usati per istanziare l’equazione è diverso da `degree + 1`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

Dovresti sollevare un’eccezione `TypeError` all’interno della nuova istruzione `if` e usare la stringa fornita per restituire un messaggio di errore personalizzato.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

Dovresti passare `2` e `3` per istanziare `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
