---
id: 662fa2e2cf27c09f21f4f5d0
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

In Python, i tipi di dati vengono riconosciuti durante l'esecuzione (quando il codice viene eseguito). Perciò, non devi specificare il tipo di dato di una variabile quando la dichiari. Tuttavia, puoi annotare una variabile per chiarire che conterrà un tipo di dato specifico con `variable: <data type> = value` o semplicemente `variable: <data type>`. Nota che l'interprete Python non applica i tipi usati per annotare le variabili, e normalmente servono strumenti esterni per farlo.

All'interno della classe `Equation`, definisci un attributo di classe `degree`. Non assegnargli un valore. Usa invece un'annotazione di tipo `int` per indicare che conterrà un numero intero all'interno delle classi concrete.

Più avanti, userai questo attributo di classe come parte del processo di validazione degli argomenti passati per istanziare gli oggetti equazione.

# --hints--

Dovresti definire un attributo di classe chiamato `degree` e annotarlo con `int` all'interno della classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
