---
id: 662f6d7c92381a3049e4c987
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Un'interfaccia non deve definire solo metodi astratti, ma può anche implementare metodi da ereditare nelle classi concrete.

Prima di occuparti dell'implementazione effettiva di `solve` e `analyze` nella classe `Equation`, definisci un metodo `__init__`. Non usare alcun decoratore su di esso.

# --hints--

Dovresti definire un metodo `__init__` nella tua classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

Il tuo metodo `__init__` dovrebbe prendere un parametro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
