---
id: 662f96576ef178927de87975
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Per essere riconosciuto come metodo astratto, un metodo dovrebbe essere decorato con il decoratore `@abstractmethod`.

Un <dfn>decoratore</dfn> si usa in Python per modificare il comportamento di una funzione. Ecco un esempio di come usare un decoratore chiamato `spam`:

```py
@spam
def foo():
    pass
```

Modifica l’istruzione di importazione per importare il decoratore `abstractmethod` e decora sia i metodi `solve` sia `analyze` della classe `Equation`. Questo genererà due eccezioni.

Una volta che una classe che eredita da `ABC` ha un metodo astratto, la classe non può più essere istanziata. Quindi, elimina l’istanza `Equation` per eliminare l’errore.

L’altro errore si verifica perché la classe `LinearEquation` deve implementare tutti i metodi astratti definiti nell’interfaccia. Assicurati di definirli anche all’interno della classe `LinearEquation`. Non devi usare il decoratore `abstractmethod` nella classe concreta.

# --hints--

Dovresti importare `abstractmethod` dal modulo `abc`.

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

Dovresti decorare con `@abstractmethod` il metodo `solve` all’interno della classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

Dovresti decorare con `@abstractmethod` il metodo `analyze` all’interno della classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

Dovresti definire un metodo chiamato `solve` all’interno della classe `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

Il tuo metodo `solve` dovrebbe prendere un parametro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

Dovresti definire un metodo chiamato `analyze` all’interno della classe `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

Il tuo metodo `solve` dovrebbe prendere un parametro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
