---
id: 6639fdcc701833a54c364211
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

Dopo aver validato i coefficienti, devi archiviarli in un attributo dell'istanza. Usa una comprensione del dizionario per creare un dizionario in cui la chiave è il grado del coefficiente e il valore corrispondente è il coefficiente, e assegnalo a un attributo chiamato `coefficients`.

Per esempio, un oggetto `LinearEquation` istanziato con `2` e `4` dovrebbe avere il seguente attributo `coefficients`: `{1: 2, 0: 4}`, perché `2` corrisponde al primo grado di `x` e `4` corrisponde al grado zero di `x`.

Crea le coppie chiave-valore nel tuo nuovo dizionario seguendo lo stesso ordine di `args`.

# --hints--

Dovresti dichiarare un attributo chiamato `coefficients` all'interno del tuo metodo `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

Dovresti usare una comprensione del dizionario per archiviare i tuoi coefficienti.

```js
({ test: () => runPython(`
import ast
node = _Node(_code).find_class("Equation").find_function("__init__").find_variable("self.coefficients")
assert isinstance(node.tree.value, ast.DictComp)
`) })
```

Il tuo attributo `coefficients` dovrebbe essere un dizionario contenente coppie chiave-valore nella forma grado-coefficiente. Ricorda di seguire lo stesso ordine in cui i coefficienti sono archiviati dentro `args`.

```js
({ test: () => runPython(`
actual1 = list(LinearEquation(1, 6).coefficients.items())
expected1 = list({1: 1, 0: 6}.items())
actual2 = list(LinearEquation(-3.5, 0).coefficients.items())
expected2 = list({1: -3.5, 0: 0}.items())
assert actual1 == expected1
assert actual2 == expected2
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass

    
lin_eq = LinearEquation(2, 3)
```
