---
id: 6639f947d3a1818c9322c64a
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

L'ultimo passaggio per validare i coefficienti è verificare che il coefficiente di grado più alto sia diverso da zero. Ricorda che il coefficiente di grado più alto dovrebbe essere passato come primo argomento quando si istanzia l'oggetto.

Aggiungi un'istruzione `if` per questo e genera un `ValueError` usando la seguente stringa per fornire un messaggio personalizzato: `'Highest degree coefficient must be different from zero'`.

# --hints--

Dovresti creare un'istruzione `if` che controlli se il primo coefficiente passato per istanziare l'equazione è uguale a zero.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]") or cond.is_equivalent("not args[0]")
`)) })
```

Dovresti generare un `ValueError` all'interno della nuova istruzione `if` e usare la stringa fornita per restituire un messaggio di errore personalizzato.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_bodies()[0].has_stmt("raise ValueError('Highest degree coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
