---
id: 663b83a28943e6aa6275a514
title: Passaggio 19
challengeType: 20
dashedName: step-19
---

# --description--

Sempre all'interno della classe `Equation`, definisci un metodo `__str__` per dare una corretta rappresentazione in stringa agli oggetti equazione che creerai.

Per ora, all'interno del metodo `__str__`, dichiara una variabile `terms` e assegnale una lista vuota. Userai questa variabile per archiviare ogni termine (coefficiente per \\( x^n \\)) della tua equazione.

Poi, dichiara una variabile `equation_string`, assegnale il risultato della concatenazione degli elementi nella lista `terms` separati da uno spazio. Infine, restituisci `equation_string`.

# --hints--

Dovresti definire un metodo `__str__` all'interno della classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

Il tuo metodo `__str__` dovrebbe prendere un parametro, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

Dovresti dichiarare una variabile `terms` e assegnarle una lista vuota all'interno del metodo `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

Dovresti dichiarare una variabile `equation_string` e assegnarle il risultato della concatenazione degli elementi in `terms` separati da uno spazio all'interno del metodo `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

Dovresti restituire `equation_string` dal tuo metodo `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
