---
id: 664e4a590b52ba8d2adff19f
title: Passaggio 30
challengeType: 20
dashedName: step-30
---

# --description--

Il discriminante di un'equazione quadratica nella forma \\( ax^2 + bx + c = 0 \\), solitamente indicato con la lettera greca maiuscola delta, è uguale a \\( Δ = b^2 - 4ac \\).

All'interno della classe `QuadraticEquation`, definisci un metodo `__init__`. Usa `super()` per chiamare il metodo `__init__` dalla classe genitore. Poi, definisci un nuovo attributo chiamato `delta`, che archivia il valore del discriminante dell'equazione.

# --hints--

Dovresti definire un metodo `__init__` all'interno della classe `QuadraticEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("__init__")`)) })
```

Il tuo metodo `__init__` dovrebbe prendere due parametri, `self` e `*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_args("self, *args")`)) })
```

Dovresti chiamare `super().__init__(*args)` all'interno del tuo metodo `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_call("super().__init__(*args)")`)) })
```

Dovresti dichiarare un attributo `delta` all'interno del tuo metodo `__init__` e assegnargli il valore del discriminante dell'equazione.

```js
({ test: () => runPython(`
eq = QuadraticEquation(2, -3, -4)
assert eq.delta == 41
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2
--fcc-editable-region--
    
--fcc-editable-region--    
    def solve(self):
        pass
    
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
