---
id: 66759e32b88fb5459b1e0234
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Il metodo `__init_subclass__` viene chiamato ogni volta che la classe che lo definisce viene sottoclasse e permette di personalizzare le classi figlie. Il metodo prende un parametro chiamato per convenzione `cls` (che indica "classe"), che rappresenta la nuova classe figlia.

Definisci un metodo `__init_subclass__` nella tua classe `Equation` e dagli un parametro `cls`.

# --hints--

Dovresti definire un metodo `__init_subclass__` con un parametro `cls` nella tua classe `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
