---
id: 6675a38a8b535e4ff3274520
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

La funzione integrata `hasattr` prende un oggetto come primo argomento e una stringa che rappresenta il nome di un attributo come secondo argomento. Restituisce un booleano che indica se l'oggetto possiede l'attributo specificato.

Ora userai il metodo `__init_subclass__` per verificare se la classe figlio ha l'attributo `degree` al momento dell'istanza.

Crea un'istruzione `if` per controllare se `cls` non ha un attributo `degree`. In tal caso, genera un `AttributeError` e usa la stringa `f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"` per fornire un messaggio personalizzato.

Dopodiché, correggi l'errore apparso nel terminale dichiarando un attributo di classe `degree` all'interno della classe `LinearEquation`. Questo attributo dovrebbe rappresentare il grado dell'equazione, cioè l'esponente del termine \\( x \\) più alto. Quindi, assegna l'intero `1` all'attributo `degree`.

# --hints--

Dovresti creare un'istruzione `if` che controlli se `cls` non ha l'attributo `degree` all'interno del metodo `__init_subclass__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_conditions()[0].is_equivalent("not hasattr(cls, 'degree')")`)) })
```

Dovresti generare un `AttributeError` usando la stringa fornita all'interno della tua istruzione `if`.

```js
({ test: () => runPython(`
raise_stmt = 'raise AttributeError(f"Cannot create \\'{cls.__name__}\\' class: missing required attribute \\'degree\\'")'
node = _Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_bodies()[0]
assert node.has_stmt(raise_stmt) 
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        pass

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    
--fcc-editable-region--
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
