---
id: 6675aaf418b41157f6ccd692
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

È ora di tornare al metodo `__init__`. A seconda del tipo di equazione, dovrai passare un numero variabile di argomenti durante l'istanza.

Aggiungi un secondo parametro `args` al metodo e usa l'operatore `*` per farlo accettare un numero variabile di argomenti.

# --hints--

Il tuo metodo `__init__` dovrebbe prendere due parametri, `self` e `*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
