---
id: 56533eb9ac21ba0edf2244b3
title: بناء محول من سيلسيوس إلى فهرنهايت
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

في هذا المختبر، ستكتب دالة تُحوّل درجة الحرارة من سيلسيوس إلى فهرنهايت. الصيغة لتحويل من سيلسيوس إلى فهرنهايت هي:

```js
fahrenheit = celsius * (9/5) + 32
```

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `convertCtoF`.
2. يجب أن تأخذ `convertCtoF` معلمة رقمية واحدة، وهي درجة الحرارة بالسيلسيوس.
3. يجب أن تُرجع `convertCtoF` رقمًا.

# --hints--

يجب أن تنشئ دالة باسم `convertCtoF`.

```js
assert.isFunction(convertCtoF);
```

يجب أن تأخذ `convertCtoF` معلمة واحدة.

```js
assert.lengthOf(convertCtoF, 1);
```

يجب أن تُرجع `convertCtoF(0)` رقمًا.

```js
assert.isNumber(convertCtoF(0));
```

يجب أن تُرجع `convertCtoF(-30)` القيمة `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

يجب أن تُرجع `convertCtoF(-10)` القيمة `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

يجب أن تُرجع `convertCtoF(0)` القيمة `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

يجب أن تُرجع `convertCtoF(20)` القيمة `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

يجب أن تُرجع `convertCtoF(30)` القيمة `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
