---
id: bd7993c9c69feddfaeb7bdef
title: Multiplica dos números decimales con JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/ce2GeHq'
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

In JavaScript, you can also perform calculations with decimal numbers, just like whole numbers.

Multipliquemos dos números decimales para obtener su producto.

# --instructions--

Cambia el `0.0` para que el producto sea igual a `5.0`.

# --hints--

La variable `product` debe ser igual a `5.0`.

```js
assert(product === 5.0);
```

Debes utilizar el operador `*`

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --after-user-code--

```js
(function(y){return 'product = '+y;})(product);
```

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
