---
id: 60eebd07ea685b0e777b5583
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Ласкаво просимо до проєкту форми реєстрації! Почніть із додавання оголошення `!DOCTYPE html` на початку документа, щоб браузер знав, який тип документа він читає.

# --hints--

Ваш код має містити референс `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Ви повинні додати пробіл після референса `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Ви повинні визначити тип документа як `html`.

```js
assert(code.match(/html/gi));
```

Ви повинні закрити оголошення `DOCTYPE` `>` після типу.

```js
assert(code.match(/html\s*>/gi));
```

Ваше оголошення `DOCTYPE` має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
