---
id: 60f027099a15b00485563dd2
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Під `DOCTYPE` додайте елемент `html` з атрибутом `lang`, значення якого дорівнює `en`, щоб мати місце для початку написання коду.

# --hints--

Ваше оголошення `DOCTYPE` має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ваш елемент `html` повинен мати початковий тег і атрибут `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш елемент `html` повинен мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Ваші теги `html` мають бути в правильному порядку.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

У вас має бути лише один елемент `html`.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
