---
id: 60f0286404aefb0562a4fdf9
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Додайте `title` та `meta` елементи всередину `head` елемента. Дайте своєму проєкту назву `Registration Form`, а також додайте атрибут `charset` зі значенням `utf-8` до вашого `meta` елемента.

# --hints--

Ваш код має містити `title` елемент.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ваш `title` елемент має бути всередині `head` елемента.

```js
assert.exists(document.querySelector('head > title'));
```

Ваш проєкт має мати назву `Registration Form`. Пам’ятайте, що важливі регістр і правопис назви.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Ви маєте створити `meta` елемент всередині `head` елемента.

```js
assert.exists(document.querySelector('head > meta'));
```

Ви маєте додати `meta` атрибут зі значенням `charset` до `utf-8` елемента.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Ваш `meta` елемент має бути порожнім елементом, він не має закриваючого тегу `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
