---
id: 60f027c87bc98f050395c139
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Далі додайте відкриваючий і закриваючий `head` та `body` теги всередині елемента `html`.

# --hints--

У вас має бути відкриваючий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

У вас має бути закриваючий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

У вас має бути відкриваючий тег `body`.

```js
assert(code.match(/<body\s*>/i));
```

У вас має бути закриваючий тег `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` і `body` мають бути братськими.

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

Елемент `head` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Елемент `body` має бути всередині елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
