---
id: 60f803d5241e6a0433a523a1
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--
Одиниця `rem` означає кореневий `em` і відноситься до розміру шрифту елемента `html`.

Оскільки елементи `label` за замовчуванням є рядковими, вони всі відображаються поруч на одній лінії, через що їхній текст важко читати. Щоб зробити так, щоб вони з’являлися на окремих лініях, додайте `display: block` до елемента `label` і додайте властивість `margin` зі значенням `0.5rem 0`, щоб відокремити їх один від одного.

# --hints--

Ви повинні використати селектор `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Ви повинні додати властивість `display` зі значенням `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Ви повинні додати властивість `margin` зі значенням `0.5rem 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
