---
id: 686e64be3cd9db25282a956d
title: Diseña una Exposición de Fotografía
challengeType: 25
dashedName: lab-photography-exhibit
demoType: onClick
---

# --description--

En este laboratorio, crearás una estructura de exhibición fotográfica usando las clases utilitarias de Tailwind CSS. 

Para las imágenes, puedes usar las siguientes URLs si quieres:

- https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg
- https://cdn.freecodecamp.org/curriculum/labs/alps.jpg
- https://cdn.freecodecamp.org/curriculum/labs/sea.jpg

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener un elemento con una clase llamada `main-container`.
2. Tu elemento `.main-container` también debe tener las siguientes clases de utilidad de Tailwind CSS:
   - la clase de utilidad correcta para crear un contenedor CSS Grid.
   - una clase de utilidad para definir el número de columnas dentro de la cuadrícula. Debes usar la clase de utilidad que utiliza el número fijo de columnas. Ej. `grid-cols-<number>`.
   - una clase de utilidad para establecer el espacio entre filas y columnas de los elementos de la cuadrícula.
3. Dentro de tu elemento `.main-container`, debe haber al menos tres elementos, cada uno con una clase llamada `card`.
4. Cada uno de tus elementos `.card` también debe tener las siguientes clases utilitarias de Tailwind CSS:
   - una clase de utilidad para establecer un borde redondeado de tu elección.
   - una clase de utilidad para establecer el padding que elijas. Debes usar este formato `p-<number>`.
5. Dentro de cada uno de tus elementos `.card`, debes tener los siguientes elementos:
   - un elemento de imagen con atributos `src` y `alt`.
   - un elemento con la clase llamada `subheading`.
   - un elemento con la clase llamada `description`.
6. Cada uno de tus elementos de imagen debe tener una clase de utilidad para establecer un borde redondeado de tu elección.
7. Cada uno de tus elementos `.subheading` debe usar clases de utilidad para establecer el peso y tamaño de fuente que elijas.
8. Cada uno de tus elementos `.description` debe usar una clase de utilidad para establecer el tamaño de fuente que elijas.

# --hints--

Debes tener un elemento con una clase llamada `main-container`.

```js
assert.exists(document.querySelector(".main-container"));
```

Tu elemento `.main-container` debe usar la clase de utilidad correcta de Tailwind CSS para crear una estructura CSS Grid.

```js
const mainContainer = document.querySelector(".main-container");
assert.isTrue(mainContainer.classList.contains("grid"));
```

Tu elemento `.main-container` debe usar una clase de utilidad para definir el número de columnas dentro de la cuadrícula.

```js
const mainContainer = document.querySelector(".main-container");
// Grid classes can vary, so this check is for the presence of any grid-cols class
const hasGridCols = [...mainContainer.classList].some(className => className.startsWith("grid-cols-"));
assert.isTrue(hasGridCols);
```

Tu elemento `.main-container` debe usar una clase de utilidad para establecer el espacio entre filas y columnas de los elementos de la cuadrícula.

```js
const mainContainer = document.querySelector(".main-container");
// Check for any gap class including gap-x and gap-y

const hasGap = [...mainContainer.classList].some(c => c.startsWith("gap-")) ||
  (classList.some(c => c.startsWith("gap-x-")) &&
   classList.some(c => c.startsWith("gap-y-")));

assert.isTrue(hasGap);
```

Tu elemento `.main-container` debe contener al menos tres elementos, cada uno con una clase llamada `card`.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isAtLeast(cards.length, 3);
```

Cada uno de tus elementos `.card` debe tener una clase de utilidad para establecer un borde redondeado.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);

const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

cards.forEach(card => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => card.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
})
```

Cada uno de tus elementos `.card` debe tener una clase de utilidad para establecer el padding.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);
const hasPadding = [...cards].every(card => {
  return [...card.classList].some(c => c.startsWith("p-"));
});
assert.isTrue(hasPadding);
```

Cada uno de tus elementos `.card` debe contener un elemento de imagen.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
```

Cada uno de tus elementos `img` debe tener un atributo `src`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("src")));
images.forEach(img => assert.isTrue(img.getAttribute("src").length > 0));
```

Cada uno de tus elementos `img` debe tener un atributo `alt`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("alt")));
images.forEach(img => assert.isTrue(img.getAttribute("alt").length > 0));
```

Cada uno de tus elementos `.card` debe contener un elemento con la clase llamada `subheading`.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isAtLeast(subheadings.length, 3);
```

Cada uno de tus elementos `.subheading` debe tener algún contenido de texto.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(subheadings);

subheadings.forEach(subheading => {
  assert.isNotEmpty(subheading.textContent.trim());
});
```

Cada uno de tus elementos `.card` debe contener un elemento con la clase llamada `description`.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
assert.isAtLeast(descriptions.length, 3);
```

Cada uno de tus elementos `.description` debe tener algún contenido de texto.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(descriptions);

descriptions.forEach(description => {
  assert.isNotEmpty(description.textContent.trim());
});
```

Cada uno de tus elementos `img` debe tener una clase de utilidad para establecer un borde redondeado.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isNotEmpty(images);
const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

images.forEach(img => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => img.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
});
```

Cada uno de tus elementos `.subheading` debe usar clases de utilidad para establecer el peso de la fuente.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedFontWeights = [
  "font-thin",
  "font-extralight",
  "font-light",
  "font-normal",
  "font-medium",
  "font-semibold",
  "font-bold",
  "font-extrabold",
  "font-black"
];

subheadings.forEach(subheading => {
  const hasFontWeight = predefinedFontWeights.some(weight => subheading.classList.contains(weight));
  assert.isTrue(hasFontWeight);
});
```

Cada uno de tus elementos `.subheading` debe usar clases de utilidad para establecer el tamaño de la fuente.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

subheadings.forEach(subheading => {
  const hasFontSize = predefinedSizes.some(size => subheading.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

Cada uno de tus elementos `.description` debe usar una clase de utilidad para establecer el tamaño de la fuente.

```js
const descriptions = document.querySelectorAll(".description");
assert.isNotEmpty(descriptions);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

descriptions.forEach(description => {
  const hasFontSize = predefinedSizes.some(size => description.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
 
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-purple-800 text-yellow-400 min-h-screen">
    <h1 class="text-center text-4xl font-bold my-12">Photography Exhibit</h1>

    <main class="main-container grid grid-cols-1 md:grid-cols-3 gap-8 px-4">
      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="Ancient Colosseum"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Colosseum, Rome</h2>
          <p class="description text-sm mb-4">
            There is so much history and beauty. The light hits it just right.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="Alpine Mountains"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">The Alps</h2>
          <p class="description text-sm mb-4">
            The mountains go on forever and the view is breathtaking.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="Calm Sea"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Seascape</h2>
          <p class="description text-sm mb-4">
            The water was still, the air was cool and the feeling is pure peace.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>
    </main>
  </body>
</html>
```
