---
id: 66c6491fe4c8e0f16845425f
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Sasa ni wakati wa kuongeza tunda lingine kwenye orodha.

Ukitegemea njia ile ile ya safu ya data kama awali, ongeza mfuatano wa herufi `"Grapes"` mwishoni mwa safu ya data `shoppingList`.

Kisha, ongeza `console.log` na fanya wito wa kitendakazi `getShoppingListMsg` ukiwa na safu ya data `shoppingList` kama hoja ndani ya `console.log` ili kuona orodha iliyosasishwa ikichapishwa kwenye konsoli.

# --hints--

Safu yako ya data `shoppingList` sasa inapaswa kuwa na vitu viwili.

```js
assert.lengthOf(shoppingList, 2);
```

Safu yako ya data `shoppingList` inapaswa kuwa na mfuatano wa herufi `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

Unapaswa kufanya wito wa kitendakazi `getShoppingListMsg(shoppingList)` ndani ya `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
