---
id: 66c64095f61166ede6e39a84
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Kwa orodha hii ya ununuzi, utatumia safu ya data kuwakilisha vitu unavyohitaji kununua.

Katika mafundisho yaliyopita, ulijifunza jinsi ya kuunda safu za data kama hizi:

```js
const fruits = ['apple', 'banana', 'orange'];
```

Katika hatua hii, unda kigezo kinachoitwa `shoppingList` na uweke thamani safu ya data tupu.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

Unapaswa kuweka thamani ya safu ya data kwenye kigezo cha `shoppingList`.

```js
assert.isArray(shoppingList);
```

Safu yako ya data inapaswa kuwa tupu.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
