---
id: 66c64210413532ee9d3bd342
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Kwa sehemu hii inayofuata ya warsha, utafanya mazoezi ya kuongeza vyakula kwenye orodha ya manunuzi.

Anza kwa kutumia `console.log` kuandika ujumbe `"It will be nice to have some fruit to eat."` kwenye konsoli.

# --hints--

Unapaswa kuandika ujumbe `"It will be nice to have some fruit to eat."` kwenye konsoli.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
