---
id: 66c644c08ebcd2ef90c750bd
title: Hatua ya 4
challengeType: 1
dashedName: step-4
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kuongeza vitu mwishoni mwa safu ya data kwa kutumia njia ya `push` kama ifuatavyo:

```js
fruits.push('pear');
```

Kwa kutumia njia ya `push`, ongeza mfuatano wa herufi `"Apples"` kwenye safu ya data `shoppingList`.

# --hints--

Safu yako ya data ya `shoppingList` inapaswa kuwa na kipengee kimoja ndani yake.

```js
assert.lengthOf(shoppingList, 1);
```

Safu yako ya data ya `shoppingList` inapaswa kuwa na mfuatano wa herufi `"Apples"` ndani yake. Hakikisha umeangalia makosa ya tahajia na herufi kubwa.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
