---
id: 66c645b345a39bf04864dc50
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Ili kuona orodha ya ununuzi iliyosasishwa, utahitaji kuandika orodha ya sasa ya ununuzi pamoja na ujumbe mfupi kwenye konsoli.

Kwa kuwa ujumbe huu utatumika mara kwa mara katika warsha, ni bora kuunda kitendakazi kinachoweza kutumika tena.

Unda kitendakazi kinachoitwa `getShoppingListMsg` kinachopokea safu ya data kama kigezo na kurudisha mfuatano wa herufi `"Current Shopping List: "` ukifuatiwa na yaliyomo katika safu ya data iliyotolewa.

Unaweza kutumia herufi za kiolezo au kuunganisha mfuatano wa herufi kwa kutumia kigezo cha `+` kuunganisha mfuatano wa herufi na safu ya data iliyopitishwa kwa kitendakazi.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `getShoppingListMsg` katika msimbo wako.

```js
assert.isFunction(getShoppingListMsg);
```

Kitendakazi chako cha `getShoppingListMsg` kinapaswa kurudisha mfuatano wa herufi.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

Kitendakazi chako cha `getShoppingListMsg` kinapaswa kupokea safu ya data na kurudisha ujumbe `"Current Shopping List: "` ukifuatiwa na yaliyomo katika safu hiyo. Hakikisha upau na makosa ya tahajia.

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
