---
id: 66c63bfa4685e4e3843fa0dc
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utaendelea kujifunza kuhusu safu za data kwa kutengeneza orodha ya ununuzi wa vyakula.

Anza kwa kuongeza `console.log` inayochapisha mfuatano wa herufi `"Grocery shopping list"` kwenye konsoli.

# --hints--

Unapaswa kuwa na `console.log` katika msimbo wako.

```js
assert.match(code, /console\.log(.*)/);
```

`console.log` yako inapaswa kutoa ujumbe `"Grocery shopping list"`. Hakikisha upitie makosa ya tahajia na herufi kubwa.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
