---
id: 66c726c34ecf1e238aa9d7d5
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kuongeza vipengele mwanzoni mwa safu ya data kwa kutumia njia ya `unshift()`. 

Hapa ni ukumbusho wa jinsi ya kutumia njia ya `unshift()`:

```js
array.unshift(item1, item2, ..., itemX);
```

Tumia njia ya `unshift()` kuongeza mfuatano wa herufi `"Vegetable Oil"` mwanzoni mwa safu ya data ya `shoppingList`.

# --hints--

Unapaswa kuwa na vitu vitatu katika safu ya data yako ya `shoppingList`.

```js
assert.lengthOf(shoppingList, 3);
```

Unapaswa kutumia njia ya `unshift()` kuongeza `"Vegetable Oil"` mwanzoni mwa safu ya data ya `shoppingList`.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
