---
id: 66c73a7798f6f62b2ae58f22
title: Hatua 16
challengeType: 1
dashedName: step-16
---

# --description--

Sasa ni wakati wa kuongeza vitu zaidi mwanzoni mwa orodha ya vitu vya dukani.

Anza kwa kuongeza tamko la `console.log` linaloandika ujumbe `"It might be nice to get a dessert."`

Chini ya tamko hilo la `console`, tumia njia sahihi ya safu ya data kuongeza mfuatano wa herufi `"Chocolate Cake"` mwanzoni mwa safu ya data `shoppingList`.

Mwishowe, ongeza `console.log` na fanya wito wa kitendakazi `getShoppingListMsg` ukiwa na safu ya data `shoppingList` kama hoja ndani ya `console.log` ili kuona orodha iliyosasishwa ikionyeshwa kwenye konsoli.

# --hints--

Unapaswa kuwa na `console.log` linaloandika ujumbe `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

Safu yako ya data ya `shoppingList` inapaswa kuwa na jumla ya vitu `6`.

```js
assert.lengthOf(shoppingList, 6);
```

Unapaswa kutumia njia ya `unshift` kuongeza mfuatano wa herufi `"Chocolate Cake"` mwanzoni mwa safu ya data `shoppingList`.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

Unapaswa kuita `getShoppingListMsg(shoppingList)` ndani ya `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
