---
id: 66c730ee6ae076281721d0b9
title: Hatua 14
challengeType: 1
dashedName: step-14
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kuondoa vitu kutoka mwisho wa safu ya data kwa kutumia njia ya `pop`. 

Hapa ni ukumbusho wa jinsi ya kutumia njia ya `pop`:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

Tumia njia ya `pop` kuondoa kipengee cha mwisho kutoka kwenye safu ya data ya `shoppingList`.

# --hints--

Safu yako ya data ya `shoppingList` sasa inapaswa kuwa na vitu `5` ndani yake.

```js
assert.lengthOf(shoppingList, 5);
```

Safu yako ya data ya `shoppingList` haipaswi tena kuwa na mfuatano wa herufi `"Potato Chips"` ndani yake.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
