---
id: 66c72b0ffbc5522525768558
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Katika hatua zilizopita, ulikagua jinsi ya kuongeza kipengee mwishoni mwa safu ya data kwa kutumia njia ya `push`.

Lakini njia ya `push` inakubali hoja nyingi, hivyo unaweza kuongeza vipengee vingi mwishoni mwa safu ya data kama ifuatavyo:

```js
array.push(item1, item2, item3);
```

Katika hatua hii, tumia njia ya `push` kuongeza mfuatano wa herufi `"Popcorn", "Beef Jerky", "Potato Chips"` hadi `shoppingList` kwenye safu ya data.

Mpangilio ni muhimu, hivyo hakikisha unaongeza vipengee kwa mpangilio vilivyoorodheshwa.

# --hints--

Safu yako ya data ya `shoppingList` inapaswa kuwa na jumla ya mfuatano wa herufi `6`.

```js
assert.lengthOf(shoppingList, 6);
```

`shoppingList` inapaswa kuwa na mfuatano wa herufi `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

`shoppingList` inapaswa kuwa na mfuatano wa herufi `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

`shoppingList` inapaswa kuwa na mfuatano wa herufi `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
