---
id: 66c72a55418cc9247b710827
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Ifuatayo, ongeza `console.log` na fanya wito wa kitendakazi `getShoppingListMsg` ukiwa na safu ya data `shoppingList` kama hoja ndani ya `console.log` ili kuona orodha iliyosasishwa ikirekodiwa kwenye konsoli.

# --hints--

Unapaswa kufanya wito wa `getShoppingListMsg(shoppingList)` ndani ya `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
