---
id: 66c72f4d0528bd268a82107b
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Sasa ni wakati wa kuandika safu ya data ya `shoppingList` iliyosasishwa kwenye konsoli.

Ongeza `console.log` nyingine na fanya wito wa kitendakazi cha `getShoppingListMsg` ukiwa na safu ya data ya `shoppingList` kama hoja ndani ya `console.log` ili kuona orodha iliyosasishwa ikichapishwa kwenye konsoli.

# --hints--

Unapaswa kufanya wito wa `getShoppingListMsg(shoppingList)` ndani ya `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 4);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");

--fcc-editable-region--

--fcc-editable-region--
```
