---
id: 66cbe2319d3845545a293a0b
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Sasa ni wakati wa kuona ujumbe uliandikwa kwenye konsoli.

Ongeza `console.log` na fanya wito wa kitendakazi `getShoppingListMsg` ukiwa na safu ya data `shoppingList` kama hoja ndani ya `console.log` ili kuona ujumbe uliandikwa kwenye konsoli.

# --hints--

Unapaswa kufanya wito wa kitendakazi `getShoppingListMsg(shoppingList)` ndani ya `console.log`.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
