---
id: 685523c05fb41d0848c4b52f
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Всередині першого елемента `div` створіть елемент `span` з текстом `Soundflow`, елемент `h1` з текстом `Discover New Music` та елемент `p` з текстом `Stream your favorite tracks and discover new artists.`

# --hints--

Ви маєте створити елемент `span` з текстом `Soundflow` всередині першого елемента `div`.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

Ви маєте створити елемент `h1` з текстом `Discover New Music` всередині першого елемента `div`.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

Ви маєте створити елемент `p` з текстом `Stream your favorite tracks and discover new artists.` всередині першого елемента `div`.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
