---
id: 68551b677a5f6898af364b85
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви дізнаєтеся більше про Tailwind CSS, використовуючи його утиліти гнучкого контейнера для створення адаптивного компонента заклику до дії (CTA).

Для вас уже підготовлено базовий шаблонний код HTML, тож спочатку потрібно створити HTML-елементи.

Почніть зі створення елемента `div`. Всередині цього `div` створіть ще два окремі елементи `div`.

# --hints--

Ви повинні створити елемент `div`.

```js
assert.exists(document.querySelector("div"));
```

Усередині елемента `div` має бути ще два елементи `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
