---
id: 68552464ef378e1051a55c28
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Всередині другого `div` створіть два елементи посилання. Перший має містити текст `Learn more`, а другий — текст `Start listening`.

Для атрибута `href` обох елементів посилання встановіть значення `#`.

# --hints--

Ви маєте створити два елементи `a` всередині другого елемента `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

Ваш перший елемент `a` повинен містити текст `Learn more` та атрибут `href` зі значенням `#` всередині другого елемента `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

Ваш другий елемент `a` повинен містити текст `Start listening` та атрибут `href` зі значенням `#` всередині другого елемента `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
