---
id: 6858ebf2552046ab3de9f17c
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

На цьому етапі ви помітите, що праворуч від вмісту в `div` занадто багато пробілу. Ви можете зменшити це, додавши ширину.

Додайте клас `md:w-1/2` до `div`, щоб зменшити ширину до 50% на середніх і більших екранах. Тобто на екранах шириною від 768 пікселів і більше.

Щоб побачити це в дії, розгорніть вкладку попереднього перегляду.

# --hints--

Ваш зовнішній елемент `div` має мати утилітний клас `md:w-1/2`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
