---
id: 6858f90c54647a48ddb8822d
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

На цьому етапі вміст, а саме два елементи `div` всередині зовнішнього `div`, буде розташований один над одним незалежно від розміру екрану.

Це тому, що у вас є клас `flex-col` у обгортковому елементі `div`. Щоб це виправити, використайте префікс `lg:`, щоб додати клас `flex-row` до `div`.

Це означає, що на пристроях із більшими екранами напрямок flex буде `row`. Щоб знову побачити це в дії, змініть ширину браузера на більшу.

# --hints--

Ваш зовнішній елемент `div` має містити утилітарний клас `lg:flex-row`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
