---
id: 6858f3a67aa8580ff37016c7
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Якщо ви знову зміните ширину вкладки попереднього перегляду на ширину, ідентичну мобільному пристрою, ви помітите, що вміст завжди розташований поруч. Це не те, що потрібно на мобільних пристроях.

Щоб це виправити, додайте клас `flex-col` до `div`. Це змінить напрямок flex на колонку, тож елементи будуть розташовуватися один над одним.

Ось як використовувати утилітний клас `flex-col`:

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

Ваш відкриваючий елемент `div` повинен мати утилітний клас `flex-col`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
