---
id: 6858ef974c8ad5dcd7abb9f1
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Вміст всередині `div` можна краще розташувати та відступити. Для решти відступів і розташування потрібно використати гнучкий контейнер.

Додайте клас `flex` до `div`, щоб два дочірні елементи `div` розмістилися поруч.

Також додайте `justify-around`, щоб простір навколо і між вмістом розподілявся рівномірно, і `items-center`, щоб увесь вміст нарешті був по центру в `div`. 

# --hints--

Ваш зовнішній елемент `div` має мати утилітарний клас `flex`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

Ваш зовнішній елемент `div` має мати утилітарний клас `justify-around`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

Ваш зовнішній елемент `div` має мати утилітарний клас `items-center`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
