---
id: 6858e6408a0999603b2a4fee
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Для `h1` додайте класи `font-bold` і `text-4xl`. Це зробить текст жирнішим, більшим і також додасть більше міжрядкового інтервалу.

Також додайте клас `my-4`, щоб збільшити верхнє та нижнє поле `h1` на 16 пікселів, або `1rem`.

# --hints--

Ваш елемент `h1` має мати утилітний клас `font-bold`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

Ваш елемент `h1` має мати утилітний клас `text-4xl`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

Ваш елемент `h1` має мати утилітний клас `my-4`.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
         --fcc-editable-region--
        <h1>Discover New Music</h1>
         --fcc-editable-region--
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
