---
id: 68597114730d7c6b8838d204
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Абзац під `h1` виглядає добре, тому йому не потрібне додаткове стилізування. Це означає, що ви можете перейти до елементів посилання всередині другого елемента `div`.

Додайте першому елементу посилання класи `bg-white`, `hover:bg-gray-100` і `text-indigo-600`. Це надасть йому білий фон, колір тексту індиго та трохи світліший фон при наведенні.

Для другого елемента посилання додайте класи `bg-pink-500`, `hover:bg-pink-600` і `text-white`.

# --hints--

Ваш перший елемент посилання має мати утилітарний клас `bg-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

Ваш перший елемент посилання має мати утилітарний клас `hover:bg-gray-100`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

Ваш перший елемент посилання має мати утилітарний клас `text-indigo-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

Ваш другий елемент посилання має мати утилітарний клас `bg-pink-500`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

Ваш другий елемент посилання має мати утилітарний клас `hover:bg-pink-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

Ваш другий елемент посилання має мати утилітарний клас `text-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
