---
id: 698b42492a7f5f737895892e
title: 建造一個鄰接列表轉矩陣轉換器
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

在這個實驗中，你將建造一個函式，將圖的鄰接列表表述轉換成鄰接矩陣。鄰接列表是一個物件，其中每個鍵代表一個節點，對應的值是一個節點陣列，表示該鍵節點所連接的節點。鄰接矩陣是一個二維陣列，位置 `[i][j]` 的值為 `1` 表示從節點 `i` 到節點 `j` 有邊，否則為 `0`。

例如，給定鄰接列表：

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

對應的鄰接矩陣將會是：

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該定義一個名為 `adjacencyListToMatrix` 的函式（程式），將鄰接列表轉換為鄰接矩陣。
2. 該函式（程式）應該接受一個表述無權重（無向或有向）圖的鄰接列表的物件作為其引數。
3. 該函式（程式）應該：
   - 將鄰接列表轉換為鄰接矩陣。
   - 列印鄰接矩陣中的每一列。
   - 傳回鄰接矩陣。

例如，`adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})` 應該列印：

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

並傳回 `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`。

# --hints--

你應該定義一個名為 `adjacencyListToMatrix` 的函式。

```js
assert.isFunction(adjacencyListToMatrix);
```

`adjacencyListToMatrix` 函式（程式）應該有一個參數。

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

該函式（程式）應正確從鄰接列表中判斷節點數量。

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

該函式（程式）應正確將矩陣值設為 `1`，以表示現有的邊。

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

該函式（程式）應該列印矩陣的每一列。

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

該函式（程式）應該傳回鄰接矩陣。

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

當給定鄰接列表 `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}` 時，該函式應該傳回 `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`。

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

當給定鄰接列表 `{0: [1], 1: [0]}` 時，該函式應該傳回 `[[0, 1], [1, 0]]`。

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

當給定鄰接列表 `{0: [], 1: [], 2: []}` 時，該函式應該傳回 `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`。

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
