---
id: 66c06d618d075c7f7f1b890a
title: 创建一个占卜机
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

在本实验室中，你将定义五个运势，并随机选择其中一个向用户显示运势。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应初始化五个变量 `fortune1`、`fortune2`、`fortune3`、`fortune4` 和 `fortune5`，它们的值是你自定义的一段字符串。 你也可以使用以下选项：

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. 你应在 1 到 5（包括 5）之间选择一个随机数，并将其赋值给变量 `randomNumber` 。

3. 你应创建一个 `selectedFortune` 变量，并根据以下规则分配适当的运势：

   - 如果 `randomNumber` 是 1，就把 `fortune1` 的值赋给 `selectedFortune`。
   - 如果 `randomNumber` 是 2，就把 `fortune2` 的值赋给 `selectedFortune`。
   - 如果 `randomNumber` 是 3，就把 `fortune3` 的值赋给 `selectedFortune`。
   - 如果 `randomNumber` 是 4，就把 `fortune4` 的值赋给 `selectedFortune`。
   - 如果 `randomNumber` 是 5，就把 `fortune5` 的值赋给 `selectedFortune`。

4. 你应将 `selectedFortune` 记录到控制台。

# --hints--

你应该用一个字符串值初始化 `fortune1` 。

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

你应该用一个字符串值初始化 `fortune2` 。

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

你应该用一个字符串值初始化 `fortune3` 。

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

你应该用一个字符串值初始化 `fortune4` 。

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

你应该用一个字符串值初始化 `fortune5` 。

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

你应使用 `Math.random()` 方法生成随机数。

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

你应该生成一个介于 1 和 5 之间的随机数，并将其赋值给变量 `randomNumber` 。

```js
assert.isNotNull(randomNumber);
assert.include([1,2,3,4,5], randomNumber);
```

你应该有一个 `selectedFortune` 变量，并根据 `randomNumber` 的值赋值。

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` 应与其运势相对应。 例如，如果 `randomNumber` 为 1，`selectedFortune` 应等于 `fortune1`，以此类推。

```js
const condition1 = randomNumber === 1 && selectedFortune === fortune1;
const condition2 = randomNumber === 2 && selectedFortune === fortune2;
const condition3 = randomNumber === 3 && selectedFortune === fortune3;
const condition4 = randomNumber === 4 && selectedFortune === fortune4;
const condition5 = randomNumber === 5 && selectedFortune === fortune5;

assert.isTrue(condition1 || condition2 || condition3 || condition4 || condition5);
```

你应将 `selectedFortune` 输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
