---
id: 6718d2d59337c822ecb697ff
title: crea un programma per la gestione di un conto bancario
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi definire una classe chiamata `BankAccount` con un costruttore. Il costruttore deve impostare il `balance` predefinito su `0` e inizializzare un array vuoto chiamato `transactions` per archiviare i record delle transazioni come oggetti.

2. Ogni transazione archiviata nell'array `transactions` deve essere un oggetto con due proprietà: `type` e `amount`. La proprietà `type` deve essere `deposit` o `withdraw`, e la proprietà `amount` deve essere l'importo depositato o prelevato.

3. Devi definire un metodo chiamato `deposit` che prende come parametro l'importo del deposito. Quando l'importo del deposito è maggiore di `0`, deve:

   - Inserire un nuovo oggetto nell'array delle transazioni con tipo `deposit` e l'importo depositato.
   - Aggiornare il saldo.
   - Restituire `"Successfully deposited $[amount]. New balance: $[balance]"`.

4. Se l'importo è minore o uguale a `0`, il metodo `deposit` deve restituire `"Deposit amount must be greater than zero."`.

5. Devi definire un metodo chiamato `withdraw` che prende come parametro un importo. Questo metodo deve aggiornare il saldo corrente in base ai prelievi. Quando l'importo da prelevare è maggiore di `0` e minore o uguale al saldo corrente, deve:

   - Inserire un nuovo oggetto nell'array delle transazioni con tipo `withdraw` e l'importo prelevato.
   - Aggiornare il saldo.
   - Restituire `"Successfully withdrew $[amount]. New balance: $[balance]"`.

6. Se l'importo da prelevare è minore o uguale a `0` o maggiore del saldo corrente, il metodo `withdraw` deve restituire `"Insufficient balance or invalid amount."`.

7. Devi definire un metodo chiamato `checkBalance` che restituisce il saldo corrente nel formato `"Current balance: $[balance]"`.

8. Devi definire un metodo chiamato `listAllDeposits` che itera attraverso l'array delle transazioni e restituisce tutti i depositi nel formato `"Deposits: amount,amount,..."`.

9. Devi definire un metodo chiamato `listAllWithdrawals` che itera attraverso l'array delle transazioni e restituisce tutti i prelievi nel formato `"Withdrawals: amount,amount,..."`.

10. Devi creare una nuova istanza di `BankAccount` chiamata `myAccount`.

11. Il tuo conto bancario `myAccount` deve avere almeno cinque transazioni.

12. Il tuo conto bancario `myAccount` deve avere almeno due depositi.

13. Il tuo conto bancario `myAccount` deve avere almeno due prelievi.

14. Il tuo conto bancario `myAccount` deve avere un saldo maggiore di `$100`.

# --hints--

Devi definire una classe chiamata `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

L'oggetto `BankAccount` deve inizialmente avere un `balance` di `0` e un array vuoto `transactions` per archiviare i record delle transazioni.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

Devi avere un metodo `deposit` che prende come parametro l'importo del deposito.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

Devi avere un metodo `withdraw` che prende come parametro l'importo del prelievo.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

Devi avere un metodo `checkBalance` che controlla il saldo corrente.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

Devi avere un metodo `listAllDeposits` che elenca tutti i depositi.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

Devi avere un metodo `listAllWithdrawals` che elenca tutti i prelievi.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` deve restituire `"Successfully deposited $100. New balance: $100"`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` deve restituire `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` deve restituire `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

Quando il saldo del conto è `100`, `BankAccount.withdraw(150)` deve restituire `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` deve restituire `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` deve restituire `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

Quando il saldo del conto è `200`, `BankAccount.withdraw(150)` deve restituire `"Successfully withdrew $150. New balance: $50"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

Quando il saldo del conto è `200`, `BankAccount.checkBalance()` deve restituire `"Current balance: $200"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

Quando depositi `10`, `35`, `90`, il metodo `listAllDeposits` deve restituire `"Deposits: 10,35,90"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

Quando prelevi `20`, `50`, `100`, il metodo `listAllWithdrawals` deve restituire `"Withdrawals: 20,50,100"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

Devi avere un'istanza di `BankAccount` chiamata `myAccount`.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

Il tuo conto bancario `myAccount` deve avere almeno cinque transazioni.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

Il tuo conto bancario `myAccount` deve avere almeno due depositi.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

Il tuo conto bancario `myAccount` deve avere almeno due prelievi.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```

Il tuo conto bancario `myAccount` deve avere un saldo maggiore di `$100`.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
