---
id: 69b83e35f19ba26ba1fa517a
title: Побудова послідовності світлофора
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

У цій лабораторній роботі ви симулюватимете конфігуровані цикли світлофора та фіксуватимете аномалії.

Ви працюватимете з об’єктами конфігурації, які описують фази світлофора. Кожен об’єкт конфігурації має такі властивості:

- `fault`: булевий прапорець, який запускає дострокове припинення, коли встановлений у `true`.
- `phases`: масив об’єктів фаз.

Кожен об’єкт фази всередині `phases` має такі властивості:

- `color`: рядок, що представляє колір світла (`"green"`, `"yellow"` або `"red"`).
- `duration`: додатне ціле число, що вказує, скільки триває фаза в секундах.

Ви можете звертатися до наданих об’єктів `config1`, `config2`, `config3` та `config4` як прикладів можливих об’єктів конфігурації.

**Примітка**: не додавайте додаткових викликів `console.log()`, оскільки це може спричинити помилки в тестах.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

# Історії користувача

1. Ви повинні мати функцію з назвою `runSequence` з двома параметрами: `config` і `cycles`, де `config` представляє об’єкт конфігурації, а `cycles` — максимальну кількість разів, яку послідовність може виконуватися перед зупинкою.

1. Ви повинні реалізувати `runSequence(config, cycles)` за допомогою циклу `for` або `while`, щоб ітерувати по кожній фазі протягом заданої кількості циклів.

1. Функція `runSequence` повинна:
    - Виводити в консоль `No phases found` і негайно повертати, якщо `config.phases` порожній.
    - Виводити в консоль `Faulted phase!` і достроково зупиняти симуляцію, якщо `config.fault` встановлено в `true`.
    - Виводити в консоль `Invalid phase detected`, якщо `duration <= 0`.
    - Виводити в консоль `Switching to [color] for [duration] s` для кожної дійсної фази. Замініть `[color]` і `[duration]` відповідними властивостями об’єкта фази.
    - Наприклад, `runSequence(config1, 1)` має вивести:

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. Ви повинні мати функцію з назвою `generateTimeline` з двома параметрами: `config` і `cycles`.

1. Функція `generateTimeline` повинна:
    - Записувати накопичений час, що минув після кожної фази протягом циклів, у масив, додаючи `duration` кожної фази до загального часу під час ітерації.
    - Обробляти всі пошкоджені та недійсні фази без валідації, навіть якщо `config.fault` дорівнює `true` або `duration <= 0`.
    - Повернути масив накопичених часових позначок.
    - Наприклад, `generateTimeline(config1, 1)` має повернути масив `[5, 7, 11]`.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

Ви не повинні вносити зміни в об’єкти `config1`, `config2`, `config3` та `config4`.

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

Ви повинні мати функцію з назвою `runSequence`.

```js
assert.isFunction(runSequence);
```

Ваша функція `runSequence` повинна приймати два параметри: `config` і `cycles`.

```js
const regex = __helpers.functionRegex('runSequence', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regex);
```

`runSequence(config1, 1)` має вивести в консоль `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` у такому порядку.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` має вивести в консоль `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`, `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` у такому порядку.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` має вивести в консоль `Switching to red for 3 s`, `Invalid phase detected`, `Switching to green for 6 s` у такому порядку.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` має вивести в консоль `Faulted phase!`.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` має вивести в консоль `No phases found`.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

Ви повинні мати функцію з назвою `generateTimeline`.

```js
assert.isFunction(generateTimeline);
```

Ваша функція `generateTimeline` повинна приймати два параметри: `config` і `cycles`.

```js
const regEx = __helpers.functionRegex('generateTimeline', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regEx);
```

`generateTimeline(config1, 1)` має повернути `[5, 7, 11]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` має повернути `[5, 7, 11, 16, 18, 22]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` має повернути `[3, 1, 7, 10, 8, 14]`.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` має повернути `[5, 7, 13]`.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` має повернути `[]`.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
