---
id: 587d8257367417b2b2512c7c
title: "Prüfen, ob ein Element in einem binären Suchbaum vorhanden ist"
challengeType: 1
forumTopicId: 301623
dashedName: check-if-an-element-is-present-in-a-binary-search-tree
---

# --description--

Jetzt, da wir eine allgemeine Vorstellung davon haben, was ein binärer Suchbaum ist, sprechen wir etwas ausführlicher darüber. Binäre Suchbäume bieten im Durchschnitt logarithmische Zeit für die häufigen Operationen Nachschlagen, Einfügen und Löschen, und im schlimmsten Fall lineare Zeit. Warum ist das so? Jede dieser Grundoperationen erfordert, dass wir ein Element im Baum finden (oder im Fall des Einfügens, wo es hingehört), und aufgrund der Baumstruktur verzweigen wir an jedem Elternknoten nach links oder rechts und schließen effektiv die Hälfte der verbleibenden Baumgröße aus. Das macht die Suche proportional zum Logarithmus der Anzahl der Knoten im Baum, was im Durchschnitt logarithmische Zeit für diese Operationen ergibt. Aber wie sieht es mit dem schlimmsten Fall aus? Betrachten Sie den Aufbau eines Baums aus den folgenden Werten, die von links nach rechts hinzugefügt werden: `10`, `12`, `17`, `25`. Nach unseren Regeln für einen binären Suchbaum fügen wir `12` rechts von `10` hinzu, `17` rechts davon und `25` rechts davon. Nun ähnelt unser Baum einer verketteten Liste, und um `25` zu finden, müssten wir alle Elemente linear durchlaufen. Daher lineare Zeit im schlimmsten Fall. Das Problem hier ist, dass der Baum unausgeglichen ist. Wir werden in den folgenden Aufgaben noch etwas genauer darauf eingehen, was das bedeutet.

# --instructions--

In dieser Aufgabe erstellen wir ein Hilfsmittel für unseren Baum. Schreiben Sie eine Methode `isPresent`, die einen ganzzahligen Wert als Eingabe erhält und einen Boolean-Wert zurückgibt, der angibt, ob dieser Wert im binären Suchbaum vorhanden ist oder nicht.

# --hints--

Die `BinarySearchTree`-Datenstruktur sollte existieren.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `isPresent` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isPresent == 'function';
  })()
);
```

Die `isPresent`-Methode sollte korrekt prüfen, ob Elemente, die dem Baum hinzugefügt wurden, vorhanden sind oder nicht.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    test.add(4);
    test.add(7);
    test.add(411);
    test.add(452);
    return (
      test.isPresent(452) &&
      test.isPresent(411) &&
      test.isPresent(7) &&
      !test.isPresent(100)
    );
  })()
);
```

`isPresent` sollte Fälle behandeln, in denen der Baum leer ist.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    return test.isPresent(5) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.isPresent = function (value) {
    var current = this.root
    while (current) {
      if (value === current.value) {
        return true;
      }
      current = value < current.value ? current.left : current.right;
    }
    return false;
  }
}
```
