---
id: 587d825b367417b2b2512c8c
title: Heap-Sort mit einem Min-Heap implementieren
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--

Da wir nun Elemente hinzufügen und entfernen können, sehen wir uns einige Anwendungen an, für die Heaps verwendet werden können. Heaps werden häufig verwendet, um Prioritätswarteschlangen zu implementieren, da sie immer ein Element mit dem größten oder kleinsten Wert an erster Stelle speichern. Außerdem werden sie verwendet, um einen Sortieralgorithmus namens Heap-Sort zu implementieren. Hier sehen wir, wie das funktioniert. Heap-Sort verwendet einen Min-Heap, das Gegenteil eines Max-Heaps. Ein Min-Heap speichert immer das Element mit dem kleinsten Wert an der Wurzelposition.

Heap-Sort funktioniert, indem ein unsortiertes Array genommen wird, jedes Element des Arrays in einen Min-Heap eingefügt wird und dann jedes Element aus dem Min-Heap in ein neues Array extrahiert wird. Die Struktur des Min-Heaps stellt sicher, dass das neue Array die ursprünglichen Elemente von klein nach groß enthält. Dies ist einer der effizientesten Sortieralgorithmen mit einer durchschnittlichen und schlechtesten Fall-Performance von O(nlog(n)).

# --instructions--

Implementieren Sie Heap-Sort mit einem Min-Heap. Sie können gerne Ihren Max-Heap-Code hier anpassen. Erstellen Sie ein Objekt `MinHeap` mit den Methoden `insert`, `remove` und `sort`. Die Methode `sort` soll ein Array aller Elemente im Min-Heap zurückgeben, sortiert von klein nach groß.

# --hints--

Die Datenstruktur `MinHeap` sollte existieren.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MinHeap` sollte eine Methode namens `insert` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MinHeap` sollte eine Methode namens `remove` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`MinHeap` sollte eine Methode namens `sort` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

Die Methode `sort` sollte ein Array zurückgeben, das alle zum Min-Heap hinzugefügten Elemente in sortierter Reihenfolge enthält.

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
