---
id: 587d8255367417b2b2512c73
title: Spread-Operator und Hinweise zur ES5-Set()-Integration verwenden
challengeType: 1
forumTopicId: 301720
dashedName: use-spread-and-notes-for-es5-set-integration
---

# --description--

Erinnern Sie sich an den ES6 Spread-Operator `...`?

`...` kann iterierbare Objekte in ES6 nehmen und in Arrays umwandeln.

Lassen Sie uns ein Set erstellen und die Spread-Funktion ausprobieren.

```js
var set = new Set([1,2,3]);
var setToArr = [...set]
console.log(setToArr) // returns [ 1, 2, 3 ]
```

# --instructions--

In dieser Aufgabe übergeben wir ein Set-Objekt an die `checkSet`-Funktion. Sie soll ein Array zurückgeben, das die Werte des Sets enthält.

Jetzt haben Sie erfolgreich gelernt, wie Sie das ES6 `Set()`-Objekt verwenden, gut gemacht!

# --hints--

`checkSet(new Set([1,2,3,4,5,6,7])` sollte `[1, 2, 3, 4, 5, 6, 7]` zurückgeben.

```js
var test = checkSet(new Set([1, 2, 3, 4, 5, 6, 7]));
assert.deepEqual(test, [1, 2, 3, 4, 5, 6, 7]);
```

# --seed--

## --seed-contents--

```js
function checkSet(set){
   // Only change code below this line

   // Only change code above this line
}
```

# --solutions--

```js
function checkSet(set){
return [...set];}
```
