---
id: 587d825c367417b2b2512c90
title: Breitensuche
challengeType: 1
forumTopicId: 301622
dashedName: breadth-first-search
---

# --description--

Bisher haben wir verschiedene Möglichkeiten kennengelernt, wie man Graphen darstellen kann. Und jetzt? Eine natürliche Frage ist, wie die Abstände zwischen zwei beliebigen Knoten im Graphen sind. Hier kommen <dfn>Graphdurchlauf-Algorithmen</dfn> ins Spiel.

<dfn>Durchlauf-Algorithmen</dfn> sind Algorithmen, um Knoten in einem Graphen zu durchlaufen oder zu besuchen. Eine Art von Durchlauf-Algorithmus ist der Breitensuche-Algorithmus.

Dieser Algorithmus startet bei einem Knoten und besucht alle seine Nachbarn, die eine Kante entfernt sind. Danach besucht er jeden ihrer Nachbarn und so weiter, bis alle Knoten erreicht wurden.

Eine wichtige Datenstruktur, die bei der Implementierung des Breitensuche-Algorithmus hilft, ist die Warteschlange. Dabei handelt es sich um ein Array, in das Sie Elemente an einem Ende hinzufügen und am anderen Ende entfernen können. Das ist auch bekannt als eine <dfn>FIFO</dfn>- oder <dfn>First-In-First-Out</dfn>-Datenstruktur.

Visuell sieht das so aus, wenn der Algorithmus arbeitet:
<img alt="animation showing the breadth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/breadth-first-search.gif'>

Die graue Schattierung zeigt einen Knoten, der in die Warteschlange eingefügt wird, und die schwarze Schattierung zeigt einen Knoten, der aus der Warteschlange entfernt wird. Sehen Sie, wie jedes Mal, wenn ein Knoten aus der Warteschlange entfernt wird (der Knoten wird schwarz), alle seine Nachbarn in die Warteschlange eingefügt werden (der Knoten wird grau).

Um diesen Algorithmus zu implementieren, benötigen Sie eine Graph-Datenstruktur und einen Knoten, bei dem Sie starten möchten.

Zuerst sollten Sie sich der Abstände vom Startknoten bewusst sein, also der Anzahl der Kanten, die entfernt sind. Sie sollten alle Abstände mit einer großen Zahl beginnen, wie `Infinity`. Das verhindert Zählfehler, falls ein Knoten vom Startknoten aus nicht erreichbar ist. Danach gehen Sie vom Startknoten zu seinen Nachbarn. Diese Nachbarn sind eine Kante entfernt, und an dieser Stelle sollten Sie eine Einheit Abstand zu den Abständen hinzufügen, die Sie verfolgen.

# --instructions--

Schreiben Sie eine Funktion `bfs()`, die eine Adjazenzmatrix-Graph-Datenstruktur (ein zweidimensionales Array) und eine Knotenbezeichnung root als Parameter erhält. Die Knotenbezeichnung ist einfach der ganzzahlige Wert des Knotens zwischen `0` und `n - 1`, wobei `n` die Gesamtzahl der Knoten im Graphen ist.

Ihre Funktion gibt ein JavaScript-Objekt mit Schlüssel-Wert-Paaren aus, wobei der Schlüssel der Knoten und der Wert der Abstand vom root ist. Wenn der Knoten nicht erreicht werden konnte, sollte der Abstand `Infinity` sein.

# --hints--

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` mit Startknoten `1` sollte `{0: 1, 1: 0, 2: 1, 3: 2}` zurückgeben

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 1);
assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: 2 });
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` mit Startknoten `1` sollte `{0: 1, 1: 0, 2: 1, 3: Infinity}` zurückgeben

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 0],
  [0, 0, 0, 0]
];
var results = bfs(graph, 1);
 assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: Infinity });
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` mit Startknoten `0` sollte `{0: 0, 1: 1, 2: 2, 3: 3}` zurückgeben

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 0);
 assert.deepEqual(results, { 0: 0, 1: 1, 2: 2, 3: 3 });
```

Der Eingabegraph `[[0, 1], [1, 0]]` mit Startknoten `0` sollte `{0: 0, 1: 1}` zurückgeben

```js
var graph = [
  [0, 1],
  [1, 0]
];
var results = bfs(graph, 0);
assert.deepEqual(results, { 0: 0, 1: 1 });
```

# --seed--

## --seed-contents--

```js
function bfs(graph, root) {
  var nodesLen = {};

  return nodesLen;
};

var exBFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(bfs(exBFSGraph, 3));
```

# --solutions--

```js
function bfs(graph, root) {
  var nodesLen = {};
  // Set all distances to infinity
  for (var i = 0; i < graph.length; i++) {
    nodesLen[i] = Infinity;
  }
  nodesLen[root] = 0; // ...except root node
  var queue = [root]; // Keep track of nodes to visit
  var current; // Current node traversing
  // Keep on going until no more nodes to traverse
  while (queue.length !== 0) {
    current = queue.shift();
    // Get adjacent nodes from current node
    var curConnected = graph[current]; // Get layer of edges from current
    var neighborIdx = []; // List of nodes with edges
    var idx = curConnected.indexOf(1); // Get first edge connection
    while (idx !== -1) {
      neighborIdx.push(idx); // Add to list of neighbors
      idx = curConnected.indexOf(1, idx + 1); // Keep on searching
    }
    // Loop through neighbors and get lengths
    for (var j = 0; j < neighborIdx.length; j++) {
      // Increment distance for nodes traversed
      if (nodesLen[neighborIdx[j]] === Infinity) {
        nodesLen[neighborIdx[j]] = nodesLen[current] + 1;
        queue.push(neighborIdx[j]); // Add new neighbors to queue
      }
    }
  }
  return nodesLen;
}
```
