---
id: 587d8258367417b2b2512c80
title: Löschen eines Blattknotens in einem binären Suchbaum
challengeType: 1
forumTopicId: 301637
dashedName: delete-a-leaf-node-in-a-binary-search-tree
---

# --description--

Dies ist die erste von drei Aufgaben, in denen wir eine schwierigere Operation in binären Suchbäumen implementieren: das Löschen. Das Löschen ist schwierig, weil das Entfernen von Knoten Verbindungen im Baum unterbricht. Diese Verbindungen müssen sorgfältig wiederhergestellt werden, um sicherzustellen, dass die Struktur des binären Baums erhalten bleibt. Bei einigen Löschungen bedeutet dies, dass der Baum umstrukturiert werden muss. Im Allgemeinen begegnen Sie beim Versuch, einen Knoten zu löschen, einem von drei Fällen: Blattknoten: Das Ziel zum Löschen hat keine Kinder. Ein Kind: Das Ziel zum Löschen hat nur ein Kind. Zwei Kinder: Das Ziel zum Löschen hat zwei Kindknoten. Das Entfernen eines Blattknotens ist einfach, wir entfernen ihn einfach. Das Löschen eines Knotens mit einem Kind ist ebenfalls relativ einfach, wir entfernen ihn und verbinden seinen Elternknoten mit dem Kind des gelöschten Knotens. Das Entfernen eines Knotens mit zwei Kindern ist jedoch schwieriger, da dadurch zwei Kindknoten entstehen, die wieder mit dem übergeordneten Baum verbunden werden müssen. Wie wir mit diesem Fall umgehen, sehen wir in der dritten Aufgabe. Außerdem müssen Sie bei der Behandlung des Löschens einige Randfälle beachten. Was, wenn der Baum leer ist? Was, wenn der zu löschende Knoten der Wurzelknoten ist? Was, wenn nur zwei Elemente im Baum sind? Für den Moment behandeln wir den ersten Fall, bei dem wir einen Blattknoten löschen.

# --instructions--

Erstellen Sie eine Methode in unserem binären Baum namens `remove`. Dort bauen wir die Logik für unsere Löschoperation auf. Zuerst sollten Sie eine Funktion innerhalb von remove erstellen, die den Knoten findet, den wir im aktuellen Baum löschen möchten. Wenn der Knoten nicht im Baum vorhanden ist, sollte `remove` `null` zurückgeben. Wenn der Zielknoten ein Blattknoten ohne Kinder ist, sollte die Referenz des Elternknotens auf ihn auf `null` gesetzt werden. Dadurch wird der Knoten effektiv aus dem Baum gelöscht. Dazu müssen Sie auch die Referenz auf den Elternknoten des zu löschenden Knotens verfolgen. Es ist außerdem hilfreich, eine Möglichkeit zu schaffen, die Anzahl der Kinder des Zielknotens zu verfolgen, da dies bestimmt, welcher Fall beim Löschen vorliegt. Die zweite und dritte Fälle behandeln wir in den nächsten Aufgaben. Viel Erfolg!

# --hints--

Die `BinarySearchTree`-Datenstruktur sollte existieren.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `remove` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Der Versuch, ein Element aus einem leeren Baum zu entfernen, sollte `null` zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

Der Versuch, ein Element zu entfernen, das nicht existiert, sollte `null` zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(30);
    return test.remove(100) == null;
  })()
);
```

Wenn der Wurzelknoten keine Kinder hat, sollte das Löschen ihn auf `null` setzen.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

Die `remove`-Methode sollte Blattknoten aus dem Baum entfernen.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
}
```

# --solutions--

```js
// solution required
```
