---
id: 587d8251367417b2b2512c65
title: Elemente aus einer verketteten Liste nach Index entfernen
challengeType: 1
forumTopicId: 301711
dashedName: remove-elements-from-a-linked-list-by-index
---

# --description--

Bevor wir zu einer anderen Datenstruktur übergehen, sammeln wir noch ein paar letzte Übungserfahrungen mit verketteten Listen.

Schreiben wir eine `removeAt`-Methode, die den `element` an einem gegebenen `index` entfernt. Die Methode soll `removeAt(index)` heißen. Um einen `element` an einem bestimmten `index` zu entfernen, müssen wir während des Durchlaufens der verketteten Liste eine laufende Zählung jedes Knotens führen.

Eine gängige Technik, um die Elemente einer verketteten Liste zu durchlaufen, verwendet einen <dfn>„Runner“</dfn> oder Wächter, der auf die Knoten zeigt, die Ihr Code vergleicht. In unserem Fall starten wir am `head` unserer Liste mit einer `currentIndex`-Variablen, die bei `0` beginnt. Der `currentIndex` sollte für jeden durchlaufenen Knoten um eins erhöht werden.

Genau wie bei unserer `remove(element)`-Methode, die <a href="/learn/coding-interview-prep/data-structures/remove-elements-from-a-linked-list" target="_blank" rel="noopener noreferrer nofollow">in einer vorherigen Lektion behandelt wurde</a>, müssen wir darauf achten, den Rest unserer Liste nicht zu isolieren, wenn wir den Knoten in unserer `removeAt(index)`-Methode entfernen. Wir halten unsere Knoten zusammenhängend, indem wir sicherstellen, dass der Knoten, der eine Referenz auf den entfernten Knoten hatte, eine Referenz auf den nächsten Knoten erhält.

# --instructions--

Schreiben Sie eine `removeAt(index)`-Methode, die einen Knoten an einem gegebenen `index` entfernt und zurückgibt. Die Methode soll `null` zurückgeben, wenn der angegebene `index` entweder negativ ist oder größer oder gleich der `length` der verketteten Liste.

**Hinweis:** Denken Sie daran, die `currentIndex` mitzuzählen.

# --hints--

Ihre `LinkedList`-Klasse sollte eine `removeAt`-Methode haben.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.removeAt === 'function';
  })()
);
```

Ihre `removeAt`-Methode sollte die `length` der verketteten Liste um eins verringern.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.removeAt(1);
    return test.size() === 2;
  })()
);
```

Ihre `removeAt`-Methode sollte das Element am angegebenen Index aus der verketteten Liste entfernen.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.add('bird');
    test.removeAt(1);
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"kitten","next":{"element":"bird","next":null}}}'
    );
  })()
);
```

Wenn nur ein Element in der verketteten Liste vorhanden ist, sollte Ihre `removeAt`-Methode das Element am angegebenen Index entfernen und zurückgeben sowie die Länge der verketteten Liste verringern.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    var removedItem = test.removeAt(0);
    return test.head() === null && test.size() === 0 && removedItem === 'cat';
  })()
);
```

Ihre `removeAt`-Methode sollte das Element des entfernten Knotens zurückgeben.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.removeAt(1) === 'dog';
  })()
);
```

Ihre `removeAt`-Methode sollte `null` zurückgeben und die verkettete Liste darf sich nicht ändern, wenn der angegebene Index kleiner als `0` ist.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(-1);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

Ihre `removeAt`-Methode sollte `null` zurückgeben und die verkettete Liste darf sich nicht ändern, wenn der angegebene Index größer oder gleich der `length` der Liste ist.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(3);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
      head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function (element) {
    this.element = element;
    this.next = null;
  };

  this.size = function () {
    return length;
  };

  this.head = function () {
    return head;
  };

  this.add = function (element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.removeAt = function (index) {
    var currentNode = head;
    var previous = head;
    var count = 0;
    if (index >= length || index < 0) {
      return null;
    }
    if (index === 0) {
      var removed = head.element;
      head = currentNode.next;
    } else {
      while (count < index) {
        previous = currentNode;
        currentNode = currentNode.next;
        count++;
      }
      var removed = previous.next.element;
      previous.next = currentNode.next;
    }
    length--;
    return removed;
  };
}
```
