---
id: 587d825a367417b2b2512c87
title: Eine doppelt verkettete Liste erstellen
challengeType: 1
forumTopicId: 301626
dashedName: create-a-doubly-linked-list
---

# --description--

Alle verketteten Listen, die wir bisher erstellt haben, sind einfach verkettete Listen. Hier erstellen wir eine <dfn>doppelt verkettete Liste</dfn>. Wie der Name schon sagt, haben Knoten in einer doppelt verketteten Liste Referenzen sowohl auf den nächsten als auch auf den vorherigen Knoten in der Liste.

Das ermöglicht es uns, die Liste in beide Richtungen zu durchlaufen, erfordert aber auch mehr Speicher, da jeder Knoten eine zusätzliche Referenz auf den vorherigen Knoten in der Liste enthalten muss.

# --instructions--

Wir haben ein `Node`-Objekt bereitgestellt und mit unserem `DoublyLinkedList` begonnen. Fügen Sie unserer doppelt verketteten Liste zwei Methoden namens `add` und `remove` hinzu. Die `add`-Methode soll das gegebene Element zur Liste hinzufügen, während die `remove`-Methode alle Vorkommen eines gegebenen Elements aus der Liste entfernen soll.

Achten Sie darauf, alle möglichen Randfälle beim Schreiben dieser Methoden zu berücksichtigen, wie zum Beispiel das Löschen des ersten oder letzten Elements. Außerdem sollte das Entfernen eines Elements aus einer leeren Liste `null` zurückgeben.

# --hints--

Die `DoublyLinkedList`-Datenstruktur sollte existieren.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

Die `DoublyLinkedList` sollte eine Methode namens `add` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.add == undefined) {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

Die `DoublyLinkedList` sollte eine Methode namens `remove` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.remove == undefined) {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Das Entfernen eines Elements aus einer leeren Liste sollte `null` zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.remove(100) == null;
  })()
);
```

Die `add`-Methode sollte Elemente zur Liste hinzufügen.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(5);
    test.add(6);
    test.add(723);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '56723';
  })()
);
```

Jeder Knoten sollte den vorherigen Knoten verfolgen.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(50);
    test.add(68);
    test.add(73);
    var result = [];
    var curr = test.tail;
    while (curr) {
      result.push(curr.data);
      curr = curr.prev;
    }
    return result.join('') == '736850';
  })()
);
```

Das erste Element sollte aus der Liste entfernt werden können.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(25);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '3560';
  })()
);
```

Das letzte Element sollte aus der Liste entfernt werden können.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(60);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '2535';
  })()
);
```


# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
