---
id: 587d8257367417b2b2512c7d
title: Minimum- und Maximumhöhe eines binären Suchbaums finden
challengeType: 1
forumTopicId: 301641
dashedName: find-the-minimum-and-maximum-height-of-a-binary-search-tree
---

# --description--

In der letzten Aufgabe haben wir ein Szenario beschrieben, in dem ein Baum unausgeglichen sein kann. Um das Konzept der Balance zu verstehen, betrachten wir eine weitere Eigenschaft eines Baums: die Höhe. Die Höhe in einem Baum stellt die Entfernung vom Wurzelknoten zu einem beliebigen Blattknoten dar. Unterschiedliche Pfade in einer stark verzweigten Baumstruktur können unterschiedliche Höhen haben, aber für einen gegebenen Baum gibt es ein Minimum und ein Maximum der Höhe. Wenn der Baum ausgeglichen ist, unterscheiden sich diese Werte höchstens um eins. Das bedeutet, dass in einem ausgeglichenen Baum alle Blattknoten auf derselben Ebene liegen oder, falls sie nicht auf derselben Ebene sind, höchstens eine Ebene auseinanderliegen.

Die Eigenschaft der Balance ist für Bäume wichtig, weil sie die Effizienz der Baumoperationen bestimmt. Wie wir in der letzten Aufgabe erklärt haben, treten im schlimmsten Fall Zeitkomplexitäten bei stark unausgeglichenen Bäumen auf. Selbstbalancierende Bäume werden häufig verwendet, um dieses Problem bei Bäumen mit dynamischen Datensätzen zu lösen. Bekannte Beispiele hierfür sind AVL-Bäume, Rot-Schwarz-Bäume und B-Bäume. Diese Bäume enthalten alle zusätzliche interne Logik, die den Baum neu ausbalanciert, wenn Einfügungen oder Löschungen einen Zustand der Unausgeglichenheit erzeugen.

**Hinweis:** Eine ähnliche Eigenschaft wie die Höhe ist die Tiefe, die angibt, wie weit ein gegebener Knoten vom Wurzelknoten entfernt ist.

# --instructions--

Schreiben Sie zwei Methoden für unseren binären Baum: `findMinHeight` und `findMaxHeight`. Diese Methoden sollen jeweils einen ganzzahligen Wert für die minimale und maximale Höhe innerhalb eines gegebenen binären Baums zurückgeben. Wenn der Knoten leer ist, weisen wir ihm eine Höhe von `-1` zu (das ist der Basisfall). Fügen Sie schließlich eine dritte Methode `isBalanced` hinzu, die `true` oder `false` zurückgibt, je nachdem, ob der Baum ausgeglichen ist oder nicht. Sie können die beiden zuvor geschriebenen Methoden verwenden, um dies zu bestimmen.

# --hints--

Die `BinarySearchTree`-Datenstruktur sollte existieren.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `findMinHeight` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMinHeight == 'function';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `findMaxHeight` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMaxHeight == 'function';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `isBalanced` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isBalanced == 'function';
  })()
);
```

Die `findMinHeight`-Methode sollte die minimale Höhe des Baums zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMinHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMinHeight() == 1;
  })()
);
```

Die `findMaxHeight`-Methode sollte die maximale Höhe des Baums zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMaxHeight() == 5;
  })()
);
```

Ein leerer Baum sollte eine Höhe von `-1` zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    return test.findMaxHeight() == -1;
  })()
);
```

Die `isBalanced`-Methode sollte `false` zurückgeben, wenn der Baum ein unausgeglichener binärer Suchbaum ist.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.isBalanced() === false;
  })()
);
```

Die `isBalanced`-Methode sollte `true` zurückgeben, wenn der Baum ein ausgeglichener binärer Suchbaum ist.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(10);
    test.add(3);
    test.add(22);
    test.add(1);
    test.add(4);
    test.add(17);
    test.add(32);
    return test.isBalanced() === true;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
  this.findMinHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMinHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMinHeight(root.left) + 1;
    }
    const lHeight = this.findMinHeight(root.left);
    const rHeight = this.findMinHeight(root.right);
    return Math.min(lHeight, rHeight) + 1;
  };
  this.findMaxHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMaxHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMaxHeight(root.left) + 1;
    }
    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    return Math.max(lHeight, rHeight) + 1;
  };
  this.isBalanced = function(root = this.root) {
    if (root === null) {
      return true;
    }

    if (root.left === null && root.right === null) {
      return true;
    }

    if (root.left === null) {
      return this.findMaxHeight(root.right) <= 0;
    }

    if (root.right === null) {
      return this.findMaxHeight(root.left) <= 0;
    }

    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    if (Math.abs(lHeight - rHeight) > 1) {
      return false;
    }
    return this.isBalanced(root.left) && this.isBalanced(root.right);
  };
}
```
