---
id: 587d8257367417b2b2512c7b
title: Ein neues Element zu einem binären Suchbaum hinzufügen
challengeType: 1
forumTopicId: 301618
dashedName: add-a-new-element-to-a-binary-search-tree
---

# --description--

Diese Reihe von Aufgaben führt in die Datenstruktur Baum ein. Bäume sind eine wichtige und vielseitige Datenstruktur in der Informatik. Natürlich stammt ihr Name daher, dass sie bei der Visualisierung sehr ähnlich aussehen wie die Bäume, die wir aus der Natur kennen. Eine Baum-Datenstruktur beginnt mit einem Knoten, der typischerweise als Wurzel bezeichnet wird, und verzweigt sich von dort aus zu weiteren Knoten, von denen jeder weitere Kindknoten haben kann, und so weiter. Die Datenstruktur wird üblicherweise mit dem Wurzelknoten oben visualisiert; man kann sich das vorstellen wie einen natürlichen Baum, der auf den Kopf gestellt wurde.

Zuerst beschreiben wir einige gebräuchliche Begriffe, die wir bei Bäumen antreffen werden. Der Wurzelknoten ist die Spitze des Baums. Datenpunkte im Baum werden Knoten genannt. Knoten mit Zweigen, die zu anderen Knoten führen, werden als Elternknoten des Knotens bezeichnet, zu dem der Zweig führt (dem Kind). Weitere kompliziertere familiäre Begriffe gelten, wie Sie es erwarten würden. Ein Teilbaum bezeichnet alle Nachfahren eines bestimmten Knotens, Zweige können als Kanten bezeichnet werden, und Blattknoten sind Knoten am Ende des Baums, die keine Kinder haben. Schließlich ist zu beachten, dass Bäume von Natur aus rekursive Datenstrukturen sind. Das heißt, alle Kinder eines Knotens sind Elternknoten ihres eigenen Teilbaums, und so weiter. Die rekursive Natur von Bäumen ist wichtig zu verstehen, wenn man Algorithmen für gängige Baumoperationen entwirft.

Zunächst besprechen wir einen bestimmten Typ von Baum, den binären Baum. Tatsächlich sprechen wir über einen bestimmten binären Baum, einen binären Suchbaum. Lassen Sie uns beschreiben, was das bedeutet. Während die Baum-Datenstruktur an einem einzelnen Knoten beliebig viele Zweige haben kann, kann ein binärer Baum nur zwei Zweige für jeden Knoten haben. Außerdem ist ein binärer Suchbaum bezüglich der Kind-Teilbäume geordnet, sodass der Wert jedes Knotens im linken Teilbaum kleiner oder gleich dem Wert des Elternknotens ist, und der Wert jedes Knotens im rechten Teilbaum größer oder gleich dem Wert des Elternknotens ist. Es ist sehr hilfreich, diese Beziehung zu visualisieren, um sie besser zu verstehen:

<div style='width: 100%; display: flex; justify-content: center; align-items: center;'><img alt="ein Beispiel für einen binären Suchbaum" style='width: 100%; max-width: 350px; background-color: var(--gray-05);' src='https://user-images.githubusercontent.com/18563015/32136009-1e665d98-bbd6-11e7-9133-63184f9f9182.png'></div>

Diese geordnete Beziehung ist nun sehr leicht zu erkennen. Beachten Sie, dass jeder Wert links von 8, dem Wurzelknoten, kleiner als 8 ist, und jeder Wert rechts größer als 8 ist. Beachten Sie auch, dass diese Beziehung für jeden der Teilbäume gilt. Zum Beispiel ist das erste linke Kind ein Teilbaum. 3 ist der Elternknoten, und es hat genau zwei Kindknoten – nach den Regeln für binäre Suchbäume wissen wir ohne hinzuschauen, dass das linke Kind dieses Knotens (und alle seine Kinder) kleiner als 3 sein wird, und das rechte Kind (und alle seine Kinder) größer als 3 (aber auch kleiner als der Wert der Wurzelstruktur), und so weiter.

Binäre Suchbäume sind sehr häufige und nützliche Datenstrukturen, weil sie im Durchschnitt logarithmische Zeit für mehrere gängige Operationen wie Nachschlagen, Einfügen und Löschen bieten.

# --instructions--

Wir fangen einfach an. Wir haben hier das Grundgerüst einer binären Suchbaum-Datenstruktur definiert sowie eine Funktion, um Knoten für unseren Baum zu erstellen. Beachten Sie, dass jeder Knoten einen linken und rechten Wert haben kann. Diese werden Kind-Teilbäumen zugewiesen, falls sie existieren. In unserem binären Suchbaum werden Sie eine Methode erstellen, um neue Werte zum Baum hinzuzufügen. Die Methode sollte `add` heißen und einen ganzzahligen Wert akzeptieren, der zum Baum hinzugefügt wird. Achten Sie darauf, die Invariante eines binären Suchbaums beizubehalten: Der Wert in jedem linken Kind sollte kleiner oder gleich dem Wert des Elternknotens sein, und der Wert in jedem rechten Kind sollte größer oder gleich dem Wert des Elternknotens sein. Hier machen wir es so, dass unser Baum keine doppelten Werte enthalten kann. Wenn wir versuchen, einen Wert hinzuzufügen, der bereits existiert, sollte die Methode `null` zurückgeben. Andernfalls, wenn das Hinzufügen erfolgreich ist, sollte `undefined` zurückgegeben werden.

**Hinweis:** Bäume sind von Natur aus rekursive Datenstrukturen!

# --hints--

Die `BinarySearchTree`-Datenstruktur sollte existieren.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `add` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

Die Add-Methode sollte Elemente gemäß den Regeln des binären Suchbaums hinzufügen.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    const expectedResult = [1, 4, 7, 8, 34, 45, 73, 87];
    const result = test.inOrder();
    return expectedResult.toString() === result.toString();
  })()
);
```

Das Hinzufügen eines Elements, das bereits existiert, sollte `null` zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    return test.add(4) == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    inOrder() {
      if (!this.root) {
        return null;
      }
      var result = new Array();
      function traverseInOrder(node) {
        node.left && traverseInOrder(node.left);
        result.push(node.value);
        node.right && traverseInOrder(node.right);
      }
      traverseInOrder(this.root);
      return result;
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.add = function(element) {
    let current = this.root;
    if (!current) {
      this.root = new Node(element);
      return;
    } else {
      const searchTree = function(current) {
        if (current.value > element) {
          if (current.left) {
            return searchTree(current.left);
          } else {
            current.left = new Node(element);
            return;
          }
        } else if (current.value < element) {
          if (current.right) {
            return searchTree(current.right);
          } else {
            current.right = new Node(element);
            return;
          }
        } else {
          return null;
        }
      };
      return searchTree(current);
    }
  };
}
```
