---
id: 587d8255367417b2b2512c74
title: Erstellen Sie eine Prioritätswarteschlangen-Klasse
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--

In dieser Aufgabe erstellen Sie eine Prioritätswarteschlange. Eine Prioritätswarteschlange ist eine spezielle Art von Warteschlange, bei der Elemente zusätzliche Informationen enthalten können, die ihre Priorität angeben. Dies kann einfach durch eine ganze Zahl dargestellt werden. Die Priorität eines Elements überschreibt die Platzierungsreihenfolge bei der Bestimmung der Reihenfolge, in der Elemente aus der Warteschlange entfernt werden. Wenn ein Element mit höherer Priorität nach Elementen mit niedrigerer Priorität eingefügt wird, wird das Element mit der höheren Priorität vor allen anderen entfernt.

Stellen Sie sich zum Beispiel vor, wir haben eine Prioritätswarteschlange mit drei Elementen:

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Hier stellt der zweite Wert (eine ganze Zahl) die Priorität des Elements dar. Wenn wir `['human', 1]` mit einer Priorität von `1` einfügen (angenommen, niedrigere Prioritäten haben Vorrang), wäre es das erste Element, das entfernt wird. Die Sammlung sähe dann so aus:

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Wir haben begonnen, eine `PriorityQueue` im Code-Editor zu schreiben. Sie müssen eine `enqueue`-Methode zum Hinzufügen von Elementen mit Priorität, eine `dequeue`-Methode zum Entfernen und Zurückgeben von Elementen, eine `size`-Methode zum Zurückgeben der Anzahl der Elemente in der Warteschlange, eine `front`-Methode zum Zurückgeben des Elements an der Vorderseite der Warteschlange und schließlich eine `isEmpty`-Methode hinzufügen, die `true` zurückgibt, wenn die Warteschlange leer ist, oder `false`, wenn sie nicht leer ist.

Die `enqueue` sollte Elemente im oben gezeigten Format akzeptieren (`['human', 1]`), wobei `1` die Priorität darstellt. `dequeue` und `front` sollten nur den Namen des Elements zurückgeben, nicht dessen Priorität.

# --hints--

Ihre `PriorityQueue`-Klasse sollte eine `enqueue`-Methode haben.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

Ihre `PriorityQueue`-Klasse sollte eine `dequeue`-Methode haben.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

Ihre `PriorityQueue`-Klasse sollte eine `size`-Methode haben.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

Ihre `PriorityQueue`-Klasse sollte eine `front`-Methode haben.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

Ihre `PriorityQueue`-Klasse sollte eine `isEmpty`-Methode haben.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Ihre `PriorityQueue`-Klasse sollte die aktuelle Anzahl der Elemente mithilfe der `size`-Methode korrekt verfolgen, während Elemente eingefügt und entfernt werden.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

Die `front`-Methode sollte das korrekte Element an der Vorderseite der Warteschlange zurückgeben, während Elemente eingefügt und entfernt werden.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

Die `isEmpty`-Methode sollte `true` zurückgeben, wenn die Warteschlange leer ist.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

Die Prioritätswarteschlange sollte Elemente mit höherer Priorität vor Elementen mit niedrigerer Priorität zurückgeben und ansonsten Elemente in der Reihenfolge „first-in-first-out“ zurückgeben.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```
