---
id: 587d8255367417b2b2512c72
title: Verwendung von .has und .size in einem ES6-Set
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

Let's look at the .has and .size methods available on the ES6 Set object.

Erstelle zunächst ein ES6-Set

```js
var set = new Set([1,2,3]);
```

Die Methode .has prüft, ob der Wert in der Menge enthalten ist.

```js
var hasTwo = set.has(2);
```

Die Methode .size gibt eine ganze Zahl zurück, die die Größe des Sets angibt

```js
var howBig = set.size;
```

# --instructions--

In dieser Übung werden wir ein Array und einen Wert an die Funktion checkSet() übergeben. Deine Funktion sollte ein ES6-Set aus dem Array-Argument erstellen. Herausfinden, ob die Menge das Wertargument enthält. Finde die Größe der Menge. Und diese beiden Werte in einem Array zurückgeben.

# --hints--

`checkSet([4, 5, 6], 3)` sollte [ false, 3 ] zurückgeben

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
