---
id: 587d8255367417b2b2512c72
title: .has und .size bei einem ES6 Set verwenden
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

Schauen wir uns die Methoden .has und .size an, die beim ES6 Set-Objekt verfügbar sind.

Erstellen Sie zuerst ein ES6 Set

```js
var set = new Set([1,2,3]);
```

Die Methode .has prüft, ob der Wert im Set enthalten ist.

```js
var hasTwo = set.has(2);
```

Die Methode .size gibt eine ganze Zahl zurück, die die Größe des Sets darstellt.

```js
var howBig = set.size;
```

# --instructions--

In dieser Aufgabe übergeben wir ein Array und einen Wert an die Funktion checkSet(). Ihre Funktion soll aus dem Array-Argument ein ES6 Set erstellen. Finden Sie heraus, ob das Set den Wert-Argument enthält. Ermitteln Sie die Größe des Sets. Und geben Sie diese beiden Werte in einem Array zurück.

# --hints--

`checkSet([4, 5, 6], 3)` sollte [ false, 3 ] zurückgeben

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
