---
id: 587d8255367417b2b2512c75
title: Eine zirkuläre Queue erstellen
challengeType: 1
forumTopicId: 301625
dashedName: create-a-circular-queue
---

# --description--

In dieser Aufgabe erstellen Sie eine zirkuläre Queue. Eine zirkuläre Queue ist eine Queue, die am Ende einer Sammlung schreibt und dann beginnt, sich am Anfang der Sammlung selbst zu überschreiben. Diese Art von Datenstruktur ist in bestimmten Situationen nützlich. Zum Beispiel kann eine zirkuläre Queue für Streaming-Medien verwendet werden. Sobald die Queue voll ist, überschreiben neue Mediendaten alte Daten.

Eine gute Möglichkeit, dieses Konzept zu veranschaulichen, ist ein Array der Länge `5`:

```js
[null, null, null, null, null]
 ^Read @ 0
 ^Write @ 0
```

Hier befinden sich der Lese- und der Schreibzeiger beide an Position `0`. Nun erhält die Queue 3 neue Einträge `a`, `b` und `c`. Unsere Queue sieht jetzt so aus:

```js
[a, b, c, null, null]
 ^Read @ 0
          ^Write @ 3
```

Während der Lesezeiger liest, kann er Werte entfernen oder behalten:

```js
[null, null, null, null, null]
                   ^Read @ 3
                   ^Write @ 3
```

Jetzt schreiben wir die Werte `d`, `e` und `f` in die Queue. Sobald der Schreibzeiger das Ende des Arrays erreicht, springt er zurück zum Anfang:

```js
[f, null, null, d, e]
                ^Read @ 3
    ^Write @ 1
```

Dieser Ansatz benötigt eine konstante Menge an Speicher, ermöglicht aber die Verarbeitung von Dateien mit viel größerer Größe.

# --instructions--

In dieser Aufgabe implementieren wir eine zirkuläre Queue. Die zirkuläre Queue sollte `enqueue` und `dequeue` Methoden bereitstellen, mit denen Sie aus der Queue lesen und in die Queue schreiben können. Die Klasse selbst sollte auch ein ganzzahliges Argument akzeptieren, mit dem Sie die Größe der Queue bei der Erstellung angeben können. Wir haben die Startversion dieser Klasse für Sie im Code-Editor vorbereitet.

Wenn Sie Elemente in die Queue einfügen, sollte sich der Schreibzeiger vorwärts bewegen und zum Anfang zurückspringen, sobald er das Ende der Queue erreicht. Die `enqueue` Methode sollte das eingefügte Element zurückgeben, wenn sie erfolgreich ist; andernfalls gibt sie `null` zurück.

Ebenso sollte sich der Lesezeiger vorwärts bewegen, wenn Sie Elemente aus der Queue entfernen. Wenn Sie ein Element aus der Queue entfernen, sollte dieses Element zurückgegeben werden. Wenn Sie kein Element entfernen können, sollten Sie `null` zurückgeben.

Der Schreibzeiger darf nicht am Lesezeiger vorbeigehen (unsere Klasse lässt Sie keine Daten überschreiben, die Sie noch nicht gelesen haben), und der Lesezeiger darf nicht über Daten hinausgehen, die Sie geschrieben haben.

# --hints--

Die `enqueue` Methode sollte Elemente zur zirkulären Queue hinzufügen.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Sie sollten keine Elemente hinter dem Lesezeiger einfügen.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    test.enqueue(13);
    test.enqueue(25);
    test.enqueue(59);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Die `dequeue` Methode sollte Elemente aus der Queue entfernen.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 && test.dequeue() === 32 && test.dequeue() === 591
    );
  })()
);
```

Nachdem ein Element entfernt wurde, sollte seine Position in der Queue auf `null` zurückgesetzt werden.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(672);
    test.dequeue();
    test.dequeue();
    var print = test.print();
    return print[0] === null && print[1] === null && print[2] === 672;
  })()
);
```

Der Versuch, hinter dem Schreibzeiger zu entfernen, sollte `null` zurückgeben und den Schreibzeiger nicht vorwärts bewegen.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 &&
      test.dequeue() === 32 &&
      test.dequeue() === 591 &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.enqueue(100) === 100 &&
      test.dequeue() === 100
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
class CircularQueue {
  constructor(size) {

    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line

    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line

    // Only change code above this line
  }
}
```

# --solutions--

```js
class CircularQueue {
  constructor(size) {
    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line
    console.log(this.write, this.max);
    if (this.queue[this.write] === null) {
      this.queue[this.write++] = item;

      if (this.write > this.max) {
        this.write = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line
    if (this.queue[this.read] !== null) {
      let item = this.queue[this.read];
      this.queue[this.read++] = null;
      if (this.read > this.max) {
        this.read = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }
}
```
