---
id: 587d8250367417b2b2512c60
title: Eine Queue-Klasse erstellen
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--

Wie Stacks sind Queues eine Sammlung von Elementen. Aber im Gegensatz zu Stacks folgen Queues dem FIFO-Prinzip (First-In First-Out). Elemente, die zu einer Queue hinzugefügt werden, werden ans Ende der Queue geschoben, und nur das Element am Anfang der Queue darf entfernt werden.

Wir könnten ein Array verwenden, um eine Queue darzustellen, aber genau wie bei Stacks wollen wir die Kontrolle über unsere Queues einschränken.

Die zwei Hauptmethoden einer Queue-Klasse sind die enqueue- und die dequeue-Methode. Die enqueue-Methode fügt ein Element ans Ende der Queue hinzu, und die dequeue-Methode entfernt und gibt das Element am Anfang der Queue zurück. Weitere nützliche Methoden sind die front-, size- und isEmpty-Methoden.

# --instructions--

Schreiben Sie eine `enqueue`-Methode, die ein Element ans Ende der Queue schiebt, eine `dequeue`-Methode, die das vorderste Element entfernt und zurückgibt, eine `front`-Methode, die uns das vorderste Element zeigt, eine `size`-Methode, die die Länge anzeigt, und eine `isEmpty`-Methode, um zu prüfen, ob die Queue leer ist.

# --hints--

Ihre `Queue`-Klasse sollte eine `enqueue`-Methode haben.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

Ihre `Queue`-Klasse sollte eine `dequeue`-Methode haben.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

Ihre `Queue`-Klasse sollte eine `front`-Methode haben.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

Ihre `Queue`-Klasse sollte eine `size`-Methode haben.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

Ihre `Queue`-Klasse sollte eine `isEmpty`-Methode haben.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Die `dequeue`-Methode sollte das vorderste Element der Queue entfernen und zurückgeben.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

Die `front`-Methode sollte den Wert des vordersten Elements der Queue zurückgeben.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

Die `size`-Methode sollte die Länge der Queue zurückgeben.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

Die `isEmpty`-Methode sollte `false` zurückgeben, wenn Elemente in der Queue sind.

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```
