---
id: 587d825a367417b2b2512c8a
title: Ein Element in einen Max-Heap einfügen
challengeType: 1
forumTopicId: 301703
dashedName: insert-an-element-into-a-max-heap
---

# --description--

Jetzt wenden wir uns einer anderen Baum-Datenstruktur zu, dem binären Heap. Ein binärer Heap ist ein teilweise geordneter binärer Baum, der die Heap-Eigenschaft erfüllt. Die Heap-Eigenschaft legt eine Beziehung zwischen Eltern- und Kindknoten fest. Sie können einen Max-Heap haben, bei dem alle Elternknoten größer oder gleich ihren Kindknoten sind, oder einen Min-Heap, bei dem das Gegenteil gilt. Binäre Heaps sind außerdem vollständige binäre Bäume. Das bedeutet, dass alle Ebenen des Baums vollständig gefüllt sind und wenn die letzte Ebene nur teilweise gefüllt ist, dann von links nach rechts.

Obwohl binäre Heaps als Baumstrukturen mit Knoten implementiert werden können, die Links- und Rechts-Referenzen enthalten, erlaubt uns die teilweise Ordnung gemäß der Heap-Eigenschaft, den Heap mit einem Array darzustellen. Die Eltern-Kind-Beziehung ist dabei entscheidend, und mit einfacher Arithmetik können wir die Kinder eines beliebigen Elternknotens und den Elternknoten eines beliebigen Kindknotens berechnen.

Betrachten Sie zum Beispiel diese Array-Darstellung eines binären Min-Heaps:

```js
[ 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Der Wurzelknoten ist das erste Element, `6`. Seine Kinder sind `22` und `30`. Wenn wir uns die Beziehung zwischen den Array-Indizes dieser Werte ansehen, dann sind für den Index `i` die Kinder `2 * i + 1` und `2 * i + 2`. Ebenso ist das Element am Index `0` der Elternknoten dieser beiden Kinder an den Indizes `1` und `2`. Allgemeiner gesagt, können wir den Elternknoten eines Knotens an einem beliebigen Index mit folgender Formel finden: `Math.floor((i - 1) / 2)`. Diese Muster gelten, egal wie groß der binäre Baum wird. Schließlich können wir eine kleine Anpassung vornehmen, um diese Arithmetik noch einfacher zu machen, indem wir das erste Element im Array überspringen. Dadurch entsteht folgende Beziehung für ein Element an einem gegebenen Index `i`:

Beispiel für eine Array-Darstellung:

```js
[ null, 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Das linke Kind eines Elements: `i * 2`

Das rechte Kind eines Elements: `i * 2 + 1`

Der Elternknoten eines Elements: `Math.floor(i / 2)`

Wenn Sie die Mathematik verstanden haben, ist die Verwendung einer Array-Darstellung sehr nützlich, weil sich die Positionen der Knoten mit dieser Arithmetik schnell bestimmen lassen und der Speicherverbrauch geringer ist, da keine Referenzen zu Kindknoten gepflegt werden müssen.

# --instructions--

Anleitung: Hier erstellen wir einen Max-Heap. Beginnen Sie damit, einfach eine `insert`-Methode zu erstellen, die Elemente zu unserem Heap hinzufügt. Während der Einfügung ist es wichtig, stets die Heap-Eigenschaft zu erhalten. Für einen Max-Heap bedeutet das, dass das Wurzelelement immer den größten Wert im Baum haben sollte und alle Elternknoten größer als ihre Kinder sein müssen. Bei einer Array-Implementierung eines Heaps wird dies typischerweise in drei Schritten erreicht:

<ol>
  <li>Fügen Sie das neue Element ans Ende des Arrays an.</li>
  <li>Wenn das Element größer als sein Elternknoten ist, tauschen Sie sie.</li>
  <li>Fahren Sie mit dem Tauschen fort, bis das neue Element entweder kleiner als sein Elternknoten ist oder Sie die Wurzel des Baums erreichen.</li>
</ol>

Fügen Sie abschließend eine `print`-Methode hinzu, die ein Array aller Elemente zurückgibt, die dem Heap hinzugefügt wurden.

# --hints--

Die `MaxHeap`-Datenstruktur sollte existieren.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` sollte eine Methode namens `insert` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` sollte eine Methode namens `print` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

Die `insert`-Methode sollte Elemente gemäß der Max-Heap-Eigenschaft hinzufügen.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    test.insert(50);
    test.insert(100);
    test.insert(700);
    test.insert(32);
    test.insert(51);
    test.insert(800);
    const result = test.print();
    const solution = JSON.stringify([null,800,51,700,32,50,100]);
    const solutionWithoutNull = JSON.stringify([800,51,700,32,50,100]);

    return (result.length == 6) ? (JSON.stringify(result) == solutionWithoutNull) : (JSON.stringify(result) == solution);
  })()
);
```

# --seed--

## --seed-contents--

```js
var MaxHeap = function() {
  // Only change code below this line
    
  // Only change code above this line
};
```

# --solutions--

```js
var MaxHeap = function() {
    // Only change code below this line
    this.heap = [];
    this.parent = index => {
      return Math.floor((index - 1) / 2);
    }
    this.insert = element => {
      this.heap.push(element);
      this.heapifyUp(this.heap.length - 1);
    }
    this.heapifyUp = index => {
      let currentIndex = index,
      parentIndex = this.parent(currentIndex);
      while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
        this.swap(currentIndex, parentIndex);
        currentIndex = parentIndex;
        parentIndex = this.parent(parentIndex);
      }
    }
    this.swap = (index1, index2) => {
      [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
    }
    this.print = () => {
      return this.heap;
    }
    // Only change code above this line
};
```
