---
id: 587d8257367417b2b2512c7e
title: Tiefensuche in einem binären Suchbaum verwenden
challengeType: 1
forumTopicId: 301719
dashedName: use-depth-first-search-in-a-binary-search-tree
---

# --description--

Wir wissen, wie man in einem binären Suchbaum nach einem bestimmten Wert sucht. Aber was, wenn wir einfach den gesamten Baum erkunden wollen? Oder was, wenn wir keinen geordneten Baum haben und einfach nach einem Wert suchen müssen? Hier stellen wir einige Baumdurchlaufmethoden vor, mit denen man Baum-Datenstrukturen erkunden kann. Zuerst kommt die Tiefensuche. Bei der Tiefensuche wird ein gegebener Teilbaum so tief wie möglich erkundet, bevor die Suche in einem anderen Teilbaum fortgesetzt wird. Es gibt drei Möglichkeiten, dies zu tun: In Reihenfolge: Beginnen Sie die Suche am linksesten Knoten und enden Sie am rechtsesten Knoten. Preorder: Erkunden Sie alle Wurzeln vor den Blättern. Postorder: Erkunden Sie alle Blätter vor den Wurzeln. Wie Sie sich denken können, wählen Sie je nach Art der Daten, die Ihr Baum speichert, und wonach Sie suchen, unterschiedliche Suchmethoden. Für einen binären Suchbaum liefert ein Inorder-Durchlauf die Knoten in sortierter Reihenfolge.

# --instructions--

Hier erstellen wir diese drei Suchmethoden für unseren binären Suchbaum. Tiefensuche ist eine von Natur aus rekursive Operation, die so lange weitere Teilbäume erkundet, wie Kindknoten vorhanden sind. Sobald Sie dieses Grundkonzept verstanden haben, können Sie einfach die Reihenfolge ändern, in der Sie die Knoten und Teilbäume erkunden, um eine der drei oben genannten Sucharten zu erzeugen. Zum Beispiel möchten wir bei der Postorder-Suche bis zu einem Blattknoten rekursiv vorgehen, bevor wir beginnen, die Knoten selbst zurückzugeben, während wir bei der Preorder-Suche zuerst die Knoten zurückgeben und dann weiter den Baum hinab rekursiv durchlaufen. Definieren Sie die Methoden `inorder`, `preorder` und `postorder` in unserem Baum. Jede dieser Methoden sollte ein Array von Elementen zurückgeben, die den Baumdurchlauf darstellen. Stellen Sie sicher, dass Sie die ganzzahligen Werte an jedem Knoten im Array zurückgeben, nicht die Knoten selbst. Geben Sie schließlich `null` zurück, wenn der Baum leer ist.

# --hints--

Die Datenstruktur `BinarySearchTree` sollte existieren.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `inorder` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.inorder == 'function';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `preorder` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.preorder == 'function';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `postorder` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.postorder == 'function';
  })()
);
```

Die Methode `inorder` sollte ein Array der Knotewerte zurückgeben, die sich aus einem Inorder-Durchlauf ergeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.inorder().join('') == '012345678910';
  })()
);
```

Die Methode `preorder` sollte ein Array der Knotewerte zurückgeben, die sich aus einem Preorder-Durchlauf ergeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.preorder().join('') == '710325469810';
  })()
);
```

Die Methode `postorder` sollte ein Array der Knotewerte zurückgeben, die sich aus einem Postorder-Durchlauf ergeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.postorder().join('') == '024653181097';
  })()
);
```

Die Methode `inorder` sollte `null` für einen leeren Baum zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    return test.inorder() == null;
  })()
);
```

Die Methode `preorder` sollte `null` für einen leeren Baum zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    return test.preorder() == null;
  })()
);
```

Die Methode `postorder` sollte `null` für einen leeren Baum zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    return test.postorder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.result = [];

  this.inorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.inorder(node.left);
    this.result.push(node.value);
    if (node.right) this.inorder(node.right);
    return this.result;
  };
  this.preorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    this.result.push(node.value);
    if (node.left) this.preorder(node.left);
    if (node.right) this.preorder(node.right);
    return this.result;
  };
  this.postorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.postorder(node.left);
    if (node.right) this.postorder(node.right);
    this.result.push(node.value);

    return this.result;
  };
}
```
