---
id: 587d8254367417b2b2512c70
title: Erstellen und Hinzufügen zu Sets in ES6
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--

Nachdem Sie ES5 durchgearbeitet haben, werden Sie nun etwas Ähnliches in ES6 machen. Das wird erheblich einfacher sein. ES6 enthält eine eingebaute Datenstruktur `Set`, sodass viele der Operationen, die Sie zuvor manuell geschrieben haben, jetzt für Sie enthalten sind. Schauen wir uns das an:

Um ein neues leeres Set zu erstellen:

```js
var set = new Set();
```

Sie können ein Set mit einem Wert erstellen:

```js
var set = new Set(1);
```

Sie können ein Set mit einem Array erstellen:

```js
var set = new Set([1, 2, 3]);
```

Sobald Sie ein Set erstellt haben, können Sie die gewünschten Werte mit der `add`-Methode hinzufügen:

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

Zur Erinnerung: Ein Set ist eine Datenstruktur, die keine doppelten Werte enthalten kann:

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set contains [1, 2, 3] only
```

# --instructions--

Geben Sie für diese Aufgabe ein Set mit den folgenden Werten zurück: `1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--

Ihr `Set` sollte nur die Werte `1, 2, 3, Taco, Cat, Awesome` enthalten.

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```
