---
id: 587d8258367417b2b2512c7f
title: Breitensuche in einem binären Suchbaum verwenden
challengeType: 1
forumTopicId: 301718
dashedName: use-breadth-first-search-in-a-binary-search-tree
---

# --description--

Hier stellen wir eine weitere Baumdurchlaufmethode vor: die Breitensuche. Im Gegensatz zu den Tiefensuche-Methoden aus der letzten Aufgabe erkundet die Breitensuche alle Knoten auf einer bestimmten Ebene innerhalb eines Baums, bevor sie zur nächsten Ebene übergeht. Typischerweise werden Warteschlangen als Hilfsdatenstrukturen beim Entwurf von Breitensuche-Algorithmen verwendet.

Bei dieser Methode beginnen wir damit, den Wurzelknoten in eine Warteschlange einzufügen. Dann starten wir eine Schleife, in der wir das erste Element aus der Warteschlange entfernen, es einem neuen Array hinzufügen und anschließend beide Kind-Teilbäume untersuchen. Wenn die Kinder nicht null sind, werden sie jeweils in die Warteschlange eingefügt. Dieser Vorgang wird fortgesetzt, bis die Warteschlange leer ist.

# --instructions--

Erstellen Sie eine Breitensuche-Methode in unserem Baum namens `levelOrder`. Diese Methode soll ein Array zurückgeben, das die Werte aller Baumknoten enthält, die in Breitensuche-Reihenfolge erkundet wurden. Achten Sie darauf, die Werte im Array zurückzugeben, nicht die Knoten selbst. Eine Ebene soll von links nach rechts durchlaufen werden. Schreiben Sie anschließend eine ähnliche Methode namens `reverseLevelOrder`, die dieselbe Suche durchführt, jedoch in umgekehrter Richtung (von rechts nach links) auf jeder Ebene.

# --hints--

Die `BinarySearchTree`-Datenstruktur sollte existieren.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `levelOrder` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.levelOrder == 'function';
  })()
);
```

Der binäre Suchbaum sollte eine Methode namens `reverseLevelOrder` haben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.reverseLevelOrder == 'function';
  })()
);
```

Die `levelOrder`-Methode sollte ein Array der Baumknotenwerte zurückgeben, die in Ebenenreihenfolge erkundet wurden.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.levelOrder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.levelOrder().join('') == '719038102546';
  })()
);
```

Die `reverseLevelOrder`-Methode sollte ein Array der Baumknotenwerte zurückgeben, die in umgekehrter Ebenenreihenfolge erkundet wurden.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.reverseLevelOrder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.reverseLevelOrder().join('') == '791108305264';
  })()
);
```

Die `levelOrder`-Methode sollte `null` für einen leeren Baum zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.levelOrder !== 'function') {
      return false;
    }
    return test.levelOrder() == null;
  })()
);
```

Die `reverseLevelOrder`-Methode sollte `null` für einen leeren Baum zurückgeben.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.reverseLevelOrder !== 'function') {
      return false;
    }
    return test.reverseLevelOrder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  this.levelOrder = (root = this.root) => {
    if(!root) return null;
    let queue = [root];
    let results = [];
    while(queue.length > 0) {
      let node = queue.shift();
      results.push(node.value);
      if(node.left) queue.push(node.left);
      if(node.right) queue.push(node.right);
    }
    return results;
  }

  this.reverseLevelOrder = (root = this.root) => {
    if(!root) return null;
    let queue = [root];
    let results = [] ;
    while ( queue.length > 0) {
      let node = queue.shift();
      results.push(node.value);
      if(node.right) queue.push(node.right);
      if(node.left ) queue.push(node.left);
    }
    return results;
  }
  // Only change code above this line
}
```
