---
id: 587d825b367417b2b2512c8e
title: Eine Hash-Tabelle erstellen
challengeType: 1
forumTopicId: 301627
dashedName: create-a-hash-table
---

# --description--

In dieser Aufgabe lernen wir Hash-Tabellen kennen. Eine Hash-Tabelle wird verwendet, um assoziative Arrays oder Abbildungen von Schlüssel-Wert-Paaren zu implementieren, wie die Objekte und Maps, die wir gerade studiert haben. Ein JavaScript-Objekt könnte zum Beispiel als Hash-Tabelle implementiert sein (die tatsächliche Implementierung hängt von der Umgebung ab, in der es ausgeführt wird). Die Funktionsweise einer Hash-Tabelle besteht darin, dass sie einen Schlüssel als Eingabe nimmt und diesen Schlüssel auf deterministische Weise in einen numerischen Wert hasht. Dieser numerische Wert wird dann als tatsächlicher Schlüssel verwendet, unter dem der zugehörige Wert gespeichert wird. Wenn Sie dann versuchen, denselben Schlüssel erneut zuzugreifen, verarbeitet die Hash-Funktion den Schlüssel, gibt dasselbe numerische Ergebnis zurück, das dann verwendet wird, um den zugehörigen Wert nachzuschlagen. Dies ermöglicht im Durchschnitt eine sehr effiziente Suchzeit von O(1).

Hash-Tabellen können als Arrays implementiert werden, wobei Hash-Funktionen Array-Indizes innerhalb eines bestimmten Bereichs erzeugen. Bei dieser Methode ist die Wahl der Array-Größe wichtig, ebenso wie die Hash-Funktion. Was passiert zum Beispiel, wenn die Hash-Funktion denselben Wert für zwei verschiedene Schlüssel erzeugt? Das nennt man eine Kollision. Eine Möglichkeit, mit Kollisionen umzugehen, besteht darin, beide Schlüssel-Wert-Paare an diesem Index zu speichern. Beim Nachschlagen eines der beiden müssten Sie dann den Bucket der Elemente durchlaufen, um den gesuchten Schlüssel zu finden. Eine gute Hash-Funktion minimiert Kollisionen, um eine effiziente Suchzeit zu gewährleisten.

Hier werden wir uns nicht mit den Details des Hashings oder der Implementierung von Hash-Tabellen beschäftigen, sondern nur ein allgemeines Verständnis dafür entwickeln, wie sie funktionieren.

# --instructions--

Lassen Sie uns die Grundfunktionalität einer Hash-Tabelle erstellen. Wir haben für Sie eine naive Hash-Funktion vorbereitet, die Sie verwenden können. Sie können einen String-Wert an die Funktion `hash` übergeben, und sie gibt einen gehashten Wert zurück, den Sie als Schlüssel zum Speichern verwenden können. Speichern Sie Elemente basierend auf diesem gehashten Wert im `this.collection`-Objekt. Erstellen Sie diese drei Methoden: `add`, `remove` und `lookup`. Die erste sollte ein Schlüssel-Wert-Paar akzeptieren, das zur Hash-Tabelle hinzugefügt wird. Die zweite sollte ein Schlüssel-Wert-Paar entfernen, wenn ein Schlüssel übergeben wird. Die dritte sollte einen Schlüssel akzeptieren und den zugehörigen Wert zurückgeben oder `null`, wenn der Schlüssel nicht vorhanden ist.

Achten Sie darauf, Ihren Code so zu schreiben, dass Kollisionen berücksichtigt werden!

**Hinweis:** Die Tests der `remove`-Methode werden erst bestehen, wenn die `add`- und `lookup`-Methoden korrekt implementiert sind.

# --hints--

Die `hash`-Funktion sollte gültig sein.

```js
let calls = 0;
const ourHash = string => {
  calls++;
  let hashCode = 0;
  for (let i = 0; i < string.length; i++) {
    hashCode += string.charCodeAt(i);
  }
  return hashCode;
};

assert.strictEqual(hash('yek'),ourHash('yek'));

assert.strictEqual(hash('key'),ourHash('key'));
assert.strictEqual(hash('key1'),ourHash('key1'));
assert.strictEqual(hash('key2'),ourHash('key2'));
assert.strictEqual(hash('key3'),ourHash('key3'));

assert.strictEqual(hash('1key'),ourHash('1key'));
assert.strictEqual(hash('ke1y'),ourHash('ke1y'));
assert.strictEqual(hash('altKey'),ourHash('altKey'));

assert.strictEqual(called,calls); 
```

Die `HashTable`-Datenstruktur sollte existieren.

```js
assert.isDefined(HashTable); 
let newHashTable = new HashTable(); 
assert.isObject(newHashTable); 
```

Die `HashTable` sollte eine `add`-Methode haben.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.add); 
```

Die `HashTable` sollte eine `lookup`-Methode haben.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.lookup); 
```

Die `HashTable` sollte eine `remove`-Methode haben.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.remove); 
```

Die `add`-Methode sollte Schlüssel-Wert-Paare hinzufügen und die `lookup`-Methode sollte die Werte zurückgeben, die mit einem gegebenen Schlüssel assoziiert sind.

```js
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
assert.strictEqual(newHashTable.lookup('key'),'value'); 
```

Die `remove`-Methode sollte einen Schlüssel als Eingabe akzeptieren und das zugehörige Schlüssel-Wert-Paar entfernen.

```js
let hashValue = hash('key');
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
newHashTable.remove('key'); 
assert.notProperty(newHashTable.collection,hashValue); 
```

Die `remove`-Methode sollte nur das korrekte Schlüssel-Wert-Paar entfernen.

```js
let newHashTable = new HashTable(); 
let hashValue = hash('key');
newHashTable.add('key','value');
newHashTable.add('yek','value');
newHashTable.add('altKey','value');
newHashTable.remove('yek');

assert.notExists(newHashTable.lookup('yek'));
assert.exists(newHashTable.lookup('altKey'));
assert.exists(newHashTable.lookup('key'));

newHashTable.remove('key');

assert.notProperty(newHashTable.collection,hashValue); 
assert.exists(newHashTable.lookup('altKey'));

```

Elemente sollten mit der Hash-Funktion hinzugefügt werden.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1', 'value1')
newHashTable.add('key2', 'value2');
newHashTable.add('key3', 'value3');
assert.strictEqual(called,3); 
```

Die Hash-Tabelle sollte Kollisionen behandeln.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1','value1');
newHashTable.add('1key', 'value2'); 
newHashTable.add('ke1y', 'value3'); 
assert.strictEqual(newHashTable.lookup('key1'),'value1');
assert.strictEqual(newHashTable.lookup('1key'),'value2');
assert.strictEqual(newHashTable.lookup('ke1y'),'value3');
```

# --seed--

## --seed-contents--

```js
var called = 0;
var hash = string => {
  called++;
  var hashed = 0;
  for (var i = 0; i < string.length; i++) {
    hashed += string.charCodeAt(i);
  }
  return hashed;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
var called = 0;
var hash = (string) => {
  called++;
  var hash = 0;
  for (var i = 0; i < string.length; i++) { hash += string.charCodeAt(i); }
  return hash;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line

  this.add = function(key, val) {
    var theHash = hash(key);
    if (!this.collection.hasOwnProperty(theHash)) {
      this.collection[theHash] = {};
    }
    this.collection[theHash][key] = val;
  }

  this.remove = function(key) {
    var theHash = hash(key);
    var hashedObj = this.collection[theHash];
    if (hashedObj.hasOwnProperty(key)) {
      delete hashedObj[key];
    }
    if (!Object.keys(hashedObj).length) {
      delete this.collection[theHash];
    }
  }

  this.lookup = function(key) {
    var theHash = hash(key);
    if (this.collection.hasOwnProperty(theHash)) {
      return this.collection[theHash][key];
    }
    return null
  }
  // Only change code above this line
};
```
