---
id: 587d8251367417b2b2512c64
title: Suche innerhalb einer verketteten Liste
challengeType: 1
forumTopicId: 301715
dashedName: search-within-a-linked-list
---

# --description--

Fügen wir unserer Klasse für verkettete Listen noch einige nützliche Methoden hinzu. Wäre es nicht praktisch, wenn wir feststellen könnten, ob unsere Liste leer ist oder nicht, so wie bei unseren `Stack`- und `Queue`-Klassen?

Wir sollten auch in der Lage sein, bestimmte Elemente in unserer verketteten Liste zu finden. Das Durchlaufen von Datenstrukturen ist etwas, worin Sie viel Übung brauchen! Erstellen wir eine `indexOf`-Methode, die ein `element` als Argument nimmt und den `index` dieses Elements in der verketteten Liste zurückgibt. Wenn das Element in der verketteten Liste nicht gefunden wird, geben Sie `-1` zurück.

Implementieren wir außerdem eine Methode, die das Gegenteil macht: eine `elementAt`-Methode, die ein `index` als Argument nimmt und das `element` am angegebenen `index` zurückgibt. Wenn kein `element` gefunden wird, geben Sie `undefined` zurück.

# --instructions--

Schreiben Sie eine `isEmpty`-Methode, die überprüft, ob die verkettete Liste leer ist, eine `indexOf`-Methode, die den `index` eines gegebenen Elements zurückgibt, und eine `elementAt`, die ein `element` an einem gegebenen `index` zurückgibt.

# --hints--

Ihre `LinkedList`-Klasse sollte eine `isEmpty`-Methode haben.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.isEmpty === 'function';
  })()
);
```

Ihre `isEmpty`-Methode sollte `false` zurückgeben, wenn mindestens ein Element in der verketteten Liste vorhanden ist.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.isEmpty() === false;
  })()
);
```

Ihre `isEmpty`-Methode sollte `true` zurückgeben, wenn keine Elemente in der verketteten Liste vorhanden sind.

```js
assert(
  (function () {
    var test = new LinkedList();
    return test.isEmpty() === true;
  })()
);
```

Ihre `LinkedList`-Klasse sollte eine `indexOf`-Methode haben.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.indexOf === 'function';
  })()
);
```

Ihre `indexOf`-Methode sollte den Index eines gegebenen Elements in der verketteten Liste zurückgeben.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('cat') === 0;
  })()
);
```

Ihre `indexOf`-Methode sollte `-1` zurückgeben, wenn das gegebene Element in der verketteten Liste nicht gefunden wird.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('pony') === -1;
  })()
);
```

Ihre `LinkedList`-Klasse sollte eine `elementAt`-Methode haben.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.elementAt === 'function';
  })()
);
```

Ihre `elementAt`-Methode sollte das Element zurückgeben, das an einem gegebenen Index in der verketteten Liste gefunden wird.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(1) === 'dog';
  })()
);
```

Ihre `elementAt`-Methode sollte `undefined` zurückgeben, wenn an einem gegebenen Index in der verketteten Liste kein Element gefunden wird.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(5) === undefined;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
      head = currentNode.next;
    } else {
      while(currentNode.element !== element) {
        previousNode = currentNode;
        currentNode = currentNode.next;
      }

      previousNode.next = currentNode.next;
    }

    length --;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
        head = currentNode.next;
    } else {
        while(currentNode.element !== element) {
            previousNode = currentNode;
            currentNode = currentNode.next;
        }

        previousNode.next = currentNode.next;
    }

    length --;
  };

  this.indexOf = function(element) {
    if (head === null) return -1

    let current = head;
    let index = 0;

    while (current.element !== element && current.next !== null) {
      current = current.next;
      index++
    }

    if (current.element !== element && current.next === null) {
      return -1
    }

    return index;
  }

  this.elementAt = function(index) {
    if (head === null) return undefined;

    let current = head;
    let currentIndex = 0;

    while (currentIndex !== index && current.next !== null) {
      current = current.next;
      currentIndex++
    }

    if (currentIndex !== index && current.next === null) {
      return undefined;
    }

    return current.element;
  }

  this.isEmpty = function() {
    return length === 0;
  }
}
```
