---
id: 587d8251367417b2b2512c63
title: Elemente aus einer verketteten Liste entfernen
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--

Die nächste wichtige Methode, die jede Implementierung einer verketteten Liste benötigt, ist eine `remove`-Methode. Diese Methode sollte das Element, das wir entfernen möchten, als Argument nehmen und dann die Liste durchsuchen, um den Knoten zu finden und zu entfernen, der dieses Element enthält.

Wann immer wir einen Knoten aus einer verketteten Liste entfernen, ist es wichtig, dass wir dabei nicht versehentlich den Rest der Liste „verwaisen“ lassen. Denken Sie daran, dass die `next`-Eigenschaft jedes Knotens auf den Knoten zeigt, der ihm in der Liste folgt. Wenn wir zum Beispiel das mittlere Element entfernen, wollen wir sicherstellen, dass die `next`-Eigenschaft des vorherigen Knotens dieses Elements auf die `next`-Eigenschaft des mittleren Elements zeigt (also auf den nächsten Knoten in der Liste!).

Das klingt vielleicht verwirrend, also kehren wir zum Beispiel der Conga-Linie zurück, damit wir ein gutes konzeptuelles Modell haben. Stellen Sie sich vor, Sie stehen in einer Conga-Linie, und die Person direkt vor Ihnen verlässt die Linie. Die Person, die gerade die Linie verlassen hat, hat ihre Hände nicht mehr an jemandem in der Linie – und Sie haben Ihre Hände nicht mehr an der Person, die gegangen ist. Sie treten einen Schritt vor und legen Ihre Hände auf die nächste Person, die Sie sehen.

Wenn das Element, das wir entfernen möchten, das `head`-Element ist, weisen wir die `head` dem zweiten Knoten der verketteten Liste zu.

# --instructions--

Schreiben Sie eine `remove`-Methode, die ein Element nimmt und es aus der verketteten Liste entfernt.

**Hinweis:** Die `length` der Liste sollte jedes Mal um eins verringert werden, wenn ein Element aus der verketteten Liste entfernt wird.

# --hints--

Ihre `LinkedList`-Klasse sollte eine `remove`-Methode haben.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

Ihre `remove`-Methode sollte die `head` auf den zweiten Knoten neu zuweisen, wenn der erste Knoten entfernt wird.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

Ihre `remove`-Methode sollte die `length` der verketteten Liste für jeden entfernten Knoten um eins verringern.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

Ihre `remove`-Methode sollte die Referenz des vorherigen Knotens des entfernten Knotens auf die `next`-Referenz des entfernten Knotens neu zuweisen.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

Ihre `remove`-Methode sollte die verkettete Liste nicht verändern, wenn das Element nicht in der verketteten Liste existiert.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```
