---
id: 587d8252367417b2b2512c67
title: Elemente an einem bestimmten Index in einer verketteten Liste hinzufügen
challengeType: 1
forumTopicId: 301619
dashedName: add-elements-at-a-specific-index-in-a-linked-list
---

# --description--

Lassen Sie uns eine Methode addAt(index,element) erstellen, die ein Element an einem bestimmten Index hinzufügt. Genau wie beim Entfernen von Elementen an einem bestimmten Index müssen wir den aktuellen Index verfolgen, während wir die verkettete Liste durchlaufen. Wenn der aktuelle Index mit dem angegebenen Index übereinstimmt, müssen wir die next-Eigenschaft des vorherigen Knotens neu zuweisen, damit sie auf den neu hinzugefügten Knoten verweist. Und der neue Knoten sollte auf den nächsten Knoten im aktuellen Index verweisen. Zurück zum Beispiel mit der Polonaise: Eine neue Person möchte sich in die Reihe einreihen, aber sie möchte in der Mitte dazukommen. Sie stehen in der Mitte der Reihe, also nehmen Sie Ihre Hände von der Person vor Ihnen weg. Die neue Person kommt herüber und legt ihre Hände auf die Person, die Sie zuvor berührt haben, und Sie haben nun Ihre Hände auf der neuen Person.

# --instructions--

Erstellen Sie eine `addAt(index,element)`-Methode, die ein Element an einem bestimmten Index hinzufügt. Geben Sie false zurück, wenn ein Element nicht hinzugefügt werden konnte. **Hinweis:** Denken Sie daran zu prüfen, ob der angegebene Index negativ ist oder länger als die Länge der verketteten Liste.

# --hints--

Ihre `addAt`-Methode sollte `head` auf den neuen Knoten neu zuweisen, wenn der angegebene Index 0 ist.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'fish');
    return test.head().element === 'fish' && test.head().next.element === 'cat';
  })()
);
```

Ihre `addAt`-Methode sollte die Länge der verketteten Liste um eins erhöhen für jeden neuen Knoten, der zur verketteten Liste hinzugefügt wird.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'cat');
    return test.size() === 3;
  })()
);
```

Ihre `addAt`-Methode sollte `false` zurückgeben, wenn ein Knoten nicht hinzugefügt werden konnte.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.addAt(4, 'cat') === false;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element) {
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function() {
    return head;
  };

  this.add = function(element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if (head === null){
        head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };
  this.addAt = function (index, element) {
    if (index > length || index < 0) {
      return false;
    }
    var newNode = new Node(element);
    var currentNode = head;
    if (index === 0) {
      head = newNode;
    } else {
      var previousNode = null;
      var i = 0;
      while (currentNode && i < index) {
        previousNode = currentNode;
        currentNode = currentNode.next;
        i++;
      }
      previousNode.next = newNode;
    }
    newNode.next = currentNode;
    length++;
  }
}
```
