---
id: 587d825b367417b2b2512c8b
title: Ein Element aus einem Max-Heap entfernen
challengeType: 1
forumTopicId: 301710
dashedName: remove-an-element-from-a-max-heap
---

# --description--

Da wir nun Elemente zu unserem Heap hinzufügen können, sehen wir uns an, wie wir Elemente entfernen können. Das Entfernen und Einfügen von Elementen erfordert ähnliche Logik. In einem Max-Heap möchten Sie normalerweise den größten Wert entfernen, was einfach bedeutet, ihn aus der Wurzel unseres Baums zu extrahieren. Dies bricht die Heap-Eigenschaft unseres Baums, daher müssen wir sie auf irgendeine Weise wiederherstellen. Typischerweise wird dies bei einem Max-Heap folgendermaßen gemacht:

<ol>
  <li>Bewegen Sie das letzte Element im Heap an die Wurzelposition.</li>
  <li>Wenn eines der Kinder der Wurzel größer ist als sie, tauschen Sie die Wurzel mit dem Kind mit dem größeren Wert.</li>
  <li>Fahren Sie mit dem Tauschen fort, bis der Elternknoten größer ist als beide Kinder oder Sie die letzte Ebene im Baum erreichen.</li>
</ol>

# --instructions--

Anleitung: Fügen Sie unserem Max-Heap eine Methode namens `remove` hinzu. Diese Methode soll den größten Wert zurückgeben, der zu unserem Max-Heap hinzugefügt wurde, und ihn aus dem Heap entfernen. Außerdem soll sie den Heap so neu anordnen, dass die Heap-Eigenschaft erhalten bleibt. Nachdem ein Element entfernt wurde, soll das nächstgrößere verbleibende Element im Heap die Wurzel werden.

# --hints--

Die `MaxHeap`-Datenstruktur sollte existieren.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` sollte eine Methode namens `print` haben.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`MaxHeap` sollte eine Methode namens `insert` haben.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` sollte eine Methode namens `remove` haben.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Die `remove`-Methode sollte das größte Element aus dem Max-Heap entfernen und dabei die Max-Heap-Eigenschaft erhalten.

```js
function isHeap(arr, i, n) {
    if( arr[i] < arr[2 * i + 1] || arr[i] < arr[2 * i + 2] ){
        return false;
    }
    if (i > (n - 1) / 2) {
        return true;
    }
    if (isHeap(arr, 2 * i + 1, n) && isHeap(arr, 2 * i + 2, n)) {
        return true;
    }
    return false;
}


assert(
  (function () {
    let test = false;

    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    
    let max = Infinity;
    const [result, vals] = [[], [9, 3, 5, 2, 15, 3, 7, 12, 7, 10, 90]];

    vals.forEach((val) => test.insert(val));

    for (let i = 0; i < vals.length; i++) {
      const curHeap = test.print();
      const arr = curHeap[0] === null ? curHeap.slice(1) : curHeap;
      
      if (!isHeap(arr, 0, arr.length - 1)) {
        return false;
      }

      const removed = test.remove();
      if (!vals.includes(removed)) return false;
      if (removed > max) return false
      max = removed;
      result.push(removed);
    }
    
    for (let i = 0; i < vals.length; i++) {
      if (!result.includes(vals[i])) {
        return false;
      }
    }
    
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
const MaxHeap = function () {
  this.heap = [];
  this.parent = index => {
    return Math.floor((index - 1) / 2);
  }
  this.insert = element => {
    this.heap.push(element);
    this.heapifyUp(this.heap.length - 1);
  }
  this.heapifyUp = index => {
    let currentIndex = index,
    parentIndex = this.parent(currentIndex);
    while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
      this.swap(currentIndex, parentIndex);
      currentIndex = parentIndex;
      parentIndex = this.parent(parentIndex);
    }
  }
  this.swap = (index1, index2) => {
    [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
  }
  this.print = () => {
    return this.heap;
  }
  // Only change code below this line

  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
