---
id: 587d8251367417b2b2512c62
title: Eine Klasse für eine verkettete Liste erstellen
challengeType: 1
forumTopicId: 301628
dashedName: create-a-linked-list-class
---

# --description--

Lassen Sie uns eine `linked list` Klasse erstellen. Jede verkettete Liste sollte mit einigen grundlegenden Eigenschaften beginnen: einer `head` (dem ersten Element in Ihrer Liste) und einer `length` (Anzahl der Elemente in Ihrer Liste). Manchmal sehen Sie Implementierungen von verketteten Listen, die eine `tail` für das letzte Element der Liste enthalten, aber vorerst bleiben wir nur bei diesen beiden. Jedes Mal, wenn wir ein Element zur verketteten Liste hinzufügen, sollte unsere `length` Eigenschaft um eins erhöht werden.

Wir möchten eine Möglichkeit haben, Elemente zu unserer verketteten Liste hinzuzufügen, also ist die erste Methode, die wir erstellen wollen, die `add` Methode.

Wenn unsere Liste leer ist, ist das Hinzufügen eines Elements zur verketteten Liste ganz einfach: Wir verpacken dieses Element einfach in eine `Node` Klasse und weisen diesen Knoten der `head` unserer verketteten Liste zu.

Aber was ist, wenn unsere Liste bereits ein oder mehrere Mitglieder hat? Wie fügen wir der Liste ein Element hinzu? Denken Sie daran, dass jeder Knoten in einer verketteten Liste eine `next` Eigenschaft hat. Um einen Knoten zur Liste hinzuzufügen, finden Sie den letzten Knoten in der Liste und zeigen Sie mit der `next` Eigenschaft dieses letzten Knotens auf unseren neuen Knoten. (Hinweis: Sie wissen, dass Sie das Ende einer verketteten Liste erreicht haben, wenn die `next` Eigenschaft eines Knotens `null` ist.)

# --instructions--

Schreiben Sie eine add-Methode, die den ersten Knoten, den Sie zur verketteten Liste hinzufügen, der `head` zuweist; danach sollte bei jedem Hinzufügen eines Knotens jeder Knoten durch die `next` Eigenschaft des vorherigen Knotens referenziert werden.

Hinweis

Die `length` Ihrer Liste sollte jedes Mal um eins erhöht werden, wenn ein Element zur verketteten Liste hinzugefügt wird.

# --hints--

Ihre `LinkedList` Klasse sollte eine `add` Methode haben.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.add === 'function';
  })()
);
```

Ihre `LinkedList` Klasse sollte `head` dem ersten hinzugefügten Knoten zuweisen.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    return test.head().element === 'cat';
  })()
);
```

Der vorherige `node` in Ihrer `LinkedList` Klasse sollte eine Referenz auf den neu erstellten Knoten haben.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('fish');
    return test.head().next.element === 'dog' && test.head().next.next.element === 'fish';
  })()
);
```

Die `size` Ihrer `LinkedList` Klasse sollte der Anzahl der Knoten in der verketteten Liste entsprechen.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.size() === 2;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() { 
  var length = 0; 
  var head = null; 

  var Node = function(element){
    this.element = element; 
    this.next = null; 
  }; 

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line
    if (head == null) {
      head = new Node(element);
    } 
    else {
      let currentNode = head;
      while (currentNode.next != null) {
        // currentNode.next will be last node of linked list after loop
        currentNode = currentNode.next;
      }
      currentNode.next = new Node(element);
    }
    length++;
    // Only change code above this line
  };
}
```
