---
id: 587d8250367417b2b2512c5f
title: Eine Stack-Klasse erstellen
challengeType: 1
forumTopicId: 301633
dashedName: create-a-stack-class
---

# --description--

Im letzten Abschnitt haben wir darüber gesprochen, was ein Stack ist und wie wir ein Array verwenden können, um einen Stack darzustellen. In diesem Abschnitt werden wir unsere eigene Stack-Klasse erstellen. Obwohl Sie Arrays verwenden können, um Stacks zu erstellen, ist es manchmal besser, die Kontrolle über unsere Stacks einzuschränken. Abgesehen von der `push` und `pop` Methode haben Stacks weitere nützliche Methoden. Fügen wir unserer Stack-Klasse eine `peek`, `isEmpty` und `clear` Methode hinzu.

# --instructions--

Schreiben Sie eine `push` Methode, die ein Element oben auf den Stack legt, eine `pop` Methode, die das Element oben auf dem Stack entfernt und zurückgibt, eine `peek` Methode, die das oberste Element im Stack betrachtet, eine `isEmpty` Methode, die überprüft, ob der Stack leer ist, und eine `clear` Methode, die alle Elemente aus dem Stack entfernt. Normalerweise haben Stacks diese Methode nicht, aber wir haben eine `print` Hilfsmethode hinzugefügt, die die Sammlung in der Konsole ausgibt.

# --hints--

Ihre `Stack` Klasse sollte eine `push` Methode haben.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.push === 'function';
  })()
);
```

Ihre `Stack` Klasse sollte eine `pop` Methode haben.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.pop === 'function';
  })()
);
```

Ihre `Stack` Klasse sollte eine `peek` Methode haben.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.peek === 'function';
  })()
);
```

Ihre `Stack` Klasse sollte eine `isEmpty` Methode haben.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.isEmpty === 'function';
  })()
);
```

Ihre `Stack` Klasse sollte eine `clear` Methode haben.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.clear === 'function';
  })()
);
```

Die `peek` Methode sollte das oberste Element des Stacks zurückgeben.

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.peek() === 'CS50' && test.peek() === 'CS50';
  })()
);
```

Die `pop` Methode sollte das oberste Element des Stacks entfernen und zurückgeben.

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.pop() === 'CS50' && test.pop() === 'CS61';
  })()
);
```

Die `isEmpty` Methode sollte true zurückgeben, wenn ein Stack keine Elemente enthält.

```js
assert(
  (function () {
    var test = new Stack();
    return test.isEmpty();
  })()
);
```

Die `clear` Methode sollte alle Elemente aus dem Stack entfernen.

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    test.clear();
    return test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Stack() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
class Stack {
  constructor() {
    this.collection = [];
  }
  print() {
    console.log(this.collection);
  }
  push(val) {
    this.collection.push(val);
  }
  pop() {
    return this.collection.pop();
  }
  peek() {
    return this.collection[this.collection.length - 1];
  }
  isEmpty() {
    return this.collection.length === 0;
  }
  clear() {
    return (this.collection.length = 0);
  }
}
```
