---
id: 587d825d367417b2b2512c96
title: Tiefensuche
challengeType: 1
forumTopicId: 301640
dashedName: depth-first-search
---

# --description--

Ähnlich wie bei der <dfn>Breitensuche</dfn> lernen wir hier einen weiteren Algorithmus zur Graphdurchquerung kennen, die sogenannte <dfn>Tiefensuche</dfn>.

Während die Breitensuche inkrementell Kantenlängen vom Startknoten aus durchsucht, folgt die <dfn>Tiefensuche</dfn> zunächst einem Pfad von Kanten so weit wie möglich nach unten.

Sobald sie das Ende eines Pfads erreicht, geht die Suche zum letzten Knoten mit einem noch nicht besuchten Kantenpfad zurück und setzt die Suche fort.

Die folgende Animation zeigt, wie der Algorithmus funktioniert. Der Algorithmus beginnt mit dem obersten Knoten und besucht die Knoten in der nummerierten Reihenfolge.

<img alt="animation showing the depth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/depth-first-search.gif'>

Beachten Sie, dass im Gegensatz zur Breitensuche bei jedem besuchten Knoten nicht alle Nachbarn besucht werden. Stattdessen wird zuerst einer seiner Nachbarn besucht und diesem Pfad so lange gefolgt, bis keine weiteren Knoten auf diesem Pfad mehr zu besuchen sind.

Um diesen Algorithmus zu implementieren, sollten Sie einen Stack verwenden. Ein Stack ist ein Array, bei dem das zuletzt hinzugefügte Element als erstes entfernt wird. Dies wird auch als <dfn>Last-In-First-Out</dfn>-Datenstruktur bezeichnet. Ein Stack ist bei Tiefensuche-Algorithmen hilfreich, weil wir beim Hinzufügen von Nachbarn zum Stack die zuletzt hinzugefügten Nachbarn zuerst besuchen und vom Stack entfernen wollen.

Eine einfache Ausgabe dieses Algorithmus ist eine Liste von Knoten, die von einem gegebenen Knoten aus erreichbar sind. Daher sollten Sie auch die besuchten Knoten im Auge behalten.

# --instructions--

Schreiben Sie eine Funktion `dfs()`, die eine ungerichtete Adjazenzmatrix `graph` und eine Knotenbezeichnung `root` als Parameter erhält. Die Knotenbezeichnung ist einfach der numerische Wert des Knotens zwischen `0` und `n - 1`, wobei `n` die Gesamtzahl der Knoten im Graphen ist.

Ihre Funktion soll ein Array aller von `root` aus erreichbaren Knoten ausgeben.

# --hints--

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` mit Startknoten `1` sollte ein Array mit `0`, `1`, `2` und `3` zurückgeben.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })(),
  [0, 1, 2, 3]
);
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` mit Startknoten `3` sollte ein Array mit `3`, `2`, `1` und `0` zurückgeben.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3, 2, 1, 0]
);
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` mit Startknoten `1` sollte ein Array mit vier Elementen zurückgeben.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })().length === 4
);
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` mit Startknoten `3` sollte ein Array mit `3` zurückgeben.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3]
);
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` mit Startknoten `3` sollte ein Array mit einem Element zurückgeben.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })().length === 1
);
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` mit Startknoten `3` sollte ein Array mit `2` und `3` zurückgeben.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [2, 3]
);
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` mit Startknoten `3` sollte ein Array mit zwei Elementen zurückgeben.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })().length === 2
);
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` mit Startknoten `0` sollte ein Array mit `0` und `1` zurückgeben.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })(),
  [0, 1]
);
```

Der Eingabegraph `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` mit Startknoten `0` sollte ein Array mit zwei Elementen zurückgeben.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })().length === 2
);
```

# --seed--

## --seed-contents--

```js
function dfs(graph, root) {

}

var exDFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(dfs(exDFSGraph, 3));
```

# --solutions--

```js
function dfs(graph, root) {
    var stack = [];
    var tempV;
    var visited = [];
    var tempVNeighbors = [];
    stack.push(root);
    while (stack.length > 0) {
        tempV = stack.pop();
        if (visited.indexOf(tempV) == -1) {
            visited.push(tempV);
            tempVNeighbors = graph[tempV];
            for (var i = 0; i < tempVNeighbors.length; i++) {
                if (tempVNeighbors[i] == 1) {
                    stack.push(i);
                }
            }
        }
    }
    return visited;
}
```
