---
id: 587d825b367417b2b2512c8d
title: Erstellen einer ES6-JavaScript-Map
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

Die neue Version von JavaScript stellt uns ein eingebautes Map-Objekt zur Verfügung, das viele der Funktionen bietet, die wir in der letzten Aufgabe selbst geschrieben haben. Dieses Map-Objekt ist zwar ähnlich wie normale JavaScript-Objekte, bietet aber einige nützliche Funktionen, die normale Objekte nicht haben. Zum Beispiel verfolgt eine ES6-Map die Einfügereihenfolge der hinzugefügten Elemente. Hier ist eine ausführlichere Übersicht über seine Methoden:

- `.has(key)` gibt true oder false zurück, je nachdem, ob ein Schlüssel vorhanden ist
- `.get(key)` gibt den Wert zurück, der mit einem Schlüssel verknüpft ist
- `.set(key, value)` legt ein neues Schlüssel-Wert-Paar fest
- `.delete(key)` entfernt ein Schlüssel-Wert-Paar
- `.clear()` entfernt alle Schlüssel-Wert-Paare
- `.keys()` gibt ein neues Map-Iterator-Objekt zurück, das alle Schlüssel in Einfügereihenfolge enthält
- `.values()` gibt ein neues Map-Iterator-Objekt zurück, das alle Werte in Einfügereihenfolge enthält
- `.entries()` gibt ein neues Map-Iterator-Objekt zurück, das alle Schlüssel-Wert-Paare in Einfügereihenfolge enthält

# --instructions--

Definieren Sie ein JavaScript-Map-Objekt und weisen Sie es einer Variablen namens myMap zu. Fügen Sie das Schlüssel-Wert-Paar `freeCodeCamp`, `Awesome!` hinzu.

# --hints--

Das `myMap`-Objekt sollte existieren.

```js
assert(typeof myMap === 'object');
```

`myMap` sollte das Schlüssel-Wert-Paar `freeCodeCamp`, `Awesome!` enthalten.

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
