---
id: 587d8254367417b2b2512c71
title: Elemente aus einem Set in ES6 entfernen
challengeType: 1
forumTopicId: 301713
dashedName: remove-items-from-a-set-in-es6
---

# --description--

Lassen Sie uns üben, wie man Elemente mit der `delete`-Methode aus einem ES6-Set entfernt.

Erstellen Sie zuerst ein ES6-Set:

```js
var set = new Set([1,2,3]);
```

Entfernen Sie nun ein Element aus Ihrem Set mit der `delete`-Methode.

```js
set.delete(1);
console.log([...set]) // should return [ 2, 3 ]
```

# --instructions--

Erstellen Sie jetzt ein Set mit den ganzen Zahlen 1, 2, 3, 4 und 5.

Entfernen Sie die Werte 2 und 5 und geben Sie dann das Set zurück.

# --hints--

Ihr Set sollte die Werte 1, 3 und 4 enthalten

```js
assert(
  (function () {
    var test = checkSet();
    return test.has(1) && test.has(3) && test.has(4) && test.size === 3;
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet(){
  // Only change code below this line
  var set = null;

  // Only change code above this line
  return set;   
}
```

# --solutions--

```js
function checkSet(){
var set = new Set([1,2,3,4,5]);
set.delete(2);
set.delete(5);
return set;}
```
