---
id: 5cc0c1b32479e176caf3b422
title: "Prüfen, ob ein Baum ein binärer Suchbaum ist"
challengeType: 1
forumTopicId: 301624
dashedName: check-if-tree-is-binary-search-tree
---

# --description--

Da Sie bereits wissen, was ein binärer Suchbaum ist, wird diese Aufgabe klären, wie Sie erkennen können, ob ein Baum ein binärer Suchbaum ist oder nicht.

Das Hauptmerkmal eines binären Suchbaums ist, dass die Knoten in einer organisierten Reihenfolge angeordnet sind. Knoten haben höchstens 2 Kindknoten (rechts und/oder links angeordnet), je nachdem, ob der Wert des Kindknotens größer oder gleich (rechts) oder kleiner (links) als der des Elternknotens ist.

# --instructions--

In dieser Aufgabe erstellen Sie ein Hilfsmittel für Ihren Baum. Schreiben Sie eine JavaScript-Methode `isBinarySearchTree`, die einen Baum als Eingabe erhält und einen Boolean zurückgibt, der angibt, ob der Baum ein binärer Suchbaum ist oder nicht. Verwenden Sie Rekursion, wann immer möglich.

# --hints--

Ihr binärer Suchbaum sollte true zurückgeben, wenn er mit `isBinarySearchTree()` überprüft wird.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    test.push(1);
    test.push(5);
    test.push(3);
    test.push(2);
    test.push(4);
    return isBinarySearchTree(test) == true;
  })()
);
```

`isBinarySearchTree()` sollte false zurückgeben, wenn er mit einem Baum überprüft wird, der kein binärer Suchbaum ist.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    test.push(1);
    test.root.left = new Node(1);
    return isBinarySearchTree(test) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype.push = function(val) {
  var root = this.root;

  if (!root) {
    this.root = new Node(val);
    return;
  }

  var currentNode = root;
  var newNode = new Node(val);

  while (currentNode) {
    if (val < currentNode.value) {
      if (!currentNode.left) {
        currentNode.left = newNode;
        break;
      } else {
        currentNode = currentNode.left;
      }
    } else {
      if (!currentNode.right) {
        currentNode.right = newNode;
        break;
      } else {
        currentNode = currentNode.right;
      }
    }
  }
};
```

# --seed--

## --seed-contents--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
}
function isBinarySearchTree(tree) {
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
}
function isBinarySearchTree(tree) {
  if (tree.root == null) {
    return null;
  } else {
    let isBST = true;
    function checkTree(node) {
      if (node.left != null) {
        const left = node.left;
        if (left.value >= node.value) {
          isBST = false;
        } else {
          checkTree(left);
        }
      }
      if (node.right != null) {
        const right = node.right;
        if (right.value < node.value) {
          isBST = false;
        } else {
          checkTree(right);
        }
      }
    }
    checkTree(tree.root);
    return isBST;
  }
};
```
