---
id: 6626b4c58c027d86478ff5eb
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Mwalimu ana furaha sana na programu uliyoitengeneza hadi sasa. Lakini sasa wanataka njia rahisi ya kuangalia kama mwanafunzi ana alama ya kupita. Alama ya kupita ni chochote ambacho si `"F"`.

Kamilisha kitendakazi `hasPassingGrade` kinachopokea alama ya mwanafunzi kama kigezo. Kitendakazi chako kinapaswa kurudisha `true` ikiwa mwanafunzi ana alama ya kupita na `false` ikiwa hana.

**Vidokezo**

- Tumia kitendakazi `getGrade` kupata alama ya mwanafunzi. Kisha angalia kama alama ni ya kupita au siyo.

# --hints--

Kitendakazi chako cha `hasPassingGrade` kinapaswa kurudisha thamani ya kweli au siyo kweli.

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

Kitendakazi chako cha `hasPassingGrade` kinapaswa kurudisha `true` ikiwa alama ni `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Kitendakazi chako cha `hasPassingGrade` kinapaswa kurudisha `false` ikiwa alama ni `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Kitendakazi chako cha `hasPassingGrade` kinapaswa kurudisha `true` kwa alama zote za kupita.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
