---
id: 662693f82c91a66be46c881b
title: Hatua ya 1
challengeType: 1
dashedName: step-1
---

# --description--

Mwalimu amemaliza kupima mitihani ya wanafunzi wake na anahitaji msaada wako kuhesabu wastani wa alama za darasa.

Kamilisha kitendakazi cha `getAverage` kinachopokea safu ya data ya alama za mtihani na kurudisha wastani wa alama hizo.

Wastani huhesabiwa kwa kuongeza alama zote kisha kugawanya kwa jumla ya alama zilizopo.

```md
average = sum of all scores / total number of scores
```

Wito kadhaa za kitendakazi zimeandaliwa ili uweze kujaribu msimbo wako.

**Vidokezo**

- Unaweza kutumia mzunguko kurudia juu ya safu ya data ya `scores` na kuongeza alama zote.
- Unaweza kutumia sifa ya `length` kupata jumla ya alama zilizopo.

# --hints--

Kitendakazi chako cha `getAverage` kinapaswa kurudisha nambari.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` inapaswa kurudisha `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` inapaswa kurudisha `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` inapaswa kurudisha `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

Kitendakazi chako cha `getAverage` kinapaswa kurudisha wastani wa alama za mitihani.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
