---
id: 6626a060c4006f793e10cb33
title: Hatua ya 2
challengeType: 1
dashedName: step-2
---

# --description--

Sasa mwalimu anahitaji msaada wako kubadilisha alama ya mwanafunzi kuwa daraja la herufi.

Kamilisha kitendakazi cha `getGrade` kinachopokea nambari `score` kama kigezo. Kitendakazi chako kinapaswa kurudisha mfuatano wa herufi unaowakilisha daraja la herufi kulingana na alama.

Hapa kuna alama na daraja zao za herufi zinazolingana:

| Eneo la alama | Daraja   |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**Vidokezo**

- Kumbuka kwamba umejifunza kuhusu tamko la masharti (`if`, `else if`, na `else`).
- Kumbuka kwamba umejifunza kuhusu waendeshaji wa kulinganisha (`>`, `<`, `>=`, `<=`, `===`).
- Kumbuka kwamba umejifunza kuhusu waendeshaji wa mantiki AND (`&&`).

# --hints--

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"A++"` ikiwa alama ni `100`.

```js
assert.strictEqual(getGrade(100), "A++");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"A"` ikiwa alama ni `94`.

```js
assert.strictEqual(getGrade(94), "A");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"B"` ikiwa alama iko kati ya `80` na `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"C"` ikiwa alama ni `78`.

```js
assert.strictEqual(getGrade(75), "C");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"D"` ikiwa alama iko kati ya `60` na `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"F"` ikiwa alama ni `57`.

```js
assert.strictEqual(getGrade(57), "F");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"F"` ikiwa alama iko kati ya `0` na `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
