---
id: 6626b8dcf5057f896f948440
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa mwalimu ana taarifa zote anazohitaji, anataka aweze kutuma ujumbe kwa mwanafunzi kuhusu matokeo.

Kamilisha kitendakazi cha `studentMsg` kwa `totalScores` na `studentScore` kama vigezo. Kitendakazi kinapaswa kurudisha mfuatano wa herufi unaowakilisha ujumbe kwa mwanafunzi.

Ikiwa mwanafunzi amefaulu kozi, mfuatano wa herufi unapaswa kufuata muundo huu:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

Ikiwa mwanafunzi ameshindwa kozi, mfuatano wa herufi unapaswa kufuata muundo huu:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

Badilisha `average-goes-here` na wastani wa alama zote.
Badilisha `grade-goes-here` na daraja la mwanafunzi.

**Vidokezo**

- Tumia kitendakazi cha `getAverage` kupata wastani wa darasa.
- Tumia kitendakazi cha `getGrade` kupata daraja la mwanafunzi.
- Tumia mchanganyiko wa mfuatano wa herufi (`+`) kujenga ujumbe.
- Kuwa makini na alama za uandishi na nafasi katika ujumbe.

# --hints--

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` inapaswa kurudisha ujumbe ufuatao: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` inapaswa kurudisha ujumbe ufuatao: `"Class average: 50.8. Your grade: A++. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

Kitendakazi chako cha `studentMsg` kinapaswa kurudisha ujumbe sahihi kulingana na alama za mwanafunzi na wastani wa darasa.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
