---
id: 61fd6ab779390f49148773bb
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Chini ya kipengele chako cha `h1`, tengeneza kipengele cha `div`. Mpe sifa ya kipengele ya `id` iliyowekwa kuwa `years`. Unataka kipengele hiki fulani kifichwe kwa kisomaji cha skrini, kwa hiyo mpe sifa ya kipengele ya `aria-hidden` iliyowekwa kuwa `true`.

# --hints--

Unapaswa kutengeneza kipengele kipya cha `div` baada ya kipengele chako cha `h1`.

```js
assert(document.querySelector('h1')?.nextElementSibling?.localName === 'div');
```

Kipengele chako kipya cha `div` kinapaswa kuwa na sifa ya kipengele ya `id` iliyowekwa kuwa `years`.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'years');
```

Kipengele chako kipya cha `div` kinapaswa kuwa na sifa ya kipengele ya `aria-hidden` iliyowekwa kuwa `true`.

```js
assert(document.querySelector('div')?.getAttribute('aria-hidden') === 'true');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
