---
id: 61fd75ea7f663457612dba02
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Ndani ya kila kipengele chako kipya cha `th`, panga ndani kipengele cha `span` chenye sifa ya kipengele `class` iliyowekwa kuwa `sr-only year`. Wape maandishi yafuatayo (kwa mpangilio): `2019`, `2020`, na `2021`.

Toa kipengele chako cha tatu cha `th` sifa ya kipengele `class` iliyowekwa kuwa `current`.

Acha kipengele cha `td` kuwa tupu. Kipengele hiki kipo tu kuhakikisha jedwali lako lina mpangilio wa safu nne na kuhusisha vichwa na safu sahihi.

# --hints--

Kila kipengele chako cha `th` kinapaswa kuwa na kipengele cha `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

Kila kipengele chako kipya cha `span` kinapaswa kuwa na sifa ya kipengele `class` iliyowekwa kuwa `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

Kipengele chako cha kwanza cha `span` kinapaswa kuwa na maandishi `2019`.

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

Kipengele chako cha pili cha `span` kinapaswa kuwa na maandishi `2020`.

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

Kipengele chako cha tatu cha `span` kinapaswa kuwa na maandishi `2021`.

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

Kipengele chako cha tatu cha `th` kinapaswa kuwa na `class` iliyowekwa kuwa `current`.

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

Kipengele chako cha `td` kinapaswa kuwa tupu.

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
