---
id: 5d822fd413a79914d39e98cc
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Bem-vindo ao projeto de Horizonte com variáveis do CSS!Comece adicionando um elemento `link` que vincule seu arquivo `styles.css` dentro do elemento `head`.

# --hints--

O código deve ter um elemento `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Você não deve mudar as tags `head` existentes.Verifique se você não excluiu a tag de fechamento.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Você deve ter um elemento nulo `link`.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
