---
id: 695774002591bbc5f8cf3e53
title: Crea una funzione per applicare uno sconto
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

In questo laboratorio scriverai una funzione che calcola il prezzo finale di un elemento dopo aver applicato uno sconto percentuale.

Per esempio, se il prezzo di un elemento è `50` e viene applicato uno sconto del `20`, l'importo dello sconto è `10` e il prezzo finale è `40`.

**Obiettivo:** Completa le storie utente qui sotto per terminare il laboratorio.

**Storie utente:**

1. Devi definire una funzione chiamata `apply_discount`.

2. La funzione `apply_discount` deve prendere esattamente due parametri: `price` e `discount`.

3. Se `price` non è un numero (`int` o `float`), la funzione deve restituire la stringa `The price should be a number`.

4. Se `discount` non è un numero (`int` o `float`), la funzione deve restituire la stringa `The discount should be a number`.

5. Se `price` è minore o uguale a `0`, la funzione deve restituire la stringa `The price should be greater than 0`.

6. Se `discount` è minore di `0` o maggiore di `100`, la funzione deve restituire la stringa `The discount should be between 0 and 100`.

7. Se entrambi gli input sono validi, la funzione deve calcolare lo sconto come percentuale del prezzo.

8. La funzione deve restituire il prezzo finale dopo aver applicato lo sconto.

# --hints--

Devi avere una funzione chiamata `apply_discount`.

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

La tua funzione `apply_discount` deve prendere due parametri: `price` e `discount`.

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

Quando `apply_discount` viene chiamata con un `price` (primo argomento) che non è un numero (`int` o `float`) deve restituire `The price should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

Quando `apply_discount` viene chiamata con un `discount` (secondo argomento) che non è un numero (`int` o `float`) deve restituire `The discount should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

Quando `apply_discount` viene chiamata con un `price` minore o uguale a `0`, deve restituire `The price should be greater than 0`.

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

Quando `apply_discount` viene chiamata con un `discount` minore di `0` o maggiore di `100`, deve restituire `The discount should be between 0 and 100`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` deve restituire `80`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` deve restituire `100`.

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` deve restituire `50`.

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

Quando `apply_discount` viene chiamata con uno sconto di `100`, deve restituire `0`.

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` deve restituire `59.6`.

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
