---
id: 6551ef504d91936d2d4e54f8
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Ein *-Argument* ist ein Objekt oder ein Ausdruck, der an eine Funktion übergeben wird — zwischen den öffnenden und schließenden Klammern hinzugefügt — wenn es aufgerufen wird:

```py
greet = 'Hello!'
print(greet)
```

Der Code im obigen Beispiel würde den String `'Hello!'`, der der Wert der Variable `greet`ist, als Argument an `print()` übergeben.

Zeige deine `text` Variable auf dem Bildschirm, indem Du die `text` -Variable als Argument an die Funktion `print()` übergibst.

# --hints--

Du solltest `text` an die Funktion `print()` übergeben, indem Du den Namen dieser Variable zwischen der öffnenden und schließenden Klammer hinzufügst. Achte darauf, dass der Funktionsaufruf am Anfang der Zeile steht.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

Dein Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
