---
id: 6551ecf597038c36637f3759
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

Variablen sind ein wesentlicher Bestandteil von Python und jeder Programmiersprache. Eine Variable ist ein Name, der auf ein Objekt referenziert oder verweist. Du kannst eine Variable deklarieren, indem Du den Variablennamen auf der linken Seite des Zuweisungsoperators `=` schreibst und auf der rechten Seite des Zuweisungsoperators den Wert angibst, der dieser Variable zugewiesen wird:

```py
variable_name = value
```

Erstelle eine Variable namens `number` und weise deiner neuen Variablen den Wert `5` zu.

# --hints--

Du solltest eine Variable mit dem Namen `number` deklarieren. Achte darauf, dass der Variablenname am Anfang der Zeile steht.

```js
assert.match(code, /^number\s*=/m)
```

Du solltest der `number`-Variablen den Wert `5` zuweisen.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

Dein Code enth채lt ung체ltige Syntax und/oder ung체ltige Einr체ckung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
