---
id: 6551f6f9ec6db270eb83f3db
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Du siehst `11` auf dem Terminal sehen, weil `'Hello World'` 11 Zeichen enthält.

Eine weitere nützliche Funktion ist `type()`, die den Datentyp einer Variable zurückgibt. Ändere die `print()` Aufrufe, um den Datentyp `text` anzuzeigen.

# --hints--

Du solltest die `type()`-Funktion aufrufen.

```js
assert.match(code, /type\s*\(.*\)/)
```

Du solltest `text` an die Funktion `type()` übergeben, indem Du diesen zwischen den Klammern einfügst.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

Du solltes der `print()`Funktion `type(text)` mitgeben, inden du ihn innerhalb der Klammern einfügst. Achte darauf, dass der Funktionsaufruf am Anfang der Zeile steht.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Du solltest keine `print(len(text))` in deinem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Dein Code enthält ungültige Syntax und/oder eine ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
