---
id: 6551f4259d66026ee28cb205
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Du kannst auch ausgehend vom Ende des Strings auf Zeichenketten zugreifen. Das letzte Zeichen hat einen Index von `-1`, das zweitletzte `-2` und so weiter.

Ändere nun den bestehenden `print()` Aufruf, um das letzte Zeichen in Ihrem String auszugeben.

# --hints--

Du solltest die `print()`-Funktion aufrufen.

```js
assert.match(code, /print\s*\(.*\)/)
```

Du solltest `text[-1]` an die Funktion `print()` übergeben, indem Du diesen zwischen den Klammern einfügst. Achte darauf, dass der Funktionsaufruf am Anfang der Zeile steht.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

Dein Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
