---
id: 6551f3bfed58796ebb1268dc
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Jedes String-Zeichen kann durch einen numerischen Index referenziert werden. Die Indexzählung beginnt bei Null. Das erste Zeichen eines Strings hat also einen Index von `0`. Beispielsweise befindet sich im String `'Hello World'` das `'H'` an der Indexstelle `0`, `'e'` im Index an Stelle `1`, und so weiter.

Auf jedes Zeichen eines Strings kann mit Klammer-Notation zugegriffen werden. Dazu musst du den Varablennamen, gefolgt von eckigen Klammern angeben und den Index des Zeichens innerhalb der eckigen Klammern angeben:

```py
text = 'Hello World'
r = text[8]
```

Anstatt jetzt `text`zu schreiben, gib nur das Zeichen im Index `6` an.

# --hints--

Du solltest die `print()`-Funktion aufrufen.

```js
assert.match(code, /print\s*\(.*\)/)
```

Du solltest `text[6]` an die Funktion `print()` übergeben, indem Du diesen zwischen den Klammern einfügen. Achte darauf, dass der Funktionsaufruf am Anfang der Zeile steht.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Dein Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
