---
id: 6551f5756c20146fc94f8675
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Du kannst mit der eingebauten Funktion `len()` auf die Anzahl der Zeichen in einem String zugreifen.

Ändere den bestehenden `print()` Aufruf, indem Du `len(text)` statt `Text[-1]` übergibst.

# --hints--


Du solltest die `len()`-Funktion aufrufen.

```js
assert.match(code, /len\s*\(.*\)/)
```

Du solltest `text` an die Funktion `len()` übergeben, indem Du diesen zwischen den Klammern einfügst.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

Du solltest der `print()`Funktion `len(text)` mitgeben, indem du es innerhalb der Klammern einfügst. Achte darauf, dass der Funktionsaufruf am Anfang der Zeile steht.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

Du solltest kein `print(text[-1])` in deinem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

Dein Code enthält ungültige Syntax und/oder eine ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
