---
id: 6553a572f7a65718f1e42e18
title: Schritt 49
challengeType: 20
dashedName: step-49
---

# --description--

Eine Funktion ist im Wesentlichen ein wiederverwendbarer Code-Baustein. Du hast bereits einige eingebaute Funktionen wie `print()`, `find()` und `len()` kennengelernt. Du kannst jedoch auch benutzerdefinierte Funktionen definieren:

```py
def function_name():
    <code>
```

Eine Funktionsdeklaration beginnt mit dem Keyword `def` gefolgt von dem Funktionsnamen — einem gültigen Variablennamen — und einem Klammernpaar. Die Deklaration endet mit einem Doppelpunkt.

Deklariere gleich nach der `shift` Variable eine Funktion namens `caesar` und rücke alle folgenden Zeilen ein, um der neuen Funktion einen Körper zu geben.

# --hints--

Du sollst das Keyword `def` verwenden, um eine neue Funktion zu deklarieren.

```js
assert.match(code, /^def\s+/m)
```

You should write `caesar` as the function name after the `def` keyword. Remember to add a space after `def`.

```js
assert.match(code, /^def\s+caesar/m)
```

You should add a pair of parentheses after the function name. Don't forget the final colon.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

You should indent all the lines after `shift = 3` so that they become your new function body.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
