---
id: 659db3e3670d3e712be82593
title: Schritt 22
challengeType: 20
dashedName: step-22
---

# --description--

Den Prozess der Lokalisierung des Buchstabens im Alphabet mehrfach zu wiederholen und die Bestimmung des verschobenen Buchstabens für jedes Zeichen im `-Text`, kann ermüdend sein. Dankenswerterweise kann dies mit einer Schleife vereinfacht werden.

Entferne vorerst alle Zeilen Code unterhalb der Deklaration der `-alphabet`-Variable.

# --hints--

Du solltest immer noch `text = 'Hello World"`in deinem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

Du solltest immer noch `shift = 3` in deinem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

Du solltest immer noch `alphabet = 'abcdefghijklmnopqrstuvwxyz'` in deinem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

Du solltest die `index` Variable und ihren Wert löschen.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

Du solltest kein `print(index)` in deinem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

Du solltest die `shifted` Variable und ihren Wert löschen.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

Du solltest kein `print(shifted)` in deinem Code haben.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

Dein Code enthält ungültige Syntax und/oder ungültige Einrückung.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
