---
id: 6998b99a8e195cd0c1f211df
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите простий трекер дикої природи, використовуючи об’єкти JavaScript.

Ось приклад об’єкта:

```js
const animal = {}; // empty object
```

Це створює порожній об’єкт.

Тепер створіть змінну з іменем `tiger` і призначте їй порожній об’єкт.

# --hints--

Ви повинні створити змінну з іменем `tiger`.

```js
assert.exists(tiger);
```

Змінній `tiger` слід призначити порожній об’єкт.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
