---
id: 6999a0c1698b93d5d41156ed
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер додайте ще одну властивість до об’єкта `tiger` з назвою `isEndangered`.

Встановіть властивість `isEndangered` рівною `true`.

# --hints--

`tiger` має мати властивість `isEndangered`.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` має бути рівним `true`.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
