---
id: 6999b0d79ea938edcb032237
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

У цьому фінальному кроці практичного заняття Wildlife Tracker ви створите функцію, яка отримує доступ до властивості об’єкта за допомогою дужкової нотації.

Ось приклад доступу до властивості за допомогою дужкової нотації:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

У цьому прикладі дужкова нотація дозволяє отримати доступ до властивості, використовуючи змінну.

Тепер створіть функцію з назвою `getProperty`.

Функція має приймати два параметри: `animal` і `propertyName`.

Поверніть значення властивості, використовуючи дужкову нотацію.

Після створення функції використайте `console.log`, щоб викликати `getProperty(tiger, "species")` і `getProperty(elephant, "age")`, щоб побачити повернені значення у консолі.

На цьому практичне заняття Wildlife Tracker завершено!

# --hints--

Ви маєте створити функцію з назвою `getProperty`.

```js
assert.isFunction(getProperty);
```

Функція `getProperty` має мати два параметри: `animal` і `propertyName`.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` має використовувати дужкову нотацію для доступу до властивості.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

Ви маєте вивести `getProperty(tiger, "species")` у консоль.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

Ви маєте вивести `getProperty(elephant, "age")` у консоль.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

Виклик `getProperty(tiger, "species")` має повертати `"Tiger"`.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

Виклик `getProperty(elephant, "age")` має повертати `12`.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` має використовувати параметри функції і працювати з будь-яким об’єктом.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
