---
id: 69999ff7f7862fdfb885ea5d
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Тепер оновіть свій об’єкт `tiger`, додавши нову властивість з назвою `age`.

Встановіть властивість `age` у значення `5`.

# --hints--

`tiger` має мати властивість `age`.

```js
assert.property(tiger, "age");
```

`tiger.age` має дорівнювати `5`.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
