---
id: 6999b0198d18d402c36cc5fe
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

У цьому кроці ви створите функцію, яка видаляє властивість з об’єкта.

Ось приклад видалення властивості за допомогою `delete`:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

У цьому прикладі властивість `canFly` видаляється з об’єкта `bird`.

Тепер створіть функцію з назвою `removeEndangeredStatus`. Функція має приймати один параметр з ім’ям `animal`.

Всередині функції видаліть властивість `isEndangered` з об’єкта `animal` за допомогою ключового слова `delete`. Поверніть оновлений об’єкт `animal`.

Після створення функції використайте `console.log`, щоб викликати `removeEndangeredStatus(tiger)` і побачити оновлений об’єкт у консолі.

# --hints--

Ви маєте створити функцію з назвою `removeEndangeredStatus`.

```js
assert.isFunction(removeEndangeredStatus);
```

Функція `removeEndangeredStatus` має мати один параметр, `animal`.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Функція `removeEndangeredStatus` має повертати оновлений об’єкт `animal`.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` має використовувати ключове слово `delete` для видалення властивості `isEndangered`.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

Ви маєте вивести `removeEndangeredStatus(tiger)` у консоль.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

Виклик `removeEndangeredStatus(tiger)` має видалити властивість `isEndangered`.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` має використовувати параметри функції і працювати з будь-яким об’єктом.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
