---
id: 6999ac171928136168654774
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

У цьому кроці ви створите функцію для доступу до властивості `species` об’єкта.

Ось невеликий приклад доступу до властивості об’єкта:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

У цьому прикладі функція приймає об’єкт як параметр і повертає його властивість `species`.

Тепер створіть функцію з назвою `getSpecies`.

Функція має приймати один параметр із назвою `animal`.

Поверніть властивість `species` об’єкта `animal` за допомогою точкової нотації.

Після створення функції використайте `console.log`, щоб викликати `getSpecies(tiger)` і побачити результат у консолі.

# --hints--

Ви маєте створити функцію з назвою `getSpecies`.

```js
assert.isFunction(getSpecies);
```

Функція `getSpecies` має мати один параметр, `animal`.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ви маєте вивести `getSpecies(tiger)` у консоль.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` має отримувати доступ до властивості species за допомогою точкової нотації.

```js
assert.match(code, /animal\.species/);
```

Виклик `getSpecies(tiger)` має повертати `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` має використовувати параметр функції і працювати з будь-яким об’єктом.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
