---
id: 6999faf886634b72be7a819b
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Для цього трекера дикої природи ви оновите об’єкт, додавши нові властивості.

Ось приклад створення об’єкта з властивістю:

```js
const animal = {
  name: "Lion" // property: value
};
```

Тепер оновіть об’єкт `tiger`, щоб він містив властивість з назвою `species`.

Встановіть її значення рівним `"Tiger"`.

# --hints--

`tiger` має бути об’єктом.

```js
assert.isObject(tiger);
```

`tiger` має мати властивість `species`.

```js
assert.property(tiger, "species");
```

`tiger.species` має бути рівним `"Tiger"`.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
