---
id: 6999b07c34c2cde21ee6e1d8
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

У цьому кроці ви створите функцію, яка перевіряє, чи має об’єкт певну властивість.

Ось приклад використання `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

У цьому прикладі функція `hasColor` використовує `hasOwnProperty` для перевірки, чи існує властивість `color` у об’єкта.

Тепер створіть функцію з назвою `hasHabitat`. Функція має приймати один параметр з ім’ям `animal`.

Використайте метод `hasOwnProperty`, щоб повернути `true`, якщо об’єкт `animal` має властивість з назвою `"habitat"`, і `false` в іншому випадку.

Після створення функції використайте `console.log`, щоб викликати `hasHabitat(tiger)` і `hasHabitat(elephant)`, щоб побачити обидва результати в консолі.

# --hints--

Ви маєте створити функцію з назвою `hasHabitat`.

```js
assert.isFunction(hasHabitat);
```

Функція `hasHabitat` має мати один параметр, `animal`.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` має використовувати метод `hasOwnProperty` для перевірки властивості `habitat`.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

Ви маєте вивести `hasHabitat(tiger)` у консоль.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

Ви маєте вивести `hasHabitat(elephant)` у консоль.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

Виклик `hasHabitat(tiger)` має повернути `true` після додавання середовища існування.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Виклик `hasHabitat(elephant)` має повернути `false`, якщо середовище існування відсутнє.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` має використовувати параметри функції і працювати з будь-яким об’єктом.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
