---
id: 6999adc72c02812a28c01e31
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

У цьому кроці ви створите функцію, яка оновлює властивість об’єкта.

Ось приклад оновлення властивості всередині функції:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

У цьому прикладі властивість `age` оновлюється на нове значення.

Тепер створіть функцію з назвою `updateAge`. Функція має приймати два параметри: `animal` і `newAge`.

Всередині функції оновіть властивість `age` об’єкта `animal` на `newAge`. Поверніть оновлений об’єкт `animal`.

Після створення функції використайте `console.log`, щоб викликати `updateAge(elephant, 12)` і побачити оновлений об’єкт `elephant` у консолі.

# --hints--

Ви маєте створити функцію з назвою `updateAge`.

```js
assert.isFunction(updateAge);
```

Функція `updateAge` має мати два параметри: `animal` і `newAge`.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

`updateAge` має використовувати точкову нотацію для оновлення властивості `age`.

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

Функція `updateAge` має повертати оновлений об’єкт `animal`.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

Ви маєте вивести `updateAge(elephant, 12)` у консоль.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

Виклик `updateAge(elephant, 12)` має оновити властивість age.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` має використовувати параметри функції і працювати з будь-яким об’єктом.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
