---
id: 699a00de9f564bb0effaa14d
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

У цьому кроці ви створите функцію для доступу до властивості `age` об’єкта.

Ось приклад:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

У цьому прикладі функція приймає об’єкт як параметр і повертає його властивість `age`.

Тепер створіть функцію з назвою `getAge`.

Функція має приймати один параметр з ім’ям `animal`.

Поверніть властивість `age` об’єкта `animal` за допомогою точкової нотації.

Після створення функції використайте виклик `console.log`, щоб викликати `getAge(tiger)` і побачити результат у консолі.

# --hints--

Ви маєте створити функцію з назвою `getAge`.

```js
assert.isFunction(getAge);
```

Функція `getAge` має мати один параметр, `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` має отримувати доступ до властивості `age` за допомогою точкової нотації.

```js
assert.match(code, /animal\.age/);
```

Ви маєте вивести `getAge(tiger)` у консоль.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

Виклик `getAge(tiger)` має повертати `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` має використовувати параметр функції і працювати з будь-яким об’єктом.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
