---
id: 69ba8da6b296391d913096a0
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер створіть другий об’єкт з назвою `elephant`.

Додайте такі властивості:

- `species` зі значенням `"Elephant"`
- `age` зі значенням `10`
- `isEndangered` зі значенням `true`

# --hints--

Ви повинні створити змінну з назвою `elephant`.

```js
assert.exists(elephant);
```

`elephant` має мати властивість `species` зі значенням `"Elephant"`.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` має мати властивість `age` зі значенням `10`.

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` має мати властивість `isEndangered` зі значенням `true`.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
