---
id: 67d83df6f82eda3868dd2a84
title: 建造一個 RPG 字元
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

在這個實驗中，你將透過建置一個小型應用程式來練習 Python 的基礎，該應用程式會為 RPG 冒險創建一個字元。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `create_character` 的函式（程式）。
1. 該函式（程式）應依序接受一個字元名稱，接著是三個屬性：力量、智力和魅力。
1. 字元名稱應該被驗證：
   - 如果字元名稱不是字串，該函式應該傳回 `The character name should be a string`。
   - 如果字元名稱是空字串，該函式應該傳回 `The character should have a name`。
   - 如果字元名稱長度超過 10 個字元，函式應該傳回 `The character name is too long`。
   - 如果字元名稱包含空格，函式應該傳回 `The character name should not contain spaces`。
1. 這些統計數據也應該被驗證：
   - 如果一個或多個統計數據不是整數，該函式應該傳回 `All stats should be integers`。
   - 如果一個或多個統計數據小於 1，該函式（程式）應該傳回 `All stats should be no less than 1`。
   - 如果一個或多個統計數值大於 4，該函式（程式）應該傳回 `All stats should be no more than 4`。
   - 如果所有屬性的總和不等於 7，該函式應該傳回 `The character should start with 7 points`。
1. 如果所有值都通過驗證，該函式應該傳回一個包含四行的字串：
   - 第一行應包含字元名稱
   - 第 2-4 行應以屬性縮寫 `STR`、`INT` 或 `CHA`（依此順序）開頭，接著一個空格，然後是與該屬性值相等數量的實心點（`●`），再以空心點（`○`）補足至 10 個。點數由編輯器提供。

以下是 `create_character('ren', 4, 2, 1)` 應該傳回的字串：

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

注意：雖然 `str` 和 `int` 是統計資料的常用縮寫，但請記得這些是 Python 的保留關鍵字，不應用作變數名稱。

# --hints--

你應該有一個名為 `create_character` 的函式（程式）。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

當 `create_character` 被呼叫且第一個引數不是字串時，應該傳回 `The character name should be a string`。

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

當 `create_character` 被呼叫且第一個引數是空字串時，它應該傳回 `The character should have a name`。

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

當 `create_character` 被呼叫且第一個引數長度超過 10 個字元時，應該傳回 `The character name is too long`。

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

當字元長度不超過 10 個字元時，`create_character` 函式（程式）不應該顯示字元過長的訊息。

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

當 `create_character` 被呼叫且第一個引數包含空格時，應該傳回 `The character name should not contain spaces`。

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

當 `create_character` 被呼叫且第二、第三或第四個引數不是整數時，應該傳回 `All stats should be integers`。

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

當 `create_character` 被呼叫且第二、第三或第四個引數小於 `1` 時，應該傳回 `All stats should be no less than 1`。

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

當 `create_character` 被呼叫且第二、第三或第四個引數大於 `4` 時，應該傳回 `All stats should be no more than 4`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

當 `create_character` 被呼叫且第二、第三或第四個引數的總和不等於 `7` 時，應該傳回 `The character should start with 7 points`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` 應該傳回 `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

當以有效值呼叫 `create_character` 時，應該輸出所需的字元屬性。

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
