---
id: 67d83df6f82eda3868dd2a84
title: 建造一個 RPG 字元
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

在這個實驗中，你將透過建置一個小型應用程式來練習 Python 的基礎，該應用程式會為 RPG 冒險創建一個字元。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `create_character` 的函式（程式）。
1. 該函式應依序接受一個角色名稱，接著是三個屬性：力量、智力和魅力。
1. 角色名稱應該被驗證：
   - 如果角色名稱不是字串，函式應傳回 `The character name should be a string`。
   - 如果角色名稱是空字串，函式應傳回 `The character should have a name`。
   - 如果角色名稱長度超過 10 個字元，函式應傳回 `The character name is too long`。
   - 如果角色名稱包含空白，函式應傳回 `The character name should not contain spaces`。
1. 屬性也應該被驗證：
   - 如果一個或多個屬性不是整數，函式應傳回 `All stats should be integers`。
   - 如果一個或多個屬性小於 1，函式應傳回 `All stats should be no less than 1`。
   - 如果一個或多個屬性大於 4，函式應傳回 `All stats should be no more than 4`。
   - 如果所有屬性的總和不等於 7，函式應傳回 `The character should start with 7 points`。
1. 如果所有值都通過驗證，函式應傳回一個包含四行的字串：
   - 第一行應包含角色名稱。
   - 第 2 至 4 行應依序以屬性縮寫 `STR`、`INT`、`CHA` 開頭，接著一個空格，然後是與該屬性值相等數量的實心點（`●`），以及補足到 10 個的空心點（`○`）。例如：如果力量值是 3，則應有 3 個實心點後接 7 個空心點。點的符號已在編輯器中提供。

以下是 `create_character('ren', 4, 2, 1)` 應該傳回的字串：

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

**注意：** 雖然 `str` 和 `int` 是常見的屬性縮寫，但 Python 也使用這些名稱代表字串和整數型別，因此最好避免將它們用作變數名稱。

# --hints--

你應該有一個名為 `create_character` 的函式（程式）。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

當 `create_character` 被呼叫且第一個引數 **不是** 字串時，應傳回 `The character name should be a string`。

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

當 `create_character` 被呼叫且第一個引數是字串時，**不應** 傳回 `The character name should be a string`。

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

當 `create_character` 被呼叫且第一個引數是空字串時，它應該傳回 `The character should have a name`。

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

當 `create_character` 被呼叫且第一個引數是 **非** 空字串時，不應該傳回 `The character should have a name`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

當 `create_character` 被呼叫且第一個引數長度超過 10 個字元時，應傳回 `The character name is too long`。

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

當角色名稱長度不超過 10 個字元時，`create_character` 函式**不應**說角色名稱太長。

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

當 `create_character` 被呼叫且第一個引數包含空白時，應傳回 `The character name should not contain spaces`。

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

當 `create_character` 被呼叫且第一個引數*不*包含空白時，*不應*傳回 `The character name should not contain spaces`。

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

當 `create_character` 被呼叫且第二、三或四個引數**不是**整數時，應傳回 `All stats should be integers`。

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

當 `create_character` 被呼叫且第二、三和四個引數都是整數時，**不應**傳回 `All stats should be integers`。

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

當 `create_character` 被呼叫且第二、三或四個引數小於 `1` 時，應傳回 `All stats should be no less than 1`。

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

當 `create_character` 被呼叫且第二、三和四個引數都不小於 `1` 時，**不應**傳回 `All stats should be no less than 1`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

當 `create_character` 被呼叫且第二、三或四個引數大於 `4` 時，應傳回 `All stats should be no more than 4`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

當 `create_character` 被呼叫且第二、三和四個引數都不大於 `4` 時，**不應**傳回 `All stats should be no more than 4`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

當傳入 `create_character` 的第二、三和四個引數**不**加總為 `7` 時，`create_character` 應傳回 `The character should start with 7 points`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

當 `create_character` 被呼叫且第二、三和四個引數加總為 `7` 時，**不應**傳回 `The character should start with 7 points`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` 應該傳回 `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

當 `create_character` 被呼叫且引數有效時，應依要求輸出角色屬性。

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
        return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
