---
id: 641da5462576784453146ec2
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

既然你已經擁有所需的資料，你就可以用它來填充 UI。但擷取的資料包含一個有 26 位作者的陣列，如果你同時將它們全部添加到頁面上，可能會導致效率不佳。

相反地，你應該一次為你的作者添加 8 位，並且有一個按鈕可以再添加 8 位，直到沒有更多的資料可顯示。

使用 `let` 創建 2 個名為 `startingIndex` 和 `endingIndex` 的變數，並分別將它們指定為數字值 `0` 和 `8`。同時，使用 `let` 創建一個 `authorDataArr` 變數，並將它設為空陣列。

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

你應該使用 `let` 來宣告名為 `startingIndex` 的變數。

```js
assert.match(code, /let\s+startingIndex/)
```

你應該將你的 `startingIndex` 變數設為 `0`。

```js
assert.match(code, /let\s+startingIndex\s*=\s*0\s*;?/)
```

你應該使用 `let` 來宣告名為 `endingIndex` 的變數。

```js
assert.match(code, /let\s+endingIndex/)
```

你應該將你的 `endingIndex` 變數設為 `8`。

```js
assert.match(code, /let\s+endingIndex\s*=\s*8\s*;?/)
```

你應該使用 `let` 來宣告名為 `authorDataArr` 的變數。

```js
assert.match(code, /let\s+authorDataArr/)
```

你應該將你的 `authorDataArr` 變數設為空陣列（`[]`）。

```js
assert.match(code, /let\s+authorDataArr\s*=\s*\[\s*\]\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
