---
id: 641d9a19bff38d34d5a5edb8
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

這個專案的所有 HTML 和 CSS 都已經為你提供。你可以查看這兩個檔案以熟悉它們。

首先使用 `.getElementById()` 方法取得 `#author-container` 和 `#load-more-btn` 元素。分別將它們指定給變數 `authorContainer` 和 `loadMoreBtn`。

這些變數不會改變，所以使用 `const` 來宣告它們。

# --hints--

你應該使用 `document.getElementById()` 來取得 `#author-container` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)author\-container\1\s*\)/);
```

你應該將 `#author-container` 元素指定到變數 `authorContainer`。別忘了使用 `const` 來宣告該變數。

```js
assert.match(code, /const\s+authorContainer\s*\=\s*document\.getElementById\(\s*('|"|`)author\-container\1\s*\)/);
```

你應該使用 `document.getElementById()` 來取得 `#load-more-btn` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)load\-more\-btn\1\s*\)/);
```

你應該將 `#load-more-btn` 元素指定給變數 `loadMoreBtn`。別忘了使用 `const` 來宣告該變數。

```js
assert.match(code, /const\s+loadMoreBtn\s*\=\s*document\.getElementById\(\s*('|"|`)load\-more\-btn\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
