---
id: 66f26c32ec6f90df01a44f60
title: تصميم نموذج اتصال
challengeType: 25
dashedName: design-a-contact-form
demoType: onClick
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك `div` بصنف `form-container` ليحتوي النموذج الخاص بك.

2. يجب أن يكون لديك عنصر `form` داخل عنصر `.form-container`.

3. داخل عنصر `form`، يجب أن يحتوي على العناصر التالية وحقول الإدخال:

- عنصر `h2` يحتوي على بعض النصوص.
   - حقل إدخال نصي للاسم مع تعيين الخاصية `type` إلى `text` ووجود الخصائص `id` و`name` و`required`.
   - حقل إدخال بريد إلكتروني مع تعيين الخاصية `type` إلى `email` ووجود الخصائص `id` و`name` و`required`.
   - مساحة نصية للرسالة مع الخصائص `id` و`name` و`required`.
   - عنصر `button` يحتوي على النص `"Submit"` والخاصية `type` بقيمة `submit`.

4. يجب أن يحتوي كل من عناصر `input` و`textarea` على عنصر `label` المقابل له مع خاصية `for` مساوية لقيمة `id` الخاصة بالعنصر.

5. يجب أن يحتوي عنصر `.form-container` على لون خلفية، وقيم لـ `border-radius` و`padding` و`width`.

6. يجب أن تحتوي عناصر `label` على هامش ولون خط.

7. يجب أن تحتوي عناصر `input` و`textarea` على قيم لـ:
   
- `width`.
   - `padding`.
   - هامش في الأسفل.

8. يجب أن يحتوي عنصر `button` على:
   
- لون خلفية
   - حجم خط مختلف عن الافتراضي.

9. يجب أن يحتوي الزر على تأثير المرور بالمؤشر الذي يغير لون الخلفية.

**ملاحظة:** تأكد من ربط ورقة الأنماط في ملف HTML الخاص بك وتطبيق CSS.

# --hints--

يجب أن يكون لديك `div` بصنف `form-container`.

```js
const el = document.querySelector('div.form-container');
assert.isNotNull(el);
```

يجب أن يكون لديك عنصر `form` داخل عنصر `.form-container`.

```js
const el = document.querySelector('.form-container form');
assert.isNotNull(el);
```

داخل `form`، يجب أن يكون لديك عنصر `h2` يحتوي على بعض النص.

```js
const el = document.querySelector('.form-container form h2');
assert.isNotEmpty(el.innerText);
// assert.isAtLeast(document.querySelector('h1').innerText.length, 1);

```

داخل `form`، يجب أن يكون لديك حقل إدخال نصي مع تعيين النوع إلى `text`.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.isNotNull(el);
```

يجب أن يحتوي حقل الإدخال الخاص بك من نوع `type=text` على خاصية `id`، و`name`، و`required`.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert.exists(el);
assert.isTrue(el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

داخل الـ `form`، يجب أن يكون لديك حقل إدخال للبريد الإلكتروني مع تعيين الـ `type` إلى `email`.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isNotNull(el);
```

يجب أن يحتوي حقل الإدخال الخاص بك من نوع `type=email` على خاصية `id`، و`name`، و`required`.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert.isTrue(el?.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

داخل `form`، يجب أن يكون لديك عنصر `textarea` مع خصائص `id` و`name` و`required`.

```js
const el = document.querySelector('.form-container form textarea');
assert.isNotNull(el);
assert.isNotNull(document.querySelector('.form-container form textarea[required]'));
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

داخل `form`، يجب أن يكون لديك عنصر `button` مع خاصية `type` بقيمة `submit`.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(el?.getAttribute('type'), 'submit');
```

يجب أن يحتوي الزر على نص `"Submit"`.

```js
const el = document.querySelector('.form-container form button');
assert.strictEqual(document.querySelector('.form-container form button').textContent, 'Submit');
```

يجب أن يكون لديك عنصر `label` لـ `#name` مع خاصية `for` تطابق قيمة `id` في الـ `input` المقابل.

```js
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value
assert.isNotEmpty(nameInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === nameInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

يجب أن يكون لديك عنصر `label` لـ `#email` مع خاصية `for` تطابق قيمة `id` في الـ `input` المقابل.

```js
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);
// check if id has a value
assert.isNotEmpty(emailInput);
let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === emailInput) {
        found = true;
        break;
    }
}

assert.isTrue(found);
```

يجب أن يكون لديك عنصر `label` لعنصر `#message` مع خاصية `for` تطابق قيمة `id` في عنصر `textarea`.

```js
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

// check if id has a value

assert.isNotEmpty(messageInput);

assert.exists(document.querySelector(`label[for="${messageInput}"]`));
```

يجب ألا تكون التسميات فارغة.

```js
const labels = document.querySelectorAll('.form-container form label');

assert.isAbove(labels.length, 0);

for (let label of labels) {
    const labelText = label.textContent.trim();
    assert.isNotEmpty(labelText);
}

```

يجب أن يحتوي عنصر `.form-container` على لون خلفية.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

يجب أن يحتوي عنصر `.form-container` على `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

يجب أن يحتوي عنصر `.form-container` على حشو.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

يجب أن يحتوي عنصر `.form-container` على عرض.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

يجب أن تحتوي المسميات على هامش ولون خط.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

يجب أن تحتوي عناصر `input` على عرض، حشو، وهامش في الأسفل.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('input'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);

```

يجب أن يحتوي عنصر `textarea` على عرض، حشو، وهامش في الأسفل.

```js
// Helper function to extract properties from multiple selectors
function getStyleProperties(selectors, properties) {
  const result = {};
  properties.forEach((property) => {
    result[property] = selectors.map((selector) => selector?.[property]).find(Boolean);
  });
  return result;
}

// Initializing selectors
const selectors = [
  new __helpers.CSSHelp(document).getStyle('textarea'),
  new __helpers.CSSHelp(document).getStyle('input, textarea'),
  new __helpers.CSSHelp(document).getStyle('textarea, input')
];

// Extract width, padding, and margin properties
const { width } = getStyleProperties(selectors, ['width']);
assert.isNotEmpty(width);

const {
  padding,
  paddingTop,
  paddingBottom,
  paddingLeft,
  paddingRight
} = getStyleProperties(selectors, ['padding', 'paddingTop', 'paddingBottom', 'paddingLeft', 'paddingRight']);

// Assert padding properties
assert.isOk(
  padding ||
  (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

const { marginBottom } = getStyleProperties(selectors, ['marginBottom']);
assert.isNotEmpty(marginBottom);
```

يجب أن يحتوي عنصر `button` على لون خلفية.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

يجب أن يحتوي عنصر `button` على حجم خط.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

يجب أن يحتوي عنصر `button` على تأثير المرور بالمؤشر الذي يغير لون الخلفية.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const btn_hover = gs('button:hover')?.backgroundColor;
assert.isNotEmpty(btn_hover);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled Form Project</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="form-container">
        <form>
            <h2>Contact Us</h2>
            <label for="name">Name:</label>
            <input type="text" id="name" name="name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" required>

            <label for="message">Message:</label>
            <textarea id="message" name="message" required></textarea>

            <button type="submit">Submit</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background: linear-gradient(45deg, #83a4d4, #b6fbff);
    font-family: Arial, sans-serif;
}

.form-container {
    background-color: white;
    border-radius: 10px;
    padding: 20px;
    width: 300px;
    text-align: center;
}

form h2 {
    margin-bottom: 20px;
    color: #333;
}

label {
    display: block;
    margin: 10px 0 5px;
    color: #555;
}

input{
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

textarea {
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

button {
    background-color: #4CAF50;
    color: white;
    padding: 10px 15px;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    font-size: 1em;
}

button:hover {
    background-color: #45a049;
}
```
