---
id: a3566b1109230028080c9345
title: すべての数字を合計するアルゴリズムを設計する
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

このラボでは、すべての数字を合計するアルゴリズムを設計する必要があります。

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. 2つの数字の配列を受け取る`sumAll`という名前の関数が必要です。
1. `sumAll([n, m])`は`n`と`m`の合計に加え、その間のすべての数字の合計を返す必要があります。最小の数字が常に最初に来るとは限りません。例えば、`sumAll([4,1])`は`10`を返すべきです。なぜなら`1`と`4`（両方を含む）の間のすべての数字の合計が`10`だからです。

# --hints--

`sumAll`という名前の関数が必要です。

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])`は数値を返す必要があります。

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])`は`10`を返す必要があります。

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])`は`10`を返す必要があります。

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])`は`45`を返す必要があります。

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])`は`45`を返す必要があります。

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
