---
id: 63ee3fe4381756f9716727f0
title: أساسيات CSS التمرين ب
challengeType: 25
dashedName: css-foundations-exercise-b
---

# --description--

**الهدف:** هناك عدة عناصر في ملف HTML المقدم، يجب عليك إضافة خاصيات صنف أو معرف إليها. ثم عليك إضافة قواعد في ملف CSS المقدم باستخدام صيغة محدد صحيحة.

## قصص المستخدم

1. يجب أن ترى خلفية `yellow` لكل العناصر ذات الأرقام الفردية في القائمة.

1. يجب أن تستخدم محدد `class` لكل العناصر ذات الأرقام الفردية في القائمة.

1. يجب أن ترى أن العنصر الثاني في القائمة يحتوي على نص `blue` وحجم خط `font-size` بقيمة `36px`.

1. يجب تعيين `font-size` ولون النص في العنصر الثاني باستخدام خاصية `id`.

1. يجب أن ترى أن العنصر الثالث في القائمة له حجم خط `font-size` بقيمة `24px`.

1. يجب تعيين `font-size` على العنصر الثالث باستخدام خاصية `class`.

1. يجب أن ترى أن العنصر الرابع في القائمة له خلفية `red`، وحجم خط `font-size` بقيمة `24px`، ووزن خط `font-weight` عريض `bold`.

1. يجب تعيين `font-size` للعنصر الرابع باستخدام خاصية `class`. يجب تعيين `font-weight` واللون باستخدام خاصية `id`.

# --hints--

يجب أن يحتوي كل عنصر فردي على خاصية `class`.

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasClass = p?.every((paragraph) => paragraph.classList.length > 0);

assert(everyPHasClass);
```

يجب أن تحتوي العناصر الفردية على خاصية `background-color` بقيمة `yellow`.

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasBackgroundColor = p?.every((paragraph) => {
  const style = getComputedStyle(paragraph);
  
  return style?.backgroundColor === 'rgb(255, 255, 0)';
})

assert.equal(everyPHasBackgroundColor, true);
```

يجب أن يحتوي العنصر الثاني على نص `blue` وحجم خط `font-size` بقيمة `36px`.

```js
const secondElementId = document.querySelectorAll('div')?.[0]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${secondElementId}`);

assert.equal(style?.color, 'blue');
assert.equal(style?.fontSize, '36px');
```

يجب أن يحتوي العنصر الثالث على نص وحجم خط `font-size` بقيمة `24px`.

```js
const thirdElement = document.querySelectorAll('p')?.[1];

assert(thirdElement?.innerText?.length > 0);

const thirdElementClasses = Array.from(thirdElement?.classList?.values());

assert(thirdElementClasses.some(thirdElementClass => {

  const style = new __helpers.CSSHelp(document).getStyle(`.${thirdElementClass}`);

  return style?.fontSize === '24px';

}))

``` 

يجب أن يحتوي العنصر الرابع على حجم خط `font-size` بقيمة `24px`.

```js
const fourthElementClass = document.querySelectorAll('div')?.[1]?.classList[0];

const style = new __helpers.CSSHelp(document).getStyle(`.${fourthElementClass}`);

assert(style?.fontSize === '24px');
```

يجب أن يحتوي العنصر الرابع على خاصية `background-color` بقيمة `red`.

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.backgroundColor === 'red');
```

يجب أن يحتوي العنصر الرابع على وزن خط `font-weight` عريض `bold`.

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.fontWeight === 'bold');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p>Number 1 - I'm a class!</p>
    <div>Number 2 - I'm one ID.</div>
    <p>Number 3 - I'm a class, but cooler!</p>
    <div>Number 4 - I'm another ID.</div>
    <p>Number 5 - I'm a class!</p>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p class="odd">Number 1 - I'm a class!</p>
    <div id="two">Number 2 - I'm one ID.</div>
    <p class="odd adjust-font-size">Number 3 - I'm a class, but cooler!</p>
    <div id="four" class="adjust-font-size">Number 4 - I'm another ID.</div>
    <p class="odd">Number 5 - I'm a class!</p>
  </body>
</html>
```

```css
.odd {
  background-color: yellow;
  font-family: Verdana, "DejaVu Sans", sans-serif;
}

.adjust-font-size {
  font-size: 24px;
}

#two {
  color: blue;
  font-size: 36px;
}

#four {
  background-color: red;
  font-weight: bold;
}
```
