---
id: 63ee3fe9381756f9716727f1
title: أساسيات CSS التمرين ج
challengeType: 25
dashedName: css-foundations-exercise-c
---

# --description--

لنُطوّر قليلاً على التمرين السابق، حيث أضفت عدة أصناف إلى عنصر واحد لتطبيق قاعدتين مختلفتين عليه.

1. يجب أن ترى خلفية سوداء ونصًا أبيض على أول عنصر `button`.
1. يجب أن ترى خلفية صفراء على ثاني عنصر `button`.
1. يجب أن تحدد حجم الخط `28px` باستخدام قائمة محددات.
1. يجب أن تحتوي قائمة الخطوط على `Helvetica` و`Times New Roman` مع `sans-serif` كنظام بديل في قائمة المحدد.
1. يجب أن ترى اسم صنف فريد على كل عنصر.
1. يجب أن يكون لديك قائمة محددات للأنماط التي يشترك بها كلا العنصرين.

# --hints--

يجب أن يكون لديك خلفية `black` على العنصر الأول.

```js
const classes = document.querySelectorAll('button')?.[0].classList;

const style = new __helpers.CSSHelp(document).getStyle(`.${classes[0]}`);

assert.equal(style?.backgroundColor, 'black');

```

يجب أن يكون لون النص في العنصر الأول `white`.

```js

const classes = document.querySelectorAll('button')?.[0].classList;

const style = new __helpers.CSSHelp(document).getStyle(`.${classes[0]}`);

assert.equal(style?.color, 'white');

```

يجب أن تحدد `font-size` لكلا العنصرين إلى `28px` باستخدام قائمة محددات.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}
assert.equal(eitherOr()?.fontSize, '28px');
```

يجب أن يكون لديك اسم صنف فريد على كل عنصر.

```js
const elementOneClasses = document.querySelectorAll('button')?.[0].classList;
const elementTwoClasses = document.querySelectorAll('button')?.[1].classList;

for(let i = 0; i < elementOneClasses.length; i++){
    assert(![...elementTwoClasses].includes(elementOneClasses[i]));
}

```

يجب أن يكون لديك قائمة محددات للأنماط التي يشترك بها كلا العنصرين.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}
assert.exists(eitherOr());
```

يجب أن تحتوي قائمة الخطوط على `Helvetica` و`Times New Roman` مع `sans-serif` كنظام بديل في قائمة المحدد.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}

assert.equal(eitherOr()?.fontFamily, 'Helvetica, "Times New Roman", sans-serif');
```

يجب أن يحتوي كل عنصر على صنف واحد فقط.

```js
const elementOneClasses = document.querySelectorAll('button')?.[0].classList;
const elementTwoClasses = document.querySelectorAll('button')?.[1].classList;

assert(elementOneClasses.length === 1 && elementTwoClasses.length === 1);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Grouping Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <button>Click Me!</button>
    <button>No, Click Me!</button>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
  <html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Grouping Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <button class="inverted">Click Me!</button>
    <button class="fancy">No, Click Me!</button>
  </body>
</html>
```

```css
.inverted,
.fancy {
  font-family: Helvetica, "Times New Roman", sans-serif;
  font-size: 28px;
}

.inverted {
  background-color: black;
  color: white;
}

.fancy {
  background-color: yellow;
}
```
