---
id: 63ee3f71381756f9716727ef
title: أساسيات CSS التمرين أ
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**الهدف:** 
في هذا التمرين، ستتدرب على إضافة CSS إلى ملف HTML باستخدام الطرق الثلاث كلها: CSS خارجي، CSS داخلي، وCSS مضمّن. يجب أن تستخدم محددات النوع فقط في هذا التمرين عند إضافة الأنماط عبر الطريقتين الخارجيّة والداخليّة. كما يجب أن تستخدم الكلمات المفتاحية للألوان (مثل "blue") بدلاً من استخدام قيم `RGB` أو `HEX`.

## قصص المستخدم

- يجب أن ترى عنصر `div` يحتوي على بعض النص.
  - يجب أن يكون له خلفية `red`، ونص `white`، وحجم خط `32px`، ومحاذاة النص في الوسط، ونص عريض `bold`.
  - يجب إضافة CSS لعنصر `div` خارجيًا، باستخدام محدد النوع.
- يجب أن ترى عنصر `p` يحتوي على بعض النص.
  - يجب أن يكون له خلفية `green`، ونص `white`، وحجم خط `18px`.
  - يجب إضافة CSS لعنصر `p` داخليًا، باستخدام محدد النوع.
- يجب أن ترى عنصر `button` يحتوي على بعض النص.
  - يجب أن يكون لعنصر `button` خلفية `orange` وحجم خط `18px`.
  - يجب إضافة CSS لعنصر `button` باستخدام الأنماط المضمّنة.

# --hints--

يجب أن يكون لديك عنصر `div` واحد يحتوي على بعض النص.

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

يجب أن يكون لديك ورقة أنماط خارجية تحتوي على أنماط عنصر `div`.

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

يجب ألا تُضاف أنماط CSS لعنصر `div` باستخدام الأنماط الداخلية أو المضمّنة.

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

يجب أن يكون لعنصر `div` خاصية `background-color` بقيمة `red` و`color` بقيمة `white`.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

يجب أن يكون لعنصر `div` خاصية `font-weight` مضبوطة على `bold`، و`font-size` مضبوطة على `32px`، و`text-align` مضبوطة على `center`.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

يجب أن يكون لديك عنصر `p` واحد ويحتوي على بعض النص.

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

يجب إضافة أنماط عنصر `p` داخليًا باستخدام عنصر `style`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

يجب أن يكون لعنصر `p` الخاص بك `font-size` بقيمة `18px` وأن يكون `color` مضبوطًا على `white`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

يجب أن يكون لديك عنصر `button` واحد يحتوي على بعض النص.

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

يجب أن يكون لعنصر `button` نمط مضمّن.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```
 
يجب أن يكون لعنصر `button` خاصية `background-color` مضبوطة على `orange`.

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

يجب أن يكون لعنصر `button` خاصية `font-size` مضبوطة على `18px`.

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
