---
id: 60a3e3396c7b40068ad69978
title: Passaggio 15
challengeType: 0
dashedName: step-15
---

# --description--

Usa i margini per regolare lo spazio esterno a un elemento.

Impostando la proprietà `margin`, dai all'elemento `.frame` un margine verticale di `20px` e un margine orizzontale di `auto`. Questo sposterà il riquadro verso il basso di 20 pixel e lo centrerà orizzontalmente nella pagina.

# --hints--

Dovresti impostare la proprietà `margin` su `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

Il tuo elemento `.frame` dovrebbe avere un valore `margin` di `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
