---
id: 60a3e3396c7b40068ad6996c
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Nel modello a riquadro CSS, ogni elemento HTML è trattato come un riquadro con quattro aree.

Immagina di ricevere un pacco dal tuo rivenditore online preferito: il contenuto è l'elemento nel pacco, o nel nostro caso, un elemento intestazione, paragrafo o immagine.

Modifica l'attributo `src` nel `<img>` da `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` a `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

L'elemento `img` dovrebbe avere un `src` di `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```
