---
id: 60a3e3396c7b40068ad6996a
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

A questo punto dovresti conoscere gli elementi base che una pagina HTML dovrebbe avere.

Prepara il tuo codice con una dichiarazione `DOCTYPE`, un elemento `html` con la lingua impostata su inglese, un elemento `head` e un elemento `body`.

# --hints--

Il tuo codice dovrebbe avere una dichiarazione `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Dovresti avere un tag di apertura `<html>` con un attributo `lang` impostato su `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Il tuo codice dovrebbe avere un elemento `head` all’interno dell’elemento `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

Il tuo codice dovrebbe avere un elemento `body` all’interno dell’elemento `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

Il tuo elemento `head` dovrebbe venire prima del tuo elemento `body`.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
