---
id: 60a3e3396c7b40068ad6996d
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Il contenuto è circondato da uno spazio chiamato spaziatura interna, simile a come il pluriball separa un elemento dal riquadro che lo contiene.

Pensa al bordo come alla scatola di cartone in cui è stato spedito il tuo elemento.

Modifica l'attributo `src` in `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

L'elemento `img` dovrebbe avere un `src` di `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
