---
id: 60a3e3396c7b40068ad69979
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Aggiungi un nuovo elemento `div` all'interno del tuo elemento `.canvas`.

Dai al nuovo `div` l'attributo `class` con valore `one`. Questo sarà il tuo primo rettangolo.

# --hints--

Dovresti creare un nuovo elemento `div`.

```js
assert(document.querySelectorAll('div').length === 3);
```

Dovresti annidare il nuovo elemento `div` all'interno del tuo elemento `.canvas`.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

Il tuo nuovo `div` dovrebbe avere un attributo `class` con valore `one`.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
