---
id: 69a56b5069ca99f7317e6e19
title: Створіть валідатор вантажного маніфесту
challengeType: 26
dashedName: lab-cargo-manifest-validator
---

# --description--

У цій лабораторній роботі ви використаєте JavaScript для нормалізації та перевірки вантажних маніфестів. Вантажний маніфест — це документ, який зазвичай містить перелік товарів, що перевозяться (наприклад, кораблем або поїздом), і включає деталі про ці товари.

Кожен вантажний маніфест буде представлений як об’єкт із такими властивостями:

- `containerId`: додатне ціле число, що ідентифікує відповідний контейнер з вантажем.
- `destination`: непорожній рядок (після обрізання пробілів), що позначає цільове місце призначення вантажу.
- `weight`: додатне число, що представляє вагу вантажу.
- `unit`: рядок, який описує одиницю виміру властивості `weight` вантажу (або `"kg"` для кілограмів, або `"lb"` для фунтів).
- `hazmat`: булеве значення, що вказує, чи потрібне поводження з небезпечними матеріалами.

Приклад об’єкта вантажного маніфесту:

```js
{
  containerId: 1,
  destination: "Monterey, California, USA",
  weight: 831,
  unit: "lb",
  hazmat: false
}
```

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні реалізувати функцію з назвою `normalizeUnits` з параметром `manifest`.
    - Функція не повинна змінювати оригінальний об’єкт manifest і завжди має повертати новий об’єкт, де `weight` нормалізовано до кілограмів, а `unit` встановлено в `"kg"`.
    - Якщо `weight` об’єкта manifest виражено у фунтах (`unit: "lb"`), функція має конвертувати `weight` у кілограми, використовуючи приблизне співвідношення `1 lb = 0.45 kg`, і відповідно оновити `unit`.
    - Якщо `weight` вже виражено в кілограмах (`unit: "kg"`), `weight` і `unit` мають залишатися без змін.
2. Ви повинні реалізувати функцію з назвою `validateManifest` з параметром `manifest`.
    - Функція не повинна змінювати оригінальний об’єкт manifest і завжди має повертати новий об’єкт.
    - Якщо вхідний manifest є дійсним (немає відсутніх або недійсних властивостей), функція має повертати порожній об’єкт.
    - Якщо вхідний manifest недійсний, функція має повертати об’єкт, що містить записи для кожної відсутньої або недійсної властивості. Відсутні властивості мають значення `"Missing"`, а недійсні — `"Invalid"`.

Приклад поверненого значення, коли вхідний об’єкт не має властивості `destination` і має недійсну властивість `weight`:

    ```js
    {
      destination: "Missing",
      weight: "Invalid"
    }
    ```

3. Реалізуйте функцію з назвою `processManifest` і параметром `manifest`. Функція має виводити в консоль:
    - Якщо об’єкт manifest є валідним, спочатку `Validation success: ${containerId}`, а потім значення `weight` у кілограмах у форматі `Total weight: ${weight} kg`. Для цього використовуйте `normalizeUnits()`.
    - Якщо об’єкт manifest невалідний, спочатку `Validation error: ${containerId}`, а потім об’єкт, який повертає виклик `validateManifest()` з цим manifest.

**Примітка:** Кожен із цих двох випадків має містити два виклики `console.log()`.

**Примітка:** Не оголошуйте `normalizeUnits`, `validateManifest` або `processManifest` за допомогою `const`, оскільки тести потребують можливості переназначення цих функцій.

# --before-each--

```js
const _validManifests = [
  {
    containerId: 4,
    destination: "Monterey, California, USA",
    weight: 831,
    unit: "lb",
    hazmat: false
  },
  {
    containerId: 5,
    destination: "Montreal, Quebec, Canada",
    weight: 151,
    unit: "kg",
    hazmat: false
  }
];

const _invalidManifests = [
  {
    containerId: -6,
    destination: 123,
    weight: 0,
    unit: "pounds",
    hazmat: true
  },
  {
    containerId: 0,
    weight: -21,
    unit: "KG"
  },
  {
    containerId: "eight",
    destination: "Guadalajara, Jalisco, Mexico",
    weight: 9001,
    hazmat: "no"
  },
  {
    destination: " ",
    weight: NaN,
    unit: "kg",
    hazmat: 1
  },
  {
    containerId: 10,
    destination: " ",
  },
];

const _testObj = {
  containerId: 1.51,
  destination: "Lavender Town",
  weight: NaN,
};

console.log = () => {};
```

# --hints--

Ви не повинні оголошувати `normalizeUnits`, `validateManifest` або `processManifest` за допомогою `const`, оскільки тести потребують можливості переназначення цих функцій.

```js
assert.doesNotThrow(() => {
  normalizeUnits = normalizeUnits;
  validateManifest = validateManifest;
  processManifest = processManifest;
});
```

Ви повинні мати функцію з назвою `normalizeUnits` з параметром `manifest`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.normalizeUnits);
const params = explorer.allFunctions.normalizeUnits.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), 'manifest');
```

Ваша функція `normalizeUnits` має повертати новий об’єкт `{ containerId: 68, destination: "Salinas", weight: 45.45, unit: "kg", hazmat: true }` при виклику з `{ containerId: 68, destination: "Salinas", weight: 101, unit: "lb", hazmat: true }`, не змінюючи оригінал.

```js
const testObj = {
  containerId: 68,
  destination: "Salinas",
  weight: 101,
  unit: "lb",
  hazmat: true
};

const expected = {
  containerId: 68,
  destination: "Salinas",
  weight: 45.45,
  unit: "kg",
  hazmat: true
};

const copy = normalizeUnits(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "normalizeUnits() did not return the expected normalized object.");
assert.notStrictEqual(copy, testObj, "normalizeUnits() should return a new object.");
assert.strictEqual(testObj.containerId, 68, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, "Salinas", "Original object should not be mutated.");
assert.strictEqual(testObj.weight, 101, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "lb", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, true, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 5, "Original object should not be mutated.");
```

Ваша функція `normalizeUnits` має повертати копію вхідного об’єкта manifest із нормалізованою до кілограмів властивістю `weight` та встановленою властивістю `unit` у значення `"kg"`. Для конвертації `weight` використовуйте приблизне співвідношення `1 lb = 0.45 kg`.

```js
_validManifests.forEach((obj, i) => {
  const result = normalizeUnits(obj);
  assert.isObject(result, "normalizeUnits() should return an object.");
  assert.notStrictEqual(result, obj, "normalizeUnits() should return a new object, not mutate the original.");
  const expectedWeight = obj.unit === "lb" ? obj.weight * 0.45 : obj.weight;
  const expected = {
    ...obj,
    weight: expectedWeight,
    unit: "kg"
  };

  assert.deepEqual(
    result,
    expected,
    `normalizeUnits() did not return the expected normalized object for manifest at index ${i}.`
  );
});
```

Ваша функція `normalizeUnits` має повертати нову копію вхідного об’єкта manifest, не змінюючи оригінал.

```js
const original = {
  containerId: 151,
  destination: "Saffron City",
  weight: 151,
  unit: "lb",
  hazmat: false
};

const expected = {
  containerId: 151,
  destination: "Saffron City",
  weight: 67.95,
  unit: "kg",
  hazmat: false
};

const copy = normalizeUnits(original);
assert.isObject(copy);
assert.deepEqual(copy, expected);
assert.notStrictEqual(copy, original, "normalizeUnits() should return a new object.");
assert.strictEqual(original.weight, 151, "Original object should not be mutated.");
assert.strictEqual(original.unit, "lb", "Original object should not be mutated.");
assert.equal(Object.keys(original).length, 5, "Original object should not be mutated.");
```

Ви повинні мати функцію з назвою `validateManifest` з параметром `manifest`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.validateManifest);
const params = explorer.allFunctions.validateManifest.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), 'manifest');
```

Ваша функція `validateManifest` має повертати новий порожній об’єкт `{}` при виклику з `{ containerId: 1, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }`.

```js
const testObj = {
  containerId: 1,
  destination: "Santa Cruz",
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for a valid manifest.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, 1, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, "Santa Cruz", "Original object should not be mutated.");
assert.strictEqual(testObj.weight, 304, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, false, "Original object should not be mutated.");
```

Ваша функція `validateManifest` має повертати порожній об’єкт `{}`, якщо вхідний об’єкт manifest є дійсним.

```js
const solution = {};
for (const obj of _validManifests) {
  assert.deepEqual(
    validateManifest(obj),
    solution,
    "validateManifest() should return an empty object for valid manifests."
  );
}
```

Ваша функція `validateManifest` має повертати новий об’єкт `{ containerId: "Missing", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` при виклику з `{}`, не змінюючи оригінал.

```js
const testObj = {};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.equal(Object.keys(testObj).length, 0, "validateManifest() should return a new object, not mutate the original.");
```

Ваша функція `validateManifest` має повертати новий об’єкт `{ containerId: "Invalid" }` при виклику з `{ containerId: null, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }`, не змінюючи оригінал.

```js
const testObj = {
  containerId: null,
  destination: "Santa Cruz",
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {
  containerId: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for null containerId.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
```

Ваша функція `validateManifest` має повертати новий об’єкт `{ containerId: "Invalid", destination: "Invalid", weight: "Invalid", unit: "Invalid", hazmat: "Invalid" }` при виклику з `{ containerId: 0, destination: 405, weight: -84, unit: "pounds", hazmat: "no" }`, не змінюючи оригінал.

```js
const testObj = {
  containerId: 0,
  destination: 405,
  weight: -84,
  unit: "pounds",
  hazmat: "no"
};

const expected = {
  containerId: "Invalid",
  destination: "Invalid",
  weight: "Invalid",
  unit: "Invalid",
  hazmat: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, 0, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, 405, "Original object should not be mutated.");
assert.strictEqual(testObj.weight, -84, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "pounds", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, "no", "Original object should not be mutated.");
```

Ваша функція `validateManifest` має повертати новий об’єкт `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` при виклику з `{ containerId: -2 }`, не змінюючи оригінал.

```js
const testObj = {
  containerId: -2,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, -2, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

Ваша функція `validateManifest` має повертати новий об’єкт `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` при виклику з `{ containerId: 3.50 }`, не змінюючи оригінал. Ви можете використовувати `Number.isInteger()` для перевірки цілих чисел.

```js
const testObj = {
  containerId: 3.50,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, 3.50, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

Ваша функція `validateManifest` має повертати новий об’єкт `{ containerId: "Missing", destination: "Invalid", weight: "Missing", unit: "Missing", hazmat: "Missing" }` при виклику з `{ destination: "  " }`, не змінюючи оригінал. Ви можете використовувати `.trim()` для видалення пробілів з рядка.

```js
const testObj = {
  destination: "  ",
};

const expected = {
  containerId: "Missing",
  destination: "Invalid",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.destination, "  ", "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

Ваша функція `validateManifest` має повертати новий об’єкт `{ containerId: "Missing", destination: "Missing", weight: "Invalid", unit: "Missing", hazmat: "Missing" }` при виклику з `{ weight: NaN }`, не змінюючи оригінал. Ви можете використовувати `Number.isNaN()` для перевірки значень `NaN`.

```js
const testObj = {
  weight: NaN,
};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Invalid",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.ok(Number.isNaN(testObj.weight), "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

Ваша функція `validateManifest` має повертати об’єкт, що описує відсутні та/або недійсні властивості, якщо вхідний об’єкт manifest недійсний.

```js
const solutions = [
  { containerId: 'Invalid', destination: 'Invalid', weight: 'Invalid', unit: 'Invalid' },
  { containerId: 'Invalid', destination: 'Missing', weight: 'Invalid', unit: 'Invalid', hazmat: 'Missing' },
  { containerId: 'Invalid', unit: 'Missing', hazmat: 'Invalid' },
  { containerId: 'Missing', destination: 'Invalid', weight: 'Invalid', hazmat: 'Invalid' },
  { destination: 'Invalid', weight: 'Missing', unit: 'Missing', hazmat: 'Missing' }
];

let i = 0;
for (const obj of _invalidManifests) {
  const result = validateManifest(obj);
  assert.deepEqual(
    result,
    solutions[i],
    "validateManifest() should return an object with errors for invalid manifests."
  );
  i += 1;
}
```

Ваша функція `validateManifest` має повертати новий об’єкт без зміни оригіналу.

```js
const invalidOriginal = {
  containerId: -151,
  destination: "",
  weight: -151,
  unit: "kilograms",
  hazmat: "idk"
};

const validOriginal = {
  containerId: 151,
  destination: "Pallet Town",
  weight: 151,
  unit: "kg",
  hazmat: false
};

const res0 = validateManifest(invalidOriginal);
assert.isObject(res0);
assert.notStrictEqual(res0, invalidOriginal, "validateManifest() should return a new object.");
assert.strictEqual(invalidOriginal.containerId, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.destination, "", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.weight, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.unit, "kilograms", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.hazmat, "idk", "Original object should not be mutated.");
assert.equal(Object.keys(invalidOriginal).length, 5, "Original object should not be mutated.");

const res1 = validateManifest(validOriginal);
assert.isObject(res1);
assert.notStrictEqual(res1, validOriginal, "validateManifest() should return a new object.");
assert.strictEqual(validOriginal.containerId, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.destination, "Pallet Town", "Original object should not be mutated.");
assert.strictEqual(validOriginal.weight, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(validOriginal.hazmat, false, "Original object should not be mutated.");
assert.equal(Object.keys(validOriginal).length, 5, "Original object should not be mutated.");
```

Ви повинні мати функцію з назвою `processManifest` з параметром `manifest`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.processManifest);
const params = explorer.allFunctions.processManifest.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), 'manifest');
```

Ваша функція `processManifest` має спочатку вивести в консоль `Validation success: 55`, а потім `Total weight: 180 kg` при виклику з `{ containerId: 55, destination: "Carmel", weight: 400, unit: "lb", hazmat: false }`.

```js
const testObj = {
  containerId: 55,
  destination: "Carmel",
  weight: 400,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);

  const expectedCalls = [
    ["Validation success: 55"],
    ["Total weight: 180 kg"]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation success message.");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the normalized weight in kg.");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Ваша функція `processManifest` має спочатку вивести повідомлення про успіх, `Validation success: ${containerId}`, якщо вхідний об’єкт manifest є дійсним.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Ваша функція `processManifest` має нормалізувати дійсний вхідний об’єкт manifest у кілограми за допомогою `normalizeUnits()` і потім вивести `Total weight: ${weight} kg`.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ];
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Ваша функція `processManifest` повинна виводити повідомлення про успішну перевірку з `containerId` об’єкта, а потім виводити його `weight` у кілограмах, якщо вхідний об’єкт маніфесту є коректним. Для цього потрібно використовувати `normalizeUnits()`, і це має бути реалізовано двома викликами `console.log()`.

```js
const originalNormalize = normalizeUnits;
let called = false;
try {
  normalizeUnits = function (...args) {
    called = true;
    return originalNormalize(...args);
  };

  processManifest(_validManifests[0]);

  assert.isTrue(
    called,
    "processManifest() should call normalizeUnits() when the manifest is valid."
  );

} catch (err) {
  assert.fail(err);
} finally {
  normalizeUnits = originalNormalize;
}

const spy = __helpers.spyOn(console, "log");

try {
  for (const obj of _validManifests) {
    processManifest(obj);
  }
  const expectedCalls = [
    ["Validation success: 4"],
    ["Total weight: 373.95 kg"],
    ["Validation success: 5"],
    ["Total weight: 151 kg"],
  ];
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Ваша функція `processManifest` має спочатку вивести `Validation error: -88`, а потім об’єкт `{ containerId: "Invalid", weight: "Invalid", unit: "Missing", hazmat: "Missing" }` при виклику з `{ containerId: -88, destination: "Soledad", weight: NaN }`.

```js
const testObj = {
  containerId: -88,
  destination: "Soledad",
  weight: NaN
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: -88"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message.");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object.");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Ваша функція `processManifest` має спочатку вивести `Validation error: undefined`, а потім об’єкт `{ containerId: "Missing", weight: "Missing", unit: "Missing" }` при виклику з `{ destination: "Watsonville", hazmat: true }`.

```js
const testObj = {
  destination: "Watsonville",
  hazmat: true
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: undefined"],
    [{
        containerId: "Missing",
        weight: "Missing",
        unit: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message.");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object.");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Ваша функція `processManifest` має спочатку вивести повідомлення про помилку, `Validation error: ${containerId}`, якщо вхідний об’єкт manifest недійсний.

```js
const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Ваша функція `processManifest` також має вивести об’єкт, який повертається викликом `validateManifest()` з оригінальним об’єктом manifest, якщо вхідний об’єкт manifest недійсний. Викликайте `console.log()` безпосередньо з поверненим об’єктом.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest() should call validateManifest() when the manifest is invalid."
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Ваша функція `processManifest` має вивести повідомлення про помилку з `containerId` об’єкта, а потім вивести об’єкт, який повертається викликом `validateManifest()` з вхідним об’єктом, якщо вхідний об’єкт manifest недійсний. Це має включати два виклики `console.log()`.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest() should call validateManifest() when the manifest is invalid."
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function normalizeUnits(manifest) {
  const normalized = {
    ...manifest
  }
  if (normalized.unit === "lb") {
    normalized.weight = normalized.weight * 0.45;
    normalized.unit = "kg";
  }
  return normalized;
}

function validateManifest(manifest) {
  const result = {}

  if (manifest.containerId === undefined) {
    result.containerId = "Missing";
  } else if (
    typeof manifest.containerId !== "number" ||
    !Number.isInteger(manifest.containerId) ||
    manifest.containerId <= 0
  ) {
    result.containerId = "Invalid";
  }
  if (manifest.destination === undefined) {
    result.destination = "Missing";
  } else if (typeof manifest.destination !== "string" || manifest.destination.trim() === "") {
    result.destination = "Invalid";
  }
  if (manifest.weight === undefined) {
    result.weight = "Missing";
  } else if (typeof manifest.weight !== "number" || manifest.weight <= 0 || Number.isNaN(manifest.weight)) {
    result.weight = "Invalid";
  }
  if (manifest.unit === undefined) {
    result.unit = "Missing";
  } else if (manifest.unit !== "kg" && manifest.unit !== "lb") {
    result.unit = "Invalid";
  }
  if (manifest.hazmat === undefined) {
    result.hazmat = "Missing";
  } else if (typeof manifest.hazmat !== "boolean") {
    result.hazmat = "Invalid";
  }

  return result;
}

function processManifest(manifest) {
  const validation = validateManifest(manifest);

  if (Object.keys(validation).length > 0) {
    console.log(`Validation error: ${manifest.containerId}`);
    console.log(validation);
  } else {
    const normalized = normalizeUnits(manifest);
    console.log(`Validation success: ${manifest.containerId}`);
    console.log(`Total weight: ${normalized.weight} kg`);
  }
}
```
