---
id: 686b9720ee1d032bd77a480a
title: Eseguire il debug di un validatore ISBN
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

L'ISBN (International Standard Book Number) è un identificatore unico assegnato ai libri commerciali. Può essere lungo 10 o 13 cifre, e l'ultima cifra è una cifra di controllo calcolata dalle altre cifre.

Camperbot ha provato a costruire un proprio validatore ISBN. Tuttavia, ha commesso alcuni errori lungo il percorso.

In questo laboratorio, correggerai il codice esistente per farlo funzionare correttamente.

Comportamento previsto:
Quando l'utente esegue il programma, verrà mostrato il prompt `Enter ISBN and length: `. L'utente può inserire il codice ISBN che vuole validare nel formato `ISBN,length`. Il codice ISBN non deve contenere trattini, seguito dalla sua lunghezza (`10` o `13`), separati da una virgola.

Esempi di input:
`1530051126,10` per codici ISBN-10.
`9781530051120,13` per codici ISBN-13.

Come trovare la cifra di controllo:

- Non è necessario conoscere la logica dettagliata del calcolo in questo laboratorio. Le funzioni `calculate_check_digit_10` e `calculate_check_digit_13` si occuperanno del calcolo e restituiranno la cifra di controllo prevista come stringa.
- La cifra di controllo per i codici ISBN-10 può essere un numero da `0` a `9` o una lettera maiuscola `X`.
- La cifra di controllo per i codici ISBN-13 può essere un numero da `0` a `9`.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi correggere il `IndentationError` nel codice attuale.
1. Anche se l'utente non inserisce un valore separato da virgola, il programma deve gestire il `IndexError` senza bloccarsi.
1. Quando l'utente non inserisce un valore separato da virgola, deve vedere il messaggio `Enter comma-separated values.` nella console e il programma deve terminare.
1. Anche se l'utente inserisce un valore non numerico per la lunghezza, il programma deve gestire il `ValueError` senza bloccarsi.
1. Quando l'utente inserisce un valore non numerico per la lunghezza, deve vedere il messaggio `Length must be a number.` nella console e il programma deve terminare.
1. Devi correggere l'errore off-by-one nella funzione `validate_isbn`.
1. Devi correggere il `TypeError` nel codice attuale che si verifica quando l'utente inserisce un codice ISBN valido.
1. Devi correggere il `IndexError` nel codice attuale quando l'utente inserisce un codice ISBN valido.
1. Anche se l'utente inserisce un codice ISBN errato con caratteri diversi dai numeri, il programma deve gestire il `ValueError` che si verifica senza bloccarsi.
1. Quando l'utente inserisce un codice ISBN errato con caratteri diversi dai numeri, deve vedere il messaggio `Invalid character was found.` nella console.
1. Quando l'utente inserisce `1530051126,10`, deve vedere il messaggio `Valid ISBN Code.`
1. Quando l'utente inserisce `9781530051120,13`, deve vedere il messaggio `Valid ISBN Code.`

**Importante**: dovrai commentare la chiamata a `main()` nello spazio globale affinché i test funzionino correttamente.

Quando completi il progetto, l'utente dovrebbe vedere i seguenti messaggi a seconda dei valori inseriti.

| Codice ISBN | Lunghezza | Messaggio | Esempio di input |
|-------------|-----------|-----------|------------------|
| Valido | Valido | `Valid ISBN Code.` | `1530051126,10` |
| Numero non valido | Valido | `Invalid ISBN Code.` | `1530051125,10` |
| Non corrisponde alla lunghezza specificata o è vuoto | Valido | `ISBN-10 code should be 10 digits long.` o `ISBN-13 code should be 13 digits long.`, a seconda della lunghezza inserita. | `9781530051120,10` o `1530051126,13` |
| Contiene caratteri non numerici (eccetto la cifra di controllo) | Valido | `Invalid character was found.` | `15-0051126,10` |
| Qualsiasi | Numero non valido | `Length should be 10 or 13.` | `1530051126,9` |
| Qualsiasi | Contiene caratteri non numerici o è vuoto | `Length must be a number.` | `1530051125,A` |
| Non separato da virgola | Non separato da virgola | `Enter comma-separated values.` | `1530051125` |

Puoi usare i seguenti valori per testare manualmente il tuo codice se vuoi.

Esempi di input per codici ISBN-10 validi:
`1530051126,10`
`9971502100,10`
`080442957X,10`

Esempi di input per codici ISBN-13 validi:
`9781530051120,13`
`9781947172104,13`

# --hints--

Dovresti commentare la chiamata alla funzione `main` per permettere agli altri test di funzionare correttamente.

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

Dovresti avere una funzione `validate_isbn`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

Dovresti avere una funzione `calculate_check_digit_10`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

Dovresti avere una funzione `calculate_check_digit_13`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

Dovresti avere una funzione `main`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

Quando l'utente inserisce un valore che non è separato da virgola, dovresti vedere il messaggio `Enter comma-separated values.` nella console.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Quando l'utente inserisce un valore non numerico per la lunghezza, dovresti vedere il messaggio `Length must be a number.` nella console.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Quando l'utente inserisce un codice ISBN errato con caratteri diversi dai numeri, dovresti vedere il messaggio `Invalid character was found.` nella console.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Quando l'utente inserisce `1530051126,10`, dovresti vedere il messaggio `Valid ISBN Code.` nella console.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Quando l'utente inserisce `9781530051120,13`, dovresti vedere il messaggio `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Quando l'utente inserisce `1530051125,10`, dovresti vedere il messaggio `Invalid ISBN Code.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

Quando l'utente inserisce `9781530051120,10`, dovresti vedere il messaggio `ISBN-10 code should be 10 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

Quando l'utente inserisce `1530051126,13`, dovresti vedere il messaggio `ISBN-13 code should be 13 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

Quando l'utente inserisce `15-0051126,10`, dovresti vedere il messaggio `Invalid character was found.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Quando l'utente inserisce `1530051126,9`, dovresti vedere il messaggio `Length should be 10 or 13.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

Quando l'utente inserisce `1530051125,A`, dovresti vedere il messaggio `Length must be a number.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Quando l'utente inserisce `1530051125`, dovresti vedere il messaggio `Enter comma-separated values.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Quando l'utente inserisce `9971502100,10`, dovresti vedere il messaggio `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Quando l'utente inserisce `080442957X,10`, dovresti vedere il messaggio `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Quando l'utente inserisce `9781947172104,13`, dovresti vedere il messaggio `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
