---
id: 69f0a37c3bf194a7c9c62492
title: 创建设备借用账本
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

你是一名软件开发人员，负责创建一个账本应用来管理公司服务台提供的 IT 硬件设备。该账本是一个对象，其键是资产标签，值是对象：

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该将账本建模为一个对象，每个设备以其资产标签为键，具有 `type`、`status`、`borrower`（包含 `name` 和 `email`）以及 `dueDate` 属性。

1. 你应该实现一个 `checkoutDevice(ledger, assetTag, borrower)` 函数，该函数接收一个 ledger 对象，并使用 `serializeLedger` 和 `loadLedger` 克隆它，而不改变原始对象。在克隆对象上，应设置 `borrower` 对象的 `name` 和 `email`，将 `status` 更新为 `"CheckedOut"`，并返回一个形式为 `{ ledger: updatedLedger, message: confirmationString }` 的对象，其中 `message` 包含资产标签和 `borrower` 对象的 `name`。

   - 如果账本中不存在 `assetTag`，函数应返回未更改的账本，并附带包含资产标签且说明未找到的 `message`。
   - 如果设备已经是 `"CheckedOut"` 状态，函数应返回未更改的账本，保留当前借用者，并附带包含资产标签且说明设备已被借出的 `message`。

1. 你应该实现一个 `checkinDevice(ledger, assetTag)` 函数，该函数接收一个 ledger 对象，并使用 `serializeLedger` 和 `loadLedger` 克隆它，而不改变原始对象。在克隆对象上，应清空 `borrower` 对象的 `name` 和 `email`，将 `dueDate` 重置为 `""`，将 `status` 更新为 `"CheckedIn"`，并返回一个形式为 `{ ledger: updatedLedger, message: confirmationString }` 的对象，其中 `message` 包含资产标签。

   - 如果账本中不存在 `assetTag`，函数应返回未更改的账本，并附带包含资产标签且说明未找到的 `message`。

1. 你应该实现一个 `listOverdueDevices(ledger, today)` 函数，该函数接收一个账本对象并返回一个逾期设备对象的数组，按 `dueDate` 升序排序。当设备的 `status` 是 `"CheckedOut"` 且 `dueDate` 严格早于 `today` 时，该设备即为逾期。`today` 和 `dueDate` 都是按月/日/年顺序的字符串，月和日可以有或没有前导零（例如，`"9/5/2025"` 和 `"09/05/2025"` 都有效）。

1. 你不应在代码中任何地方使用 JavaScript 的 `Date` 对象。

1. 你应该实现一个 `serializeLedger(ledger)` 函数，将账本对象转换为 JSON 字符串，以及一个 `loadLedger(json)` 函数，将 JSON 字符串转换回 JavaScript 对象。

# --hints--

你应该创建一个名为 `checkoutDevice` 的函数。

```js
assert.isFunction(checkoutDevice);
```

你应该创建一个名为 `checkinDevice` 的函数。

```js
assert.isFunction(checkinDevice);
```

你应该创建一个名为 `listOverdueDevices` 的函数。

```js
assert.isFunction(listOverdueDevices);
```

你应该创建一个名为 `serializeLedger` 的函数。

```js
assert.isFunction(serializeLedger);
```

你应该创建一个名为 `loadLedger` 的函数。

```js
assert.isFunction(loadLedger);
```

你的 `checkoutDevice` 函数不应修改原始账本。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

你的 `checkoutDevice` 函数应使用 `serializeLedger` 和 `loadLedger` 来克隆 ledger。

```js
const source = __helpers.removeJSComments(checkoutDevice.toString());
assert.match(source, /\bserializeLedger\s*\(/);
assert.match(source, /\bloadLedger\s*\(/);
```

你的 `checkoutDevice` 函数应设置借用者的 `name` 和 `email`，并将 `status` 设置为 `"CheckedOut"`。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

你的 `checkoutDevice` 函数应返回包含资产标签和借用者 `name` 的确认消息。

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

你的 `checkoutDevice` 函数应在返回的账本中保留无关设备和未更改的字段。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

当资产标签未找到时，你的 `checkoutDevice` 函数应返回未更改的账本和包含资产标签的消息。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

当设备已被借出时，你的 `checkoutDevice` 函数应返回未更改的账本，保留当前借用者，并返回包含资产标签的消息。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

你的 `checkinDevice` 函数不应修改原始账本。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

你的 `checkinDevice` 函数应使用 `serializeLedger` 和 `loadLedger` 来克隆 ledger。

```js
const source = __helpers.removeJSComments(checkinDevice.toString());
assert.match(source, /\bserializeLedger\s*\(/);
assert.match(source, /\bloadLedger\s*\(/);
```

你的 `checkinDevice` 函数应清空借用者的 `name` 和 `email`，将 `dueDate` 重置为 `""`，并将 `status` 设置为 `"CheckedIn"`。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

你的 `checkinDevice` 函数应返回包含资产标签的确认消息。

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

你的 `checkinDevice` 函数应在返回的账本中保留无关设备和未更改的字段。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

当资产标签未找到时，你的 `checkinDevice` 函数应返回未更改的账本和包含资产标签的消息。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

你的 `listOverdueDevices` 函数应返回一个数组。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

你的 `listOverdueDevices` 函数应只返回 `status` 为 `"CheckedOut"` 的逾期设备。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

你的 `listOverdueDevices` 函数应返回按 `dueDate` 升序排序的逾期设备。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

你的 `listOverdueDevices` 函数应能正确排序，即使月或日没有前导零。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

你不应在代码中任何地方使用 JavaScript 的 `Date` 对象。

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

你的 `serializeLedger` 函数应返回表示传入账本的 JSON 字符串。

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

你的 `loadLedger` 函数应返回 JSON 字符串所表示的 JavaScript 对象。

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = loadLedger(serializeLedger(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = loadLedger(serializeLedger(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
