---
id: 68e2c945cc1d8e778152be31
title: Spieler-Interface erstellen
challengeType: 27
dashedName: build-a-player-interface
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine abstrakte Klasse namens `Player` definieren, die von der `abc.ABC`-Klasse erbt.
1. Die `Player`-Klasse sollte eine `__init__`-Methode haben, die Folgendes setzt:

   - Das `moves`-Attribut auf eine leere Liste.
   - Das `position`-Attribut auf `(0, 0)`.
   - Das `path`-Attribut auf eine Liste, die die Anfangsposition enthält.

1. Die `Player`-Klasse sollte eine Methode namens `make_move` haben, die:

   - `random.choice` verwendet, um eine zufällige Bewegung aus dem `moves`-Attribut (definiert in der konkreten Klasse) zu erhalten.
   - Die Werte der ausgewählten Bewegung zur aktuellen Position addiert und das `position`-Attribut aktualisiert.
   - Das neue `position`-Tuple zum `path`-Attribut hinzufügt.
   - Das neue `position` zurückgibt.

1. Die `Player`-Klasse sollte eine abstrakte Methode namens `level_up` haben, die in konkreten Klassen implementiert wird.
1. Sie sollten eine `Pawn`-Klasse definieren, die von der `Player`-Klasse erbt.
1. Die `Pawn`-Klasse sollte `super()` verwenden, um die `__init__`-Methode der Elternklasse aufzurufen und dann das `moves`-Attribut auf eine Liste von Tuples zu setzen, die `x, y`-Koordinaten darstellen.
1. Jedes Koordinaten-Tuple sollte eine Bewegung von `1` Einheit in den folgenden Richtungen darstellen: nach oben, nach unten, nach links, nach rechts.
1. Die `Pawn`-Klasse sollte eine konkrete `level_up`-Methode implementieren, indem sie weitere Bewegungen zum `moves`-Attribut hinzufügt. Die hinzugefügten Bewegungen sollten die vier diagonalen Bewegungen darstellen (zum Beispiel `1` Einheit nach unten plus `1` Einheit nach links).

**Hinweis:** Standardbibliotheksmodule sollten ohne Aliasnamen importiert werden. Tests, die sich auf die `Player`-Klasse beziehen, schlagen fehl, bis die `Pawn`-Klasse instanziierbar wird.

# --hints--

Sie sollten eine Klasse namens `Player` haben.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Player")`) })
```

Die `Player`-Klasse sollte von der `ABC`-Klasse des `abc`-Moduls erben.

```js
({ test: () => runPython(`
from abc import ABC
import ast
cls = _Node(_code).find_class("Player")
assert cls.inherits_from("ABC") or cls.inherits_from("abc.ABC")
`) })
```

Die `Player`-Klasse sollte eine `__init__`-Methode haben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("__init__")`) })
```

Die `Player`-`__init__`-Methode sollte einen einzelnen Parameter `self` haben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("__init__").has_args("self")`) })
```

Die `Player`-`__init__`-Methode sollte `self.moves` auf eine leere Liste setzen.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.moves = []") or init.has_stmt("self.moves = list()")
`) })
```

Die `Player`-`__init__`-Methode sollte `self.position` auf `(0, 0)` setzen.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.position = (0, 0)")
`) })
```

Die `Player`-`__init__`-Methode sollte `self.path` auf eine Liste setzen, die die Anfangsposition enthält.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.path = [self.position]")
`) })
```

Die `Player`-Klasse sollte eine `make_move`-Methode haben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("make_move")`) })
```

Die `Player`-`make_move`-Methode sollte einen einzelnen Parameter `self` haben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("make_move").has_args("self")`) })
```

Die `Player`-`make_move`-Methode sollte `random.choice` verwenden, um eine zufällige Bewegung aus dem `moves`-Attribut zu erhalten.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").block_has_call("choice", "make_move")`) })
```

Die `Player`-`make_move`-Methode sollte das `position`-Attribut aktualisieren, indem sie die Koordinaten der zufällig ausgewählten Bewegung hinzufügt.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.position == (1, 2)
  p.make_move()
  assert p.position == (2, 4)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Die `Player`-`make_move`-Methode sollte das neue `position`-Tuple zum `path`-Attribut hinzufügen.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.path == [(0, 0), (1, 2)]
  p.make_move()
  assert p.path == [(0, 0), (1, 2), (2, 4)]

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Die `Player`-`make_move`-Methode sollte das aktualisierte `position`-Attribut zurückgeben.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  assert p.make_move() == (1, 2)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Die `Player`-Klasse sollte eine `level_up`-Methode haben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("level_up")`) })
```

Die `Player`-`level_up`-Methode sollte einen einzelnen Parameter `self` haben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("level_up").has_args("self")`) })
```

Die `Player`-`level_up`-Methode sollte eine abstrakte Methode sein.

```js
({ test: () => runPython(`
target = _Node(_code).find_class("Player").find_function("level_up")
assert target.has_decorators("abstractmethod") or target.has_decorators("abc.abstractmethod")
`) })
```

Die `Player`-Klasse sollte eine abstrakte Klasse sein.

```js
({ test: () => runPython(`
try:
  Player()
except TypeError as e:
  assert str(e) == "Can't instantiate abstract class Player with abstract method level_up"
else:
  assert False, "Player class should not be instantiable"
`) })
```

Sie sollten eine Klasse namens `Pawn` haben.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Pawn")`) })
```

Die `Pawn`-Klasse sollte von der `Player`-Klasse erben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").inherits_from("Player")`) })
```

Die `Pawn`-Klasse sollte eine `__init__`-Methode haben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("__init__")`) })
```

Die `Pawn`-`__init__`-Methode sollte einen einzelnen Parameter `self` haben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_args("self")`) })
```

Die `Pawn`-`__init__`-Methode sollte die `__init__`-Methode der Elternklasse mit der `super`-Funktion aufrufen.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_stmt("super().__init__()")`) })
```

Die `Pawn`-`__init__`-Methode sollte das `moves`-Attribut auf eine Liste von Tuples setzen, die `x, y`-Koordinaten darstellen, wobei jedes Koordinaten-Tuple eine Bewegung von `1` Einheit in den folgenden Richtungen darstellt: nach oben, nach unten, nach links, nach rechts.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]
assert isinstance(p.moves, list)
assert len(p.moves) == 4
assert all(move in p.moves for move in moves)
`) })
```

Die `Pawn`-Klasse sollte eine `level_up`-Methode haben.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("level_up")`) })
```

Die `Pawn`-`level_up`-Methode sollte die vier diagonalen Bewegungen von `1` Einheit zum `moves`-Attribut hinzufügen.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1), (1, 1), (1, -1), (-1, 1), (-1, -1)]
p.level_up()
assert isinstance(p.moves, list)
assert len(p.moves) == 8
assert all(move in p.moves for move in moves)
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
from abc import ABC, abstractmethod
from random import choice


class Player(ABC):
    def __init__(self):
        self.moves = []
        self.position = (0, 0)
        self.path = [self.position]

    def make_move(self):
        x, y = self.position
        move = choice(self.moves)
        new_x = move[0] + x
        new_y = move[1] + y
        self.position = (new_x, new_y)
        self.path.append(self.position)
        return self.position

    @abstractmethod
    def level_up(self):
        pass


class Pawn(Player):
    def __init__(self):
        super().__init__()
        self.moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]

    def level_up(self):
        self.moves.extend([(1, 1), (1, -1), (-1, 1), (-1, -1)])
```
