---
id: 69757cc0faae0152c1418aad
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie den Umgang mit Zahlen und mathematischen Operationen, um einen Rechnungsteiler zu erstellen. Dieses Tool berechnet, wie viel jede Person nach Hinzufügen der Essenskosten und eines Trinkgelds schuldet.

Um zu beginnen, benötigen Sie eine Möglichkeit, den Gesamtbetrag zu verfolgen, während Kosten hinzugefügt werden.

Erstellen Sie eine Variable mit dem Namen `running_total` und weisen Sie ihr den Wert `0` zu.

# --hints--

Sie sollten eine Variable mit dem Namen `running_total` haben.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

Sie sollten der Variablen `0` den ganzzahligen Wert `running_total` zuweisen.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
