---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Jeder Kurs hat einen Preis. Sie müssen diese Beträge in Variablen speichern, um sie später zu verwenden. Da diese Beträge Cent enthalten, verwenden Sie den Typ float, der zur Darstellung von Dezimalzahlen dient. Hier ist ein Beispiel für eine Variable mit einem float-Wert:

```py
change = 2.35
```

Erstellen Sie vier Variablen: `appetizers` mit dem Wert `37.89`, `main_courses` mit dem Wert `57.34`, `desserts` mit dem Wert `39.39` und `drinks` mit dem Wert `64.21`. 

# --hints--

Sie sollten eine Variable namens `appetizers` definieren und ihr den Wert `37.89` zuweisen.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

Sie sollten eine Variable namens `main_courses` definieren und ihr den Wert `57.34` zuweisen.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

Sie sollten eine Variable namens `desserts` definieren und ihr den Wert `39.39` zuweisen.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

Sie sollten eine Variable namens `drinks` definieren und ihr den Wert `64.21` zuweisen.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
