---
id: 6976cc0b0c135686a4620b94
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Als Nächstes müssen Sie die Anzahl der Personen berücksichtigen, die die Rechnung teilen. Speichern Sie diesen Wert in einer Variablen, wie Sie es im vorherigen Schritt getan haben.

Erstellen Sie eine Variable mit dem Namen `num_of_friends` und weisen Sie ihr den Wert von `4` zu. Diese Variable wird später im Workshop verwendet, um die endgültige Aufteilung zu berechnen.

# --hints--

Sie sollten eine Variable mit dem Namen `num_of_friends` haben.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

Sie sollten die Ganzzahl `4` Ihrer Variablen `num_of_friends` zuweisen.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
