---
id: 6977a8bbcbdd557886bd0718
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Mit dem jetzt enthaltenen Trinkgeld haben Sie den Gesamtbetrag für die gesamte Gruppe. Sie müssen ermitteln, wie viel jede Person schuldet, indem Sie die Gesamtrechnung durch die Anzahl der Freunde teilen.

In Python verwenden Sie den Schrägstrich `/`, um eine Division durchzuführen. Zum Beispiel:

```py
half = 10 / 2
```

Erstellen Sie eine Variable namens `final_bill` und weisen Sie ihr das Ergebnis der Division von `running_total` durch `num_of_friends` zu.

Verwenden Sie abschließend die `print()`-Funktion, um den String `Bill per person:` gefolgt von einem Leerzeichen und dem Wert von `final_bill` auszugeben.

# --hints--

Sie sollten eine Variable namens `final_bill` definieren.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

Sie sollten den `/`-Operator verwenden, um `running_total` durch `num_of_friends` zu teilen und das Ergebnis Ihrer `final_bill`-Variable zuzuweisen.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

Sie sollten den String `Bill per person:` gefolgt von einem Leerzeichen und der `final_bill`-Variable ausgeben.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
