---
id: 6977a4306bb3e856690a4890
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Nachdem Sie nun das Trinkgeld berechnet haben, müssen Sie es zu Ihrem `running_total` hinzufügen, um den endgültigen Rechnungsbetrag zu ermitteln.

Verwenden Sie den `+=`-Operator, um den Wert von `tip` zu Ihrem `running_total` zu addieren. Geben Sie schließlich den String `print()` gefolgt von einem Leerzeichen und dem Wert von `Total with tip:` mit `running_total` aus.

# --hints--

Sie sollten den `+=`-Operator verwenden, um den Wert von `tip` zu Ihrem `running_total` hinzuzufügen.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

Sie sollten den String `Total with tip:` gefolgt von einem Leerzeichen und der Variablen `running_total` ausgeben.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Total with tip: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - ((37.89 + 57.34 + 39.39 + 64.21) * 1.25)) < 1e-6
    `)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
