---
id: 69779dcb44016eccc05c15a4
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Der Service war ausgezeichnet, daher entscheidet sich die Gruppe, 25 % Trinkgeld zu geben. Um einen Prozentsatz in Python zu berechnen, können Sie den Gesamtbetrag mit dem Dezimalwert des Prozentsatzes multiplizieren.

Zum Beispiel multiplizieren Sie, um 10 % eines Werts zu finden, diesen mit `0.10` unter Verwendung des `*`-Operators:

```py
tax = total * 0.10
```

Erstellen Sie eine Variable namens `tip` und weisen Sie ihr das Ergebnis der Multiplikation von `running_total` mit `0.25` zu.

Verwenden Sie abschließend `print()`, um den String `Tip amount:` gefolgt von einem Leerzeichen und dem Wert Ihrer `tip`-Variable auszugeben.

# --hints--

Sie sollten eine Variable namens `tip` haben.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

Ihre `tip`-Variable sollte das Ergebnis von `running_total * 0.25` sein.

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

Sie sollten den String `Tip amount:` gefolgt von einem Leerzeichen und der `tip`-Variable ausgeben.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
