---
id: 660ef5105b8ba095307a0e50
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora vuelve a imprimir tu variable `character` en la consola. Debes ver la cadena de texto `"Hello"`, luego la cadena de texto `"World"` en la consola.

# --hints--

Debes usar `console.log` por segunda vez.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Debes imprimir `character` en la consola por segunda vez.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


Tu nuevo `console.log()` debería venir después de tu reasignación.

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
