---
id: 660f50a21fe7645252804f2b
title: Paso 112
challengeType: 1
dashedName: step-112
---

# --description--

¿Qué pasaría si tuvieras una forma de alternar entre una piríamida inversa y una piríamida estándar?

Comienza declarando una variable `inverted` y asignándole el valor `true`. No estás cambiando esta variable en tu código, pero necesitarás usar `let` para que nuestras pruebas puedan modificarla más tarde.

# --hints--

Debes declarar una variable `inverted` con `let`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

Debes inicializar `inverted` con el valor `true`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
