---
id: 660f17294346b7d69e79db3d
title: Paso 32
challengeType: 1
dashedName: step-32
---

# --description--

Ahora estás listo para comenzar a construir tu generador de pirámide. Tu constante `character` servirá como el bloque de construcción de la pirámide.

Es posible que `"Hello"` no funcione muy bien para eso. Cambia el valor de `character` para que sea el carácter `"#"`.

# --hints--

Tu constante `character` debe ser un símbolo #.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
