---
id: 660f374d532dc41189cc9cc2
title: Paso 69
challengeType: 1
dashedName: step-69
---

# --description--

Ahora es el momento de hacer un poco de matemáticas. Consideremos una pirámides de tres filas. Si queremos que se centre, se parecería a algo como:

```md
··#··
·###·
#####
```

Los espacios vacíos han sido reemplazados por interpuntos, o puntos medios, para la legibilidad. Si extrapolas el patrón, puedes ver que los espacios al principio y al final de una fila siguen un patrón.

Actualiza tus cadenas de texto vacías para que se repitan `rowCount - rowNumber` veces.

Abre la consola para ver el resultado.

# --hints--

Debes llamar a `.repeat()` en tus cadenas `" "` para repetirlas `rowCount - rowNumber` veces.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
