---
id: 660f46460f9c36330ebc07d8
title: Paso 88
challengeType: 1
dashedName: step-88
---

# --description--

El operador de igualdad puede conducir a un comportamiento extraño en JavaScript. Por ejemplo, `"0" == 0` es verdadero, aunque uno es una cadena y uno es un número.

El operador de <dfn>igualdad estricta</dfn> `===` se utiliza para comprobar si dos valores son iguales y comparten el mismo tipo. Como norma general, este es el operador de igualdad que debes utilizar siempre. Con el operador de igualdad estricta, `"0" === 0` se convierte en falso, porque aunque pueden tener el mismo valor de cero, no son del mismo tipo.

Actualiza tu condición `done == count` y utiliza el operador de igualdad estricta.

# --hints--

Tu condición `if` debe usar una igualdad estricta.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
