---
id: 660f505d02b2bd513a1c3468
title: Paso 111
challengeType: 1
dashedName: step-111
---

# --description--

Tu pirámide ya no está invertida. Esto se debe a que está agregando nuevas filas al **final** del array.

Actualiza el cuerpo del bucle para añadir nuevas filas al comienzo del array.

# --hints--

Debes usar el método `unshift` de tu array `rows`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

Debes pasar una llamada `padRow()` como argumento para tu método `.unshift()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

Debes pasar `i` como primer argumento a tu llamada a `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

Debes pasar `count` como segundo argumento a tu llamada a `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
