---
id: 660f207334fabaeac3269c38
title: Paso 41
challengeType: 1
dashedName: step-41
---

# --description--

Para manipular la cadena de texto `result`, usarás un tipo de bucle diferente. Específicamente, un bucle `for...of`, que iterará sobre cada elemento en un objeto iterable y lo asignará temporalmente a una variable.

La sintaxis de un bucle `for...of` se ve así:

```js
for (const value of iterable) {

}
```

Ten en cuenta que puedes usar `const` porque la variable sólo existe para una sola iteración, no durante todo el bucle.

Crea un bucle `for...of` para iterar a través de tu array `rows`, asignando cada valor a una variable `row`.

# --hints--

Debes usar otra palabra reservada en tu `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

Debes declarar una constante `row`.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

Debes usar `const` para declarar tu constante `row`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Tu bucle `for...of` debe declarar tu constante `row`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

Tu variable `row` debe extraerse de `rows` usando la palabra reservada `of`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

El cuerpo de tu bucle `for...of` debería estar vacío.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
