---
id: 660f09a2694b59c3a10ee304
title: Pasos 23
challengeType: 1
dashedName: step-23
---

# --description--

Por ahora, elimina tu primer console.log y tu `rows[rows.length - 1]`. Deja la segunda sentencia `rows` para más tarde.

# --hints--

Debes eliminar tu llamada de `console.log(rows[0])`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

Debes eliminar tu reasignación `rows[rows.length - 1]`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

No debes eliminar tu llamada de `console.log(rows)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
