---
id: 660f2eccfe3f820304af1b39
title: Paso 64
challengeType: 1
dashedName: step-64
---

# --description--

Para saber cómo formatear una fila, su función `padRow` necesitará saber en qué número de fila se encuentra y cuántas filas en total se están generando.

La mejor manera de hacerlo es creando parámetros de función para ello. Dale a tu función `padRow` un parámetro `rowNumber` y `rowCount`. Los parámetros múltiples se separan con una coma:

```js
function name(first, second) {

}
```

# --hints--

Tu función `padRow` debe tener un parámetro `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

Debes agregar una coma después del parámetro `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

Tu función `padRow` debe tener un parámetro `rowCount`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
