---
id: 660f0f980e98e8cf77f1ce31
title: Paso 30
challengeType: 1
dashedName: step-30
---

# --description--

Cambia la declaración `rows` para que se le asigne un array vacío de nuevo.

También, cambia tu cadena de texto `'Hello'` para usar comillas dobles de nuevo. Generalmente, no importa cuál de los dos prefieras, pero querrás ser coherente en esa elección a lo largo de tu proyecto.

# --hints--

Tu array `rows` debe estar vacío.

```js
assert.empty(rows);
```

Tu cadena de texto `"Hello"` debería tener comillas dobles.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
