---
id: 660f46b9c417a8341729a3ab
title: Paso 69
challengeType: 1
dashedName: step-89
---

# --description--

Cuando `done` haya alcanzado el valor de `count`, queremos que el bucle deje de ejecutarse.

Dentro del cuerpo de tu `if`, asigna el booleano `false` a tu variable `continueLoop`.

# --hints--

El cuerpo de tu `if` debe asignar `false` a tu variable `continueLoop`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{\s*continueLoop\s*=\s*false;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done === count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
