---
id: 660f061d259bbebc37461080
title: Paso 20
challengeType: 1
dashedName: step-20
---

# --description--

Puedes acceder a los valores dentro de un arreglo usando el valor del <dfn>índice</dfn>. El índice es el número que representa la posición del valor en el array, comenzando por `0` para el primer valor.

Puedes acceder al valor usando <dfn>notación de corchetes</dfn>, como `array[0]`.

Usa `console.log` y notación de corchetes para imprimir el primer valor en tu array `rows`.

# --hints--

Debes tener un `console.log()` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Debes acceder al array `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

Debes usar notación de corchetes en tu arreglo `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

Debes usar la notación de corchetes para acceder al primer elemento de tu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

Debes imprimir en consola el primer elemento de tu arreglo `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
