---
id: 660f4455f457ef2e3ec6920f
title: Paso 80
challengeType: 1
dashedName: step-80
---

# --description--

Ahora la cadena de texto ya no se imprime, porque `false` no es `true`. Pero, ¿qué pasa con otros valores?

Intenta cambiar la condición a la cadena de texto `"false"`.

# --hints--

Tu declaración `if` debe tener la cadena de texto `"false"` como condición.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*('|")false\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (false) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
