---
id: 660f4a83373de83ea101685f
title: Paso 100
challengeType: 1
dashedName: step-100
---

# --description--

Como vas a hacer el ciclo en la dirección opuesta, tu ciclo debe ejecutarse mientras `i` sea mayor que `0`. Puedes usar el operador <dfn>greater than</dfn> `>` para esto.

Establece la condición de tu bucle para ejecutarlo cuando `i` es mayor que `0`.

# --hints--

Tu bucle `for` debería ejecutarse cuando `i` sea mayor que `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
