---
id: 660f34626216270c682e2f7b
title: Paso 66
challengeType: 1
dashedName: step-66
---

# --description--

La <dfn>llamada a una función</dfn> permite utilizarla. Puede que no hayas sido consciente de ello, pero los métodos como `.push()` que has estado usando han sido llamadas a funciones.

Una función es llamada haciendo referencia al nombre de la función, y agregando `()`. Así es como se llama a la función `test`:

```js
test();
```

Reemplaza el `character.repeat(i + 1)` en tu llamada `.push()` con una llamada a tu función `padRow`. No añada ningún argumento todavía.

# --hints--

No debes usar `i + 1` en tu llamada `push`.

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

No debes usar `character.repeat` en tu llamada `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

Debes llamar a `padRow` en tu llamada `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

No Debes tener argumentos en tu llamada `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
