---
id: 660f45ccf4ca5c31f253005a
title: Paso 86
challengeType: 1
dashedName: step-86
---

# --description--

Ahora mismo, si cambias `continueLoop` a verdadero, tu bucle `while` se ejecutará para siempre. Esto se llama un <dfn>bucle infinito</dfn>, y Debes tener cuidado y evitarlos. Un bucle infinito puede bloquear tu sistema, requiriendo un reinicio completo para escapar.

Para evitar esto, empieza usando el operador de incremento para aumentar el valor de la variable `done` dentro de tu bucle.

# --hints--

Tu bucle `while` incrementará la variable `done`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
