---
id: 660f0c34aad72dc712b97624
title: Paso 27
challengeType: 1
dashedName: step-27
---

# --description--

Otro método esencial para este proyecto es el método `.pop()`. Este método elimina el último elemento de un array y <dfn>devuelve</dfn> ese elemento.

Cuando un método devuelve un valor, puedes pensar en él como si te devolviera el valor, haciendo que este valor esté disponible para su uso en otras partes de tu código.

Crea una nueva variable llamada `popped` y asígnale el resultado de `rows.pop()`. Luego, imprime `popped` en la consola.

# --hints--

Debes declarar una variable llamada `popped`.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

Debes usar `let` para declarar tu variable llamada `popped`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

Debes llamar al método `.pop()`.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

Debes llamar al método `.pop()` en tu array `rows`. 

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

Debes imprimir en consola tu variable `popped`. 

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
