---
id: 660f1bf673487ae0bb25b900
title: Paso 37
challengeType: 1
dashedName: step-37
---

# --description--

Tu bucle ahora debería ejecutarse ocho veces. Dentro del cuerpo del bucle, imprime el valor del iterador `i` y observa lo que sucede.

# --hints--

Debes usar `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

Debes imprimir el valor de `i`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

Debes imprimir el valor de `i` en tu bucle `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
