---
id: 660f415b76859a2736771607
title: Paso 78
challengeType: 1
dashedName: step-78
---

# --description--

Tu pirámide ha desaparecido de nuevo. Eso está bien - eso es de esperar.

Antes de crear tu nuevo bucle, necesitas aprender acerca de las condiciones `if`. Una <dfn>condición `if`</dfn> te permite ejecutar un bloque de código sólo cuando se cumple dicha condición. Utilizan la siguiente sintaxis:

```js
if (condition) {
  logic
}
```

Crea una sentencia `if` con el booleano `true` como la condición. En el cuerpo, imprime la cadena de texto `"Condition is true"`.

# --hints--

Debes crear una sentencia `if`.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

Tu sentencia `if` debe tener `true` como condición.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

El cuerpo de tu `if` debe mostrar en consola `"Condition is true"`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
