---
id: 660f540c2176ea5dec01306d
title: Paso 118
challengeType: 1
dashedName: step-118
---

# --description--

¡Buen trabajo! Experimenta con diferentes valores para tus variables `character`, `count` e `inverted`.

Cuando estés listo para pasar a tu próximo proyecto, ajusta `character` a `"!"`, `count` a `10`, y `inverted` a `false` para continuar.

¡Felicitaciones por completar tu primer proyecto JavaScript!

# --hints--

Debes establecer `character` a `"!"`.

```js
assert.equal(character, "!");
```

Debes establecer `count` a `10`.

```js
assert.equal(count, 10);
```

Debes establecer `inverted` a `false`.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
