---
id: 660f039ff313dbb696b007ca
title: Paso 19
challengeType: 1
dashedName: step-19
---

# --description--

Cuando un array contiene distintos valores, o <dfn>elementos</dfn>, esos valores deben estar separados por comas. Aquí hay un array que contiene dos cadenas de texto:

```js
let array = ["first", "second"];
```

Cambia tu variable `rows` para que sea un array con las cadenas de texto `"Naomi"`, `"Quincy"`y `"CamperChan"`. El orden de los valores de un array es importante, así que sigue ese orden. Recuerda que las cadenas de texto son sensibles a mayúsculas y minúsculas.

# --hints--

El primer elemento de tu array debe ser la cadena de texto `"Naomi"`.

```js
assert.equal(rows[0], "Naomi");
```

El segundo elemento de tu array debe ser la cadena de texto `"Quincy"`.

```js
assert.equal(rows[1], "Quincy");
```

El tercer elemento de tu array debe ser la cadena de texto `"CamperChan"`.

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
