---
id: 660f0a55847d6cc485f29ba5
title: Paso 26
challengeType: 1
dashedName: step-26
---

# --description--

Un <dfn>método</dfn> en JavaScript es una función asociada a ciertos valores u objetos. Un ejemplo que ya has encontrado es el método `.log()`, que es parte del objeto `console`.

Los arrays tienen sus propios métodos, y el primero que aprenderás es el método `.push()`. Esto te permite "empujar" un valor al final de un array. Aquí hay un ejemplo para añadir el número `12` al final de un array:

```js
array.push(12);
```

Usa `.push()` para añadir la cadena de texto `"freeCodeCamp"` al final de tu array `rows`. Añade este código antes de tu `console.log` para que puedas ver el cambio que hiciste en tu array.

# --hints--

Debes usar `.push()` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

Debes usar el método `.push()` de tu array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Debes pasar la cadena de texto `"freeCodeCamp"` a tu método `.push()`. 

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
