---
id: 6610c2d8d67563174fcf96dc
title: Paso 52
challengeType: 1
dashedName: step-52
---

# --description--

Cuando tienes un valor que está escrito explícitamente en tu código, como la cadena de texto `"Hello!"` en tu función, se considera como <dfn>codificado(escrito a mano)</dfn>. La codificación manual de un valor dentro de una función podría no hacerlo tan reutilizable como te gustaría.

En su lugar, puedes definir <dfn>parámetros</dfn> para la función. Los parámetros son variables especiales que le dan un valor a la función cuando esta es llamada, y pueden ser usados en una función para cambiar dinámicamente el resultado del código de la misma.

Para añadir un parámetro a tu función, necesitas añadir un nombre de variable dentro de los paréntesis. Por ejemplo, esta función `demo` tiene un parámetro `name`:

```js
function demo(name) {

}
```

`name` suena como un parámetro útil, así que continúa añadiéndolo a tu función `padRow`.

# --hints--

Tu función `padRow` debe tener un parámetro `name`.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
