---
id: 6610b741b54b90f0c0fb3d58
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Cuando los nombres de variables tienen más de una palabra, hay convenciones específicas para nombrarlas como la capitalización de las palabras. En JavaScript, la convención a usar es <dfn>camel case</dfn>.

Camel case significa que la primera palabra en el nombre es totalmente minúscula, pero las siguientes palabras son todas capitalizadas. Aquí hay algunos ejemplos del camel case:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

Usa el camel case para declarar una nueva variable `secondCharacter`.

# --hints--

Debes declarar una variable `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

Debes usar `let` para declarar tu variable `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

No debes asignar un valor a tu variable `secondCharacter`. No olvides el punto y coma.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
