---
id: 6610b8017d1671f2814e8c77
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Cuando declaras una variable sin inicializarla, esta se considera <dfn>no inicializada</dfn>. Actualmente, tu variable `secondCharacter` no está inicializada.

Agrega un `console.log()` para ver cuál es el valor de tu variable `secondCharacter`.

# --hints--

Debes añadir un segundo `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Debes imprimir en consola tu variable `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
