---
id: 6610c21b3ef82015573ffbbe
title: Paso 51
challengeType: 1
dashedName: step-51
---

# --description--

Tu variable `call` tiene un valor `undefined`, ¡Aunque lo hayas definido! Esto es porque tu función `padRow` no retorna un valor. Por defecto, las funciones devuelven `undefined` como su valor.

Para devolver algo más, necesitas usar la palabra reservada `return`. Aquí hay un ejemplo de una función que retorna la cadena de texto `"Functions are cool!"`:

```js
function demo() {
  return "Functions are cool!";
}
```

Usa la palabra reservada `return` para que tu función devuelva la cadena de texto `"Hello!"`.

# --hints--

Tu función `padRow` debe usar la palabra reservada `return`.

```js
assert.match(padRow.toString(), /return/);
```

Tu función `padRow` debería devolver la cadena de texto `"Hello!"`.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
