---
id: 6610c48c4ea0891afa7c4696
title: Pasaje 54
challengeType: 1
dashedName: step-54
---

# --description--

Si vuelves a abrir la consola, verás que tu función `padRow` devuelve `undefined`, ¡Aunque hayas definido un valor de retorno! Esto se debe a que a los parámetros se les debe dar un valor cuando **llamas** a la función.

Cuando le pasas un valor al llamado de la función, este valor se denomina <dfn>argumento</dfn>. Aquí hay un ejemplo de llamar a una función `demo` y pasar `"Naomi"` como argumento para el parámetro `name`.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

Pasa tu propio nombre como argumento para el parámetro `name` en tu llamada a la función `padRow`. Recuerda que tu nombre es una cadena de texto, así que necesitará usar comillas.

# --hints--

Debes pasar una cadena de texto a tu llamada `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
