---
id: 6610c538372aa61cc0f5b122
title: Paso 57
challengeType: 1
dashedName: step-57
---

# --description--

Las variables en JavaScript están disponibles en un <dfn>scope</dfn> específico. En otras palabras, el lugar donde se declara una variable determina en qué parte del código se puede utilizar.

El primer ámbito (scope) de aplicación es el ámbito global. Las variables que se declaran fuera de cualquier "bloque" como una función o en un bucle `for` están en el ámbito <dfn>global</dfn>. Tus variables `character`, `count`y `rows` están todas en el ámbito global.

Cuando una variable está en el ámbito global, una función puede acceder a ella en su definición. Aquí hay un ejemplo de una función usando una variable `title` global:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

Este ejemplo devolvería `"Professor Naomi"`. Actualiza tu función `padRow` para devolver el valor de concatenar tu variable `character` al principio del parámetro `name`.

# --hints--

Tu función `padRow` debe concatenar la variable `character` antes de `name`.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

Tu función `padRow` debería devolver el resultado de `character + name`.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
