---
id: 6610c83b52583e245a079217
title: Paso 61
challengeType: 1
dashedName: step-61
---

# --description--

Ahora tu variable `call` tiene el valor `"Testing"`. Pero tu función ya no está usando el parámetro `name`.

Elimina el parámetro `name` de tu declaración de función, luego elimina tu cadena de texto `"CamperChan"` de la llamada `padRow`.

Además, elimina ambos `console.log` de la función `padRow`.
# --hints--

Tu función `padRow` no debe tener un parámetro `name`.

```js
assert.notMatch(padRow.toString(), /name/);
```

No debes pasar `"CamperChan"` a tu llamada `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

Aún debes llamar a tu función `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

No debes tener un `console.log` antes de tu palabra reservada `return`.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

No Debes tener un `console.log` después de tu palabra reservada `return`.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
