---
id: 664599653fcd6e97104f9261
title: Paso 83
challengeType: 1
dashedName: step-83
---

# --description--

A veces querrás ejecutar un código diferente cuando todas las condiciones `if...else if` sean `false`. Puedes hacer esto agregando un bloque `else`. 

Un bloque `else` solo evaluará si las condiciones en el `if` y `else if` no se cumplen.

Aquí se añade el bloque `else` al bloque `else if`. 

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

Agrega un bloque `else` al bloque `else if`. Dentro del bloque `else`, imprime en consola la cadena de texto `"This is the else block"`.

Para ver los resultados en la consola, puedes cambiar manualmente el `<` en la sentencia `else if` a `>`. Eso hará que la condición sea `false` y se ejecutará el bloque `else`.

# --hints--

Debes tener un bloque `else`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

Tu bloque `else` debería mostrar en consola la cadena de texto `"This is the else block"`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

