---
id: 66643c93e05093c728abdbe9
title: Paso 55
challengeType: 1
dashedName: step-55
---

# --description--

Antes de continuar, tómate un momento para revisar cómo funcionan las funciones.

Declara una función llamada `addTwoNumbers`. Esta función debería tomar dos argumentos y devolver la suma de estos dos argumentos.

Su función no debería usar valores de codificación-rígida. Un ejemplo de una función de codificación-rígida podría ser:

```js
function sayName(firstName, lastName) {
  return "John Doe";
}

sayName("Camper", "Cat");
```

Esta función devolvería `"John Doe"` independientemente de los argumentos pasados ​​a los parámetros `firstName` y `lastName`, por lo que `"John Doe"` se considera un valor de codificación-rígida.

Declara una variable `sum` y asígnele el valor que devuelve el llamado a la función `addTwoNumbers` con `5` y `10` como sus argumentos. Imprime la variable `sum` en la consola.

# --hints--

Debes tener una función llamada `addTwoNumbers`.

```js
assert.isFunction(addTwoNumbers);
```

Su función `addTwoNumbers` debe tener 2 parámetros.

```js
assert.lengthOf(addTwoNumbers, 2);
```

Su función debe devolver la suma de los dos parámetros.

```js
assert.strictEqual(addTwoNumbers(5,10), 15);
```

Su función no debe devolver un valor de codificación-rígida. Es decir, debería funcionar con dos argumentos numéricos cualquiera.

```js
assert.strictEqual(addTwoNumbers(3, 5), 8);
```

Debes declarar una variable `sum`.

```js
assert.isDefined(sum);
```

Su variable `sum` debe tener el valor `15`.

```js
assert.strictEqual(sum, 15);
```

Debes asignar a `sum` el valor de llamar a la función `addTwoNumbers` con `5` y `10` como argumentos.

```js
assert.match(code, /sum\s*=\s*addTwoNumbers\s*\(/);
```

Debes imprimir en consola tu variable `sum`.

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
