---
id: 66eaddd04a9e533fba689001
title: تصميم بطاقة منشور مدونة
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

في هذا المختبر، ستتدرب على تنسيق الخلفيات والحدود من خلال إنشاء بطاقة منشور مدونة.

**الهدف:** أتمم قصص المستخدم أدناه واجعل جميع الاختبارات تجتاز لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر `div` بصنف `blog-post-card` ليحتوي كل عناصر البطاقة.
2. داخل عنصر `.blog-post-card`، يجب أن يكون لديك عنصر صورة مع خاصية `alt` صحيحة ونص، وصنف `post-img`. يمكنك استخدام `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` كقيمة لخاصية `src` للصورة.
3. يجب أن يكون لديك عنصر `div` بصنف `post-content` داخل `.blog-post-card` يحتوي على التالي:
    - عنصر `h2` بصنف `post-title` ونص لعنوان منشور المدونة.
    - عنصر `p` بصنف `post-excerpt` ونص يلخص منشور المدونة.
    - عنصر `a` بصنف `read-more` مع النص `Read More`.
4. يجب أن تضبط لون خلفية عنصر `body` على قيمة غير الأبيض.
5. يجب أن تطبق التنسيقات التالية على عنصر `.blog-post-card`:
    - خلفية بيضاء.
    - زوايا منحنية.
    - عرض تختاره.
    - محاذاة نص تختارها.
6. يجب تنسيق عنصر `.post-img` بحيث تملأ الصورة عرض البطاقة بالكامل ولها حدود سفلية.
7. يجب تنسيق عنصر `.post-content` ليحتوي على حشو داخل البطاقة.
8. يجب أن يكون لعناصر `.post-title` و `.post-excerpt` لون نص غير الافتراضي وهوامش من جميع الجهات.
9. يجب تنسيق عنصر `.read-more` مثل زر ويحتوي على:
    - لون نص غير الافتراضي.
    - لون خلفية.
    - هوامش من جميع الجهات.
    - حشو من جميع الجهات.
    - زوايا منحنية.
    - خاصية `display` مضبوطة على `inline-block`.
    - تأثير المرور بالمؤشر يغير لون الخلفية.

**ملاحظة:** تأكد من ربط ورقة الأنماط في ملف HTML وتطبيق CSS الخاص بك.

# --hints--

يجب أن يكون لديك عنصر `div` بصنف `blog-post-card`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

يجب أن يكون لديك عنصر `img` مع `class` بقيمة `post-img`. تأكد من أن الصورة تحتوي على خاصية `alt` مع نص وخاصية `src` بقيمة.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

يجب أن يكون لديك عنصر `div` بصنف `post-content`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

يجب أن يكون لديك عنصر `h2` بصنف `post-title`. تأكد من وجود نص لعنوان منشور المدونة.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

يجب أن يكون لديك عنصر `p` بصنف `post-excerpt`. تأكد من وجود نص يلخص منشور المدونة.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

يجب أن يكون لديك عنصر `a` مع `class` من نوع `read-more`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

يجب أن يحتوي عنصر `.read-more` على النص `Read More`.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

يجب أن يحتوي عنصر `.blog-post-card` على خاصية `border-radius` بقيمة (يجب ألا تكون `0` أو قيمة سالبة).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

يجب أن يحتوي عنصر `.blog-post-card` على خلفية بيضاء.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

يجب أن يحتوي عنصر `body` على لون خلفية غير الأبيض.

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

يجب استهداف `.blog-post-card` وضبط خاصية `width` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

يجب استهداف `.blog-post-card` وضبط خاصية `text-align` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

يجب استهداف `.post-content` وضبط خاصية `padding` الخاصة به.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

يجب أن يحتوي عنصر `.read-more` على تأثير المرور بالمؤشر.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

يجب استهداف `.read-more` وضبط خاصية `color` الخاصة به.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

يجب استهداف `.read-more` وضبط خاصية `background-color` الخاصة به.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

يجب استهداف `.read-more` وضبط خاصية `margin` الخاصة به.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

يجب استهداف `.read-more` وضبط خاصية `display` الخاصة به.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

يجب استهداف `.read-more` وضبط خاصية `border-radius` الخاصة به.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

يجب استهداف `.read-more` وضبط خاصية `padding` الخاصة به.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

يجب أن يملأ عنصر `.post-img` عرض البطاقة `width`.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

يجب أن يحتوي عنصر `.post-img` على قيمة `border-bottom`.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

يجب أن تحتوي عناصر `.post-title` و `.post-excerpt` على هوامش وألوان نص غير افتراضية.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
