---
id: 68de7360e1832ea180577a8a
title: Construye una calculadora del N-ésimo número de Fibonacci
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `fibonacci`.
2. Debes definir una lista llamada `sequence` dentro de la función `fibonacci`, y debe inicializarse con los valores `[0, 1]`.
3. La función `fibonacci` debe aceptar un parámetro, un entero positivo `n`.
4. Llamar a `fibonacci(n)` debe usar un enfoque de programación dinámica para calcular y devolver el número `n`-ésimo de la secuencia de Fibonacci, donde cada número es la suma de los dos números anteriores.
5. Cada número calculado en la posición `n` de la secuencia de Fibonacci debe almacenarse en la lista `sequence` en el índice `n - 1`.

# --hints--

Debes tener una función llamada `fibonacci`.

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

Tu función `fibonacci` debe tomar un solo parámetro.

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

Debes tener una lista llamada `sequence` dentro de la función `fibonacci` inicializada en `[0, 1]`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)` debería devolver `0`.

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)` debería devolver `1`.

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)` debería devolver `1`.

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)` debería devolver `2`.

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)` debería devolver `5`.

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)` debería devolver `55`.

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)` debería devolver `610`.

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

No debes usar recursión en tu código.

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```

