---
id: 680900675ae3d54ee19590c7
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Erstellen Sie ein weiteres `label` mit dem Text `Real Name` und einem `input` vom Typ `password` darin. Geben Sie dem `input` eine `value` von `realName`.

Verwenden Sie dann das `onChange`-Ereignis, um `realName` auf die gleiche Weise zu aktualisieren, wie Sie `heroName` aktualisiert haben.

# --hints--

Sie sollten ein `label`-Element erstellen.

```js
const labels = document.querySelectorAll("label");
assert.lengthOf(labels, 2);
```

Ihr `label`-Element sollte `Real Name` als Text haben.

```js
const labels = document.querySelectorAll("label");
assert.equal(labels[1]?.textContent.trim(), "Real Name");
```

Sie sollten ein `input` vom Typ `password` innerhalb Ihres `label`-Elements erstellen.

```js
const inputEl = document.querySelectorAll("label > input");
assert.equal(inputEl[1]?.getAttribute("type"), "password");
```

Ihr `input`-Element sollte seinen Wert auf die `realName`-Statusvariable setzen.

```js
assert.match(code, /value\s*=\s*\{\s*realName\s*\}/)
```

Sie sollten das `onChange`-Ereignis verwenden, um den Wert der `realName`-Statusvariable zu setzen, damit Sie im `input`-Feld tippen können.

```js
assert.match(code, /onChange\s*=\s*\{\s*(e|\(\s*e\s*\))\s*=>\s*setRealName\s*\(\s*e\s*\.\s*target\s*\.\s*value\s*\)\s*\}/)
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {
  
  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);
  
  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          <label>
            Hero Name
            <input
              type='text'
              value={heroName}
              onChange={e => setHeroName(e.target.value)}
            />
          </label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </div>
      </form>
    </div>
  )
};
```
