---
id: 680900675ae3d54ee19590c9
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Innerhalb Ihres `select`-Elements müssen Sie Optionen erstellen, aus denen der Benutzer seine Stromquelle auswählen kann. Das `select` sollte die folgenden sieben Optionen enthalten:

| Stromquelle                  |
|-------------------------------|
| `Select one`                    |
| `Bitten by a strange creature`  |
| `Radioactive exposure`          |
| `Science experiment`           |
| `Alien heritage `               |
| `Ancient artifact discovery`    |
| `Other`                         |

Die erste Option (`Select one`) sollte als `value`-Attribut einen leeren String haben (d. h. `value=''`). Dies dient als Platzhalter in der Benutzeroberfläche.

# --hints--

Sie sollten das erste `option`-Element mit dem Text `Select one` innerhalb Ihres `select`-Elements erstellen.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.textContent, "Select one");
```

Sie sollten dem ersten `option`-Element ein `value`-Attribut mit einem leeren String geben.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.getAttribute('value'), '');
```

Sie sollten ein zweites `option`-Element mit dem Text `Bitten by a strange creature` innerhalb Ihres `select`-Elements erstellen.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[1]?.textContent, "Bitten by a strange creature");
```

Sie sollten ein drittes `option`-Element mit dem Text `Radioactive exposure` innerhalb Ihres `select`-Elements erstellen.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[2]?.textContent, "Radioactive exposure");
```

Sie sollten ein viertes `option`-Element mit dem Text `Science experiment` innerhalb Ihres `select`-Elements erstellen.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[3]?.textContent, "Science experiment");
```

Sie sollten ein fünftes `option`-Element mit dem Text `Alien heritage` innerhalb Ihres `select`-Elements erstellen.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[4]?.textContent, "Alien heritage");
```

Sie sollten ein sechstes `option`-Element mit dem Text `Ancient artifact discovery` innerhalb Ihres `select`-Elements erstellen.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[5]?.textContent, "Ancient artifact discovery");
```

Sie sollten ein siebtes `option`-Element mit dem Text `Other` innerhalb Ihres `select`-Elements erstellen.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[6]?.textContent, "Other");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          <label>
            Hero Name
            <input
              type='text'
              value={heroName}
              onChange={e => setHeroName(e.target.value)}
            />
          </label>
          <label>
            Real Name
            <input
              type='password'
              value={realName}
              onChange={e => setRealName(e.target.value)}
            />
          </label>
        </div>
          <label className='section column'>
            How did you get your powers?
            --fcc-editable-region--
            <select>
              
            </select>
            --fcc-editable-region--
          </label>
      </form>
    </div>
  )
};
```
