---
id: 6899b593d232e07617f74055
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

`button` 요소 안에 `span` 요소를 만들고 그 안에 `😊`를 넣으세요.

`span` 요소에 `role` 속성을 값 `img`로, 그리고 `aria-hidden` 속성을 값 `true`로 지정하세요.

그다음, 텍스트 `span`를 포함하는 `class`의 `count` 속성을 가진 두 번째 `0/10` 요소를 만드세요.

# --hints--

`button` 요소 안에는 두 개의 `span` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("button > span"), 2);
```

첫 번째 `span` 요소는 텍스트 내용이 `😊`여야 합니다.

```js
assert.equal(document.querySelectorAll("button > span")[0]?.textContent.trim(), "😊");
```

첫 번째 `span` 요소는 값이 `role`인 `img` 속성을 가져야 합니다.

```js
assert.equal(document.querySelectorAll("button > span")[0]?.getAttribute("role"), "img");
```

첫 번째 `span` 요소는 값이 `aria-hidden`인 `true` 속성을 가져야 합니다.

```js
assert.equal(document.querySelectorAll("button > span")[0]?.getAttribute("aria-hidden"), "true");
```

두 번째 `span` 요소는 값이 `class`인 `count` 속성을 가져야 합니다.

```js
assert.equal(document.querySelectorAll("button > span")[1]?.className, "count");
```

두 번째 `span` 요소는 텍스트 내용이 `0/10`이어야 합니다.

```js
assert.equal(document.querySelectorAll("button > span")[1]?.textContent.trim(), "0/10");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>

    <h1 class="title">How are you feeling today?</h1>

    <p class="description">
      Click on the buttons below to rate your emotions.
    </p>
    <div class="btn-container">
--fcc-editable-region--
      <button id="happy-btn" class="emoji-btn" aria-label="Happy face emoji">
        
      </button>
--fcc-editable-region--
    </div>
  </main>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```
