---
id: 689b17e7e63b8e2c13839e98
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

이제 실제 카운팅 기능을 구현할 시간입니다.

현재 카운트를 단순히 기록하는 대신, 값을 1씩 증가시키고 화면에 표시를 업데이트하세요.

`console.log`를 제거하세요.

`.count` 안에 있는 첫 번째 숫자가 이미 10인지 확인하세요(텍스트는 `10/10`). 만약 그렇다면 아무 작업도 하지 마세요.

숫자가 10보다 작으면 현재 카운트를 1만큼 증가시키고 `.count` 요소의 텍스트를 업데이트하세요. 예를 들어, 현재 카운트가 2라면 3으로 증가시키고 `.count` 요소의 텍스트는 `3/10`가 되어야 합니다.

# --hints--

버튼을 클릭하면 카운트가 0에서 1로 증가해야 합니다.

```js
const spy = __helpers.spyOn(console, "log");
const countEl = document.querySelector("#happy-btn .count");
countEl.textContent = "0/10";
document.querySelector("#happy-btn").click();
assert.equal(countEl.textContent, "1/10");
```

버튼을 여러 번 클릭하면 카운트가 제대로 증가해야 합니다.

```js
const countEl = document.querySelector("#happy-btn .count");
countEl.textContent = "3/10";
document.querySelector("#happy-btn").click();
assert.equal(countEl.textContent, "4/10");
```

카운트가 이미 10일 때는 클릭해도 더 이상 증가하지 않아야 합니다.

```js
const countEl = document.querySelector("#happy-btn .count");
countEl.textContent = "10/10";
document.querySelector("#happy-btn").click();
assert.equal(countEl.textContent, "10/10");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>

    <h1 class="title">How are you feeling today?</h1>

    <p class="description">
      Click on the buttons below to rate your emotions.
    </p>
    <div class="btn-container">
      <button id="happy-btn" class="emoji-btn" aria-label="Happy face emoji">
        <span role="img" aria-hidden="true">😊</span>
        <span class="count">0/10</span>
      </button>
    </div>
  </main>
  <script src="./script.js"></script>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```

```js
--fcc-editable-region--
const happyBtn = document.querySelector("#happy-btn");

happyBtn.addEventListener("click", () => {
  const countEl = happyBtn.querySelector(".count");
  const currCount = +countEl.textContent.split("/")[0];
  console.log("Current count:", currCount);
})


--fcc-editable-region--
```
