---
id: 587d825c367417b2b2512c8f
title: Merge-Sort implementieren
challengeType: 1
forumTopicId: 301614
dashedName: implement-merge-sort
---

# --description--

Ein weiterer gängiger Sortieralgorithmus mittlerer Komplexität ist Merge-Sort. Wie Quick-Sort verwendet Merge-Sort ebenfalls eine Divide-and-Conquer-Strategie und arbeitet rekursiv, um ein Array zu sortieren. Er nutzt die Tatsache, dass es relativ einfach ist, zwei Arrays zu sortieren, solange jedes für sich bereits sortiert ist. Aber wir starten nur mit einem Array als Eingabe – wie kommen wir also zu zwei sortierten Arrays daraus? Nun, wir können die ursprüngliche Eingabe rekursiv halbieren, bis wir den Basisfall eines Arrays mit nur einem Element erreichen. Ein Array mit einem einzelnen Element ist von Natur aus sortiert, sodass wir dann mit dem Zusammenführen beginnen können. Dieses Zusammenführen löst die rekursiven Aufrufe auf, die das ursprüngliche Array geteilt haben, und erzeugt schließlich ein endgültiges sortiertes Array mit allen Elementen. Die Schritte von Merge-Sort sind also:

**1)** Teilen Sie das Eingabe-Array rekursiv in zwei Hälften, bis ein Teilarray mit nur einem Element entsteht.

**2)** Führen Sie jedes sortierte Teilarray zusammen, um das endgültige sortierte Array zu erzeugen.

Merge-Sort ist eine effiziente Sortiermethode mit einer Zeitkomplexität von *O(nlog(n))*. Dieser Algorithmus ist beliebt, weil er leistungsfähig und relativ einfach zu implementieren ist.

Nebenbei bemerkt: Dies wird der letzte Sortieralgorithmus sein, den wir hier behandeln. Später im Abschnitt über Baum-Datenstrukturen werden wir Heap-Sort beschreiben, eine weitere effiziente Sortiermethode, die einen binären Heap in ihrer Implementierung benötigt.

**Anleitung:** Schreiben Sie eine Funktion `mergeSort`, die ein Array ganzer Zahlen als Eingabe erhält und ein Array dieser Zahlen in aufsteigender Reihenfolge zurückgibt. Eine gute Möglichkeit, dies zu implementieren, ist, eine Funktion zu schreiben, zum Beispiel `merge`, die für das Zusammenführen zweier sortierter Arrays zuständig ist, und eine andere Funktion, zum Beispiel `mergeSort`, die für die Rekursion verantwortlich ist, die einzelne Arrays erzeugt, die in das Zusammenführen eingespeist werden. Viel Erfolg!

# --hints--

`mergeSort` sollte eine Funktion sein.

```js
assert.isFunction(mergeSort);
```

`mergeSort` sollte ein sortiertes Array (aufsteigend) zurückgeben.

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    mergeSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`mergeSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` sollte ein Array zurückgeben, das bis auf die Reihenfolge unverändert ist.

```js
assert.sameMembers(
  mergeSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`mergeSort` sollte die eingebaute `.sort()`-Methode nicht verwenden.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    mergeSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function mergeSort(array) {
  if (array.length === 1) {
    return array;
  } else {
    const splitIndex = Math.floor(array.length / 2);
    return merge(
      mergeSort(array.slice(0, splitIndex)),
      mergeSort(array.slice(splitIndex))
    );
  }

  // Merge two sorted arrays
  function merge(array1, array2) {
    let merged = [];
    while (array1.length && array2.length) {
      if (array1[0] < array2[0]) {
        merged.push(array1.shift());
      } else if (array1[0] > array2[0]) {
        merged.push(array2.shift());
      } else {
        merged.push(array1.shift(), array2.shift());
      }
    }

    // After looping ends, one array is empty, and other array contains only
    // values greater than all values in `merged`
    return [...merged, ...array1, ...array2];
  }
}

mergeSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]);
```
