---
id: 587d825a367417b2b2512c89
title: Quick Sort implementieren
challengeType: 1
forumTopicId: 301615
dashedName: implement-quick-sort
---

# --description--

Hier steigen wir in einen fortgeschrittenen Sortieralgorithmus ein: Quick Sort. Quick Sort ist ein effizienter, rekursiver Divide-and-Conquer-Ansatz zum Sortieren eines Arrays. Bei dieser Methode wird ein Pivot-Wert im ursprünglichen Array ausgewählt. Das Array wird dann in zwei Teilarrays aufgeteilt, die Werte enthalten, die kleiner beziehungsweise größer als der Pivot-Wert sind. Anschließend kombinieren wir das Ergebnis der rekursiven Aufrufe des Quick-Sort-Algorithmus auf beiden Teilarrays. Dies wird fortgesetzt, bis der Basisfall eines leeren oder ein Element enthaltenden Arrays erreicht ist, das wir zurückgeben. Das Zurückgeben der rekursiven Aufrufe liefert uns das sortierte Array.

Quick Sort ist eine sehr effiziente Sortiermethode, die im Durchschnitt eine Leistung von *O(nlog(n))* bietet. Außerdem ist sie relativ einfach zu implementieren. Diese Eigenschaften machen sie zu einer beliebten und nützlichen Sortiermethode.

**Anleitung:** Schreiben Sie eine Funktion `quickSort`, die ein Array ganzer Zahlen als Eingabe erhält und ein Array dieser Zahlen in aufsteigender Reihenfolge zurückgibt. Obwohl die Wahl des Pivot-Werts wichtig ist, reicht für unsere Zwecke hier jeder Pivot aus. Zur Vereinfachung kann das erste oder letzte Element verwendet werden.

# --hints--

`quickSort` sollte eine Funktion sein.

```js
assert(typeof quickSort == 'function');
```

`quickSort` sollte ein sortiertes Array (aufsteigend) zurückgeben.

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    quickSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`quickSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` sollte ein Array zurückgeben, das bis auf die Reihenfolge unverändert ist.

```js
assert.sameMembers(
  quickSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quickSort` sollte die eingebaute `.sort()`-Methode nicht verwenden.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quickSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function quickSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function quickSort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quickSort(lesser), ...equal, ...quickSort(greater)];
  }
}
```
