---
id: 587d8259367417b2b2512c85
title: Selection Sort implementieren
challengeType: 1
forumTopicId: 301616
dashedName: implement-selection-sort
---

# --description--

Hier implementieren wir Selection Sort. Selection Sort funktioniert, indem es den kleinsten Wert in einer Liste auswählt und ihn mit dem ersten Wert der Liste vertauscht. Dann beginnt es an der zweiten Position, wählt den kleinsten Wert in der verbleibenden Liste aus und vertauscht ihn mit dem zweiten Element. Es fährt fort, die Liste zu durchlaufen und Elemente zu vertauschen, bis es das Ende der Liste erreicht hat. Nun ist die Liste sortiert. Selection Sort hat in allen Fällen eine quadratische Zeitkomplexität.

**Anleitung**: Schreiben Sie eine Funktion `selectionSort`, die ein Array ganzer Zahlen als Eingabe erhält und ein Array dieser Zahlen in sortierter Reihenfolge von klein nach groß zurückgibt.

# --hints--

`selectionSort` sollte eine Funktion sein.

```js
assert.isFunction(selectionSort);
```

`selectionSort` sollte ein sortiertes Array (von klein nach groß) zurückgeben.

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}

assert.isTrue(
  isSorted(
    selectionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`selectionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` sollte ein Array zurückgeben, das bis auf die Reihenfolge unverändert ist.

```js
assert.sameMembers(
  selectionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` sollte die eingebaute `.sort()`-Methode nicht verwenden.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function selectionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
