---
id: 587d8259367417b2b2512c86
title: Insertion Sort implementieren
challengeType: 1
forumTopicId: 301613
dashedName: implement-insertion-sort
---

# --description--

Die nächste Sortiermethode, die wir uns ansehen, ist Insertion Sort. Diese Methode funktioniert, indem sie am Anfang der Liste ein sortiertes Array aufbaut. Sie beginnt das sortierte Array mit dem ersten Element. Dann betrachtet sie das nächste Element und tauscht es rückwärts in das sortierte Array, bis es an der richtigen Stelle einsortiert ist. Sie fährt fort, die Liste zu durchlaufen und neue Elemente rückwärts in den sortierten Teil einzufügen, bis sie das Ende erreicht. Dieser Algorithmus hat im Durchschnitt und im schlimmsten Fall eine quadratische Zeitkomplexität.

**Anleitung:** Schreiben Sie eine Funktion `insertionSort`, die ein Array ganzer Zahlen als Eingabe erhält und ein Array dieser Zahlen in aufsteigender Reihenfolge zurückgibt.

# --hints--

`insertionSort` sollte eine Funktion sein.

```js
assert.isFunction(insertionSort);
```

`insertionSort` sollte ein sortiertes Array zurückgeben (aufsteigend).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    insertionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`insertionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` sollte ein Array zurückgeben, das bis auf die Reihenfolge unverändert ist.

```js
assert.sameMembers(
  insertionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort([5, 4, 33, 2, 8])` sollte `[2, 4, 5, 8, 33]` zurückgeben.

```js
assert.deepEqual(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` sollte die eingebaute `.sort()`-Methode nicht verwenden.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function insertionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
