---
id: 61abc7ebf3029b56226de5b6
title: Binäre Suche implementieren
challengeType: 1
forumTopicId: 487618
dashedName: implement-binary-search
---

# --description--

Die binäre Suche ist ein Algorithmus mit einer Effizienz von **O(log(n))** zur Suche in einem sortierten Array, um ein Element zu finden. Er arbeitet mit den folgenden Schritten:

1. Finden Sie das mittlere `value` eines sortierten Arrays. Wenn `value == target`, geben Sie `true` zurück (Der Wert wurde gefunden und die Suche ist abgeschlossen).
1. Wenn das mittlere `value < target`, durchsuchen Sie die rechte Hälfte des Arrays beim nächsten Vergleich.
1. Wenn das mittlere `value > target`, durchsuchen Sie die linke Hälfte des Arrays beim nächsten Vergleich.
1. Wenn nach der Durchsuchung des gesamten Arrays der Wert nicht vorhanden ist, geben Sie `false` zurück (Das Array wurde durchsucht und der Wert ist nicht im Array).

Wie Sie sehen, halbieren Sie das Array sukzessive, was Ihnen die log(n)-Effizienz gibt. Für diese Aufgabe möchten wir, dass Sie Ihren Weg zeigen – wie Sie zum Zielwert gelangt sind... den Pfad, den Sie genommen haben!

# --instructions--

Schreiben Sie eine Funktion `binarySearch`, die den Algorithmus der binären Suche auf einem Array implementiert und den Pfad zurückgibt, den Sie genommen haben (jeden mittleren Wert beim Vergleich), um das Ziel in einem Array zu finden.

Die Funktion erhält ein sortiertes Array von ganzen Zahlen und einen Zielwert als Eingabe. Sie gibt ein Array zurück, das (in Reihenfolge) den mittleren Wert enthält, den Sie bei jeder Halbierung des ursprünglichen Arrays gefunden haben, bis Sie den Zielwert gefunden haben. Der Zielwert sollte das letzte Element des zurückgegebenen Arrays sein. Wenn der Wert nicht gefunden wird, geben Sie den String `Value Not Found` zurück.

Zum Beispiel würde `binarySearch([1,2,3,4,5,6,7], 5)` `[4,6,5]` zurückgeben.

Für diese Aufgabe müssen Sie beim Halbieren unbedingt `Math.floor()` bei der Division verwenden: `Math.floor(x/2)`. Dies sorgt für einen konsistenten, testbaren Pfad.

**Hinweis:** Das folgende Array wird in den Tests verwendet:

```js
const testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
```

# --hints--

`binarySearch` sollte eine Funktion sein.

```js
assert.isFunction(binarySearch);
```

`binarySearch(testArray, 0)` sollte `[13, 5, 2, 0]` zurückgeben.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 0), [13, 5, 2, 0]);
```

`binarySearch(testArray, 1)` sollte `[13, 5, 2, 0, 1]` zurückgeben.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 1), [13, 5, 2, 0, 1]);
```


`binarySearch(testArray, 2)` sollte `[13, 5, 2]` zurückgeben.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 2), [13, 5, 2]);
```

`binarySearch(testArray, 6)` sollte den String `Value Not Found` zurückgeben.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.strictEqual(binarySearch(_testArray, 6), 'Value Not Found');
```

`binarySearch(testArray, 11)` sollte `[13, 5, 10, 11]` zurückgeben.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 11), [13, 5, 10, 11])
```

`binarySearch(testArray, 13)` sollte `[13]` zurückgeben.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 13), [13]);
```

`binarySearch(testArray, 70)` sollte `[13, 19, 22, 49, 70]` zurückgeben.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 70), [13, 19, 22, 49, 70]);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let arrayPath = [];
  return arrayPath;
}
```



# --solutions--

```js
let binarySearch = (searchList, value) => {
  let arrayPath = [];

  // set initial L - M - R
  let left = 0;
  let right = searchList.length - 1;
  let middle = Math.floor(right / 2);

  // if first comparison finds value
  if (searchList[middle] == value) {
    arrayPath.push(searchList[middle]);
    return arrayPath;
  }

  while (searchList[middle] !== value) {
    // add to output array
    arrayPath.push(searchList[middle]);

    // not found
    if (right < left) {
      return 'Value Not Found';
    }
    // value is in left or right portion of array
    // update L - M - R
    if (searchList[middle] > value) {
      right = middle - 1;
      middle = left + Math.floor((right - left) / 2);
    } else {
      left = middle + 1;
      middle = left + Math.floor((right - left) / 2);
    }

    // if found update output array and exit
    if (searchList[middle] == value) {
      arrayPath.push(searchList[middle]);

      break;
    }
  }
  return arrayPath;
};
```
