---
id: a3f503de51cfab748ff001aa
title: Paare bilden
challengeType: 1
forumTopicId: 301617
dashedName: pairwise
---

# --description--

Gegeben ist ein Array `arr`. Finden Sie Elementpaare, deren Summe dem zweiten Argument `arg` entspricht, und geben Sie die Summe ihrer Indizes zurück.

Sie dürfen mehrere Paare verwenden, die dieselben numerischen Elemente, aber unterschiedliche Indizes haben. Jedes Paar sollte die niedrigstmöglichen verfügbaren Indizes verwenden. Sobald ein Element verwendet wurde, kann es nicht erneut verwendet werden, um mit einem anderen Element ein Paar zu bilden. Zum Beispiel erzeugt `pairwise([1, 1, 2], 3)` ein Paar `[2, 1]` unter Verwendung der 1 an Index 0 statt der 1 an Index 1, weil 0+2 &lt; 1+2 ist.

Zum Beispiel gibt `pairwise([7, 9, 11, 13, 15], 20)` `6` zurück. Die Paare, die 20 ergeben, sind `[7, 13]` und `[9, 11]`. Wir können dann das Array mit ihren Indizes und Werten aufschreiben.

<div style='margin-left: 2em;'>

| Index | 0 | 1 | 2  | 3  | 4  |
| ----- | - | - | -- | -- | -- |
| Wert  | 7 | 9 | 11 | 13 | 15 |

</div>

Im Folgenden nehmen wir ihre entsprechenden Indizes und addieren sie.

<div style='margin-left: 2em;'>

7 + 13 = 20 → Indizes 0 + 3 = 3  
9 + 11 = 20 → Indizes 1 + 2 = 3  
3 + 3 = 6 → Rückgabe `6`

</div>

# --hints--

`pairwise([1, 4, 2, 3, 0, 5], 7)` sollte 11 zurückgeben.

```js
assert.deepEqual(pairwise([1, 4, 2, 3, 0, 5], 7), 11);
```

`pairwise([1, 3, 2, 4], 4)` sollte 1 zurückgeben.

```js
assert.deepEqual(pairwise([1, 3, 2, 4], 4), 1);
```

`pairwise([1, 1, 1], 2)` sollte 1 zurückgeben.

```js
assert.deepEqual(pairwise([1, 1, 1], 2), 1);
```

`pairwise([0, 0, 0, 0, 1, 1], 1)` sollte 10 zurückgeben.

```js
assert.deepEqual(pairwise([0, 0, 0, 0, 1, 1], 1), 10);
```

`pairwise([], 100)` sollte 0 zurückgeben.

```js
assert.deepEqual(pairwise([], 100), 0);
```

# --seed--

## --seed-contents--

```js
function pairwise(arr, arg) {
  return arg;
}

pairwise([1,4,2,3,0,5], 7);
```

# --solutions--

```js
function pairwise(arr, arg) {
  var sum = 0;
  arr.forEach(function(e, i, a) {
    if (e != null) {
      var diff = arg-e;
      a[i] = null;
      var dix = a.indexOf(diff);
      if (dix !== -1) {
        sum += dix;
        sum += i;
        a[dix] = null;
      }
    }
  });
  return sum;
}

pairwise([1,4,2,3,0,5], 7);
```
