---
id: 685dbd19200ad656d48b77ff
title: Musik-Shopping-Cart-Seite erstellen
challengeType: 25
dashedName: lab-music-shopping-cart-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten ein Element mit einem `id` namens `shopping-cart-container` haben.
2. Ihr `#shopping-cart-container`-Element sollte die korrekte Utility-Klasse verwenden, um das Element auf ein Flexbox-Layout einzustellen.
3. Ihr `#shopping-cart-container`-Element sollte die Richtung der Flex-Elemente für kleinere Geräte auf Spalte und für größere Geräte auf Reihe setzen. Denken Sie daran, dass Tailwind CSS den Mobile-First-Ansatz verwendet und Sie das `lg:`-Breakpoint-Präfix verwenden, um größere Geräte anzusprechen.
4. Innerhalb Ihres `#shopping-cart-container`-Elements sollten Sie ein Element mit einem `id` namens `products-container` haben.
5. Ihr `#products-container`-Element sollte mindestens zwei Kindelemente haben, die jeweils eine Klasse namens `card` besitzen.
6. Jedes `.card`-Element sollte die folgenden verschachtelten Elemente enthalten:
   - Ein `h2`-Element mit Text, der den Produktnamen darstellt, und einer Utility-Klasse, die die Schriftgröße nur mit vordefinierten Größenklassen wie `text-sm`, `text-md`, `text-lg`, `text-xl`, `text-2xl` usw. einstellt.
   - Ein Element mit einer Klasse namens `quantity` und Text für die Anzahl der Warenkorb-Artikel dieses Produkts.
   - Ein Element mit einer Klasse namens `price` und Text für den Preis.
   - Ein `button` mit einer Klasse namens `remove-button` und dem Text `Remove`. Ihr `button` sollte Utility-Klassen für eine vordefinierte rote Hintergrundfarbe Ihrer Wahl und eine andere vordefinierte rote Hintergrundfarbe für den Hover-Zustand haben. Beispiele für vordefinierte rote Hintergrundfarben sind `bg-red-500`, `bg-red-600` usw. Beispiele für vordefinierte Hover-Rot-Hintergrundfarben sind `hover:bg-red-500`, `hover:bg-red-600` usw.
7. Innerhalb Ihres `#shopping-cart-container`-Elements sollten Sie ein Element mit einem `id` namens `order-summary-container` haben.
8. Ihr `#order-summary-container`-Element sollte die folgenden Stile haben:
   - Eine Utility-Klasse Ihrer Wahl für vordefinierte abgerundete Ecken. Beispiele sind `rounded`, `rounded-lg`, `rounded-full` usw.
   - Eine Utility-Klasse Ihrer Wahl, um die Rahmenbreite auf allen Seiten einzustellen.
9. Ihr `#order-summary-container`-Element sollte die folgenden verschachtelten Elemente enthalten:
   - Ein `h2`-Element mit dem Text `Order Summary` und einer Utility-Klasse Ihrer Wahl, die die Schriftgröße einstellt.
   - Ein Element mit dem Text `Total:` und `id`, das auf `total` gesetzt ist. Dieses Element sollte auch Utility-Klassen verwenden, um die Schriftstärke auf eine vordefinierte Tailwind-CSS-Schriftstärke und die Schriftgröße Ihrer Wahl einzustellen.
   - Ein Element mit einem `id`, das auf `total-amount` gesetzt ist, um die Summe aller Artikel im Warenkorb anzuzeigen.
   - Einen Link mit dem Text `Checkout` und der korrekten Utility-Klasse zum Zentrieren des Texts. Der `href`-Wert sollte auf `#` gesetzt sein. Ihr Link sollte außerdem eine Utility-Klasse haben, um die Hintergrundfarbe auf eine vordefinierte Tailwind-CSS-Blauton Ihrer Wahl einzustellen, sowie eine andere blaue Hintergrundfarbe für den Hover-Zustand.

# --hints--

Sie sollten ein Element mit einem `id` haben, das auf `shopping-cart-container` gesetzt ist.

```js
assert.exists(document.getElementById("shopping-cart-container"));
```

Ihr `#shopping-cart-container`-Element sollte die `flex`-Utility-Klasse verwenden, um das Element auf ein Flexbox-Layout einzustellen.

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.isTrue(shoppingCartContainer.classList.contains("flex"));
```

Ihr `#shopping-cart-container`-Element sollte die korrekte Utility-Klasse verwenden, um die Richtung der Flex-Elemente auf Spalte zu setzen.

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.isTrue(shoppingCartContainer.classList.contains("flex-col"));
```

Ihr `#shopping-cart-container`-Element sollte das `lg:`-Breakpoint-Präfix gefolgt von der korrekten Utility-Klasse verwenden, um die Richtung der Flex-Elemente für größere Geräte auf Reihe zu setzen.

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.isTrue(shoppingCartContainer.classList.contains("lg:flex-row"));
```

Sie sollten ein Element mit einem `id` haben, das auf `products-container` gesetzt ist, innerhalb Ihres `#shopping-cart-container`-Elements.

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.exists(shoppingCartContainer);
const productsContainer = document.getElementById("products-container");
assert.exists(productsContainer);
assert.isTrue(shoppingCartContainer.contains(productsContainer));
```

Ihr `#products-container`-Element sollte mindestens zwei Kindelemente haben, die jeweils eine Klasse namens `card` besitzen.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");
assert.isAtLeast(cards.length, 2);
```

Jedes `.card`-Element sollte ein `h2`-Element mit Text enthalten, der den Produktnamen darstellt.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const productName = card.querySelector("h2");
  assert.exists(productName);
  assert.isNotEmpty(productName.textContent);
});
```

Jedes `h2`-Element innerhalb des `.card`-Elements sollte eine Utility-Klasse verwenden, die die Schriftgröße nur mit vordefinierten Größenklassen wie `text-sm`, `text-base`, `text-lg`, `text-xl`, `text-2xl` usw. einstellt.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

cards?.forEach((card) => {
  const productName = card.querySelector("h2");
  assert.exists(productName);
  const hasPredefinedSize = predefinedSizes.some(size => productName.classList.contains(size));
  assert.isTrue(hasPredefinedSize);
});
```

Jedes `.card`-Element sollte ein Element mit einer Klasse namens `quantity` haben.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const quantity = card.querySelector(".quantity");
  assert.exists(quantity);
});
```

Jedes `.quantity`-Element sollte Text enthalten, der die Anzahl der Warenkorb-Artikel dieses Produkts darstellt.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const quantity = card.querySelector(".quantity");
  assert.exists(quantity);
  assert.isNotEmpty(quantity.textContent.trim());
});
```

Jedes `.card`-Element sollte ein Element mit einer Klasse namens `price` haben.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");  

cards?.forEach((card) => {
  const price = card.querySelector(".price");
  assert.exists(price);
});
```

Jedes `.price`-Element sollte Text enthalten, der den Preis darstellt.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const price = card.querySelector(".price");
  assert.exists(price);
  assert.isNotEmpty(price.textContent.trim());
});
```

Jedes `.card`-Element sollte ein `button`-Element mit einer Klasse namens `remove-button` und dem Text `Remove` haben.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const removeButton = card.querySelector("button.remove-button");
  assert.exists(removeButton);
  assert.equal(removeButton.textContent.trim(), "Remove");
});
```

Jedes `button`-Element sollte eine Utility-Klasse für eine rote Hintergrundfarbe Ihrer Wahl haben.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");
const predefinedRedColors = [
  "bg-red-50",
  "bg-red-100",
  "bg-red-200",
  "bg-red-300",
  "bg-red-400",
  "bg-red-500",
  "bg-red-600",
  "bg-red-700",
  "bg-red-800",
  "bg-red-900",
  "bg-red-950"
];

cards?.forEach((card) => {
  const removeButton = card.querySelector("button.remove-button");
  assert.exists(removeButton);
  const hasRedBackground = predefinedRedColors.some(color => removeButton.classList.contains(color));
  assert.isTrue(hasRedBackground);
});
```

Jedes `button`-Element sollte eine andere rote Hintergrundfarbe für den Hover-Zustand haben.

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

const predefinedRedHoverColors = [
  "hover:bg-red-50",
  "hover:bg-red-100",
  "hover:bg-red-200",
  "hover:bg-red-300",
  "hover:bg-red-400",
  "hover:bg-red-500",
  "hover:bg-red-600",
  "hover:bg-red-700",
  "hover:bg-red-800",
  "hover:bg-red-900",
  "hover:bg-red-950"
];

cards?.forEach((card) => {
  const removeButton = card.querySelector("button.remove-button");
  assert.exists(removeButton);

  const classList = removeButton.classList;

  const hoverBgClass = [...classList].find(cls => cls.startsWith("hover:bg-red-"));
  const bgClass = [...classList].find(cls => cls.startsWith("bg-red-") && !cls.startsWith("hover:"));

  assert.notEqual(bgClass, hoverBgClass.replace("hover:", ""));
});
```

Sie sollten ein Element mit einem `id` namens `order-summary-container` innerhalb Ihres `#shopping-cart-container`-Elements haben.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
assert.exists(orderSummaryContainer);
```

Ihr `#order-summary-container`-Element sollte eine Utility-Klasse Ihrer Wahl für abgerundete Ecken haben.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

const hasRoundedCorners = predefinedRoundedCorners.some(corner => orderSummaryContainer.classList.contains(corner));
assert.isTrue(hasRoundedCorners);
```

Ihr `#order-summary-container`-Element sollte eine Utility-Klasse Ihrer Wahl haben, um die Rahmenbreite auf allen Seiten einzustellen.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");

// border classes can vary in width, so this check is for the presence of any border class
const hasBorder = [...orderSummaryContainer.classList].some(className => className.startsWith("border-") || className === "border");
assert.isTrue(hasBorder);
```

Ihr `#order-summary-container`-Element sollte ein `h2`-Element mit dem Text `Order Summary` haben.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const orderSummaryTitle = orderSummaryContainer.querySelector("h2");
assert.exists(orderSummaryTitle);
assert.equal(orderSummaryTitle.textContent.trim(), "Order Summary");
```

Ihr `h2` innerhalb des `#order-summary-container`-Elements sollte eine Utility-Klasse Ihrer Wahl verwenden, die die Schriftgröße einstellt.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

const orderSummaryTitle = orderSummaryContainer.querySelector("h2");
assert.exists(orderSummaryTitle);
const hasPredefinedSize = predefinedSizes.some(size => orderSummaryTitle.classList.contains(size));
assert.isTrue(hasPredefinedSize);
```

Ihr `#order-summary-container`-Element sollte ein Element mit dem Text `Total:` und `id` haben, das auf `total` gesetzt ist.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total");
assert.exists(totalElement);
assert.equal(totalElement.textContent.trim(), "Total:");
```

Ihr `#total`-Element sollte eine Utility-Klasse verwenden, um die Schriftstärke einzustellen.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total");
const predefinedFontWeights = [
  "font-thin",
  "font-extralight",
  "font-light",
  "font-normal",
  "font-medium",
  "font-semibold",
  "font-bold",
  "font-extrabold",
  "font-black"
];

const hasFontWeight = predefinedFontWeights.some(weight => totalElement.classList.contains(weight));
assert.isTrue(hasFontWeight);
```

Ihr `#total`-Element sollte eine Utility-Klasse verwenden, um die Schriftgröße einzustellen.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total");
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

const hasPredefinedSize = predefinedSizes.some(size => totalElement.classList.contains(size));
assert.isTrue(hasPredefinedSize);
```

Ihr `#order-summary-container`-Element sollte ein Element mit dem `id` haben, das auf `total-amount` gesetzt ist, um die Summe aller Artikel im Warenkorb anzuzeigen.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total-amount");
assert.exists(totalElement);
assert.isNotEmpty(totalElement.textContent.trim());
```

Ihr `#order-summary-container`-Element sollte einen Link mit dem Text `Checkout` und einem `#` für den `href`-Wert haben.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
assert.exists(checkoutLink);
assert.equal(checkoutLink.textContent.trim(), "Checkout");
assert.equal(checkoutLink.getAttribute("href"), "#");
```

Ihr `Checkout`-Link sollte eine Utility-Klasse zum Zentrieren des Texts haben.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
const hasTextCentering = checkoutLink.classList.contains("text-center");
assert.isTrue(hasTextCentering);
```

Ihr `Checkout`-Link sollte eine Utility-Klasse haben, um die Hintergrundfarbe auf eine vordefinierte Tailwind-CSS-Blauton Ihrer Wahl einzustellen.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
const predefinedBlueColors = [
  "bg-blue-50",
  "bg-blue-100",
  "bg-blue-200",
  "bg-blue-300",
  "bg-blue-400",
  "bg-blue-500",
  "bg-blue-600",
  "bg-blue-700",
  "bg-blue-800",
  "bg-blue-900",
  "bg-blue-950"
];

const hasBlueBackground = predefinedBlueColors.some(color => checkoutLink.classList.contains(color));
assert.isTrue(hasBlueBackground);
```

Ihr `Checkout`-Link sollte eine andere blaue Hintergrundfarbe für den Hover-Zustand haben.

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
const predefinedBlueHoverColors = [
  "hover:bg-blue-50",
  "hover:bg-blue-100",
  "hover:bg-blue-200",
  "hover:bg-blue-300",
  "hover:bg-blue-400",
  "hover:bg-blue-500",
  "hover:bg-blue-600",
  "hover:bg-blue-700",
  "hover:bg-blue-800",
  "hover:bg-blue-900",
  "hover:bg-blue-950"
];

const classList = checkoutLink.classList;
const hoverBgClass = [...classList].find(cls => cls.startsWith("hover:bg-blue-"));
const bgClass = [...classList].find(cls => cls.startsWith("bg-blue-") && !cls.startsWith("hover:"));

assert.notEqual(bgClass, hoverBgClass.replace("hover:", ""));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music Shopping Cart Page</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music Shopping Cart Page</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-900 text-white">
    <h1 class="text-center text-3xl my-8">Shopping Cart Page</h1>
    <main
      id="shopping-cart-container"
      class="flex flex-col lg:flex-row items-center gap-8"
    >
      <div id="products-container" class="w-full md:w-2/3 px-4">
        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">Ludwig Supraphonic Snare Drum Chrome</h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">2</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$1398.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>

        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">Shure SM58S Mic With Switch Standard</h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">1</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$119.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>

        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">
            Fender American Professional II Stratocaster
          </h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">1</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$1700.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>

        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">
            Moog Subsequent 37 Analog Synthesizer
          </h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">1</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$1799.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>
      </div>

      <div
        id="order-summary-container"
        class="rounded-lg border border-gray-200 p-8 bg-gray-800 w-full md:w-1/3 mt-8 lg:mt-0"
      >
        <h2 class="text-2xl mb-4">Order Summary</h2>
        <ul>
          <li>
            <span id="subtotal" class="font-semibold text-lg">Subtotal:</span>
            <span>$5,016.00</span>
          </li>
          <li class="my-6">
            <span id="shipping" class="font-semibold text-lg">Shipping:</span>
            <span>Calculated in checkout</span>
          </li>
          <li>
            <span id="total" class="font-semibold text-lg">Total:</span>
            <span id="total-amount">$5,016.00</span>
          </li>
        </ul>
        <a
          href="#"
          class="text-center rounded-lg mx-auto mt-4 block px-5 py-2.5 font-medium bg-blue-600 hover:bg-blue-700"
          >Checkout</a
        >
      </div>
    </main>
  </body>
</html>
```
