---
id: bad87fee1348bd9aedf08817
title: Створення мертвих посилань за допомогою символу хешу
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMdkytL'
forumTopicId: 18230
dashedName: make-dead-links-using-the-hash-symbol
---

# --description--

Іноді ви хочете додати `a` елементи на свій вебсайт, перш ніж знати, куди вони будуть вести.

Це також зручно, коли ви змінюєте поведінку посилання за допомогою `JavaScript`, про що ми дізнаємося пізніше.

# --instructions--

Поточне значення атрибута `href` — це посилання, яке веде на "`https://www.freecatphotoapp.com`". Замініть значення атрибута `href` на `#`, також відомий як символ хешу, щоб створити мертве посилання.

Наприклад: `href="#"`

# --hints--

Ваш елемент `a` має бути мертвим посиланням зі значенням атрибута `href`, встановленим на "#".

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'),'#');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#" target="_blank">cat photos</a>.</p>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
