---
id: bad87fee1348bd9aedf08812
title: Додайте зображення на свій вебсайт
challengeType: 0
forumTopicId: 16640
dashedName: add-images-to-your-website
---

# --description--

Ви можете додати зображення на свій вебсайт, використовуючи елемент `img`, і вказати URL конкретного зображення за допомогою атрибута `src`.

Приклад цього:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg">
```

Зверніть увагу, що `img` — це порожній елемент.

Всі елементи `img` **повинні** мати атрибут `alt`. Текст всередині атрибута `alt` використовується читачами екрана для покращення доступності і відображається, якщо зображення не завантажується.

**Примітка:** Якщо зображення є виключно декоративним, найкращою практикою є використання порожнього атрибута `alt`.

Ідеально, щоб атрибут `alt` не містив спеціальних символів, якщо це не потрібно.

Давайте додамо атрибут `alt` до нашого прикладу `img` вище:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg" alt="freeCodeCamp logo">
```

# --instructions--

Спробуємо додати зображення на наш вебсайт:

Всередині існуючого елемента `main` вставте елемент `img` перед існуючими елементами `p`.

Тепер встановіть атрибут `src`, щоб він вказував на URL `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`.

Нарешті, не забудьте додати атрибут `img` з відповідним текстом до вашого елемента `alt`.

# --hints--

На вашій сторінці має бути елемент зображення.

```js
assert.exists(document.querySelector('img'));
```

Ваше зображення має мати атрибут `src`, який вказує на зображення кошеняти.

```js
const url = document.querySelector('img').getAttribute('src');
assert.match(url,/^https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg$/i);
```

Атрибут `alt` вашого елемента зображення не повинен бути порожнім.

```js
assert.exists(document.querySelector('img').getAttribute('alt'));
assert.isNotEmpty(document.querySelector('img').getAttribute('alt'));
assert.match(__helpers.removeWhiteSpace(code),/<(?:img|IMG)\S*alt=(['"])(?!\1|>)\S+\1\S*\/?>/)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>


  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
