---
id: bad87fee1348bd9aede08817
title: Вкласти елемент посилання в абзац
challengeType: 0
forumTopicId: 18244
dashedName: nest-an-anchor-element-within-a-paragraph
---

# --description--

Ви можете вкладати посилання всередину інших текстових елементів.

```html
<p>
  Here's a <a target="_blank" href="https://www.freecodecamp.org"> link to www.freecodecamp.org</a> for you to follow.
</p>
```

Розберемо приклад. Звичайний текст обгорнутий у елемент `p`:

```html
<p> Here's a ... for you to follow. </p>
```

Далі йде *елемент посилання* `<a>` (який потребує кінцевого тегу `</a>`):  

```html
<a> ... </a>
```

`target` — це атрибут елемента посилання, який вказує, де відкривати посилання. Значення `_blank` означає відкривати посилання у новій вкладці. `href` — це атрибут елемента посилання, який містить URL-адресу посилання:  

```html
<a href="https://www.freecodecamp.org" target="_blank"> ... </a>
```

Текст `link to www.freecodecamp.org` всередині елемента `a` називається <dfn>текстом посилання</dfn> і відображає посилання, на яке можна натиснути:

```html
<a href=" ... " target="...">link to freecodecamp.org</a>
```

Фінальний вихідний результат прикладу виглядатиме так:  

Ось <a href="https://www.freecodecamp.org" target="_blank">посилання на www.freecodecamp.org</a>, за яким ви можете перейти.

# --instructions--

Вкладіть існуючий елемент `a` всередину нового елемента `p`. Не створюйте новий елемент посилання. Новий абзац повинен містити текст `View more cat photos`, де `cat photos` — це посилання, а решта — звичайний текст.

# --hints--

У вас має бути лише один елемент `a`.

```js
assert.lengthOf(document.querySelectorAll('a'), 1 );
```

Елемент `a` повинен направляти на "`https://www.freecatphotoapp.com`".

```js
assert.lengthOf(document.querySelectorAll('a[href="https://www.freecatphotoapp.com"]'),1);
```

Ваш елемент `a` повинен мати текст посилання `cat photos`

```js
assert.match(document.querySelector('a').textContent,/cat\sphotos/gi)
```

Ви повинні створити новий елемент `p`. Усього в коді HTML має бути щонайменше 3 теги `p`.

```js
assert.lengthOf(document.querySelectorAll('p'),3)
```

Ваш елемент `a` має бути вкладений у новий елемент `p`.

```js
const anchorParent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode;
assert.strictEqual(anchorParent.tagName,"P")
```

Ваш елемент `p` повинен містити текст `View more ` (з пробілом після нього).

```js
const textContent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode.textContent;
assert.match(textContent,/View\smore\s/gi);
```

Ваш елемент `a` <em>не</em> повинен містити текст `View more`.

```js
assert.notMatch(document.querySelector('a').textContent,/View\smore/gi); 
```

Кожен ваш елемент `p` повинен мати кінцевий тег.

```js
assert.match(code,/<\/p>/g);
assert.match(code,/<p/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

Кожен ваш елемент `a` повинен мати кінцевий тег.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>View more <a target="_blank" href="https://www.freecatphotoapp.com">cat photos</a></p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
