---
id: bad87fee1348bd9aecf08801
title: Вступ до елементів HTML5
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBkZGpt7'
forumTopicId: 301097
dashedName: introduction-to-html5-elements
---

# --description--

HTML5 вводить більш описові HTML теги. Серед них є `main`, `header`, `footer`, `nav`, `video`, `article`, `section` та інші.

Ці теги надають описову структуру вашому HTML, роблять його легшим для читання та допомагають із пошуковою оптимізацією (SEO) і доступністю. `main` тег HTML5 допомагає пошуковим системам і іншим розробникам знаходити основний вміст вашої сторінки.

Приклад використання: `main` елемент із двома вкладеними дочірніми елементами всередині нього:

```html
<main> 
  <h1>Hello World</h1>
  <p>Hello Paragraph</p>
</main>
```

**Примітка:** Багато нових тегів HTML5 і їхні переваги розглядаються в розділі «Прикладна доступність».

# --instructions--

Створіть другий `p` елемент із таким текстом kitty ipsum: `Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.`

Потім створіть `main` елемент і вкладете лише два `p` елементи всередину `main` елемента.

# --hints--

У вас має бути 2 `p` елементи з текстом Kitty Ipsum.

```js
assert.lengthOf(document.querySelectorAll('p'),2);
```

Кожен із ваших `p` елементів повинен мати кінцевий тег.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

Ваш `p` елемент має містити перші кілька слів наданого додаткового тексту `kitty ipsum`.

```js
assert.match(document.querySelectorAll('p')[1].textContent,/Purr\s+jump\s+eat/gi);
```

Ваш код повинен мати один `main` елемент.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`main` елемент має мати два абзаци як дочірні елементи.

```js
const main = document.querySelector('main');
const children = main.querySelectorAll("p"); 
assert.lengthOf(children,2);
```

Відкриваючий тег `main` має йти перед першим тегом абзацу.

```js
assert.match(code,(/<main>\s*?<p>/g));
```

Кінцевий тег `main` має йти після другого закриваючого тегу абзацу.

```js
assert.match(code,(/<\/p>\s*?<\/main>/g));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
