---
id: bad87fee1348bd9aedf08802
title: Прибрати коментарі в HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBmG9T7'
forumTopicId: 18329
dashedName: uncomment-html
---

# --description--

Коментування — це спосіб залишати коментарі для інших розробників у своєму коді, не впливаючи на вихідні дані, які бачить кінцевий користувач.

Коментування також зручно, щоб зробити код неактивним, не видаляючи його повністю.

Коментарі в HTML починаються з `<!--` і закінчуються `-->`

# --instructions--

Приберіть коментарі з ваших елементів `h1`, `h2` і `p`.

# --hints--

Ваш елемент `h1` має бути видимим на сторінці після того, як ви приберете коментарі.

```js
assert.notEmpty(document.querySelectorAll('h1'));
```

Ваш елемент `h2` має бути видимим на сторінці після того, як ви приберете коментарі.

```js
assert.notEmpty(document.querySelectorAll('h2'));
```

Ваш елемент `p` має бути видимим на сторінці після того, як ви приберете коментарі.

```js
assert.notEmpty(document.querySelectorAll('p'));
```

На сторінці не повинно бути видимих закриваючих тегів коментарів (наприклад, `-->`).

```js
const elements = document.querySelectorAll('*');
const potentialComments = Array.from(elements).filter(el => el.textContent.includes('-->'));
assert.notExists(potentialComments[1])
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
