---
id: bad87fee1348bd9aede08835
title: Вкласти багато елементів у один елемент div
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cNW4kC3'
forumTopicId: 18246
dashedName: nest-many-elements-within-a-single-div-element
---

# --description--

Елемент `div`, також відомий як елемент ділення, є загальним контейнером для інших елементів.

Елемент `div` ймовірно є найпоширенішим HTML-елементом.

Як і будь-який інший елемент, що не є самозакриваючим, ви можете відкрити елемент `div` тегом `<div>` і закрити його на іншій лінії тегом `</div>`.

# --instructions--

Вкладіть ваші списки "Речі, які люблять коти" та "Топ 3 речі, які коти ненавидять" усередину одного елемента `div`.

Підказка: Спробуйте розмістити відкриваючий тег `div` над елементом `p` "Речі, які люблять коти" і закриваючий тег `div` після закриваючого тегу `ol`, щоб обидва списки були всередині одного `div`.

# --hints--

Ваші елементи `p` мають бути вкладені всередину елемента `div`.

```js
const div = document.querySelector('div');
const children = div.querySelectorAll('p');
assert.isAbove(children.length,1);
```

Ваш елемент `ul` має бути вкладений всередину елемента `div`.

```js
const div = document.querySelector('div');
const children = div.querySelectorAll('ul');
assert.notEmpty(children);
```

Ваш елемент `ol` має бути вкладений всередину елемента `div`.

```js
const div = document.querySelector('div');
const children = div.querySelectorAll('ol');
assert.notEmpty(children);
```

Ваш елемент `div` має мати кінцевий тег.

```js
assert.match(code,/<\/div>/g);
assert.strictEqual(code.match(/<\/div>/g).length,code.match(/<div>/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>

  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving" checked> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving" checked> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
