---
id: bad87fee1348bd9aedf08804
title: Перетворити HTML на коментарі
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cGyGbca'
forumTopicId: 16782
dashedName: comment-out-html
---

# --description--

Пам’ятайте, що щоб почати коментар, потрібно використати `<!--`, а щоб закінчити коментар — `-->`.

Тут вам потрібно закінчити коментар перед тим, як починається ваш `h2` елемент.

# --instructions--

Перетворіть на коментарі ваші `h1` та `p` елементи, але не `h2` елемент.

# --hints--

Ваш `h1` елемент має бути перетворений на коментарі, щоб він не був видимим на сторінці.

```js
assert.isEmpty(document.querySelectorAll('h1'));
```

Ваш `h2` елемент не повинен бути перетворений на коментарі, щоб він був видимим на сторінці.

```js
assert.isNotEmpty(document.querySelectorAll('h2'));
```

Ваш `p` елемент має бути перетворений на коментарі, щоб він не був видимим на сторінці.

```js
assert.isEmpty(document.querySelectorAll('p'));
```

Кожен ваш коментар має бути закритий `-->`.

```js
assert.isAbove(code.match(/[^fc]-->/g).length, 1);
```

Ви не повинні змінювати порядок елементів `h1`, `h2` або `p` у коді.

```js
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[0],'<h1>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[1],'<h2>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[2],'<p>');
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<!--<h1>Hello World</h1>-->
<h2>CatPhotoApp</h2> 
<!--<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p> -->
```
