---
id: bad87fee1348bd9aedf08834
title: Створіть множину радіокнопок
challengeType: 0
forumTopicId: 16822
dashedName: create-a-set-of-radio-buttons
---

# --description--

Ви можете використовувати <dfn>радіокнопки</dfn> для питань, де хочете, щоб користувач дав лише одну відповідь із кількох варіантів.

Радіокнопки — це тип `input`.

Кожна ваша радіокнопка може бути вкладена у власний елемент `label`. Обгорнувши елемент `input` всередині елемента `label`, ви автоматично зв’яжете вхідні дані радіокнопки з елементом мітки, що її оточує.

Усі пов’язані радіокнопки повинні мати однаковий атрибут `name`, щоб створити групу радіокнопок. Створюючи таку групу, вибір будь-якої однієї радіокнопки автоматично зніме вибір з інших кнопок у тій самій групі, гарантуючи, що користувач надасть лише одну відповідь.

Ось приклад радіокнопки:

```html
<label> 
  <input type="radio" name="indoor-outdoor">Indoor 
</label>
```

Вважається найкращою практикою встановлювати атрибут `for` на елементі `label` зі значенням, що співпадає зі значенням атрибута `id` елемента `input`. Це дозволяє допоміжним технологіям створити зв’язок між міткою та пов’язаним елементом `input`. Наприклад:

```html
<input id="indoor" type="radio" name="indoor-outdoor">
<label for="indoor">Indoor</label>
```

Ми також можемо вкладати елемент `input` всередину тегів `label`:

```html
<label for="indoor"> 
  <input id="indoor" type="radio" name="indoor-outdoor">Indoor 
</label>
```

# --instructions--

Додайте пару радіокнопок до своєї форми, кожну вкладену у власний елемент `label`. Одна має мати опцію `indoor`, а інша — опцію `outdoor`. Обидві повинні мати однаковий атрибут `name` зі значенням `indoor-outdoor`, щоб створити групу радіокнопок.

# --hints--

На вашій сторінці має бути два елементи `radio` кнопок.

```js
assert.lengthOf(document.querySelectorAll('input[type="radio"]'),2);
```

Ваші радіокнопки повинні мати атрибут `name` зі значенням `indoor-outdoor`.

```js
assert.lengthOf([...document.querySelectorAll('input[type="radio"]')].filter(input => input.name === "indoor-outdoor"),2);
```

Кожен із двох елементів радіокнопок має бути вкладений у власний елемент `label`.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="radio"]:only-child'),2);
```

Кожен із ваших елементів `label` має мати кінцевий тег.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length);
```

Одна з ваших радіокнопок має мати мітку `indoor`.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label =>  label.match(/indoor/gi));
assert.isTrue(hasMatch); 
```

Одна з ваших радіокнопок має мати мітку `outdoor`.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label => label.match(/outdoor/gi));
assert.isTrue(hasMatch); 
```

Кожен із ваших елементів радіокнопок має бути доданий всередину тегу `form`.

```js
const labelParent = document.querySelector('label').parentNode;
assert.strictEqual(labelParent.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
   <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
