---
id: bad87fee1348bd9aedf08830
title: додати текст заповнювача до текстового поля
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cKdJDhg'
forumTopicId: 16647
dashedName: add-placeholder-text-to-a-text-field
---

# --description--

Текст заповнювача — це те, що відображається у вашому `input` елементі до того, як користувач введе будь-які вхідні дані.

Ви можете створити текст заповнювача так:

```html
<input type="text" placeholder="this is placeholder text">
```

**Примітка:** Пам’ятайте, що `input` є порожнім елементом.

# --instructions--

Встановіть значення `placeholder` вашого текстового `input` на "cat photo URL".

# --hints--

Ви повинні додати атрибут `placeholder` до існуючого текстового `input` елемента.

```js
assert.notEmpty(document.querySelectorAll('input[placeholder]'));
```

Ви повинні встановити значення вашого атрибута `placeholder` на `cat photo URL`.

```js
assert.exists(document.querySelector('input'));
assert.exists(document.querySelector('input').getAttribute('placeholder'));
const placeholder = document.querySelector('input').getAttribute('placeholder');
assert.match(placeholder,/cat\s+photo\s+URL/gi);
```

Готовий `input` елемент не повинен мати кінцевого тегу.

```js
assert.notMatch(code,/<input.*\/?>.*<\/input>/gi);
```

Готовий `input` елемент повинен мати правильний синтаксис.

```js
assert.notEmpty(document.querySelectorAll('input[type=text]'));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```
