---
id: bad87fee1348bd9aedf08833
title: Заповніть пропуск текстом-заповнювачем
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cgR7Dc7'
forumTopicId: 18178
dashedName: fill-in-the-blank-with-placeholder-text
---

# --description--

Веб-розробники традиційно використовують <dfn>lorem ipsum text</dfn> як текст-заповнювач. Lorem ipsum текст випадково взято з відомого уривка Цицерона з Стародавнього Риму.

Lorem ipsum текст використовували як текст-заповнювач друкарі з XVI століття, і ця традиція продовжується в інтернеті.

Ну, 5 століть — це вже досить довго. Оскільки ми створюємо застосунок CatPhotoApp, давайте використаємо текст під назвою "kitty ipsum".

# --instructions--

Замініть текст всередині вашого `p` елемента на перші кілька слів цього kitty ipsum тексту: `Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.`

# --hints--

Ваш `p` елемент має містити перші кілька слів наданого "kitty ipsum" тексту.

```js
assert.match(document.querySelector('p').textContent,/Kitty(\s)+ipsum/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Hello Paragraph</p>
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff</p>
```
