---
id: bad87fee1348bd9aedf08816
title: Посилання на зовнішні сторінки за допомогою елементів посилання
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c8EkncB'
forumTopicId: 18226
dashedName: link-to-external-pages-with-anchor-elements
---

# --description--

Ви можете використовувати `a` (*елементи посилання*) для створення посилань на вміст поза межами вашої вебсторінки.

`a` елементам потрібна адреса призначення, яка називається `href` атрибутом. Також їм потрібен текст посилання. Ось приклад:

```html
<a href="https://www.freecodecamp.org">this links to freecodecamp.org</a>
```

Тоді ваш браузер відобразить текст `this links to freecodecamp.org` як посилання, на яке можна натиснути. І це посилання направить вас на вебадресу `https://www.freecodecamp.org`.

# --instructions--

Створіть `a` елемент, який направляє на `https://www.freecatphotoapp.com` і має текст посилання "cat photos".

# --hints--

Ваш `a` елемент має містити текст посилання `cat photos`.

```js
assert.match(document.querySelector('a').textContent,/cat photos/gi);
```

Вам потрібен `a` елемент, який направляє на `https://www.freecatphotoapp.com`

```js
assert.match(document.querySelector('a').getAttribute('href'),/^https?:\/\/(www\.)?freecatphotoapp\.com\/?$/i);
```

Ваш `a` елемент має мати кінцевий тег.

```js
assert.match(code,/<\/a>/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>



  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <a href="https://www.freecatphotoapp.com">cat photos</a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
