---
id: bad88fee1348bd9aedf08816
title: Направити на внутрішні розділи сторінки за допомогою елементів посилання
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cyrDRUL'
forumTopicId: 301098
dashedName: link-to-internal-sections-of-a-page-with-anchor-elements
---

# --description--

`a` (*елементи посилання*) також можна використовувати для створення внутрішніх посилань, щоб переходити до різних розділів на сторінці.

Щоб створити внутрішнє посилання, призначте атрибуту `href` посилання символ решітки `#` плюс значення атрибута `id` елемента, на який ви хочете направити посилання всередині сторінки, зазвичай далі вниз по сторінці. Потім потрібно додати той самий атрибут `id` до елемента, на який ви посилаєтеся. `id` — це атрибут, який унікально описує елемент.

Нижче наведено приклад внутрішнього посилання-анкера та його цільового елемента:

```html
<a href="#contacts-header">Contacts</a>
...
<h2 id="contacts-header">Contacts</h2>
```

Коли користувачі натискають на `Contacts` посилання, їх буде перенаправлено до розділу сторінки з елементом заголовка **Contacts**.

# --instructions--

Змініть своє зовнішнє посилання на внутрішнє, змінивши атрибут `href` на `#footer` і текст з `cat photos` на `Jump to Bottom`.

Видаліть атрибут `target="_blank"` з тегу посилання, оскільки він викликає відкриття пов’язаної сторінки у новій вкладці вікна.

Потім додайте атрибут `id` зі значенням `footer` до елемента `<footer>` внизу сторінки.

# --hints--

На вашій сторінці має бути лише один тег посилання.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

На вашій сторінці має бути лише один тег `footer`.

```js
assert.lengthOf(document.querySelectorAll('footer'),1);
```

Тег `a` має мати атрибут `href`, встановлений у "#footer".

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'), '#footer');
```

Тег `a` не повинен мати атрибуту `target`.

```js
assert.notExists(document.querySelector('a').getAttribute('target'));
```

Текст `a` має бути "Jump to Bottom".

```js
assert.match(document.querySelector('a').textContent,/Jump to Bottom/gi);
```

Тег `footer` має мати атрибут `id`, встановлений у "footer".

```js
assert.strictEqual(document.querySelector('footer').getAttribute('id'),'footer');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer>Copyright Cat Photo App</footer>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="#footer">Jump to Bottom</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer id="footer">Copyright Cat Photo App</footer>
```
