---
id: bad87fed1348bd9aedf08833
title: Видалити HTML-елементи
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ckK73C9'
forumTopicId: 17559
dashedName: delete-html-elements
---

# --description--

У нашого телефону мало вертикального простору.

Давайте видалимо непотрібні елементи, щоб почати створювати наш застосунок CatPhotoApp.

# --instructions--

Видаліть свій `h1` елемент, щоб ми могли спростити наш вигляд.

# --hints--

Ваш `h1` елемент слід видалити.

```js
assert.notMatch(code,/<h1>/gi);
assert.notMatch(code,/<\/h1>/gi);
```

Ваш `h2` елемент має бути на сторінці.

```js
assert.match(code,/<h2>[\w\W]*<\/h2>/gi);
```

Ваш `p` елемент має бути на сторінці.

```js
assert.match(code,/<p>[\w\W]*<\/p>/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2><p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
