---
id: bad87fee1348bd9aedf08801
title: Повідомлення за допомогою елемента абзац
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ceZ7DtN'
forumTopicId: 18202
dashedName: inform-with-the-paragraph-element
---

# --description--

Елемент `p` є рекомендованим елементом для тексту абзацу на вебсайтах. `p` — це скорочення від "paragraph".

Ви можете створити елемент абзацу так:

```html
<p>I'm a p tag!</p>
```

# --instructions--

Створіть елемент `p` під вашим елементом `h2` і додайте йому текст `Hello Paragraph`.

**Примітка:** За конвенцією всі HTML-теги пишуться в нижньому регістрі, наприклад `<p></p>`, а не `<P></P>`.

# --hints--

Ваш код має містити дійсний елемент `p`.

```js
assert.lengthOf(document.querySelectorAll('p'),1);
```

Ваш елемент `p` має містити текст `Hello Paragraph`.

```js
assert.match(document.querySelector('p').textContent,/hello(\s)+paragraph/gi);
```

Ваш елемент `p` має мати кінцевий тег.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
```

# --solutions--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
<p>Hello Paragraph</p>
```
