---
id: bad87fee1348bd9aedf08827
title: Створіть невпорядкований список з маркерами
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cDKVPuv'
forumTopicId: 16814
dashedName: create-a-bulleted-unordered-list
---

# --description--

HTML має спеціальний елемент для створення <dfn>невпорядкованих списків</dfn>, або списків у вигляді маркерів.

Невпорядковані списки починаються з відкриваючого `<ul>` елемента, за яким йде будь-яка кількість `<li>` елементів. Нарешті, невпорядковані списки закриваються `</ul>`.

Наприклад:

```html
<ul>
  <li>milk</li>
  <li>cheese</li>
</ul>
```

створить список у вигляді маркерів з `milk` та `cheese`.

# --instructions--

Видаліть останні два `p` елементи та створіть невпорядкований список з трьох речей, які люблять коти, внизу сторінки.

# --hints--

Створіть `ul` елемент.

```js
assert.isNotEmpty(document.querySelectorAll('ul'));
```

У вашому `li` елементі має бути три `ul` елементи.

```js
assert.lengthOf(document.querySelectorAll('ul li'),3);
```

Ваш `ul` елемент повинен мати кінцевий тег.

```js
assert.match(code,/<\/ul>/gi);
assert.match(code,/<ul/gi);
assert.strictEqual(code.match(/<\/ul>/gi).length,code.match(/<ul/gi).length);
```

Ваші `li` елементи повинні мати кінцеві теги.

```js
assert.match(code,/<\/li>/gi);
assert.match(code,/<li[\s>]/gi);
assert.strictEqual(code.match(/<\/li>/gi).length,code.match(/<li[\s>]/gi).length);
```

Ваші `li` елементи не повинні містити порожній рядок або лише пробіли.

```js
assert.isEmpty([...document.querySelectorAll('ul li')].filter((item) => item.textContent.trim() === ""));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <ul>
    <li>milk</li>
    <li>mice</li>
    <li>catnip</li>
  </ul>
</main>
```
