---
id: bd7123c8c441eddfaeb5bdef
title: Привітання елементам HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cE8Gpt2'
forumTopicId: 18276
dashedName: say-hello-to-html-elements
---

# --description--

Ласкаво просимо до завдань з кодування HTML на freeCodeCamp. Вони проведуть вас крок за кроком через розробку вебсторінок.

Спочатку ви почнете зі створення простої вебсторінки за допомогою HTML. Ви можете редагувати код у своєму редакторі коду, який вставлений у цю вебсторінку.

Чи бачите код у своєму редакторі коду, який каже `<h1>Hello</h1>`? Це HTML-елемент.

Більшість HTML-елементів мають початковий тег і кінцевий тег.

Початкові теги виглядають так:

```html
<h1>
```

Кінцеві теги виглядають так:

```html
</h1>
```

Єдина різниця між початковими та кінцевими тегами — це коса риска після відкриваючої дужки кінцевого тегу.

У кожному завданні є тести, які ви можете запускати в будь-який час, натиснувши на кнопку "Run tests". Коли ви пройдете всі тести, вам запропонують відправити своє рішення і перейти до наступного завдання з кодування.

# --instructions--

Щоб пройти тест у цьому завданні, змініть текст вашого `h1` елемента на `Hello World`.

# --hints--

Ваш `h1` елемент повинен містити текст `Hello World`.

```js
assert.match(document.querySelector('h1').textContent,/hello(\s)+world/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello</h1>
```

# --solutions--

```html
<h1>Hello World</h1>
```
