---
id: 60f803d5241e6a0433a523a1
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

En los videos de la clase anterior, aprendiste cómo trabajar con unidades `rem`. Recuerda que la unidad `rem` significa raíz `em`, y es relativa al tamaño de fuente del elemento `html`.

Como los elementos `label` son en línea por defecto, todos se muestran uno al lado del otro en la misma línea, dificultando la lectura de su texto.

Para que aparezcan en líneas separadas, añade `display: block` al elemento `label`, y agrega un `margin` de `0.5rem 0`, para separarlos entre sí.

# --hints--

Debes usar el selector `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Debes añadir una propiedad `display` con el valor `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Debes añadir una propiedad `margin` con el valor `0.5rem 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```

