---
id: 63b6088bb56e2d1cac364043
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Continuando el patrón, crea un `fieldset` para `dinner` con los mismos elementos anidados.

# --hints--

Debes crear un tercer elemento `fieldset` en tu `form`.

```js
assert.exists(document.querySelectorAll('form fieldset')[2]);
```

Tu tercer elemento `fieldset` debe venir después del segundo elemento `fieldset`.

```js
assert.equal(document.querySelectorAll('form fieldset')[2]?.previousElementSibling?.tagName, "FIELDSET");
```

Tu tercer elemento `fieldset` debe tener un `id` establecido a `dinner`.

```js
assert.equal(document.querySelectorAll('form fieldset')[2]?.id, "dinner");
```

Tu tercer elemento `fieldset` debe contener un elemento `legend`.

```js
assert.exists(document.querySelectorAll('form fieldset')[2]?.querySelector('legend'));
```

Tu nuevo elemento `legend` debe tener el texto `Dinner`.

```js
assert.equal(document.querySelectorAll('form fieldset')[2]?.querySelector('legend')?.innerText, "Dinner");
```

Tu tercer elemento `fieldset` debe contener un elemento `div`.

```js
assert.exists(document.querySelectorAll('form fieldset')[2]?.querySelector('div'));
```

Tu nuevo elemento `div` debe tener una `class` establecida a `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[2]?.querySelector('div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
