---
id: 646c5e727e56e743c9aed4a1
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

CSS 위치 지정은 브라우저에서 요소를 어떻게 배치할지 설정할 수 있게 합니다. `position` 속성에 `static`, `absolute`, `relative`, `sticky` 또는 `fixed` 중 하나를 설정할 수 있습니다.

요소의 `position` 속성을 설정하면 `top`, `right`, `left` 또는 `bottom` 속성 중 하나 이상에 픽셀 또는 백분율 값을 설정해 요소를 이동할 수 있습니다.

`static`은 모든 요소의 기본 위치 지정 방식입니다. 요소에 이 속성을 할당하면 `top`, `right`, `left` 또는 `bottom`로 이동할 수 없습니다.

`.cat-head`에 `position` 속성으로 `static`를 할당한 다음, `top`와 `left` 속성에 각각 `100px` 값을 설정하세요.


# --hints--

`.cat-head` 선택자에 `position` 속성을 `static`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'static')
```

`.cat-head` 선택자에 `top` 속성을 `100px`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px')
```

`.cat-head` 선택자에 `left` 속성을 `100px`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
