---
id: 646c5ffef5598d449b52ec12
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

이제 고양이 귀 작업을 하세요. 오른쪽 귀와 왼쪽 귀가 각각 있고, 각 귀 안에는 안쪽 귀가 있습니다.

`.cat-head` 요소 안에 클래스 `div`를 가진 `cat-ears` 요소를 만드세요.

# --hints--

기존에 있는 클래스 `div`를 가진 `cat-head` 요소는 변경하지 마세요.

```js
assert.lengthOf(document.querySelectorAll('div.cat-head'), 1);
```

`div` 요소 안에 `.cat-head` 요소를 하나 만들어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.cat-head div'), 1);
```

`div` 요소는 클래스 `cat-ears`를 가져야 합니다.

```js
assert.equal(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class'),'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">
        
      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
