---
id: 646c586be7180e40ddf74ff6
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

`body` 요소에 `background-color`의 `#c9d2fc`을(를) 지정하세요.

# --hints--

`body` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` 선택자에 `background-color`으로 설정된 `#c9d2fc` 속성이 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(201, 210, 252)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```

