---
id: 646ceb843412c74edee27a79
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

클래스 선택자를 사용해 `.cat-right-ear` 요소에 `height`와 `width` 속성을 `100px`으로 설정하세요. `background-color`를 `white`로 설정하세요. 왼쪽과 오른쪽 테두리는 `35px solid blue`으로 설정하세요. 위쪽과 아래쪽 테두리는 `35px solid red`으로 설정하세요.

# --hints--

`.cat-right-ear` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

`.cat-right-ear` 선택자에 `height` 속성을 `100px`로 설정하세요. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

`.cat-right-ear` 선택자에 `width` 속성을 `100px`로 설정하세요. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

`.cat-right-ear` 선택자에 `background-color` 속성을 `white`로 설정하세요. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

`.cat-right-ear` 선택자에 `border-left` 속성을 `35px solid blue`로 설정하세요. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

`.cat-right-ear` 선택자에 `border-right` 속성을 `35px solid blue`로 설정하세요. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

`.cat-right-ear` 선택자에 `border-top` 속성을 `35px solid red`로 설정하세요. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

`.cat-right-ear` 선택자에 `border-bottom` 속성이 `35px solid red`로 설정되어 있어야 합니다. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
