---
id: 646cf1206cac5f51804f49cf
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

왼쪽 귀를 적절히 위치시키려면 CSS transform을 사용해 특정 각도로 회전시킬 수 있습니다.

`transform` 속성은 레이아웃을 변경하거나 주변 요소에 영향을 주지 않고 요소의 모양, 위치, 크기를 수정할 수 있게 해줍니다. `translate()`, `rotate()`, `scale()`, `skew()`, `matrix()` 같은 기능들이 있습니다.

`transform` 속성을 `rotate(-45deg)`로 설정하고 어떤 변화가 있는지 확인하세요.

# --hints--

`transform` 요소의 `.cat-left-ear` 속성을 `rotate(-45deg)`로 설정해야 합니다. 세미콜론을 잊지 마세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.transform, 'rotate(-45deg)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  --fcc-editable-region--
  
  --fcc-editable-region--
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
