---
id: 646cf2249f02ca5233d9af7c
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

이제 고양이의 오른쪽 귀를 작업할 수 있습니다. 이미 HTML이 준비되어 있습니다.

클래스 선택자를 사용해 `.cat-right-ear` 요소에 좌우 테두리를 각각 `35px solid transparent`로 지정하세요. 또한, 아래쪽 테두리는 `70px solid #5e5e5e`로 설정하세요.

# --hints--

`.cat-right-ear` 선택자가 있어야 합니다.

```js 
assert.exists(new __helpers.CSSHelp(document).getStyle('.cat-right-ear'))
```

`.cat-right-ear` 선택자에 `border-left` 속성을 `35px solid transparent`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid transparent')
```

`.cat-right-ear` 선택자에 `border-right` 속성을 `35px solid transparent`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid transparent')
```

`.cat-right-ear` 선택자에 `border-bottom` 속성을 `70px solid #5e5e5e`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

--fcc-editable-region--

--fcc-editable-region--
```
