---
id: 646c48df8674cf2b91020ecc
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

기본 보일러플레이트가 제공되었습니다. `link` 요소 안에 `head` 요소를 추가하는 것부터 시작하세요. 그 `link` 요소에 `rel` 속성을 `stylesheet`로, `href`를 `./styles.css`로 설정하세요.

# --hints--

`link` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("head>link"));
```

`link` 요소에는 `rel`와 `href` 속성이 각각 `stylesheet`과 `./styles.css`로 설정되어야 합니다.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
