---
id: 646c59e942f35541923104bf
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

HTML로 돌아가서 `main` 요소를 만드세요. 그 `main` 요소 안에, 클래스가 `div`인 `cat-head` 요소를 추가하세요.

# --hints--

`main` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

`div` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

`div` 요소에 클래스 `cat-head`가 있어야 합니다.

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

`div` 요소가 `main` 태그 안에 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
