---
id: 6476f7a4827bcc61682f2347
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

`absolute` 속성에 `position` 값을 사용하면 요소가 문서의 일반 흐름에서 벗어나고, 그 위치는 `top`, `right`, `bottom`, `left` 속성에 의해 결정됩니다.

`position` 요소의 `.cat-head` 속성을 `absolute`로 설정한 다음, `top`와 `left` 속성을 임의의 양수 픽셀 값으로 설정하세요.

<!-- **참고**: 여기서 `top`, `left`, `bottom`, `right` 속성을 실험할 수 있지만, 테스트는 `top`의 `300px`와 `400px`의 left에 대해서만 통과합니다. -->

# --hints--

`.cat-head` 선택자에 `position` 속성을 `absolute`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

`.cat-head` 선택자에 `top` 속성을 `px` 내의 임의 숫자로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

`.cat-head` 선택자에 `left` 속성을 `px` 내의 임의 숫자로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
