---
id: 6476f5c17f99146071ee884c
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

아무 일도 일어나지 않는 것을 볼 수 있습니다. `.cat-head` 요소는 각각 `top`인 `left`과 `100px`를 설정했음에도 움직이지 않았습니다. 하지만 `relative` 위치 지정은 그렇지 않습니다.

`relative` 값을 사용하면 요소가 여전히 문서의 일반 흐름에 따라 위치하지만 `top`, `left`, `bottom`, 그리고 `right` 값이 활성화됩니다.

`static` 대신 `.cat-head`에 `relative` 위치를 지정하고 `top`과 `left` 속성은 그대로 두세요.

# --hints--

`.cat-head` 선택자에 `position` 속성을 `relative`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

`.cat-head` 선택자에 `top` 속성을 `100px`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

`.cat-head` 선택자에 `left` 속성을 `100px`로 설정하세요. 세미콜론을 꼭 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
