---
id: 6476fc5cf14b276e6d04e82a
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

클래스 선택자를 사용해 `.box` 요소에 너비 `200px`, 높이 `600px`, 배경색 `#000`을 지정하세요. 또한 `position`를 `absolute`로, `top`를 `800px`로, `left`를 `650px`로 지정하세요.

# --hints--

`.box` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.box'));
```

`.box` 선택자에 `width` 속성을 `200px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.width, '200px');
```

`.box` 선택자에 `height` 속성을 `600px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.height, '600px');
```

`.box` 선택자에 `background-color` 속성을 `#000`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor, 'rgb(0, 0, 0)');
```

`.box` 선택자에 `position` 속성을 `absolute`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.position, 'absolute');
```

`.box` 선택자에 `left` 속성을 `650px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.left, '650px');
```

`.box` 선택자에 `top` 속성을 `800px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.top, '800px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

