---
id: 6476fb7cbaafa36d65e9cf35
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

`fixed`는 사용자가 페이지에서 어디로 스크롤하든 요소를 고정할 수 있게 하는 `position` 속성 값입니다.

`fixed` 위치 지정이 어떻게 작동하는지 보려면 마크업을 좀 더 해야 합니다. HTML에서 클래스가 `div`인 `box` 요소를 만드세요.

# --hints--

`div` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

`div` 요소에 클래스 `box`이 있어야 합니다.

```js
assert.equal(document.querySelectorAll('div')[1]?.getAttribute('class'), 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
