---
id: 64770351e8586671ec0911f0
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

`sticky` 위치 지정은 `relative`과 `fixed` 위치 지정의 혼합형입니다. 스크롤 위치에 따라 요소가 포함된 요소나 뷰포트 내 특정 위치에 **고정**되도록 허용합니다.

`position`의 `.cat-head` 속성 값을 `sticky`로 변경하고, `top`를 `0`로 설정한 다음 `left`와 그 값을 제거하세요.

**참고**: `sticky`가 어떻게 작동하는지 보려면 `.cat-head` 요소 앞뒤에 텍스트를 몇 개 배치해야 합니다. 그 후 아래로 스크롤하면 `.cat-head`가 상단에 고정되어 그 자리에 머무르는 것을 볼 수 있습니다.

# --hints--

`.cat-head` 선택자에 `position` 속성을 `sticky`로 설정해야 합니다. 세미콜론을 꼭 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'sticky');
```

`.cat-head` 선택자에 `top` 속성을 `0`로 설정해야 합니다. 세미콜론을 꼭 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

코드에 `left` 속성과 그 값을 포함하지 않아야 합니다.

```js
assert.notMatch(code, /left:\s*100px;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
