---
id: 6676a8a8d9861319ee901a90
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

이제 고양이 귀 작업을 시작할 수 있습니다.

검토 코드를 정리하여 `.cat-right-ear` 선택자와 그 모든 속성을 제거하세요.

클래스 선택자를 사용해 `.cat-left-ear` 요소에 좌우 테두리를 각각 `35px solid transparent`로 지정하세요. 또한, 아래쪽 테두리는 `70px solid #5e5e5e`로 설정하세요.

# --hints--


`.cat-right-ear` 선택자에는 `height` 속성이 없어야 합니다.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
}
```

`.cat-right-ear` 선택자에는 `width` 속성이 없어야 합니다.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
}
```


`.cat-right-ear` 선택자에는 `border-left` 속성이 없어야 합니다.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
}
```

`.cat-right-ear` 선택자에는 `border-right` 속성이 없어야 합니다.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
}
```

`.cat-right-ear` 선택자에는 `border-top` 속성이 없어야 합니다.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
}
```

`.cat-right-ear` 선택자에는 `border-bottom` 속성이 없어야 합니다.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
}
```

`.cat-right-ear` 선택자가 없어야 합니다.

```js
assert.notExists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

`.cat-left-ear` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-left-ear'))
```

`.cat-left-ear` 선택자에는 `border-left` 속성이 `35px solid transparent`로 설정되어 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft, '35px solid transparent')
```

`.cat-left-ear` 선택자에는 `border-right` 속성이 `35px solid transparent`로 설정되어 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight, '35px solid transparent')
```

`.cat-left-ear` 선택자에는 `border-bottom` 속성이 `70px solid #5e5e5e`로 설정되어 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom, '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
