---
id: 6676a8b01e56ec1a1e07c202
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

이제 두꺼운 테두리가 있는 흰색 사각형이 생긴 것을 확인하세요. 테두리는 대각선 모서리를 가지고 있어 삼각형을 만들 때 사용할 수 있습니다. 고양이 귀를 만들 때 필요합니다.

같은 클래스 선택자 내에서 `height`와 `width`을 `0`로 조정하세요. 왼쪽, 오른쪽, 위쪽 테두리를 `transparent`로 설정하세요.

`background-color` 속성을 제거하면 삼각형이 보일 것입니다.

# --hints--


`.cat-right-ear` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

`.cat-right-ear` 선택자에는 `height` 속성이 `0`로 설정되어 있어야 합니다. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '0px')
```

`.cat-right-ear` 선택자에는 `width` 속성이 `0`로 설정되어 있어야 합니다. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '0px')
```

`.cat-right-ear` 선택자에는 `border-left` 속성이 `35px solid transparent`로 설정되어 있어야 합니다. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid transparent')
```

`.cat-right-ear` 선택자에는 `border-right` 속성이 `35px solid transparent`로 설정되어 있어야 합니다. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid transparent')
```

`.cat-right-ear` 선택자에는 `border-top` 속성이 `35px solid transparent`로 설정되어 있어야 합니다. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid transparent')
```


`.cat-right-ear` 선택자에는 `border-bottom` 속성이 `35px solid red`로 설정되어 있어야 합니다. 세미콜론을 잊지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

`.cat-right-ear` 선택자에는 `background-color` 속성이 없어야 합니다.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
