---
id: 66ec4c8e9878d8441956516f
title: بناء جدول فهرس الكتب
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ عنصر `table` يعرض معلومات الكتب.
1. يجب أن يحتوي جدولك على عنصر رأس الجدول مع صف واحد بداخله.
1. يجب أن يحتوي الصف في عنصر رأس الجدول على أربعة عناصر رأس جدول، بالنصوص `Title`، `Author`، `Genre`، و`Publication Year`، بهذا الترتيب.
1. يجب أن يحتوي جدولك على عنصر جسم الجدول مع خمسة صفوف على الأقل.
1. يجب أن يحتوي كل صف في جسم الجدول على أربعة عناصر بيانات جدول تعرض عنوان الكتاب، المؤلف، النوع، وسنة النشر.
1. يجب أن يحتوي جدولك على عنصر تذييل الجدول مع صف واحد بداخله.
1. يجب أن يحتوي الصف في عنصر تذييل الجدول على عنصر بيانات جدول يمتد عبر أربعة أعمدة ويحمل النص `Total Books: N` حيث يجب استبدال `N` بعدد الكتب في جدولك.

# --hints--

يجب أن يكون لديك عنصر `table` واحد.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

يجب أن يكون لديك عنصر `thead` واحد داخل عنصر `table` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

يجب أن يكون لديك عنصر `tr` واحد داخل عنصر `thead` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

يجب أن يكون لديك أربعة عناصر `th` داخل صف عنصر `thead` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

يجب أن تحتوي عناصر `th` الأربعة على النصوص `Title`، `Author`، `Genre`، و`Publication Year`، بهذا الترتيب.

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent.trim(), 'Title');
assert.equal(ths[1]?.textContent.trim(), 'Author');
assert.equal(ths[2]?.textContent.trim(), 'Genre');
assert.equal(ths[3]?.textContent.trim(), 'Publication Year');
```

يجب أن يكون لديك عنصر `tbody` واحد داخل عنصر `table` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

يجب أن يحتوي عنصر `tbody` على خمسة صفوف على الأقل.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

يجب أن يحتوي كل صف في عنصر `tbody` على أربعة عناصر `td` بالضبط كأبناء.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

يجب أن يحتوي كل عنصر `td` في جسم الجدول على نص يعرض معلومات الكتاب.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.trim().length, 1);
});
```

يجب أن يكون لديك عنصر `tfoot` واحد داخل عنصر `table` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

يجب أن يكون لديك عنصر `tr` واحد بالضبط داخل عنصر `tfoot` الخاص بك.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

يجب أن يحتوي عنصر `td` في صف عنصر `tfoot` على الخاصية `colspan` مضبوطة على `4`.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

يجب أن يحتوي عنصر `td` في صف عنصر `tfoot` على النص `Total Books: N` حيث `N` هو عدد الكتب في جدولك.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
