---
id: 614385513d91ae5c251c2052
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви створите сторінку журналу, використовуючи пресреліз із попередньої версії навчальної програми freeCodeCamp.

Почніть із додавання `link` елемента для `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` стилю шрифту, `link` для стилю `https://use.fontawesome.com/releases/v5.8.2/css/all.css` FontAwesome і `link` для вашого `./styles.css` стилю.

Ваш стиль шрифту завантажить три окремі шрифти: `Anton`, `Baskervville` і `Raleway`.

# --hints--

Ваш код має містити три `link` елементи.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Ваші `link` елементи повинні бути всередині вашого `head` елемента.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

Ваші три `link` елементи повинні мати атрибут `rel` зі значенням `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Один із ваших елементів зв’язку повинен мати множину `href` зі значенням `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Один із ваших елементів зв’язку повинен мати множину `href` зі значенням `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Один із ваших `link` елементів повинен мати множину `href` зі значенням `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
