---
id: 6165d3b702a5d92ad970b30c
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Після вашого елемента `img` додайте елемент `h1` з множиною `class`, встановленою на `hero-title`, і текстом, встановленим на `OUR NEW CURRICULUM`, а потім елемент `p` з множиною `class`, встановленою на `hero-subtitle`, і текстом, встановленим на `Our efforts to restructure our curriculum with a more project-based focus`.

# --hints--

Ви повинні створити елемент `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ваш елемент `h1` має йти після вашого елемента `img`.

```js
assert.equal(document.querySelector('h1')?.previousElementSibling?.localName, 'img');
```

Ваш елемент `h1` має мати множину `class`, встановлену на `hero-title`.

```js
assert.equal(document.querySelector('h1')?.className, 'hero-title');
```

Ваш елемент `h1` має мати текст, встановлений на `OUR NEW CURRICULUM`.

```js
assert.equal(document.querySelector('h1')?.textContent, 'OUR NEW CURRICULUM');
```

Ви повинні створити новий елемент `p`.

```js
assert.exists(document.querySelector('p'));
```

Ваш елемент `p` має йти після вашого елемента `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Ваш елемент `p` має мати множину `class`, встановлену на `hero-subtitle`.

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

Ваш елемент `p` має мати текст, встановлений на `Our efforts to restructure our curriculum with a more project-based focus`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
