---
id: 68b94f23f7b3ee14078c89c5
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将使用 `iframe` 元素来显示一个视频。基本的超文本标记语言 公式化 已为你准备好。

首先创建一个带有文本 `iframe Video Display` 的 `h1` 元素。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'))
```

你的 `h1` 元素应包含文本 `iframe Video Display`。

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
