---
id: 68ba9c4f1688914d72876096
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

其中一个属性是 `allow`。它就像一个权限列表，告诉浏览器 `iframe` 被允许使用哪些特色。

这是一个带有 `allow` 属性的 `iframe` 元素：

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

为你的 `allow` 属性添加值 `accelerometer`、`autoplay` 和 `clipboard-write`。

`accelerometer` 让 `iframe` 使用运动传感器，从而能够检测设备倾斜和旋转等情况。`autoplay` 让视频自动开始播放，`clipboard-write` 允许 `iframe` 将数据写入用户的剪贴板。

# --hints--

你的 `iframe` 元素应具有设置为 `accelerometer` 的 `allow` 属性。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

你的 `iframe` 元素的 `allow` 属性应包含 `autoplay` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

你的 `iframe` 元素的 `allow` 属性应包含 `clipboard-write` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
