---
id: 68ba9c4f1688914d72876093
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

现在，创建一个 `iframe` 元素。暂时不要往里面放任何内容。

# --hints--

你应该有一个 `iframe` 元素。

```js
assert.exists(document.querySelector('iframe'))
```

你的 `iframe` 元素应为空。

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
