---
id: 68ba9c4f1688914d72876097
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

为 `allow` 属性中的现有值添加 `encrypted-media`、`gyroscope` 和 `web-share`。

这三者将允许使用加密媒体扩展来保护视频，授予访问设备的运动和方向传感器的权限，并允许通过设备的本地对话框共享 iframe 内容。 

# --hints--

你的 `iframe` 元素的 `allow` 属性应包含 `encrypted-media` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

你的 `iframe` 元素的 `allow` 属性应包含 `gyroscope` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

你的 `iframe` 元素的 `allow` 属性应包含 `web-share` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
