---
id: 68ba9c4f1688914d72876094
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在关于 `iframe` 元素的第一课中，你学到了它是一个替换元素，就像 `img` 一样。这意味着它也可以使用 `width` 和 `height` 属性来确定它的高度和宽度。

为你的 `iframe` 元素设置 `width` 为 `560` 和 `height` 为 `315`。

# --hints--

你的 `iframe` 元素应具有设置为 `560` 的 `width` 属性。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

你的 `iframe` 元素应具有设置为 `315` 的 `height` 属性。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
