---
id: 68ba9c4f1688914d72876095
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

`iframe` 元素也接受一个 `src` 属性，其值指示要显示的资源的 URL 或路径。

为你的 `iframe` 元素添加一个 `src` 属性，值为 `https://www.youtube.com/embed/I0_951_MPE0`。

此时，你应该能看到页面上显示的视频，但你还需要添加一些属性。

# --hints--

你的 `iframe` 元素应具有设置为 `https://www.youtube.com/embed/I0_951_MPE0` 的 `src` 属性。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
