---
id: 68ba9c4f1688914d72876098
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

下一个你将添加的属性是 `referrerpolicy`。它是决定当你的页面连接到另一个页面时你共享多少细节的规则。

为你的元素添加 `referrerpolicy` 属性并将其设置为 `strict-origin-when-cross-origin`。这会在同一站点共享完整地址，在其他站点只共享站点名称，在不安全站点则不共享任何信息。

# --hints--

你的 `iframe` 元素应具有设置为 `strict-origin-when-cross-origin` 的 `referrerpolicy` 属性。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
