---
id: 683ecb010f4c6bafd428802c
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você validará um conjunto de dados médicos para garantir que ele esteja em conformidade com um conjunto de regras.

Declare uma variável chamada `medical_records` e atribua a ela uma lista vazia.Nos passos seguintes, você vai usá-lo para armazenar seus dados médicos.

# --hints--

Você deve declarar uma variável chamada `medical_records`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

Você deve atribuir uma lista vazia à sua variável `medical_records`.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

