---
id: 6842d659a4ae7973dabab3b0
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Complete o dicionário adicionando os seguintes três pares chave-valor:

- Uma chave `diagnosis` com o valor de `Hypertension`
- Uma chave `medications` com o valor da lista `['Lisinopril']`
- Uma chave `last_visit_id` com o valor de `V2301`.

# --hints--

Você deve ter uma chave `diagnosis` com o valor da string `Hypertension` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

Você deve ter uma chave `medications` com o valor da lista `['Lisinopril']` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

Você deve ter uma chave `last_visit_id` com o valor da string `V2301` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

Você deve ter uma chave `patient_id` com o valor da string `P1001` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

Você deve ter uma chave `age` com o valor do inteiro `34` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Você deve ter uma chave `gender` com o valor da string `Female` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
