---
id: 6842ba0086201089ea268c92
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Seguindo a mesma estrutura que você usou nas etapas anteriores, a lista `medical_records` foi preenchida para você com os dados de outros pacientes.Fique à vontade para dar uma olhada.

A seguir você começará a escrever a função para validar o conjunto de dados.Crie uma função chamada `validate` com um único parâmetro `data`.

Você quer garantir que seus dados sejam uma lista ou uma tupla.Portanto, dentro da função `validate`, declare uma variável chamada `is_sequence` e atribua a ela uma chamada para `isinstance`.Passe `data` como o primeiro argumento e uma tupla contendo `list` e `tuple` como o segundo argumento.

# --hints--

Você deve ter uma função chamada `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

Sua função `validate` deve ter um único parâmetro `data`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

Você deve ter uma variável chamada `is_sequence` dentro da sua função `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

Você deve atribuir `isinstance(data, (list, tuple))` à sua variável `is_sequence`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]
--fcc-editable-region--

--fcc-editable-region--
```
