---
id: 6842d5b537c64266420faf62
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Adicione uma chave `age` com o valor do inteiro `34` e uma chave `gender` com o valor da string `Female` ao seu dicionário. Não esqueça da vírgula entre os pares chave-valor.

# --hints--

Sua lista `medical_records` deve conter um único dicionário.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Você deve ter uma chave `age` com o valor do inteiro `34` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Você deve ter uma chave `gender` com o valor da string `Female` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

Você deve ter uma chave `patient_id` com o valor da string `P1001` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
