---
id: 6842aee5bcaf4ab8883f8191
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Como você aprendeu em lições anteriores, um dicionário é uma estrutura de dados que armazena pares chave-valor, onde as chaves são tipos de dados imutáveis (como strings) e únicas dentro de um dicionário:

```py
person = {
    'name': 'John',
    'age': 33
}
```

A lista `medical_records` armazenará dicionários, cada um deles representando um paciente. Adicione um dicionário com a chave `patient_id` e valor da string `P1001` à lista `medical_records`.

# --hints--

Sua lista `medical_records` deve conter um único dicionário.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Você deve ter uma chave `patient_id` com o valor da string `P1001` no seu dicionário.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
