---
id: 6842ff2b3c45aa4a89785ecb
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Logo após sua declaração `if`, declare uma variável chamada `is_invalid` e defina-a como `False`.Mais tarde, você usará isso como uma flag para executar uma instrução condicional.

# --hints--

Você deve ter uma variável chamada `is_invalid` na sua função.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_invalid")
`) })
```

Você deve atribuir `False` à sua variável `is_invalid`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_variable("is_invalid").is_equivalent("is_invalid = False")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False

--fcc-editable-region--
    
--fcc-editable-region--
```
