---
id: 5dc174fcf86c76b9248c6eb2
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将通过创建一个猫咪照片应用程序继续使用基本的 HTML 元素，如标题、段落和列表。

通过添加一个文本为 `CatPhotoApp` 的 `h1` 元素，开始工作坊。

# --hints--

代码中应包含文本 `CatPhotoApp`。 你可能需要检查你的拼写。

```js
assert.match(code, /catphotoapp/i);
```

你的 `h1` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert.exists(document.querySelector('h1'));
```

Your `h1` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
assert.notMatch(code, /<\/?H1>/);
```

你的 `h1` 元素应该有一个结束标签。 结束标签的语法为：`</elementName>`。

```js
assert.match(code, /<\/h1\>/);
```

你的 `h1` 元素的文本应该是 `CatPhotoApp`。 你要么遗漏了文本或拼写有误，要么没有把文本放在 `h1` 元素的开始和结束标签之间。

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
