---
id: 5dc1798ff86c76b9248c6eb3
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `h1` 元素下方，添加带有以下文本的 `h2` 元素：

`Cat Photos`

# --hints--

你的 `h1` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素应该有一个结束标签。 结束标签的语法为：`</elementName>`。

```js
assert.match(code, /<\/h1\>/);
```

你应该只有一个 `h1` 元素。 删除多余的。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `h1` 元素的文本应该是 “CatPhotoApp”。 你要么遗漏了文本，要么拼写有误。

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

Your `h2` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
assert.notMatch(code, /<\/?H2>/);
```

你的 `h2` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert.exists(document.querySelector('h2'));
```

你的 `h2` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert.match(code, /<\/h2\>/);
```

你的 `h2` 元素的文本应该是 `Cat Photos`。 仅把文本 `Cat Photos` 放在 `h2` 的开始和结束标签之间。

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

你的 `h2` 元素应该放在 `h1` 元素下面。 `h1` 元素更重要，必须位于 `h2` 元素之上。

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```
