---
id: 5dc17d3bf86c76b9248c6eb4
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `h2` 元素下方创建一个 `p` 元素，并为其添加以下文本：

`Everyone loves cute cats online!`

# --hints--

你的 `p` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert.exists(document.querySelector('p'));
```

Your `p` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
assert.notMatch(code, /<\/?P>/);
```

你的 `p` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert.match(code, /<\/p\>/);
```

你的 `p` 元素的文本应为 `Everyone loves cute cats online!`。你要么遗漏了文本，要么输入有误。

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

你的 `p` 元素应该在 `h2` 元素下面。 你把它们的顺序写错了。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```
