---
id: 5dc17dc8f86c76b9248c6eb5
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

注释能让你在不影响浏览器显示的情况下留下信息。 它也能让你的代码失效。 HTML 中的注释是以 `<!--` 开始，包含任意行数的文本，并以 `-->` 结束。

这是一个内容为 `TODO: Remove h1` 的注释示例：

```html
<!-- TODO: Remove h1 -->
```

在 `p` 元素上方添加注释：

`TODO: Add link to cat photos`

# --hints--

你的注释应该以 `<!--` 开始。 你在定义注释的开头时缺少了一个或多个字符。

```js
assert.match(code, /<!--/);
```

你的注释应该以 `-->` 结束。  你在定义注释的结尾时缺少了一个或多个字符。

```js
assert.match(code, /-->/);
```

你的代码不应该有多余的开始/结束注释字符。 你有一个多余的 `<!--` 或 `-->` 显示在浏览器中。

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g)?.length, 2)
assert.isBelow(noSpaces.match(/-->/g)?.length, 2);
```

你的注释应该在 `p` 元素上面。 你把它们的顺序写错了。

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

你的注释应该包含文本 `TODO: Add link to cat photos`。

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```
