---
id: 5dc23f9bf86c76b9248c6eba
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

你可以通过使用 `img` 元素来为你的网站添加图像。 `img` 元素只有一个开始标签，没有结束标签。 一个没有结束标签的元素被称为<dfn>空元素</dfn>。

在 `p` 元素下面添加一个 `img` 元素。 此时，在浏览器中还不会显示任何图像。

# --hints--

你的 `img` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素不应该有结束标签 `</img>`。

```js
assert.notMatch(code, /<\/img\>/);
```

你应该只有一个 `img` 元素。 删除任何多余的。

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

你的 `img` 元素应该放在 `p` 元素下面。 你把它们的顺序写错了。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```
