---
id: 5dc24165f86c76b9248c6ebc
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

所有的 `img` 元素都应该有一个 `alt` 属性。 `alt` 属性的文本用于屏幕阅读器以改进可访问性，以及在图像加载失败时显示。 

这是一个带有 `alt` 属性的 `img` 元素的例子：

```html
<img src="cat.jpg" alt="A cat">
```

在 `img` 元素中，添加一个 `alt` 属性，文本为：

`A cute orange cat lying on its back`

# --hints--

你的代码应该有一个 `img` 元素。 你删除了前一步骤中的 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素应该有一个 `alt` 属性。 确保在开始标签的名称后面有一个空格，且所有的属性名称前面也要有空格。

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

你的 `img` 元素的 `alt` 属性值被设置成了其他内容，而非 `A cute orange cat lying on its back`。 确保将 `alt` 属性的值放在引号中。

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
