---
id: 5dc2385ff86c76b9248c6eb7
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 有一些元素能够标识不同的内容区域。 这些元素能让你的 HTML 易于阅读，并有助于搜索引擎优化（SEO）和可访问性。

`main` 元素用于表示 HTML 文档正文的主要内容。 `main` 元素里的内容应该是文档中唯一的，不应该在文档的其他部分重复。

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

通过在 `h1` 元素之前添加 `<main>` 开始标签和在 `p` 元素之后添加 `</main>` 结束标签来标识此页面的主要部分。

# --hints--

你的 `main` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert.exists(document.querySelector('main'));
```

Your `main` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

你的 `main` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert.match(code, /<\/main\>/);
```

你的 `main` 元素的开始标签应该在 `body` 元素的下方。 你把它们的顺序写错了。

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

你的 `main` 元素的开始标签应该在 `h1` 元素的上方。 你把它们的顺序写错了。

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

你的 `main` 元素的结束标签应该在 `p` 元素的下方。 你把它们的顺序写错了。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```
