---
id: 5dc24614f86c76b9248c6ebd
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

你可以使用锚元素（`a`）链接到另一个页面。 

这是一个链接到 `https://www.freecodecamp.org` 的示例：

```html
<a href="https://www.freecodecamp.org"></a>
```

在段落之后添加一个链接到 `https://freecatphotoapp.com` 的锚元素。 此时，链接不会显示在预览中。

# --hints--

你的锚元素（`a`）应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert.exists(document.querySelector('a'));
```

你的锚元素（`a`）应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert.match(code, /<\/a\>/);
```

你的锚元素（`a`）应该在 `p` 元素下方。 你把它们的顺序写错了。

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

你的锚元素（`a`）缺少 `href` 属性。 确保在开始标签的名称后面有一个空格，且所有的属性名称前面也要有空格。

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

你的锚元素（`a`）应该链接到 `https://freecatphotoapp.com`。 你可能忽略了 URL 或者有拼写错误。

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

虽然你将锚元素（`a`）的 `href` 属性指向了正确的链接，但建议始终以引号来包裹属性的值。

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
