---
id: 5dc23991f86c76b9248c6eb8
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在上一步中，你将 `h1`、`h2`、注释和 `p` 元素放在 `main` 元素中。 这叫作*嵌套*。 嵌套元素应该在其所嵌套元素的右边再放两个空格。 这种空格被称为缩进，用于使 HTML 更容易阅读。

这是一个嵌套和缩进的示例：

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

下面代码中的 `h1` 元素、`h2` 元素和注释比 `main` 元素多缩进两个空格。 使用键盘上的空格键在 `p` 元素前面添加两个空格，使其正确缩进。

# --hints--

你的代码应该有一个 `h2` 元素，文本为 `Cat Photos`。 你可能不小心删除了它，它缺少开始和结束标记，或者文本已经更改。

```js
assert.exists(document.querySelector('h2'));
assert.match(code, /<\/h2\>/);
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(),'cat photos')
```

你不应添加示例中的 `ul` 或 `li` 元素。

```js
assert.notExists(document.querySelector('ul'));
assert.notExists(document.querySelector('li'));
```

你不应更改 `h2` 元素所在行的缩进。 它的开始标签前应该有 6 个空格。 可以重置该步骤来恢复原始缩进。

```js
assert.match(code.toLowerCase(), /<\/h1\>\s*\n\s{6}<h2>/);
```

你的代码应该有一个注释。 你从之前的步骤中删除了注释。

```js
assert.match(code, /<!--.*-->/);
```

注释的文本应该是 `TODO: Add link to cat photos`。 不要更改注释的文本或空格。

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i);
```

你不应更改注释元素所在行的缩进。 它的开始标签前应该有 6 个空格。 可以重置该步骤来恢复原始缩进。

```js
assert.match(
  code
    .toLowerCase(),
    /<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/
);
```

你的代码应该有一个 `p` 元素。 你删除了前一步骤中的 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

`p` 元素中的文本应该是 `Everyone loves cute cats online!`。不要改变 `p` 元素中的文本、空格和标点符号。

```js
assert.match(
  document
    .querySelector('p')
    ?.innerText?.trim().toLowerCase(),
    /^everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

开始的 `p` 标签应该有与 `h2` 和注释元素一致的缩进。 它的开始标签前应该有 6 个空格。

```js
assert.match(code.toLowerCase(), /-->\s*\n\s{6}<p>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
