---
id: 5ddbd81294d8ddc1510a8e56
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

链接的文本必须放在锚元素（`a`）的开始和结束标签之间。

这是一个链接文本是 `click here to go to freeCodeCamp.org` 的示例：

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

将锚文本 `cat photos` 添加到锚元素。 这将成为链接的文本。

# --hints--

你的锚元素（`a`）应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert.exists(document.querySelector('a'));
```

你的锚元素（`a`）应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert.match(code, /<\/a\>/);
```

你的锚元素（`a`）的文本应该是 `cat photos`。 确保将链接文本放在锚元素（`a`）的开始和结束标签之间。

```js
assert.equal(
  document.querySelector('a')?.innerText?.trim().toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
