---
id: 5dfa2407b521be39a3de7be1
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

要在新标签页中打开链接，你可以在锚元素（`a`）上使用 `target` 属性。 

`target` 属性指定打开链接文档的位置。`target="_blank"` 会在新建标签（页）或窗口中打开链接文档。

这是一个带有 `target` 属性的 `a` 元素的基本语法：

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

向锚元素（`a`）的开始标签添加一个值为 `_blank` 的 `target` 属性，用以在新的标签页中打开链接。

# --hints--

你的 `p` 元素应该嵌套一个文本为 `cat photos` 的锚（`a`）元素。 你可能删除了它或者拼写有误。

```js
const anchor = document.querySelectorAll('p > a')[1];
assert.equal(
    anchor?.innerText?.toLowerCase().trim().replace(/\s+/g, ' '), 'cat photos'
);
```

你的锚元素（`a`）没有 `target` 属性。 确保在开始标签的名称后面有一个空格，且所有的属性名称前面也要有空格。

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

`target` 属性的值应为 `_blank`。 你要么遗漏了属性值，要么拼写有误。 记住，属性值应该放在引号中。

```js
assert.equal(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
