---
id: 5dfa37b9eacea3f48c6300b0
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

要创建一个无序项目列表，你可以使用 `ul` 元素。

在带有 `Things cats love:` 文本的 `h3` 元素之后，添加一个无序列表（`ul`）元素。 注意，此时不会显示任何内容。

# --hints--

你的 `ul` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert.exists(document.querySelector('ul'));
```

你的 `ul` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert.match(code, /<\/ul>/);
```

`ul` 元素应位于第二个 `section` 元素的结束标签之上。

```js
const secondSectionLastElemNode =
  document.querySelectorAll('main > section')[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName, 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
