---
id: 5dfb6250eacea3f48c6300b2
title: 步骤 24
challengeType: 0
dashedName: 步骤-24
---

# --description--

在无序列表之后，添加一个新图像，它具有 `src` 属性，值为：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

而它的 `alt` 属性值为：

`A slice of lasagna on a plate.`

# --hints--

在结束 `</ul>` 标签之后应该有一个 `img` 元素。

```js
assert.equal(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

新图像没有 `alt` 属性。 检查开始标记的名称后面是否有空格和/或所有属性名称之前是否有空格。

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

新图像的 `alt` 值应为 `A slice of lasagna on a plate.`。确保将 `alt` 属性的值放在引号中。

```js
assert.match(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

新图像没有 `src` 属性。 检查开始标签的名称后面是否有空格和/或所有属性名称之前是否有空格。

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

新图像的 `src` 值应为 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`。 确保将 `src` 属性的值放在引号中。

```js
assert.equal(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

虽然你已将新图像的 `src` 设置为正确的 URL，但建议始终将属性的值用引号括起来。

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
