---
id: 5ef9b03c81a63668521804d0
title: 步骤 27
challengeType: 0
dashedName: step-27
---

# --description--

要强调一个特定的单词或短语，你可以使用 `em` 元素。

通过将 `figcaption` 元素中的单词 `love` 包裹在强调 `em` 元素中来强调它。

# --hints--

你的强调 `em` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert.exists(document.querySelector('em'));
```

你的强调 `em` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert.match(code, /<\/em\>/);
```

你要么删除了 `figcaption` 元素，要么它缺少开始或结束标签。

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

你的强调 `em` 元素应包裹文本 `love`。 你要么省略了文本，要么有拼写错误。

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.trim().toLowerCase(), 'love'
);
```

`figcaption` 的文本应该是 `Cats love lasagna`。 检查是否有拼写错误，以及 `em` 元素的开始和结束标签周围是否有必要的空格。

```js
assert.match(
  document
    .querySelector('figcaption')
    ?.innerText?.trim().replace(/\s+/gi, ' '),
    /^cats love lasagna\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
