---
id: 5ef9b03c81a63668521804ea
title: 步骤 38
challengeType: 0
dashedName: step-38
---

# --description--

注意，到目前为止添加到页面中的所有内容都位于`body`元素中。 所有应该呈现到页面的页面内容元素都放在 `body` 元素中。 但是，其他重要信息在 `head` 元素中。

`head` 元素用于包含关于文档的元数据，如它的标题、样式表链接和脚本。 元数据是不直接显示在页面上的关于页面的信息。

在 `body` 元素上方添加一个 `head` 元素。

# --hints--

你要么删除了 `body` 元素，要么它缺少开始标签或结束标签。

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

你的 `head` 元素应有一个起始标签。起始标签具有以下语法：`<elementName>`。

```js
assert.match(code, /\<head\>/);
```

你的 `head` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert.match(code, /\<\/head\>/);
```

你的 `head` 元素应位于 `body` 元素开始标签的上方。 你把它放在别的地方了。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
