---
id: 5ef9b03c81a63668521804d1
title: 步骤 28
challengeType: 0
dashedName: step-28
---

# --description--

在 `figure` 元素之后，添加另一个 `h3` 元素，文本为：

`Top 3 things cats hate:`

# --hints--

在第二个 `section` 元素的结束标签上方应该有一个 `h3` 元素。 确保它有一个开始和结束标签。

```js
assert.equal(
  document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName,
    'H3'
);
assert.lengthOf(code.match(/<\/h3\>/g),  2);
```

新的 `h3` 元素应该包含文本 `Top 3 things cats hate:`。 确保在文本末尾包含冒号。

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild.innerText?.trim().toLowerCase()
    .replace(/\s+/g, ' '), 'top 3 things cats hate:'
);
```

新的 `h3` 元素上方应该有一个 `figure`。 你可能不小心删除了 `figure` 元素。

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  ?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName, 'H3');
assert.equal(secondSectionLastElemNode.previousElementSibling.nodeName, 'FIGURE');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
