---
id: 5ef9b03c81a63668521804e9
title: 步骤 37
challengeType: 0
dashedName: step-37
---

# --description--

通过将现有的 `freeCodeCamp.org` 文本包含在锚元素（`a`）中，使其成为链接。 `href` 属性应该被设置为 `https://www.freecodecamp.org`。

# --hints--

你的锚元素（`a`）应该嵌套在 `footer` 中的 `p`  元素中。 确保为锚元素（`a`）添加一个开始和结束标签。

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

你的锚元素（`a`）应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

你的锚（`a`）元素应该有一个 `href` 属性，值为 `https://www.freecodecamp.org`。 你要么遗漏了属性/值，要么拼写有误。

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

链接的文本应该是 `freeCodeCamp.org`。 你要么遗漏了文本，要么拼写有误。

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

嵌套锚（`a`）元素后，浏览器中唯一可见的 `p` 元素内容应该是 `No Copyright - freeCodeCamp.org`。 仔细检查 `p` 和嵌套锚元素的文本、空格或标点符号。

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
