---
id: 5ef9b03c81a63668521804ee
title: 步骤 41
challengeType: 0
dashedName: step-41
---

# --description--

所有页面都应该以 `<!DOCTYPE html>` 开头。 这个特殊的字符串被称为<dfn>声明</dfn>，确保浏览器尝试满足行业规范。

`<!DOCTYPE html>` 告诉浏览器该文档是一个 HTML5 文档，即 HTML 的最新版本。

将此声明添加为代码的第一行。

# --hints--

你的代码应以 `<!DOCTYPE html>` 声明开头。 你要么遗漏了声明，要么拼写有误，要么它不是第一行代码。

```js
assert.match(code, /\<\s*!DOCTYPE\s+html\s*\>/);
```

你的 `<!DOCTYPE html>` 必须位于文档的顶部。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /^\<\!DOCTYPEhtml\>\<html/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

<html lang="en">
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
