---
id: 5dc174fcf86c76b9248c6eb2
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML 元素有像 `<h1>` 這樣的開啟頁籤和像 `</h1>` 這樣的關閉頁籤。元素會顯示的文字位於其開啟頁籤和關閉頁籤之間。

將下方 `h1` 元素的 `Hello World` 文字更改為 `CatPhotoApp`，並在瀏覽器預覽中觀察變化。

完成後，按下「Check Your Code」按鈕以查看是否正確。

# --hints--

程式碼中應該包含文字 `CatPhotoApp`。你可能想要檢查你的拼寫。

```js
assert(code.match(/catphotoapp/i));
```

你的 `h1` 元素應該有一個開啟標記。開啟標記具有這種語法：`<elementName>`。

```js
assert(document.querySelector('h1'));
```

你的 `h1` 頁籤應該使用小寫。依慣例，所有 HTML 頁籤皆以小寫書寫。

```js
assert.notMatch(code, /<\/?H1>/);
```

你的 `h1` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/h1\>/));
```

你有多於一個 `h1` 元素。請移除多餘的 `h1` 元素。

```js
assert(document.querySelectorAll('h1').length === 1);
```

你的 `h1` 元素的文字應該是 `CatPhotoApp`。你要麼遺漏了文字、有拼寫錯誤，要麼文字不在 `h1` 元素的開啟和關閉標籤之間。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

