---
id: 5dc17d3bf86c76b9248c6eb4
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

`p` 元素用於在網站上創建一段文字。在你的 `h2` 元素下方創建一個 `p` 元素，並給它以下文字：

`Everyone loves cute cats online!`

# --hints--

你的 `p` 元素應該有一個開啟標記。開啟標記具有以下語法：`<elementName>`。

```js
assert(document.querySelector('p'));
```

你的 `p` 標籤應該使用小寫。依慣例，所有 HTML 標籤皆以小寫書寫。

```js
assert.notMatch(code, /<\/?P>/);
```

你的 `p` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/p\>/));
```

你的 `p` 元素的文字應該是 `Everyone loves cute cats online!` 你要麼遺漏了文字，要麼有拼寫錯誤。

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

你的 `p` 元素應該在 `h2` 元素下面。你把它們的順序弄錯了。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```

