---
id: 5dc17dc8f86c76b9248c6eb5
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

註解允許你留下訊息而不影響瀏覽器顯示。它也允許你使程式碼變為非作用中。HTML 中的註解以 `<!--` 開始，包含任意數量的文字列，並以 `-->` 結束。 

這是一個帶有 `TODO: Remove h1` 的註解範例：

```html
<!-- TODO: Remove h1 -->
```

在 `p` 元素上方添加一則包含此文字的註解：

`TODO: Add link to cat photos`

# --hints--

你的註解應該以 `<!--` 開頭。你缺少定義註解開始的 1 個或多個字元。

```js
assert(code.match(/<!--/));
```

你的註解應該以 `-->` 結尾。你缺少一個或多個定義註解結尾的字元。

```js
assert(code.match(/-->/));
```

你的程式碼不應該有多餘的開啟／關閉註解字元。瀏覽器中顯示了多餘的 `<!--` 或 `-->`。

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

你的註解應該在 `p` 元素上方。你放錯順序了。

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

你的評論應該包含文字 `TODO: Add link to cat photos`。

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```

