---
id: 5dc17dc8f86c76b9248c6eb5
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

註釋能讓你在不影響瀏覽器顯示的情況下留下信息。 它也能讓你的代碼失效。 HTML 中的註釋是以 `<!--` 開始，包含任意行數的文本，並以 `-->` 結束。

這是一個內容爲 `TODO: Remove h1` 的註釋示例：

```html
<!-- TODO: Remove h1 -->
```

在 `p` 元素上方添加註釋：

`TODO: Add link to cat photos`

# --hints--

你的註釋應該以 `<!--` 開始。 你在定義註釋的開頭時缺少了一個或多個字符。

```js
assert(code.match(/<!--/));
```

你的註釋應該以 `-->` 結束。  你在定義註釋的結尾時缺少了一個或多個字符。

```js
assert(code.match(/-->/));
```

你的代碼不應該有多餘的開始/結束註釋字符。 你有一個多餘的 `<!--` 或 `-->` 顯示在瀏覽器中。

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

你的註釋應該在 `p` 元素上面。 你把它們的順序寫錯了。

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

你的註釋應該包含文本 `TODO: Add link to cat photos`。

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```

