---
id: 5dc23991f86c76b9248c6eb8
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在前一步驟中，你將 `h1`、`h2`、註解和 `p` 元素放在 `main` 元素內。這稱為*巢狀*。巢狀元素應該比它們所巢狀的元素向右多縮排兩個空格。這種間距稱為縮排，用來讓 HTML 更容易閱讀。

這裡是一個巢狀和縮排的範例：

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

在下面的程式碼中，`h1` 元素、`h2` 元素和註解的縮排比 `main` 元素多兩個空格。使用鍵盤上的空白鍵，在 `p` 元素前面再添加兩個空格，使其縮排正確。

# --hints--

你的程式碼應該有一個 `h2` 元素，裡面的文字是 `Cat Photos`。你可能不小心刪除了它，因為它缺少開啟和關閉的標籤，或者文字已經改變。

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

你不應該添加範例中的 `ul` 或 `li` 元素。

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

你不應該更改包含 `h2` 元素的行的縮排。它的開啟標記應該從該行開頭開始向右縮排 6 個空格。你可以重新開始此步驟以還原原始縮排。

```js
assert(code.toLowerCase().match(/<\/h1\>\s*\n\s{6}<h2>/));
```

你的程式碼應該有註解。你在先前的步驟中移除了註解。

```js
assert(code.match(/<!--.*-->/));
```

評論的文字應該是 `TODO: Add link to cat photos`。請勿更改評論的文字或間距。

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

你不應該更改包含你的註解元素的行的縮排。它的開啟頁籤應該從該行開頭開始向右縮排 6 個空格。你可以重新開始此步驟以還原原始縮排。

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

你的程式碼應該有一個 `p` 元素。你已經在先前的步驟中移除了 `p` 元素。

```js
assert(document.querySelector('p'));
```

`p` 元素的文字應該是 `Everyone loves cute cats online!` 不要更改 `p` 元素的文字、間距或標點符號。

```js
assert.match(
  document
    .querySelector('p')
    .innerText.toLowerCase(),
    /everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

開啟的 `p` 頁籤應該有與你的 `h2` 和註解元素相符的縮排。它的開啟頁籤應該從行首開始向右縮排 6 個空格。

```js
assert(code.toLowerCase().match(/-->\s*\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

