---
id: 5dc24073f86c76b9248c6ebb
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>屬性</dfn> 是用在元素起始標籤內用來控制元素行為的特殊字詞。`img` 元素中的 `src` 屬性指定影像的 URL（影像所在的位置）。

這是一個 `img` 元素的範例，其 `src` 屬性指向 freeCodeCamp 標誌：

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

在現有的 `img` 元素內，為你的 `src` 屬性添加此 URL：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

你的程式碼應該有一個 `img` 元素。你可能已經移除了 `img` 元素，或者你沒有用引號將 `src` 屬性的值括起來。

```js
assert(document.querySelector('img'));
```

你的 `img` 元素應該有一個 `src` 屬性。你要麼遺漏了該屬性，要麼有拼寫錯誤。請確保元素名稱和屬性名稱之間有一個空格。

```js
assert(document.querySelector('img').src);
```

你的 `img` 元素的 `src` 屬性應該設為 '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'。你要麼遺漏了 URL，要麼有拼寫錯誤。URL 的大小寫很重要。

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

雖然你已將 `img` 元素的 `src` 設定為正確的 URL，但建議始終用引號將屬性的值括起來。

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

你的 `img` 元素應該以 `>` 或 `/>` 結尾。

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

