---
id: 5dc24073f86c76b9248c6ebb
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>屬性</dfn>是位於元素的開始標籤中的特殊詞彙，用於控制元素的行爲。 在 `img` 元素中的 `src` 屬性指定圖像的 URL（圖像所在的位置）。

下面是一個 `img` 元素的 `src` 屬性指向 freeCodeCamp logo 的例子：

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

在現有的 `img` 元素中，添加一個 `src` 屬性，URL 如下：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

你的代碼應該有一個 `img` 元素。 你可能已經刪除了 `img` 元素，或者你沒有用引號包圍 `src` 屬性的值。

```js
assert(document.querySelector('img'));
```

你的 `img` 元素應該有一個 `src` 屬性。 你要麼遺漏了屬性，要麼拼寫有誤。 確保元素名稱和屬性名稱之間有一個空格。

```js
assert(document.querySelector('img').src);
```

你的 `img` 元素的 `src` 屬性應設置爲 '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'。 你要麼遺漏了 URL，要麼拼寫有誤。 URL 的大小寫很重要。

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

雖然你已經將 `img` 元素的 `src` 設置爲正確的 URL，但建議始終將屬性的值用引號括起來。

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

你的 `img` 元素應該以 `>` 或 `/>` 結束。

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

