---
id: 5ddbd81294d8ddc1510a8e56
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

聯結的文字必須放置在錨點（`a`）元素的開啟與關閉頁籤之間。

這是一個帶有文字 `click here to go to freeCodeCamp.org` 的聯結範例：

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

為你的錨點元素添加錨點文字 `cat photos`。這將成為聯結的文字。

# --hints--

你的錨點（`a`）元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert(document.querySelector('a'));
```

你的錨點（`a`）元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/a\>/));
```

你的錨點（`a`）元素的文字應該是 `cat photos`。請確保將聯結文字放在錨點（`a`）元素的起始標籤和結束標籤之間。

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

