---
id: 5dfa37b9eacea3f48c6300b0
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

要創建一個無序項目列表，你可以使用 `ul` 元素。 

在帶有 `Things cats love:` 文字的 `h3` 元素之後，添加一個無序列表（`ul`）元素。請注意，此時不會顯示任何內容。

# --hints--

你的 `ul` 元素應該有一個開啟標記。開啟標記具有此語法：`<elementName>`。

```js
assert(document.querySelector('ul'));
```

你的 `ul` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/ul>/));
```

`ul` 元素應該位於第二個 `section` 元素的結束標籤之上。

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')?.[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName , 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

