---
id: 5dfa3589eacea3f48c6300ae
title: 第 20 步
challengeType: 0
dashedName: step-20
---

# --description--

在第二個 `section` 元素中，為你的 `h2` 元素添加文字 `Cat Lists`。

# --hints--

你的 `section` 元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert(
  document.querySelectorAll('section').length === 2 &&
    code.match(/<\/section>/g).length === 2
);
```

你的 `h2` 元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert(document.querySelectorAll('h2').length === 2);
```

你的 `h2` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/h2\>/g).length === 2);
```

你的第二個 `h2` 元素應該位於第二個 `section` 元素的結束標籤正上方。它的位置不正確。

```js
const secondSection = document.querySelectorAll('section')[1];
assert(secondSection.lastElementChild.nodeName === 'H2');
```

第二個 `h2` 元素應該有文字 `Cat Lists`。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase() === 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

