---
id: 5dfa2407b521be39a3de7be1
title: 步驟 15
challengeType: 0
dashedName: step-15
---

# --description--

要在新標籤頁中打開鏈接，你可以在錨元素（`a`）上使用 `target` 屬性。

`target` 屬性指定打開鏈接文檔的位置。 `target="_blank"` 是在新標籤頁或窗口中打開鏈接的文檔。

這是一個帶有 `target` 屬性的 `a` 元素的基本語法：

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

向錨元素（`a`）的開始標籤添加一個值爲 `_blank` 的 `target` 屬性，用以在新的標籤頁中打開鏈接。

# --hints--

你的 `p` 元素應該嵌套一個文本爲 `cat photos` 的錨元素（`a`）。 你可能刪除了它或者拼寫有誤。

```js
const anchor = document.querySelectorAll('p > a');
assert.isNotEmpty(anchor);
assert.strictEqual(anchor[1]?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos')
```

你的錨元素（`a`）沒有 `target` 屬性。 檢查開始標籤的名稱後面是否有空格和/或所有屬性名稱之前是否有空格。

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

`target` 屬性的值應爲 `_blank`。 你要麼遺漏了屬性值，要麼拼寫有誤。 記住，屬性值應該放在引號中。

```js
assert.strictEqual(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

