---
id: 5dfb6a35eacea3f48c6300b4
title: 步驟 26
challengeType: 0
dashedName: step-26
---

# --description--

`figure` 元素中包含的影像使用 `figcaption` 元素來添加說明文字。 

這是一個帶有標題為 `A cute cat` 的 `figcaption` 元素範例：

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

在 `figure` 元素中嵌套的影像之後，添加一個 `figcaption` 元素，並將文字設為：

`Cats love lasagna.`

# --hints--

你的 `figcaption` 元素應該有一個開啟標記。開啟標記具有以下語法：`<elementName>`。

```js
assert(document.querySelector('figcaption'));
```

你的 `figcaption` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/figcaption\>/));
```

`figcaption` 元素應該巢狀在 `figure` 元素中。

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

千層麵的 `img` 元素應該巢狀在 `figure` 元素中。

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

嵌套在 `figure` 元素中的 `figcaption` 元素應該位於 `img` 元素下方。你放錯順序了。

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

你的 `figcaption` 元素的文字應該是 `Cats love lasagna.` 你要麼遺漏了文字，要麼有拼寫錯誤。

```js
assert(
  document.querySelector('figcaption').innerText.trim().match(/^Cats love lasagna\.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          
        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

