---
id: 5dfb655eeacea3f48c6300b3
title: 第 25 步
challengeType: 0
dashedName: step-25
---

# --description--

`figure` 元素表述獨立內容，並允許你將影像與標題關聯。

將你剛剛添加的影像嵌套在 `figure` 元素中。

# --hints--

你的 `figure` 元素應該有一個開啟標籤。開啟標籤具有以下語法：`<elementName>`。

```js
assert(document.querySelector('figure'));
```

你的 `figure` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/figure\>/));
```

在第二個 `section` 元素的結束標籤正上方應該有一個 `figure` 元素。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.lastElementChild?.nodeName, 'FIGURE');
```

千層麵的 `img` 元素應該巢狀在 `figure` 元素中。

```js
assert.strictEqual(document.querySelector('figure > img')?.src.toLowerCase(), 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
      
        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
      
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

