---
id: 5ef9b03c81a63668521804d0
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--

To place emphasis on a specific word or phrase, you can use the `em` element.

Emphasize the word `love` in the `figcaption` element by wrapping it in an emphasis `em` element.

# --hints--

你的強調元素 `em` 應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert(document.querySelector('em'));
```

你的強調元素 `em` 應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert(code.match(/<\/em\>/));
```

你要麼刪除了 `figcaption` 元素，要麼它缺少開始或結束標籤。

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

Your emphasis `em` element should surround the text `love`. 你要麼遺漏了文本，要麼拼寫有誤。

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

The `figcaption`'s text should be `Cats love lasagna`. Check for typos and that the necessary spaces are present around the `em` element's opening and closing tags.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

