---
id: 5ef9b03c81a63668521804df
title: 第 48 步
challengeType: 0
dashedName: step-48
---

# --description--

`id` 屬性用於識別特定的 HTML 元素。每個 `id` 屬性的值必須在整個頁面中與所有其他 `id` 值唯一。

這是一個帶有 `id` 屬性的 `input` 元素範例：

```html
<input id="email">
```

為你的圓鈕添加一個值為 `indoor` 的 `id` 屬性。當元素有多個屬性時，屬性的順序不重要。

# --hints--

你的圓鈕應該仍然位於 `label` 元素的開啟和關閉標籤之間。

```js
const labelChildNodes = [...document.querySelectorAll('form > label')?.[0]?.childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

你的圓鈕應該有一個 `id` 屬性。請檢查開啟標籤名稱之後是否有空格，和／或所有屬性名稱之前是否有空格。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('id'));
```

你的圓鈕應該有一個 `type` 屬性。請檢查開啟標籤名稱之後是否有空格，和／或所有屬性名稱之前是否有空格。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

你的圓鈕應該有一個 `id` 屬性，其值為 `indoor`。你要麼遺漏了該值，要麼有拼寫錯誤。請記得屬性值應該用引號括起來。

```js
assert.match(document.querySelector('input')?.id, /^indoor$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

