---
id: 5ef9b03c81a63668521804e5
title: 步驟 63
challengeType: 0
dashedName: step-63
---

# --description--

為了讓核取方塊預設被勾選或圓鈕預設被選取，你需要為它添加 `checked` 屬性。 

這是一個帶有 `checked` 屬性的圓鈕範例：

```html
<input checked type="radio" name="meal" value="breakfast"> Breakfast
```

不需要為 `checked` 屬性設定值。相反地，只要將字 `checked` 添加到 `input` 元素中，並確保它與其他屬性之間有空格。

將第一個圓鈕和第一個核取方塊設為預設選取。

# --hints--

請確定仍有兩個圓鈕和三個核取方塊巢狀在各自的 `fieldset` 元素中。

```js
assert.isTrue(
  document.querySelectorAll('input[type="radio"]').length === 2 &&
     document.querySelectorAll('fieldset > input[type="checkbox"]').length === 3
);
```

你的第一個圓鈕，`id` 設為 `indoor`，應該有 `checked` 屬性。

```js
assert.isTrue(document.querySelector('input[type="radio"]')?.hasAttribute('checked'));
```

第二個圓鈕不應該有 `checked` 屬性。

```js
assert.isFalse(document.querySelectorAll('input[type="radio"]')?.[1]?.hasAttribute('checked'));
```

第一個核取方塊缺少 `checked` 屬性。

```js
assert.isTrue(document.querySelector('input[type="checkbox"]')?.hasAttribute('checked'));
```

第二個核取方塊不應該有 `checked` 屬性。

```js
assert.isFalse(document.querySelectorAll('input[type="checkbox"]')?.[1]?.hasAttribute('checked'));
```

第三個核取方塊不應該有 `checked` 屬性。

```js
assert.isFalse(document.querySelectorAll('input[type="checkbox"]')?.[2]?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic"> Energetic</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

