---
id: 5ef9b03c81a63668521804d8
title: 步驟 39
challengeType: 0
dashedName: step-39
---

# --description--

`input` 元素允許你透過多種方式從網頁表單收集資料。像 `img` 元素一樣，`input` 元素是 <dfn>void element</dfn>，不需要閉合標籤。

將 `input` 元素巢狀於 `form` 元素中。

# --hints--

你的 `form` 元素應該有正確順序的開啟標籤和關閉標籤。你可能缺少一個或兩個必要的標籤，或是標籤的順序錯誤。

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

你的 `form` 元素的開啟標記應該只包含一個 `action` 屬性。請移除你可能輸入的其他內容。

```js
assert([...document.querySelector('form').attributes].length < 2);
```

你應該創建一個 `input` 元素。檢查語法。

```js
assert(document.querySelector('input'));
```

你的 `input` 元素應該有一個開啟的頁籤，但不應該有關閉的頁籤。

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

你的 `input` 元素應該巢狀在 `form` 元素內。

```js
assert(document.querySelector('form > input'));
```

你的 `form` 應該只包含 `input` 元素。請移除 `form` 元素標籤之間的任何 HTML 元素或文字。

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

