---
id: 5ef9b03c81a63668521804db
title: 步驟 43
challengeType: 0
dashedName: step-43
---

# --description--

為防止使用者在缺少必要資訊時提交你的表單，你需要為 `input` 元素添加 `required` 屬性。

以下是一個帶有 `required` 屬性的輸入欄位範例：

```html
<input type="text" name="firstName" required>
```

不需要為 `required` 屬性設定值。相反地，只要將字 `required` 添加到 `input` 元素中，並確保它與其他屬性之間有空格。

# --hints--

你已刪除你的 `input` 元素或它的語法無效。所有屬性的值都應該用引號括起來。

```js
assert(document.querySelectorAll('input').length);
```

你的 `form` 應該只包含 `input` 元素。請移除 `form` 元素內的任何其他 HTML 元素或文字。

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

你的 `input` 元素應該有一個 `required` 屬性。請記得，你只要在 `input` 元素的標記內添加字 `required` 即可。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('required'));
```

不應該為 `required` 屬性指定值。

```js
assert.strictEqual(document.querySelector('input')?.getAttribute('required'), '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

