---
id: 5ef9b03c81a63668521804da
title: 步驟 44
challengeType: 0
dashedName: step-44
---

# --description--

`button` 元素用於創建可點擊的按鈕。

在 `input` 元素下方添加一個帶有 `Submit` 文字的 `button` 元素。點擊沒有任何屬性的表單按鈕的預設行為是將表單提交到表單的 `action` 屬性所指定的位置。

# --hints--

你的 `button` 元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert(document.querySelector('button'));
```

你的 `button` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/button\>/));
```

你的 `button` 元素的文字應該是 `Submit`。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
assert(document.querySelector('button').innerText.toLowerCase() === 'submit');
```

你的 `button` 元素應該在 `input` 元素下方。你把它們的順序弄錯了。

```js
const collection = [...document.querySelectorAll('input, button')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('INPUT') < collection.indexOf('BUTTON'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

