---
id: 5ef9b03c81a63668521804e2
title: 步驟 56
challengeType: 0
dashedName: step-56
---

# --description--

表單通常會使用核取方塊來處理可能有多個答案的問題。`input` 元素中 `type` 屬性設定為 `checkbox` 即會創建核取方塊。

在你剛剛添加的 `legend` 元素下，添加一個 `input`，其 `type` 屬性設為 `checkbox`，並將其文字設為：

` Loving`

# --hints--

你的核取方塊的 `input` 元素應該有一個開啟標籤，但不要有關閉標籤。

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

你應該只為你的核取方塊添加一個輸入元素。移除任何多餘的元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

你的新 `input` 元素應該位於帶有文字 `What's your cat's personality?` 的 `legend` 元素下方。你把它們的順序弄錯了。

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

你的新 `input` 元素沒有 `type` 屬性。請檢查開啟標記名稱後是否有空格。

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

你的新 `input` 元素應該有一個 `type` 屬性，其值為 `checkbox`。你要麼遺漏了該值，要麼有拼寫錯誤。請記得屬性值應該用引號括起來。

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

雖然你已將新的 `input` 元素的 `type` 屬性設為 `checkbox`，但建議始終用引號將屬性的值括起來。

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

`Loving` 文字應該位於你的核取方塊的正右方。請確保元素與文字之間有一個空格。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>
            
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

