---
id: 5ef9b03c81a63668521804e3
title: 步驟 60
challengeType: 0
dashedName: step-60
---

# --description--

在你剛剛添加的核取方塊後面再添加一個。`id` 屬性值應該是 `lazy`，而 `name` 屬性值應該與最後一個核取方塊相同。

同時在新的核取方塊右側添加一個 `label` 元素，並將文字設為 `Lazy`。請確保使用 `for` 屬性將該 `label` 元素與新的核取方塊關聯。

# --hints--

你需要添加一個新的核取方塊。

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'), 2);
```

你的新核取方塊應該有一個 `id` 屬性，值為 `lazy`，以及一個 `name` 屬性，值為 `personality`。請檢查開啟標記名稱後是否有空格，和／或所有屬性名稱前是否有空格。

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'lazy' && checkbox.getAttribute('name') === 'personality'
  )
);
```

你的新核取方塊應該在第一個之後。你把它們放錯順序了。

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('loving') < checkboxes.indexOf('lazy'));
```

在你新的核取方塊右側，應該有一個帶有文字 `Lazy` 的 `label` 元素。

```js
const nextElementSibling = document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling;
assert.isOk(nextElementSibling?.nodeName === 'LABEL' && nextElementSibling?.innerText?.replace(/\s+/g, '')?.match(/^Lazy$/i)
);
```

新的 `label` 應該有一個 `for` 屬性，其值與新核取方塊的 `id` 屬性相同。你要麼遺漏了該值，要麼有拼寫錯誤。

```js
assert.strictEqual(
  document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling?.getAttribute('for'),
    'lazy'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

