---
id: 5ef9b03c81a63668521804e9
title: 步驟 66
challengeType: 0
dashedName: step-66
---

# --description--

將現有的 `freeCodeCamp.org` 文字轉換成連結，方法是將它包裹在一個錨點（`a`）元素中。`href` 屬性應設為 `https://www.freecodecamp.org`。

# --hints--

你的錨點（`a`）元素應該巢狀在 `footer` 裡的 `p` 元素內。請確保為錨點（`a`）元素添加開啟標籤和關閉標籤。

```js
assert(document.querySelectorAll('footer > p > a').length);
```

你的錨點（`a`）元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.lengthOf(aElemClosingTags, 4);
```

你的錨點（`a`）元素應該有一個 `href` 屬性，且其值為 `https://www.freecodecamp.org`。你可能遺漏了該屬性／值，或有拼寫錯誤。

```js
const nestedAnchor = document.querySelector('footer > p > a');
assert.strictEqual(nestedAnchor?.getAttribute('href')?.toLowerCase(), 'https://www.freecodecamp.org');
```

聯結的文字應該是 `freeCodeCamp.org`。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
const nestedAnchor = document.querySelector('footer > p > a');
assert.strictEqual(
  nestedAnchor?.innerText?.toLowerCase()?.replace(/\s+/g, ' '),
  'freecodecamp.org'
);
```

在巢狀的錨點（`a`）元素後，瀏覽器中唯一可見的 `p` 元素內容應該是 `No Copyright - freeCodeCamp.org`。請再次確認 `p` 和巢狀錨點元素的文字、間距或標點符號。

```js
const pText = document.querySelector('footer > p')?.innerText?.toLowerCase()?.replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving" checked> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic">Energetic</label>
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```

