---
id: 5ef9b03c81a63668521804d4
title: 步驟 34
challengeType: 0
dashedName: step-34
---

# --description--

`strong` 元素用於顯示某些文字具有重要性或緊急性。

在你剛剛添加的 `figcaption` 中，將 `hate` 包裹在 `strong` 元素中以顯示其強烈的重要性。

# --hints--

你的 `strong` 元素應該有一個開啟標記。開啟標記具有此語法：`<elementName>`。

```js
assert(document.querySelector('strong'));
```

你的 `strong` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/strong\>/));
```

你的 `strong` 元素應該包圍文字 `hate`，該文字位於 `Cats hate other cats.` 中。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
assert(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.trim().toLowerCase() === 'hate'
);
```

`figcaption` 的文字應該是 `Cats hate other cats.`，請檢查是否有拼字錯誤以及 `strong` 元素開啟和關閉標籤周圍是否有必要的空格。

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim()
      .match(/cats hate other cats\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
--fcc-editable-region--
        
          <figcaption>Cats hate other cats.</figcaption>  
        
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

