---
id: 5ef9b03c81a63668521804d9
title: 步驟 42
challengeType: 0
dashedName: step-42
---

# --description--

佔位文字用於提示使用者應該在輸入欄位中輸入何種類型的資訊。 

這是一個 `input` 元素的範例，其 placeholder 設定為 `e.g. Jane Doe`：

```html
<input type="text" placeholder="e.g. Jane Doe">
```

為你的 `input` 元素添加佔位符文字 `cat photo URL`。

# --hints--

你已刪除你的 `input` 元素或它的語法無效。所有屬性的值都應該用引號括起來。

```js
assert(document.querySelectorAll('input').length);
```

你的 `form` 應該只包含 `input` 元素。請移除 `form` 元素內的任何其他 HTML 元素或文字。

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

你的 `input` 元素沒有 `placeholder` 屬性。請檢查開啟標記名稱後是否有空格，和／或所有屬性名稱前是否有空格。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('placeholder'));
```

你的 `input` 元素應該有一個 `placeholder` 屬性，其值為 `cat photo URL`。你要麼遺漏了該值，要麼有拼寫錯誤。請記得屬性值應該用引號括起來。

```js
assert(
  document.querySelector('input')
    ?.getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

雖然你已將 `input` 元素的 `placeholder` 屬性設為 `cat photo URL`，但建議始終用引號將屬性值括起來。

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

