---
id: 5ef9b03c81a63668521804dc
title: 步驟 46
challengeType: 0
dashedName: step-46
---

# --description--

你可以為只想從多個選項中選擇一個答案的問題使用圓鈕。

這是一個圓鈕，文字設定為 `cat` 的範例：

```html
<input type="radio"> cat
```

請記住，`input` 元素是空元素。

在文字輸入之前，添加一個圓鈕，並將文字設定為：

`Indoor`

# --hints--

你應該為你的圓鈕創建一個 `input` 元素。檢查語法。

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

你的 `input` 元素應該有一個開啟的頁籤，但不應該有關閉的頁籤。

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

你應該只為你的圓鈕、單選按鈕添加一個 `input` 元素。移除任何多餘的部分。

```js
assert(document.querySelectorAll('form > input'), 2);
```

你的新 `input` 元素應該位於現有 `input` 且 `type` 屬性設為 `text` 的元素上方。你將它們放錯順序了。

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

你的新 `input` 元素沒有 `type` 屬性。請檢查開啟標記名稱後是否有空格。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

你的新 `input` 元素應該只有一個 `type` 屬性。移除任何多餘的。

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

你的新 `input` 元素應該有一個 `type` 屬性，且其值為 `radio`。你要麼遺漏了該值，要麼有拼寫錯誤。請記得屬性值應該用引號括起來。

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

雖然你已將新的 `input` 元素的 `type` 屬性設為 `radio`，但建議始終用引號將屬性值括起來。

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

`radio` 按鈕的 `Indoor` 文字應該位於它之後，而不是之前。

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

`Indoor` 文字應該位於你的 `radio` 按鈕的正右方。你要麼遺漏了該文字，要麼有拼寫錯誤。

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

