---
id: 5ef9b03c81a63668521804d0
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

若要強調特定的字或片語，你可以使用 `em` 元素。 

在 `figcaption` 元素中，將字 `love` 包裹在強調的 `em` 元素中以強調它。

# --hints--

你的強調 `em` 元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert(document.querySelector('em'));
```

你的強調 `em` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/em\>/));
```

你已刪除 `figcaption` 元素或它缺少開啟或關閉標籤。

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

你的強調 `em` 元素應該包圍文字 `love`。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

`figcaption` 的文字應該是 `Cats love lasagna`。請檢查是否有拼字錯誤，以及 `em` 元素開啟和關閉標籤周圍是否有必要的空格。

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

