---
id: 5ef9b03c81a63668521804dd
title: 步驟 47
challengeType: 0
dashedName: step-47
---

# --description--

`label` 元素用於協助將 `input` 元素的文字與該 `input` 元素本身關聯起來（特別是對於像螢幕閱讀器這類輔助技術）。 

這是一個帶有 `radio` 按鈕的 `label` 元素範例：

```html
<label><input type="radio"> cat</label>
```

在範例中，點擊字串 `"cat"` 也會選取 `radio` 按鈕。

將你的 `radio` 按鈕巢狀放入 `label` 元素中。

# --hints--

你應該確保圓鈕仍然存在。

```js
assert(document.querySelector('input[type="radio"]'));
```

`Indoor` 文字應該位於你的 `radio` 按鈕的正右方。請確保該元素與文字之間有一個空格。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
const radioInputElem = document.querySelector('input');
assert.match(
  radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Indoor/i
);
```

你的 `label` 元素應該有一個開啟標籤。開啟標籤的語法為：`<elementName>`。

```js
assert(document.querySelector('label'));
```

你的 `label` 元素應該有一個結束標記。結束標記在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/label\>/));
```

你的圓鈕和其文字應全部位於 `label` 元素的開啟和關閉標籤之間。

```js
const labelChildNodes = [...document.querySelector('form > label').childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="radio"> Indoor
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

