---
id: 5efb23e70dc218d6c85f89b1
title: 步驟 40
challengeType: 0
dashedName: step-40
---

# --description--

你可以使用 `type` 屬性創建許多種類的輸入。你可以輕鬆創建密碼欄、重設按鈕，或讓使用者從他們的電腦中選擇檔案的控制元件。

為你的 `input` 元素添加 `type` 屬性且其值為 `text`，以創建一個文字欄位來取得使用者的文字輸入。

# --hints--

你已刪除你的 `input` 元素或它的語法無效。如果你有添加屬性，請確保它們的值被引號包圍。

```js
assert(document.querySelectorAll('input').length);
```

你的 `form` 元素應該只包含 `input` 元素。請移除 `form` 元素標籤之間的任何多餘 HTML 元素或文字。

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

你的 `input` 元素沒有 `type` 屬性且其值為 `text`。請檢查開啟標記名稱後是否有空格和／或所有屬性名稱前是否有空格。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

你的 `input` 元素應該有一個 `type` 屬性，其值為 `text`。你要麼遺漏了該值，要麼有拼寫錯誤。請記住，屬性值應該用引號括起來。

```js
assert.match(
  document.querySelector('input')?.getAttribute('type').replace(/\s+/g, ' '),
  /^text$/i
);
```

雖然你已將 `input` 元素的 `type` 屬性設定為 `text`，但建議始終用引號將屬性值括起來。

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

