---
id: 5efb2c990dc218d6c85f89b2
title: 步驟 45
challengeType: 0
dashedName: step-45
---

# --description--

即使你將按鈕添加在文字輸入的下方，它們仍會在頁面上並排顯示。那是因為 `input` 和 `button` 元素都是 <dfn>行內元素</dfn>，不會出現在新行。

你添加的按鈕預設會提交表單。然而，依賴預設行為可能會造成混淆。為你的 `button` 添加 `type` 屬性，並將值設為 `submit`，以明確表示它是提交按鈕。

# --hints--

你的 `button` 元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert(document.querySelector('button'));
```

你的 `button` 元素應該有一個結束標籤。結束標籤在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/button\>/));
```

你的 `button` 元素沒有 `type` 屬性。請檢查開啟標記名稱後是否有空格。

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('type'));
```

你的 `button` 元素應該有一個 `type` 屬性，且其值為 `submit`。你要麼遺漏了該值，要麼有拼寫錯誤。請記得屬性值應該用引號括起來。

```js
assert.match(
  document.querySelector('button')?.getAttribute('type'),
  /^submit$/i
);
```

雖然你已將 `button` 元素的 `type` 屬性設為 `submit`，但建議始終用引號包住屬性的值。

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

