---
id: 5efc4f528d6a74d05e68af74
title: 步驟 58
challengeType: 0
dashedName: step-58
---

# --description--

還有另一種方法可以將 `input` 元素的文字與該元素本身關聯。你可以將文字巢狀在 `label` 元素中，並新增一個 `for` 屬性，其值與 `input` 元素的 `id` 屬性相同。

給定如下的 `input` 元素：

```html
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

一個與此 `input` 元素相關的 `label` 元素範例如下：

```html
<label for="breakfast">Breakfast</label>
```

將文字 `Loving` 與核取方塊關聯，方法是將文字 `Loving` 僅巢狀於 `label` 元素中，並為其指定適當的 `for` 屬性。

# --hints--

你應該確定核取方塊元素仍然存在，且你沒有對它做任何更改。它應該保持與起始程式碼中相同。

```js
const checkboxElementRegex = /\s*<input\s+id\s*=\s*"loving"\s*type\s*=\s*"checkbox"\s*>/;
assert.match(code, checkboxElementRegex);
```

你的核取方塊仍應該有一個 `id` 屬性，值為 `loving`。預期為 `--fcc-expected--`，但找到 `--fcc-actual--`。

```js
const checkboxElementId = document.querySelector('input[type="checkbox"]')?.getAttribute('id');
assert.strictEqual(checkboxElementId, 'loving');
```

你應該在核取方塊元素後添加一個新的 `label` 元素。確保它同時具有開啟和關閉的標籤。

```js
assert.lengthOf(document.querySelectorAll('label'), 3);
assert.lengthOf(code.match(/<\/label\>/g), 3);
```

你的 checkbox 後面應該放置 `label` 元素。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElement = checkboxElement?.nextElementSibling?.nodeName;
assert.strictEqual(checkboxSiblingElement, 'LABEL');
```

你的 `label` 元素應該有文字 `Loving`。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementText = checkboxElement?.nextElementSibling?.innerText;

assert.strictEqual(checkboxSiblingElementText, 'Loving');
```

`Loving` 標籤文字應該只在你的新 `label` 元素內。請確保你沒有在 `label` 元素外重複該文字。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingNode = checkboxElement?.nextSibling?.nodeValue ?? '';
const labelElement = checkboxElement?.nextElementSibling;
const labelSiblingNode = labelElement?.nextSibling?.nodeValue ?? '';

assert.notInclude(checkboxSiblingNode, 'Loving');
assert.notInclude(labelSiblingNode, 'Loving');
```

你的 `label` 元素沒有 `for` 屬性。請確保開啟的 `label` 頁籤名稱和 `for` 屬性之間有一個空格。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.isNotNull(checkboxSiblingElementAttr);
```

你的 `label` 元素應該有一個 `for` 屬性，值為 `loving`。預期為 `--fcc-expected--`，但找到 `--fcc-actual--`。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.strictEqual(checkboxSiblingElementAttr, 'loving');
```

你的核取方塊和你的新 `label` 元素之間應該有一個空格。

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

