---
id: 5f05a1d8e233dff4a68508d8
title: 步驟 49
challengeType: 0
dashedName: step-49
---

# --description--

在第一個圓鈕下方創建另一個圓鈕。將它巢狀在一個 `label` 元素內，並以 `Outdoor` 作為該 `label` 的文字。給該圓鈕一個 `id` 屬性，值為 `outdoor`。

# --hints--

只有原始的 `Indoor` 圓鈕、單選按鈕 `input` 應該設定 `id` 為 `indoor`。如果需要，你可以重新開始此步驟以取回原始的 HTML。

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

你不應該對 `Indoor` 圓鈕進行任何更改。如果需要，你可以重新啟動該步驟以取回原始的 HTML。

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

你應該在一個新的 `label` 元素中嵌套恰好一個新的 `input` 元素。確保你的新 `label` 同時有開啟和關閉的標籤。

```js
assert(document.querySelectorAll('label input').length === 2);
```

你不應該為任一開啟的 `label` 頁籤添加任何屬性。

```js
assert(code.match(/<label\s*>/g).length === 2);
```

你的新 `input` 應該有一個 `id` 屬性，值為 `outdoor`。 

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

你的新 `input` 應該有一個 `type` 屬性，值為 `radio`。

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

你不應該添加任何新的元素，除了巢狀在 `label` 中的 `input`。

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

你的新 `input` 左側不應該有任何文字。

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

你新的圓鈕標籤文字必須完全是 `Outdoor`。這包括大小寫。 

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

你的新圓鈕和標籤應該緊接在 `Indoor` 圓鈕和標籤的下方／後方。它們之間不應該有其他頁籤。

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

