---
id: 5f0d4ab1b435f13ab6550052
title: 步驟 54
challengeType: 0
dashedName: step-54
---

# --description--

接著，你將為表單添加一些新的 `input` 元素，所以請在目前的 `fieldset` 元素正下方再添加一個 `fieldset` 元素。

# --hints--

你的新 `fieldset` 元素應該有一個開啟標記。開啟標記的語法為：`<elementName>`。

```js
assert(document.querySelectorAll('fieldset').length >= 2);
```

你應該只添加一個開啟的 `fieldset` 標籤。請移除任何多餘的標籤。

```js
assert(document.querySelectorAll('fieldset').length === 2);
```

你的新 `fieldset` 元素應該有一個結束標記。結束標記在 `<` 字元後面緊接著有一個 `/`。

```js
assert(code.match(/<\/fieldset>/g).length >= 2);
```

你應該只添加一個結束的 `fieldset` 頁籤。請移除任何多餘的。

```js
assert(code.match(/<\/fieldset>/g).length === 2);
```

第二個 `fieldset` 元素不應該巢狀在第一個 `fieldset` 元素中。

```js
const childrenOf1stFieldset = [
  ...document.querySelector('form > fieldset').children
];
const foundElems = childrenOf1stFieldset.filter((child) => {
  return child.nodeName === 'FIELDSET';
});
assert(foundElems.length === 0);
```

兩個 `fieldset` 元素都應該在文字欄位上方。它們的順序錯亂。

```js
const formChildren = document.querySelector('form')?.children;
assert.isTrue(
  formChildren?.[0]?.nodeName === 'FIELDSET' &&
  formChildren?.[1]?.nodeName === 'FIELDSET' &&
  formChildren?.[2]?.nodeName === 'INPUT' &&
  formChildren?.[2]?.getAttribute('type') === 'text'
);
```

你的新 `fieldset` 元素應該位於現有的 `fieldset` 元素下方。你將它們放錯順序了。

```js
const fieldsetChildren = [...document.querySelectorAll('fieldset')].map(
  (elem) => elem.children
);
assert(fieldsetChildren[0].length > fieldsetChildren[1].length);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

