---
id: 5f1a80975fc4bcae0edb3497
title: 步驟 51
challengeType: 0
dashedName: step-51
---

# --description--

如果你選取 `Indoor` 圓鈕並提交表單，該按鈕的表單資料將基於其 `name` 和 `value` 屬性。由於你的圓鈕沒有 `value` 屬性，表單資料將包含 `indoor-outdoor=on`，當你有多個按鈕時這並不實用。

為你的兩個圓鈕添加 `value` 屬性。為方便起見，將按鈕的 `value` 屬性設為與其 `id` 屬性相同的值。

# --hints--

兩個圓鈕仍應位於開啟和關閉的 `label` 元素標籤之間。

```js
const labelChildNodes = [...document.querySelectorAll('form > label')].map(
  (node) => node.childNodes
);
assert(
  labelChildNodes.filter((childNode) => childNode[0].nodeName === 'INPUT')
    .length === 2
);
```

兩個圓鈕都應該有一個 `value` 屬性。請檢查開啟標籤名稱後是否有空格，和／或所有屬性名稱前是否有空格。

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(radioButtons.every((btn) => btn.hasAttribute('value')));
```

兩個圓鈕都應該有一個 `type` 屬性。請檢查開啟標籤名稱後是否有空格，和／或所有屬性名稱前是否有空格。

```js
const radioButtons = [...document.querySelectorAll('input')];
assert(radioButtons.every((btn) => btn.hasAttribute('type')));
```

`Indoor` 圓鈕的 `value` 屬性應該設為 `indoor`。你要麼遺漏了該值，要麼有拼寫錯誤。請記得屬性值應該用引號括起來。

```js
const indoorRadioButton = document.querySelector('#indoor');
assert(indoorRadioButton.getAttribute('value').match(/^indoor$/));
```

`Outdoor` 圓鈕的 `value` 屬性應該設為 `outdoor`。你要麼遺漏了該值，要麼有拼寫錯誤。請記得屬性值應該用引號括起來。

```js
const outdoorRadioButton = document.querySelector('#outdoor');
assert(outdoorRadioButton.getAttribute('value').match(/^outdoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
          <label><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

