---
id: 671144cdcc01d73f7dd79dc9
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

現在你已經添加了聯結，你可以移除註解。

# --hints--

你的程式碼中不應該有任何註解。

```js
assert.notMatch(code, /<!--/);
assert.notMatch(code, /-->/);
```

文字 `TODO: Add link to cat photos` 不應該再出現。

```js
assert.notMatch(code, /TODO:?/i)
assert.notMatch(code, /Add\s*link\s*to\s*cat\s*photos/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
--fcc-editable-region--
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
