---
id: 671141feba228a35cefba82d
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

將現有的文字 `cute cats` 轉換成一個指向以下位置的錨點元素：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

在第一個 `p` 元素中應該有一個新的錨點元素。

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

錨點文字應該是 `cute cats`。

```js
assert.strictEqual(document.querySelector('a')?.innerText.trim(), "cute cats");
```

新錨點元素的 `href` 屬性應該是 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`。

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

你的程式碼中應該只有一個 `cute cats` 的實體。  

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

`p` 元素的文字仍應該是 `Everyone loves cute cats online!`。

```js
assert.strictEqual(document.querySelector('p')?.innerText.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
