---
id: 7cf9b03d81a65668421804c3
title: 步驟 41
challengeType: 0
dashedName: step-41
---

# --description--

為了讓表單的資料能被 `action` 屬性指定的位置存取，你必須為文字欄指定 `name` 屬性並設值，以表述所提交的資料。 

這是一個帶有 `name` 屬性的 `input` 元素範例：

```html
<input type="text" name="name">
```

為你的文字欄添加值為 `catphotourl` 的 `name` 屬性。

# --hints--

你已刪除你的 `input` 元素或它的語法無效。所有屬性的值都應該用引號括起來。

```js
assert(document.querySelectorAll('input').length);
```

你的 `form` 應該只包含 `input` 元素。請移除 `form` 元素內的任何其他 HTML 元素或文字。

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

你的 `input` 元素沒有 `name` 屬性。請檢查開啟標記名稱後是否有空格，和／或所有屬性名稱前是否有空格。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

你的 `input` 元素應該有一個 `name` 屬性，其值為 `catphotourl`。你要麼遺漏了該值，要麼有拼寫錯誤。請記住，屬性值應該用引號括起來。

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

雖然你已將 `input` 元素的 `name` 屬性設為 `catphotourl`，但建議始終用引號將屬性值括起來。

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

