---
id: 65ef1d5e3d2927c5e0f4997b
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

Python에서 `is` 키워드는 객체 동일성을 검사합니다. 이는 두 변수가 메모리에서 같은 객체를 가리키는지 확인하는 데 사용됩니다. `is`와 달리 동등 연산자(`==`)는 두 객체가 메모리에서 같은 객체인지와 상관없이 값이 같은지를 판단합니다.

for 루프 밖에서 `if` 문을 만들어 `root is None`인지 확인하세요. 만약 그렇다면, 메시지 `'Failed to converge within {max_iterations} iterations.'`를 출력하세요. 메시지는 f-string을 사용해 형식을 맞추는 것을 기억하세요.

# --hints--

`if`인지 확인하는 `root is None` 문을 만들어야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

`f'Failed to converge within {max_iterations} iterations.'` 함수에 인수 `print`를 전달해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
