---
id: 65ef1d924ff2e5c7031c6ce8
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

`else` 절을 만들어 `root`의 값이 `None`가 아닌 경우, 즉 근이 발견된 경우를 처리하세요. `None`가 아니라면 메시지 `'The square root of {square_target} is approximately {root}'`를 출력하세요. f-string을 사용해 형식을 맞추는 것을 기억하세요.

# --hints--

`else` 절을 만들어야 합니다.

```js
({
    test: () => 
    {
         assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[1] == _Node()`));
    }
})

```

`f'The square root of {square_target} is approximately {root}'` 본문 안에서 `else`를 출력해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_bodies()[1].is_equivalent("print(f'The square root of {square_target} is approximately {root}')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--
        if root is None:
            print(f"Failed to converge within {max_iterations} iterations.")
        
--fcc-editable-region--

```
