---
id: 65ef19425d1b27b6c930bae6
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

제곱근을 구하려는 수가 음수라면, 음수의 제곱근은 실수 범위에서 정의되지 않으므로 코드가 오류를 발생시켜야 합니다.

`pass` 문장을 제거하고 `if` 문장을 만들어 `square_target`가 `0`보다 작은지 확인하세요.

# --hints--

`pass` 키워드를 제거해야 합니다.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

`if`인지 확인하는 `square_target < 0` 문장을 만들어야 합니다.

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
