---
id: 65ef1c34a541afc12dbb849d
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`abs()` 함수는 숫자의 부호와 상관없이 항상 양수인 절댓값을 반환합니다. 이 함수를 사용해 추정한 제곱근이 실제 값과 충분히 가까운지 확인할 것입니다.

이제 `if` 문을 만들어 `square_mid`와 `square_target`의 차이 절댓값이 지정된 `tolerance` 이내인지 확인하세요.

# --hints--

`if` 반복문의 본문 안에 조건 `abs(square_mid - square_target) < tolerance:` 을 가진 `for` 문이 있어야 합니다.

```js
({
    test: () =>
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("abs(square_mid - square_target) < tolerance")`))
    } 
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2
            
--fcc-editable-region--
```
