---
id: 65ef198fde24dfb7ff675b42
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

`raise` 문장은 특정 예외를 강제로 발생시키도록 합니다. `raise` 키워드 다음에 예외 유형이 오며, 사용자 지정 오류 메시지를 제공할 수 있습니다:

```py
raise ValueError("Invalid value")
```

위 코드를 실행하면 `ValueError`가 발생하고 메시지 `"Invalid value"`가 사용자에게 표시됩니다.

만약 `square_target`가 `0`보다 작으면, 실수 제곱근을 계산할 수 없습니다. 따라서 메시지 `ValueError`와 함께 `'Square root of negative number is not defined in real numbers'`를 발생시키세요. `pass` 키워드는 반드시 제거해야 합니다.

# --hints--

`pass` 키워드를 제거해야 합니다.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

`ValueError` 본문 안에서 메시지 `'Square root of negative number is not defined in real numbers'`와 함께 `if`를 발생시켜야 합니다.

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
