---
id: 65ef181dad3d96b14a21a03e
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 프로젝트에서는 이분법을 사용해 주어진 수의 근사 제곱근을 구할 것입니다.

이분법은 실수 함수의 근을 찾는 기법입니다. 제곱근이 존재하는 구간을 점점 좁혀가면서 지정한 허용 오차 내의 값으로 수렴할 때까지 진행합니다.

먼저 `square_root_bisection`라는 이름의 함수를 만드세요. 매개변수는 일단 비워두세요.

# --hints--

매개변수 없이 `square_root_bisection` 함수를 정의해야 합니다.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
