---
id: 65ef1be2b958cabf712296cd
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

이제 현재 구간의 중간값을 찾아 그 중간값의 제곱과 목표 값을 비교하면서 구간을 반복해서 좁혀 나가세요.

이를 위해 `else` 블록 안에 `for` 반복문을 만들어 `max_iterations`번까지 실행하세요.

반복문에는 `range` 함수를 사용하세요. 이 함수는 순회할 수 있는 수열을 생성합니다. 구문은 `range(start, stop, step)`이며, 여기서 `start`는 시작 정수(포함), `stop`는 마지막 정수(미포함), `step`는 수열에서 이전 수와 다음 수의 차이입니다.

또한 `_`를 반복 변수로 사용하세요. `_`는 자리 표시자 역할을 하며, 변수를 사용해야 하지만 실제 값은 필요 없을 때 유용합니다.

# --hints--

`for` 반복문을 만들어 `range(max_iterations)`를 순회하세요. 반복 변수로 `_`를 사용하세요.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
