---
id: 65ef1a720d6e65bbb0d22463
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--
다음으로, `square_target`이 `1` 또는 `0`와 다르면서 양수인 경우를 처리할 것입니다.

이 경우를 처리하는 `else` 절을 만드세요. 

# --hints--

`else` 절을 만들고 `pass` 키워드를 추가하는 것을 잊지 마세요.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[2] == _Node()`))
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')
--fcc-editable-region--


--fcc-editable-region--
```
