---
id: 65ef190c6b51e9b5a5f7ed29
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

`square_root_bisection` 함수에 다음 매개변수를 주세요:

- `square_target`: 제곱근을 구하려는 숫자입니다.
- `tolerance` (선택 사항): 근사값의 제곱과 실제 목표값 사이의 허용 오차입니다(기본값은 `1e-7`). 허용 오차 `1e-7`는 해가 실제 값에서 `0.0000001` 이내로 정확함을 의미하며, 정확도와 성능의 균형을 맞추는 좋은 기본값입니다.
- `max_iterations` (선택 사항): 수행할 최대 반복 횟수입니다(기본값은 `100`). 이 한도 내에 수렴하지 않으면 해를 찾지 못한 것으로 간주합니다.
  
# --hints--

함수는 `square_target`, `tolerance = 1e-7`, `max_iterations = 100` 매개변수를 가져야 합니다. 순서가 중요합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```
