---
id: 65ef19ec318e2fb9066381e3
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`square_target`가 `0`이거나 `1`일 때 각각의 경우를 만드세요.

먼저 `if`가 `square_target`와 같은지 확인하는 `1` 문장을 만드세요.

# --hints--

`if`인지 확인하는 `square_target == 1` 문장을 만드세요.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[0].is_equivalent("square_target == 1")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')

--fcc-editable-region--
```
