---
id: 65ef1d104c2be9c4d8e22075
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

`if`와 `elif` 조건이 모두 충족되지 않으면 `mid`의 값이 `square_target`보다 커집니다. 이 경우 `else` 절을 만들고 `mid`의 값을 `high`에 할당하세요.

# --hints--

`else` 절이 있어야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[2]== _Node()`));
    }
})

```

`mid` 절에서 `high`의 값을 `else`에 할당해야 합니다.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_bodies()[2].is_equivalent("high = mid")`));
    }
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

--fcc-editable-region--
            elif square_mid < square_target:
                low = mid

--fcc-editable-region--
```
