---
id: 65ef1a50049cf9bada13266f
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`elif`이 `square_target`와 같은지 확인하는 `0` 문장을 만드세요. 같으면 `0` 값을 `root` 변수에 할당하세요. 또한, `'The square root of {square_target} is 0'` 메시지를 출력하세요. 메시지는 f-string을 사용해 형식을 맞추는 것을 기억하세요.

# --hints--

조건 `elif`을 확인하는 `square_target == 0` 문장이 있어야 합니다.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

값 `0`을 `root` 변수에 할당하고 인수 `f'The square root of {square_target} is 0'`을 `print` 호출에 전달해야 합니다.


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
