---
id: 698b42492a7f5f737895892e
title: Construye un convertidor de Lista de Adyacencia a Matriz
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

En este laboratorio, construirás una función que convierte una representación de lista de adyacencia de un grafo en una matriz de adyacencia. Una lista de adyacencia es un objeto donde cada clave representa un nodo, y el valor correspondiente es un arreglo de nodos a los que el nodo clave está conectado. Una matriz de adyacencia es un arreglo 2D donde la entrada en la posición `[i][j]` es `1` si hay una arista del nodo `i` al nodo `j`, y `0` en caso contrario.

Por ejemplo, dada la lista de adyacencia:

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

La matriz de adyacencia correspondiente sería:

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes definir una función llamada `adjacencyListToMatrix` para convertir una lista de adyacencia en una matriz de adyacencia.
2. La función debe tomar un objeto que represente la lista de adyacencia de un grafo no ponderado (ya sea no dirigido o dirigido) como argumento.
3. La función debería:
   - Convierte la lista de adyacencia en una matriz de adyacencia.
   - Imprime cada fila en la matriz de adyacencia.
   - Devuelve la matriz de adyacencia.

Por ejemplo, `adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})` debería imprimir:

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

y devuelve `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`.

# --hints--

Debes definir una función llamada `adjacencyListToMatrix`.

```js
assert.isFunction(adjacencyListToMatrix);
```

La función `adjacencyListToMatrix` debe tener un parámetro.

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

La función debe determinar correctamente el número de nodos a partir de la lista de adyacencia.

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

La función debe establecer correctamente los valores de la matriz a `1` para las aristas existentes.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

La función debe imprimir cada fila de la matriz.

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

La función debe devolver la matriz de adyacencia.

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

Cuando se da la lista de adyacencia `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}`, la función debe devolver `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`.

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Cuando se da la lista de adyacencia `{0: [1], 1: [0]}`, la función debe devolver `[[0, 1], [1, 0]]`.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Cuando se proporciona la lista de adyacencia `{0: [], 1: [], 2: []}`, la función debería devolver `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
