---
id: 587d781e367417b2b2512acc
title: Закріплення елемента у вікні браузера за допомогою фіксованого позиціювання
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDNUR'
forumTopicId: 301061
dashedName: lock-an-element-to-the-browser-window-with-fixed-positioning
---

# --description--

Наступна схема макету, яку пропонує CSS, — це позиція `fixed`, яка є типом абсолютного позиціювання, що закріплює елемент відносно вікна браузера. Подібно до абсолютного позиціювання, вона використовується з CSS-властивостями компенсування і також виводить елемент із нормального потоку документа. Інші елементи більше не "бачать", де він розташований, що може вимагати деяких правок макету в інших місцях.

Головна відмінність між позиціями `fixed` і `absolute` полягає в тому, що елемент із фіксованою позицією не рухається, коли користувач прокручує сторінку.

# --instructions--

Навігаційна панель у коді має ідентифікатор `navbar`. Змініть її `position` на `fixed` і компенсуйте 0 пікселів від `top` та 0 пікселів від `left`. Після того, як ви додасте код, прокрутіть вікно попереднього перегляду, щоб побачити, як навігація залишається на місці.

# --hints--

Елемент `#navbar` повинен мати встановлену позицію `position` зі значенням `fixed`.

```js
const navbarElement = document.querySelector('#navbar');
const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.position, 'fixed');
```

Ваш код має використовувати CSS-компенсацію `top` зі значенням 0 пікселів для елемента `#navbar`.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
assert.equal(navbarStyle?.top, '0px');
```

Ваш код має використовувати CSS-компенсацію `left` зі значенням 0 пікселів для елемента `#navbar`.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.left, '0px');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {



    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```

# --solutions--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {
    position: fixed;
    top: 0;
    left: 0;
    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```
