---
id: 587d78a3367417b2b2512acf
title: Змінити позицію перекривних елементів за допомогою властивості z-index
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM94aHk'
forumTopicId: 301046
dashedName: change-the-position-of-overlapping-elements-with-the-z-index-property
---

# --description--

Коли елементи розташовані так, що перекриваються (тобто використовуючи `position: absolute | relative | fixed | sticky`), елемент, який йде пізніше в HTML-розмітці, за замовчуванням буде зверху над іншими елементами. Однак властивість `z-index` може визначати порядок, у якому елементи накладаються один на одного. Вона має бути цілим числом (тобто цілим числом, а не десятковим), і вищі значення властивості `z-index` елемента піднімають його вище в стеку, ніж ті, що мають нижчі значення.

# --instructions--

Додайте властивість `z-index` до елемента з класом `first` (червоний прямокутник) і встановіть її значення рівним 2, щоб він перекривав інший елемент (синій прямокутник).

# --hints--

Елемент з класом `first` повинен мати значення `z-index` рівне 2.

```js
const firstElement = document.querySelector('.first');
const firstStyle = window.getComputedStyle(firstElement); 
assert.equal(firstStyle?.zIndex, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;

  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>
```

# --solutions--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;
    z-index: 2;
  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>
<div class="first"></div>
<div class="second"></div>
```
