---
id: 587d78a8367417b2b2512ae3
title: "Анімувати елементи безперервно, використовуючи нескінченну кількість анімацій"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDVfq'
forumTopicId: 301041
dashedName: animate-elements-continually-using-an-infinite-animation-count
---

# --description--

Попередні завдання охоплювали, як використовувати деякі властивості анімації та правило `@keyframes`. Ще одна властивість анімації — це `animation-iteration-count`, яка дозволяє контролювати, скільки разів ви хочете повторити цикл анімації. Ось приклад:

```css
animation-iteration-count: 3;
```

У цьому випадку анімація зупиниться після трьох запусків, але можна зробити так, щоб анімація працювала безперервно, встановивши це значення в `infinite`.

# --instructions--

Щоб м’яч безперервно підстрибував праворуч у циклі, змініть властивість `animation-iteration-count` на `infinite`.

# --hints--

Властивість `animation-iteration-count` має мати значення `infinite`.

```js
const ballElement = document.querySelector('#ball');
const ballStyle = window.getComputedStyle(ballElement); 
assert.equal(ballStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: 3;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```
