---
id: 587d78a7367417b2b2512ae0
title: Використання CSS-анімації для зміни стану кнопки при наведенні
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4vZAa'
forumTopicId: 301073
dashedName: use-css-animation-to-change-the-hover-state-of-a-button
---

# --description--

Ви можете використовувати CSS `@keyframes` для зміни кольору кнопки у стані наведення.

Ось приклад зміни ширини зображення при наведенні:

```html
<style>
  img {
    width: 30px;
  }
  img:hover {
    animation-name: width;
    animation-duration: 500ms;
  }

  @keyframes width {
    100% {
      width: 40px;
    }
  }
</style>

<img src="https://cdn.freecodecamp.org/curriculum/applied-visual-design/google-logo.png" alt="Google's Logo" />
```

# --instructions--

Зверніть увагу, що `ms` означає мілісекунди, де 1000ms дорівнює 1s.

Використайте CSS `@keyframes`, щоб змінити `background-color` `button` елемента так, щоб він став `#4791d0` при наведенні користувача. Правило `@keyframes` має містити лише запис для `100%`.

# --hints--

Правило @keyframes має використовувати `animation-name` колір фону.

```js
assert.match(code, /@keyframes\s+?background-color\s*?{/g);
```

Потрібно одне правило під `@keyframes`, яке змінює `background-color` на `#4791d0` при 100%.

```js
assert.match(code, /100%\s*?{\s*?background-color:\s*?#4791d0;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }


</style>

<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }

  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
