---
id: 587d7791367417b2b2512ab3
title: Створення візуального балансу за допомогою властивості text-align
challengeType: 0
videoUrl: 'https://scrimba.com/c/c3b4EAp'
forumTopicId: 301053
dashedName: create-visual-balance-using-the-text-align-property
---

# --description--

Цей розділ навчальної програми зосереджений на прикладному візуальному дизайні. Перша група завдань базується на заданому макеті картки, щоб показати кілька основних принципів.

Текст часто займає велику частину веб-вмісту. CSS має кілька варіантів вирівнювання тексту за допомогою властивості `text-align`.

`text-align: justify;` розподіляє текст так, щоб кожен рядок мав однакову ширину.

`text-align: center;` центрує текст.

`text-align: right;` вирівнює текст по правому краю.

А `text-align: left;` (за замовчуванням) вирівнює текст по лівому краю.

# --instructions--

Вирівняйте текст у тегу `h4`, який містить слово "Google", по центру. Потім вирівняйте по ширині абзац, який містить інформацію про те, як була заснована Google.

# --hints--

Ваш код повинен використовувати властивість text-align у тегу `h4`, щоб встановити її в `center`.

```js
const h4Element =document.querySelector('h4')
const h4Style = window.getComputedStyle(h4Element);
assert.equal(h4Style?.textAlign, 'center');
```

Ваш код повинен використовувати властивість text-align у тегу `p`, щоб встановити її в `justify`.

```js
const pElement =document.querySelector('p')
const pStyle = window.getComputedStyle(pElement);
assert.equal(pStyle?.textAlign, 'justify');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {

  }
  p {

  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
  }
  p {
    text-align: justify;
  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
