---
id: 587d78a3367417b2b2512ad0
title: Центрування елемента по горизонталі за допомогою властивості margin
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJqU4'
forumTopicId: 301043
dashedName: center-an-element-horizontally-using-the-margin-property
---

# --description--

Ще одна техніка позиціювання — це центрування блокового елемента по горизонталі. Один зі способів зробити це — встановити його `margin` у значення auto.

Цей метод також працює для зображень. Зображення за замовчуванням є рядковими елементами, але їх можна змінити на блокові, якщо встановити властивість `display` у `block`.

# --instructions--

Центруйте `div` на сторінці, додавши властивість `margin` зі значенням `auto`.

# --hints--

`div` має мати встановлену `margin` зі значенням `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('div')?.margin ,'auto');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;

  }
</style>
<div></div>
```

# --solutions--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;
    margin: auto;
  }
</style>
<div></div>
```
