---
id: 587d781e367417b2b2512acb
title: Закріплення елемента до батьківського за допомогою абсолютного позиціонування
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJ7c3'
forumTopicId: 301060
dashedName: lock-an-element-to-its-parent-with-absolute-positioning
---

# --description--

Наступним варіантом для CSS властивості `position` є `absolute`, яка закріплює елемент відносно його батьківського контейнера. На відміну від позиції `relative`, це виводить елемент із нормального потоку документа, тому навколишні елементи його ігнорують. Для регулювання позиції використовують CSS властивості зміщення (top або bottom і left або right).

Одна особливість абсолютного позиціонування полягає в тому, що воно закріплюється відносно найближчого *позиціонованого* предка. Якщо ви забудете додати правило позиціонування до батьківського елемента (зазвичай це робиться за допомогою `position: relative;`), браузер продовжить шукати вгору по ланцюжку і врешті-решт за замовчуванням використає тег `body`.

# --instructions--

Закріпіть елемент `#searchbar` у верхньому правому куті його батьківського елемента `section`, оголосивши його `position` як `absolute`. Задайте йому зміщення `top` і `right` по 50 пікселів кожне.

# --hints--

Елемент `#searchbar` має мати властивість `position`, встановлену в `absolute`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.position, 'absolute');
```

Ваш код має використовувати CSS зміщення `top` на 50 пікселів для елемента `#searchbar`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.top, '50px');
```

Ваш код має використовувати CSS зміщення `right` на 50 пікселів для елемента `#searchbar`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.right, '50px');
```

# --seed--

## --seed-contents--

```html
<style>
  #searchbar {



  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```

# --solutions--

```html
<style>
  #searchbar {
    position: absolute;
    top: 50px;
    right: 50px;
  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```
