---
id: 587d78a9367417b2b2512ae8
title: "Дізнайтеся, як працюють криві Безьє"
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bDrs8'
forumTopicId: 301058
dashedName: learn-how-bezier-curves-work
---

# --description--

Останнє завдання познайомило вас із властивістю `animation-timing-function` та кількома ключовими словами, які змінюють швидкість анімації протягом її тривалості. CSS пропонує ще один варіант, окрім ключових слів, який дає ще точніший контроль над тим, як відбувається анімація, — це використання кривих Безьє.

У CSS-анімаціях криві Безьє використовуються з функцією `cubic-bezier`. Форма кривої представляє, як відбувається анімація. Крива розташована у системі координат 1 на 1. Вісь X цієї системи координат — це тривалість анімації (уявіть це як часову шкалу), а вісь Y — це зміна в анімації.

Функція `cubic-bezier` складається з чотирьох основних точок, які розташовані на цій сітці 1 на 1: `p0`, `p1`, `p2` і `p3`. `p0` і `p3` задані для вас — це початкова та кінцева точки, які завжди розташовані відповідно в початку координат (0, 0) і (1, 1). Ви задаєте значення x і y для двох інших точок, і те, де ви їх розміщуєте на сітці, визначає форму кривої, за якою відбуватиметься анімація. Це робиться в CSS, оголошуючи значення x і y "якірних" точок `p1` і `p2` у формі: `(x1, y1, x2, y2)`. Підсумовуючи, ось приклад кривої Безьє в CSS-коді:

```css
animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
```

У наведеному вище прикладі значення x і y еквівалентні для кожної точки (x1 = 0.25 = y1 і x2 = 0.75 = y2), що, якщо пригадати геометрію, дає лінію, яка простягається від початку координат до точки (1, 1). Ця анімація — це лінійна зміна елемента протягом тривалості анімації і є еквівалентною використанню ключового слова `linear`. Іншими словами, вона змінюється з постійною швидкістю.

# --instructions--

Для елемента з ідентифікатором `ball1` змініть значення властивості `animation-timing-function` з `linear` на його еквівалентне значення функції `cubic-bezier`. Використайте значення точок, наведені в прикладі вище.

# --hints--

Значення властивості `animation-timing-function` для елемента з ідентифікатором `ball1` має бути лінійно-еквівалентною функцією `cubic-bezier`.

```js
const ballOne = document.querySelector('#ball1');
const ballOneStyle = window.getComputedStyle(ballOne); 
assert.equal(
  ballOneStyle?.animationTimingFunction,
    'cubic-bezier(0.25, 0.25, 0.75, 0.75)'
);
```

Значення властивості `animation-timing-function` для елемента з ідентифікатором `ball2` не повинно змінюватися.

```js
const ballTwoElement = document.querySelector('#ball2');
const ballTwoStyle = window.getComputedStyle(ballTwoElement); 
const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(
  ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
