---
id: 587d78a6367417b2b2512adb
title: Використання властивості CSS transform skewX для нахилу елемента по осі X
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLP8Sr'
forumTopicId: 301074
dashedName: use-the-css-transform-property-skewx-to-skew-an-element-along-the-x-axis
---

# --description--

Наступна функція властивості `transform` — це `skewX()`, яка нахиляє вибраний елемент по осі X (горизонтальній) на заданий кут.

У наведеному коді елемент абзацу нахиляється на -32 градуси по осі X.

```css
p {
  transform: skewX(-32deg);
}
```

# --instructions--

Нахиліть елемент з ідентифікатором `bottom` на 24 градуси по осі X, використовуючи властивість `transform`.

# --hints--

Елемент з ідентифікатором `bottom` має бути нахилений на 24 градуси по осі X.

```js
assert.match(code, /#bottom\s*?{\s*?.*?\s*?transform:\s*?skewX\(24deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;

  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
