---
id: 587d781d367417b2b2512ac8
title: Налаштуйте стан наведення для тегу посилання
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakRGcm'
forumTopicId: 301035
dashedName: adjust-the-hover-state-of-an-anchor-tag
---

# --description--

У цьому завданні йдеться про використання псевдокласів. Псевдоклас — це ключове слово, яке можна додати до селекторів, щоб вибрати певний стан елемента.

Наприклад, стиль тегу посилання можна змінити для його стану наведення за допомогою селектора псевдокласу `:hover`. Ось CSS, який змінює `color` тегу посилання на червоний під час наведення:

```css
a:hover {
  color: red;
}
```

# --instructions--

У редакторі коду є CSS-правило, яке задає чорний колір для всіх тегів `a`. Додайте правило, щоб коли користувач наводить курсор на тег `a`, `color` став синім.

# --hints--

Тег посилання `color` має залишатися чорним, додайте CSS-правила лише для стану `:hover`.

```js
const anchorElement = document.querySelector("a"); 
const anchorStyle = window.getComputedStyle(anchorElement);
assert.equal(anchorStyle?.color, 'rgb(0, 0, 0)');
```

Тег посилання має мати `color` синього кольору при наведенні.

```js
assert.match(code,
    /a:hover\s*?{\s*?color:\s*?(blue|rgba\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?,\s*?1\s*?\)|#00F|rgb\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?\))\s*?;\s*?}/gi
  );
```

# --seed--

## --seed-contents--

```html
<style>
  a {
    color: #000;
  }



</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```

# --solutions--

```html
<style>
  a {
    color: #000;
  }
  a:hover {
    color: rgba(0,0,255,1);
  }
</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```
