---
id: 587d78a7367417b2b2512ae2
title: створення візуального напрямку за допомогою поступового зникнення елемента зліва направо
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJqqAE'
forumTopicId: 301054
dashedName: create-visual-direction-by-fading-an-element-from-left-to-right
---

# --description--

У цьому завданні ви зміните `opacity` анімованого елемента так, щоб він поступово зникав, коли досягає правої сторони екрану.

У показаній анімації круглий елемент із градієнтним фоном рухається вправо до позначки 50% анімації відповідно до правила `@keyframes`.

# --instructions--

Цільтеся на елемент з ідентифікатором `ball` і додайте множину властивостей `opacity` зі значенням 0.1 на `50%`, щоб елемент зникав під час руху вправо.

# --hints--

Правило `keyframes` для зникнення має встановити властивість `opacity` на 0.1 на 50%.

```js
assert.match(code,/@keyframes fade\s*?{\s*?50%\s*?{\s*?(?:left:\s*?60%;\s*?opacity:\s*?0?\.1;|opacity:\s*?0?\.1;\s*?left:\s*?60%;)/gi
);
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;

    }
  }

</style>

<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;
      opacity: 0.1;
    }
  }
</style>
<div id="ball"></div>
```
