---
id: 587d78a4367417b2b2512ad2
title: Дізнайтеся про третинні кольори
challengeType: 0
forumTopicId: 301057
dashedName: learn-about-tertiary-colors
---

# --description--

Комп’ютерні монітори та екрани пристроїв створюють різні кольори, комбінуючи кількості червоного, зеленого та синього світла. Це відоме як адитивна модель кольору RGB у сучасній теорії кольору. Червоний (R), зелений (G) і синій (B) називають первинними кольорами. Змішування двох первинних кольорів створює вторинні кольори бірюзовий (G + B), пурпуровий (R + B) і жовтий (R + G). Ви бачили ці кольори у завданні Доповнюючі кольори. Ці вторинні кольори є доповнюючими до первинного кольору, який не використовувався при їх створенні, і розташовані навпроти цього первинного кольору на колірному колі. Наприклад, пурпуровий створюється з червоного та синього і є доповнюючим до зеленого.

Третинні кольори — це результат поєднання первинного кольору з одним із його сусідніх вторинних кольорів. Наприклад, у моделі кольору RGB червоний (первинний) і жовтий (вторинний) утворюють помаранчевий (третинний). Це додає шість кольорів до простого колірного кола, загалом дванадцять.

Існують різні методи вибору кольорів, які дають гармонійне поєднання в дизайні. Один із прикладів, що може використовувати третинні кольори, називається схемою розщепленого доповнюючого кольору. Ця схема починається з базового кольору, а потім поєднує його з двома кольорами, що сусідять із його доповнюючим. Ці три кольори забезпечують сильний візуальний контраст у дизайні, але є більш м’якими, ніж використання двох доповнюючих кольорів.

Ось три кольори, створені за схемою розщепленого доповнюючого кольору:

<table><thead><tr><th>Колір</th><th>Hex код</th></tr></thead><thead></thead><tbody><tr><td>помаранчевий</td><td>#FF7F00</td></tr><tr><td>бірюзовий</td><td>#00FFFF</td></tr><tr><td>малиновий</td><td>#FF007F</td></tr></tbody></table>

# --instructions--

Змініть властивість `background-color` класів `orange`, `cyan` і `raspberry` на відповідні кольори. Обов’язково використовуйте hex коди, а не назви кольорів.

# --hints--

Елемент `div` з класом `orange` повинен мати властивість `background-color` зі значенням помаранчевий.

```js
const orangeElement = document.querySelector('.orange');
const orangeStyle = window.getComputedStyle(orangeElement);
assert.equal(orangeStyle?.backgroundColor, 'rgb(255, 127, 0)');
```

Елемент `div` з класом `cyan` повинен мати властивість `background-color` зі значенням бірюзовий.

```js
const cyanElement = document.querySelector('.cyan');
const cyanStyle = window.getComputedStyle(cyanElement);
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

Елемент `div` з класом `raspberry` повинен мати властивість `background-color` зі значенням малиновий.

```js
const raspberryElement = document.querySelector('.raspberry');
const raspberryStyle = window.getComputedStyle(raspberryElement);
assert.equal(raspberryStyle?.backgroundColor, 'rgb(255, 0, 127)');
```

Усі значення `background-color` для класів кольорів повинні бути hex кодами, а не назвами кольорів.

```js
assert.notMatch(code, /background-color:\s(orange|cyan|raspberry)/);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .raspberry {
    background-color: #000000;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>

<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #FF7F00;
  }

  .cyan {
    background-color: #00FFFF;
  }

  .raspberry {
    background-color: #FF007F;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>
<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```
