---
id: 587d78a5367417b2b2512ad8
title: створення текстури шляхом додавання тонкого шаблону як фонового зображення
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQdwJC8'
forumTopicId: 301052
dashedName: create-texture-by-adding-a-subtle-pattern-as-a-background-image
---

# --description--

Один зі способів додати текстуру та інтерес до фону і зробити його більш помітним — це додати тонкий шаблон. Головне — баланс, адже не хочеться, щоб фон занадто виділявся і відволікав від переднього плану. Властивість `background` підтримує функцію `url()` для направлення на зображення обраної текстури або шаблону. Адреса посилання береться в лапки всередині круглих дужок.

# --instructions--

Використовуючи url `https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png`, встановіть властивість `background` для цілої сторінки за допомогою селектора `body`.

# --hints--

Ваш елемент `body` має мати властивість `background`, встановлену в `url()` з наведеним посиланням.

```js
assert.match(code,/background(-image)?:\s*?url\(\s*("|'|)https:\/\/cdn-media-1\.freecodecamp\.org\/imgr\/MJAkxbh\.png\2\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {

  }
</style>
```

# --solutions--

```html
<style>
  body {
    background: url("https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png");
  }
</style>
```
