---
id: 587d78a5367417b2b2512ada
title: Використання властивості CSS transform scale для масштабування елемента при наведенні
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLPJuM'
forumTopicId: 301077
dashedName: use-the-css-transform-scale-property-to-scale-an-element-on-hover
---

# --description--

Властивість `transform` має різноманітні функції, які дозволяють масштабувати, переміщувати, обертати, нахиляти тощо ваші елементи. Коли її використовують із псевдокласами, такими як `:hover`, що вказують на певний стан елемента, властивість `transform` легко додає інтерактивність вашим елементам.

Ось приклад, який масштабує елементи абзацу у 2.1 раза від їхнього початкового розміру, коли користувач наводить на них курсор:

```css
p:hover {
  transform: scale(2.1);
}
```

**Примітка:** Застосування трансформації до елемента `div` також вплине на всі дочірні елементи, що містяться в div.

# --instructions--

Додайте CSS-правило для стану `hover` елемента `div` і використайте властивість `transform`, щоб масштабувати елемент `div` у 1.1 раза від його початкового розміру, коли користувач наводить на нього курсор.

# --hints--

Розмір елемента `div` має масштабуватися у 1.1 раза, коли користувач наводить на нього курсор.

```js
assert.match(code, /div:hover\s*?{\s*?transform:\s*?scale\(1\.1\);/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }



</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }
  div:hover {
    transform: scale(1.1);
  }
</style>
<div></div>
```
