---
id: 587d781a367417b2b2512ab7
title: "Використовуйте тег strong, щоб зробити текст жирним"
challengeType: 0
videoUrl: 'https://scrimba.com/c/ceJNBSb'
forumTopicId: 301080
dashedName: use-the-strong-tag-to-make-text-bold
---

# --description--

Щоб зробити текст жирним, ви можете використати тег `strong`. Його часто застосовують, щоб привернути увагу до тексту та показати, що він важливий. За допомогою тегу `strong` браузер застосовує CSS `font-weight: bold;` до елемента.

# --instructions--

Обгорніть тегом `strong` текст `Stanford University`, який знаходиться всередині тегу `p` (крапку не включайте).

# --hints--

Ваш код має додати один тег `strong` до розмітки.

```js
assert.lengthOf(document.querySelectorAll('strong'),1);
```

Тег `strong` має бути всередині тегу `p`.

```js
const paragraphElement = document.querySelector('p');
const paragraphChildren = paragraphElement?.querySelectorAll(`:scope ${'strong'}`);
assert.lengthOf(paragraphChildren, 1);
```

Тег `strong` має обгорнути слова `Stanford University`.

```js
assert.match(document.querySelector('strong')?.textContent, /^Stanford University\.?$/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at <strong>Stanford University</strong>.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
