---
id: 587d78a7367417b2b2512ae1
title: створення руху за допомогою анімації CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7amZfW'
forumTopicId: 301051
dashedName: create-movement-using-css-animation
---

# --description--

Коли елементи мають задану `position`, таку як `fixed` або `relative`, властивості компенсації CSS `right`, `left`, `top` і `bottom` можна використовувати в правилах анімації для створення руху.

Як показано в прикладі нижче, ви можете зсунути елемент вниз, а потім вгору, встановивши властивість `top` ключового кадру `50%` на 50px, але залишивши її рівною 0px для першого (`0%`) і останнього (`100%`) ключового кадру.

```css
@keyframes rainbow {
  0% {
    background-color: blue;
    top: 0px;
  }
  50% {
    background-color: green;
    top: 50px;
  }
  100% {
    background-color: yellow;
    top: 0px;
  }
}
```

# --instructions--

Додайте горизонтальний рух до анімації `div`. Використовуючи властивість компенсації `left`, додайте її до правила `@keyframes`, щоб rainbow починався з 0 пікселів на `0%`, рухався до 25 пікселів на `50%` і закінчувався на -25 пікселів на `100%`. Не замінюйте властивість `top` в редакторі — анімація має містити і вертикальний, і горизонтальний рух.

# --hints--

Правило `@keyframes` для `0%` має використовувати компенсацію `left` 0px.

```js
assert.match(code,/[^50]0%\s*?{[\s\S]*?left:\s*?0px(;[\s\S]*?|\s*?)}/gi);
```

Правило `@keyframes` для `50%` має використовувати компенсацію `left` 25px.

```js
assert.match(code,/50%\s*?{[\s\S]*?left:\s*?25px(;[\s\S]*?|\s*?)}/gi);
```

Правило `@keyframes` для `100%` має використовувати компенсацію `left` -25px.

```js
assert.match(code,/100%\s*?{[\s\S]*?left:\s*?-25px(;[\s\S]*?|\s*?)}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;

    }
    50% {
      background-color: green;
      top: 50px;

    }
    100% {
      background-color: yellow;
      top: 0px;

    }
  }
</style>

<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;
      left: 0px;
    }
    50% {
      background-color: green;
      top: 50px;
      left: 25px;
    }
    100% {
      background-color: yellow;
      top: 0px;
      left: -25px;
    }
  }
</style>
<div id="rect"></div>
```
