---
id: 587d78a8367417b2b2512ae7
title: Змінити час анімації за допомогою ключових слів
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKvwCM'
forumTopicId: 301045
dashedName: change-animation-timing-with-keywords
---

# --description--

У CSS-анімаціях властивість `animation-timing-function` контролює, як швидко анімований елемент змінюється протягом усього часу анімації. Якщо анімація — це автомобіль, що рухається з точки А в точку Б за певний час (ваш `animation-duration`), то `animation-timing-function` визначає, як автомобіль прискорюється і гальмує під час поїздки.

Існує кілька заздалегідь визначених ключових слів для популярних варіантів. Наприклад, значення за замовчуванням — `ease`, яке починається повільно, прискорюється посередині, а потім знову сповільнюється в кінці. Інші варіанти включають `ease-out`, що швидко починається, а потім уповільнюється, `ease-in`, що повільно починається, а в кінці прискорюється, або `linear`, що застосовує постійну швидкість анімації протягом усього часу.

# --instructions--

Для елементів з ідентифікаторами `ball1` і `ball2` додайте властивість `animation-timing-function` для кожного, і встановіть `#ball1` у `linear`, а `#ball2` у `ease-out`. Зверніть увагу на різницю в русі елементів під час анімації, але вони закінчують разом, оскільки мають однаковий `animation-duration` у 2 секунди.

# --hints--

Значення властивості `animation-timing-function` для елемента з ідентифікатором `ball1` має бути `linear`.

```js
const ballOne =document.querySelector('#ball1'); 
const ballOneStyle = window.getComputedStyle(ballOne); 

const ball1Animation = __helpers.removeWhiteSpace(ballOneStyle?.animationTimingFunction);
assert.isTrue(ball1Animation == 'linear' || ball1Animation == 'cubic-bezier(0,0,1,1)');
```

Значення властивості `animation-timing-function` для елемента з ідентифікатором `ball2` має бути `ease-out`.

```js
const ballTwo = document.querySelector('#ball2'); 
const ballTwoStyle = window.getComputedStyle(ballTwo); 

const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)');
```

# --seed--

## --seed-contents--

```html
<style>

  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;

  }
  #ball2 {
    left:56%;

  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left:56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
