---
id: 587d781c367417b2b2512ac2
title: Встановіть розмір шрифту для кількох заголовків
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpQNT3'
forumTopicId: 301067
dashedName: set-the-font-size-for-multiple-heading-elements
---

# --description--

Властивість `font-size` використовується, щоб вказати, наскільки великий текст у певному елементі. Це правило можна застосовувати до кількох елементів, щоб створити візуальну узгодженість тексту на сторінці. У цьому завданні ви встановите значення для всіх тегів `h1` до `h6`, щоб збалансувати розміри заголовків.

# --instructions--

У тегах `style` встановіть властивість `font-size` для:

- тегу `h1` на 68px.
- тегу `h2` на 52px.
- тегу `h3` на 40px.
- тегу `h4` на 32px.
- тегу `h5` на 21px.
- тегу `h6` на 14px.

# --hints--

Ваш код має встановити властивість `font-size` для тегу `h1` на 68 пікселів.

```js
 const fontSizeOfh1 = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh1 ,'68px');
```

Ваш код має встановити властивість `font-size` для тегу `h2` на 52 пікселі.

```js
 const fontSizeOfh2 = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh2 ,'52px');
```

Ваш код має встановити властивість `font-size` для тегу `h3` на 40 пікселів.

```js
 const fontSizeOfh3 = new __helpers.CSSHelp(document).getStyle('h3')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh3 ,'40px');
```

Ваш код має встановити властивість `font-size` для тегу `h4` на 32 пікселі.

```js
 const fontSizeOfh4 = new __helpers.CSSHelp(document).getStyle('h4')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh4 , '32px');
```

Ваш код має встановити властивість `font-size` для тегу `h5` на 21 піксель.

```js
 const fontSizeOfh5 = new __helpers.CSSHelp(document).getStyle('h5')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh5 ,'21px');
```

Ваш код має встановити властивість `font-size` для тегу `h6` на 14 пікселів.

```js
 const fontSizeOfh6 = new __helpers.CSSHelp(document).getStyle('h6')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh6 , '14px');
```

# --seed--

## --seed-contents--

```html
<style>






</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
  }
  h2 {
    font-size: 52px;
  }
  h3 {
    font-size: 40px;
  }
  h4 {
    font-size: 32px;
  }
  h5 {
    font-size: 21px;
  }
  h6 {
    font-size: 14px;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
