---
id: 587d78a4367417b2b2512ad5
title: Налаштувати тон кольору
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDJvT7'
forumTopicId: 301038
dashedName: adjust-the-tone-of-a-color
---

# --description--

Опція `hsl()` у CSS також полегшує налаштування тону кольору. Змішування білого з чистим відтінком створює відтінок цього кольору, а додавання чорного зробить тінь. Альтернативно, тон утворюється додаванням сірого або одночасним створенням відтінку та тіні. Нагадаємо, що 's' і 'l' у `hsl()` означають насиченість і світлість відповідно. Відсоток насиченості змінює кількість сірого, а відсоток світлості визначає, скільки білого або чорного міститься в кольорі. Це корисно, коли у вас є базовий відтінок, який вам подобається, але потрібні різні варіації цього відтінку.

# --instructions--

Усі елементи мають `background-color` за замовчуванням зі значенням `transparent`. Наш елемент `nav` наразі має вигляд `cyan` фону, бо елемент позаду нього має `background-color`, встановлений на `cyan`. Додайте `background-color` до елемента `nav`, щоб він використовував той самий `cyan` відтінок, але мав `80%` насиченість і `25%` світлість для зміни тону та тіні.

# --hints--

Елемент `nav` має мати `background-color` відрегульованого бірюзового тону за допомогою властивості `hsl()`.

```js
// Computed style of hsl(180, 80%, 25%) results in rgb(13,115,115)
assert.equal(
  new __helpers.CSSHelp(document).getStyle('nav').getPropVal('background-color', true), 
  'rgb(13,115,115)'
)
```

# --seed--

## --seed-contents--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {

  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>

<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```

# --solutions--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {
    background-color: hsl(180, 80%, 25%);
  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```
