---
id: 587d78a6367417b2b2512adc
title: Використання властивості CSS transform skewY для нахилу елемента вздовж осі Y
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZ2uB'
forumTopicId: 301075
dashedName: use-the-css-transform-property-skewy-to-skew-an-element-along-the-y-axis
---

# --description--

Оскільки функція `skewX()` нахиляє вибраний елемент вздовж осі X на заданий кут, не дивно, що властивість `skewY()` нахиляє елемент вздовж осі Y (вертикальної).

# --instructions--

Нахиліть елемент з ідентифікатором `top` на -10 градусів вздовж осі Y, використовуючи властивість `transform`.

# --hints--

Елемент з ідентифікатором `top` має бути нахилений на -10 градусів вздовж осі Y.

```js
assert.match(code, /#top\s*?{\s*?.*?\s*?transform:\s*?skewY\(-10deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;

  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;
    transform: skewY(-10deg);
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
