---
id: 587d78a5367417b2b2512ad6
title: Створення поступового лінійного градієнта CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4dpt9'
forumTopicId: 301047
dashedName: create-a-gradual-css-linear-gradient
---

# --description--

Застосування кольору до HTML-елементів не обмежується одним плоским відтінком. CSS надає можливість використовувати колірні переходи, відомі як градієнти, на елементах. Це здійснюється через властивість `background` за допомогою функції `linear-gradient()`. Ось загальний синтаксис:

```css
background: linear-gradient(gradient_direction, color 1, color 2, color 3, ...);
```

Перший аргумент визначає напрямок, з якого починається колірний перехід — його можна вказати у градусах, де `90deg` створює горизонтальний градієнт (зліва направо), а `45deg` — діагональний градієнт (знизу зліва до верху справа). Наступні аргументи визначають порядок кольорів, які використовуються у градієнті.

Приклад:

```css
background: linear-gradient(90deg, red, yellow, rgb(204, 204, 255));
```

# --instructions--

Використайте `linear-gradient()` для властивості `div` елемента `background` і встановіть напрямок 35 градусів, щоб змінити колір від `#CCFFFF` до `#FFCCCC`.

# --hints--

Елемент `div` повинен мати властивість `linear-gradient` `background` з вказаним напрямком і кольорами.

```js
const divElement = document.querySelector('div');
const divStyle = window.getComputedStyle(divElement); 
assert.match(divStyle?.background, /linear-gradient\(35deg, rgb\(204, 255, 255\), rgb\(255, 204, 204\)\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;

  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;
    background: linear-gradient(35deg, #CCFFFF, #FFCCCC);
  }
</style>
<div></div>
```
