---
id: 587d781b367417b2b2512abc
title: Налаштуйте властивість колір фону тексту
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDqwA6'
forumTopicId: 301032
dashedName: adjust-the-background-color-property-of-text
---

# --description--

Замість того, щоб налаштовувати загальний фон або колір тексту, щоб зробити передній план легко читабельним, ви можете додати `background-color` до елемента, який містить текст, який хочете виділити. У цьому завданні використовується `rgba()` замість `hex` кодів або звичайного `rgb()`.

<blockquote>rgba означає:<br>  r = червоний<br>  g = зелений<br>  b = синій<br>  a = альфа/рівень прозорості</blockquote>

Значення RGB можуть варіюватися від 0 до 255. Значення альфа може бути від 1, що означає повну непрозорість або суцільний колір, до 0, що означає повну прозорість або прозорий. `rgba()` чудово підходить у цьому випадку, оскільки дозволяє налаштовувати прозорість. Це означає, що вам не потрібно повністю закривати фон.

Для цього завдання ви використаєте `background-color: rgba(45, 45, 45, 0.1)`. Він створює темно-сірий колір, який майже прозорий через низьке значення прозорості 0.1.

# --instructions--

Щоб текст виділявся більше, налаштуйте `background-color` елемента `h4` на задане значення `rgba()`.

Також для `h4` видаліть властивість `height` і додайте `padding` зі значенням 10px.

# --hints--

Ваш код має додати властивість `background-color` до елемента `h4` зі значенням `rgba(45, 45, 45, 0.1)`.

```js
assert(
  /(background-color|background):rgba\(45,45,45,0?\.1\)(;?}|;)/gi.test(
    code.replace(/\s/g, '')
  )
);
```

Ваш код має додати властивість `padding` до елемента `h4` і встановити її в 10 пікселів.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.equal(h4style?.paddingTop, "10px");
assert.equal(h4style?.paddingRight, "10px");
assert.equal(h4style?.paddingBottom, "10px");
assert.equal(h4style?.paddingLeft, "10px");
```

Властивість `height` на елементі `h4` має бути видалена.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.notEqual(h4style?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;


  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    padding: 10px;
    background-color: rgba(45, 45, 45, 0.1);
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
