---
id: 587d78a6367417b2b2512ade
title: Створення складнішої форми за допомогою CSS та HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpz4fr'
forumTopicId: 301050
dashedName: create-a-more-complex-shape-using-css-and-html
---

# --description--

Одна з найпопулярніших форм у світі — це форма серця, і в цьому завданні ви створите її, використовуючи чистий CSS. Але спершу потрібно зрозуміти псевдоелементи `::before` та `::after`. `::before` створює псевдоелемент, який є першим дочірнім елементом вибраного елемента; `::after` створює псевдоелемент, який є останнім дочірнім елементом вибраного елемента. У наведеному прикладі псевдоелемент `::before` використовується, щоб додати прямокутник до елемента з класом `heart`:

```css
.heart::before {
  content: "";
  background-color: yellow;
  border-radius: 25%;
  position: absolute;
  height: 50px;
  width: 70px;
  top: -50px;
  left: 5px;
}
```

Для того, щоб псевдоелементи `::before` та `::after` працювали правильно, вони повинні мати визначену властивість `content`. Ця властивість зазвичай використовується для додавання таких речей, як фото або текст, до вибраного елемента. Коли псевдоелементи `::before` та `::after` використовуються для створення форм, властивість `content` все одно потрібна, але вона встановлена в порожній рядок. У наведеному вище прикладі елемент з класом `heart` має псевдоелемент `::before`, який створює жовтий прямокутник з висотою та шириною `50px` та `70px` відповідно. Цей прямокутник має заокруглені кути завдяки 25% `border-radius` і розташований абсолютно на відстані `5px` зліва та `50px` вище за верхню межу елемента.

# --instructions--

Перетворіть елемент на екрані на форму серця. У селекторі `.heart::after` змініть властивість `background-color` на `pink`, а `border-radius` встановіть на 50%.

Далі націльтеся на елемент з класом `heart` (просто `heart`) і заповніть властивість `transform`. Використайте функцію `rotate()` зі значенням -45 градусів.

Нарешті, у селекторі `.heart::before` встановіть властивість `content` у порожній рядок.

# --hints--

Властивість `background-color` селектора `.heart::after` має бути `pink`.

```js
const heartAfter = code.match(/\.heart::after\s*{[\s\S]+?[^\}]}/g)?.[0];
assert.match(heartAfter, /({|;)\s*background-color\s*:\s*pink\s*(;|})/g);
```

Властивість `border-radius` селектора `.heart::after` має бути 50%.

```js
assert.lengthOf(code.match(/border-radius\s*?:\s*?50%/gi),2);
```

Властивість `transform` для класу `heart` має використовувати функцію `rotate()`, встановлену на -45 градусів.

```js
assert.match(code,/transform\s*?:\s*?rotate\(\s*?-45deg\s*?\)/gi);
```

Властивість `content` селектора `.heart::before` має бути порожнім рядком.

```js
assert.match(code,/\.heart::before\s*?{\s*?content\s*?:\s*?("|')\1\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: ;
  }
  .heart::after {
    background-color: blue;
    content: "";
    border-radius: 25%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: ;
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
  }
  .heart::after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: "";
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```
