---
id: 587d781e367417b2b2512aca
title: перемістити позиційований відносно елемент за допомогою CSS-компенсації
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bQEA4'
forumTopicId: 301065
dashedName: move-a-relatively-positioned-element-with-css-offsets
---

# --description--

CSS-компенсації `top` або `bottom`, а також `left` або `right` вказують браузеру, на скільки відступити елемент від його звичайного розташування у нормальному потоці документа. Ви компенсуєте елемент від певного місця, що переміщує елемент у протилежному напрямку від вказаного боку. Як ви бачили в останньому завданні, використання компенсації `top` перемістило `h2` вниз. Аналогічно, використання компенсації `left` переміщує елемент вправо.

# --instructions--

Використайте CSS-компенсації, щоб перемістити `h2` на 15 пікселів вправо і на 10 пікселів вгору.

# --hints--

Ваш код має використовувати CSS-компенсацію, щоб відносно позиціонувати `h2` на 10px вгору. Іншими словами, перемістіть його на 10px від `bottom` від його звичайного розташування.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.bottom, '10px');
```

Ваш код має використовувати CSS-компенсацію, щоб відносно позиціонувати `h2` на 15px вправо. Іншими словами, перемістіть його на 15px від `left` від його звичайного розташування.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.left , '15px');
```

# --seed--

## --seed-contents--

```html
<head>
<style>
  h2 {
    position: relative;


  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<head>
<style>
  h2 {
    position: relative;
    left: 15px;
    bottom: 10px;
  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
