---
id: 587d7791367417b2b2512ab4
title: Налаштуйте ширину елемента за допомогою властивості width
challengeType: 0
videoUrl: 'https://scrimba.com/c/cvVLPtN'
forumTopicId: 301039
dashedName: adjust-the-width-of-an-element-using-the-width-property
---

# --description--

Ви можете вказати ширину елемента за допомогою властивості `width` у CSS. Значення можуть бути задані у відносних одиницях довжини (наприклад, `em`), абсолютних одиницях довжини (наприклад, `px`) або у відсотках від батьківського елемента, що містить цей елемент. Ось приклад, який змінює ширину зображення на 220px:

```css
img {
  width: 220px;
}
```

# --instructions--

Додайте властивість `width` до всієї картки і встановіть її в абсолютне значення 245px. Використайте клас `fullCard` для вибору елемента.

# --hints--

Ваш код має змінити властивість `width` картки на 245 пікселів, використовуючи селектор класу `fullCard`.

```js
const fullCard = code.match(/\.fullCard\s*{[\s\S]+?[^}]}/g);
const fullCardElement = document.querySelector('.fullCard');
const fullCardStyle = window.getComputedStyle(fullCardElement); 
assert.match(code,/\.fullCard\s*{[\s\S]+?[^}]}/g); 

assert.match(fullCard?.[0],/width\s*:\s*245px\s*(;|})/gi);
assert.equal(fullCardStyle?.maxWidth, 'none');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
  }
  p {
    text-align: justify;
  }
  .links {
    margin-right: 20px;
    text-align: left;
  }
  .fullCard {

    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
  }
  p {
    text-align: justify;
  }
  .links {
    margin-right: 20px;
    text-align: left;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
