---
id: 587d78a9367417b2b2512ae9
title: Використання кривої Безьє для переміщення графіки
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bnRCK'
forumTopicId: 301071
dashedName: use-a-bezier-curve-to-move-a-graphic
---

# --description--

У попередньому завданні йшлося про ключове слово `ease-out`, яке описує анімацію, що спочатку прискорюється, а потім уповільнюється наприкінці анімації. Праворуч показано різницю між ключовим словом `ease-out` (для синього елемента) та ключовим словом `linear` (для червоного елемента). Подібні прогресії анімації до ключового слова `ease-out` можна досягти, використовуючи кастомну кубічну криву Безьє.

Загалом, зміна опорних точок `p1` і `p2` призводить до створення різних кривих Безьє, які контролюють, як анімація розвивається в часі. Ось приклад кривої Безьє з використанням значень, що імітують стиль `ease-out`:

```css
animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
```

Пам’ятайте, що всі функції `cubic-bezier` починаються з `p0` у точці (0, 0) і закінчуються `p3` у точці (1, 1). У цьому прикладі крива рухається швидше по осі Y (починається з 0, переходить до `p1` значення y 0, а потім до `p2` значення y 1), ніж по осі X (0 на початку, потім 0 для `p1`, і до 0.58 для `p2`). В результаті зміна анімованого елемента відбувається швидше, ніж час анімації для цього сегмента. Наприкінці кривої співвідношення між змінами значень x і y змінюється — значення y рухається від 1 до 1 (без змін), а значення x рухаються від 0.58 до 1, що робить зміни анімації повільнішими порівняно з тривалістю анімації.

# --instructions--

Щоб побачити ефект цієї кривої Безьє в дії, змініть властивість `animation-timing-function` елемента з ідентифікатором `red` на функцію `cubic-bezier` з відповідними значеннями x1, y1, x2, y2, встановленими як `0, 0, 0.58, 1`. Це змусить обидва елементи проходити анімацію подібним чином.

# --hints--

Значення властивості `animation-timing-function` елемента з ідентифікатором `red` має бути функцією `cubic-bezier` з відповідними значеннями x1, y1, x2, y2, встановленими як 0, 0, 0.58, 1.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.equal(
  redStyle?.animationTimingFunction, 'cubic-bezier(0, 0, 0.58, 1)'
);
```

Елемент з ідентифікатором `red` більше не повинен мати властивість `animation-timing-function` зі значенням `linear`.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.notEqual(redStyle?.animationTimingFunction, 'linear');
```

Значення властивості `animation-timing-function` для елемента з ідентифікатором `blue` не повинно змінюватися.

```js
const blueElement = document.querySelector('#blue');
const blueStyle = window.getComputedStyle( blueElement);
const blueBallAnimation = __helpers.removeWhiteSpace(
  blueStyle?.animationTimingFunction
);
assert.isTrue(
  blueBallAnimation == 'ease-out' ||
    blueBallAnimation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```

# --solutions--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```
