---
id: 587d78a3367417b2b2512ace
title: Відсунути елементи вліво або вправо за допомогою властивості float
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDqu2'
forumTopicId: 301066
dashedName: push-elements-left-or-right-with-the-float-property
---

# --description--

Наступний інструмент позиціювання насправді не використовує `position`, але встановлює властивість `float` елемента. Плаваючі елементи вилучаються з нормального потоку документа і відсуваються або до `left`, або до `right` свого батьківського елемента. Зазвичай його використовують разом із властивістю `width`, щоб вказати, скільки горизонтального простору потребує плаваючий елемент.

# --instructions--

Надана розмітка добре підходить для двоколоночного макета, де елементи `section` і `aside` розташовані поруч. Дайте елементу `#left` властивість `float` зі значенням `left`, а елементу `#right` — властивість `float` зі значенням `right`.

# --hints--

Елемент з ідентифікатором `left` повинен мати значення властивості `float` рівне `left`.

```js
const leftElement = document.querySelector('#left');
const leftStyle = window.getComputedStyle(leftElement);
assert.equal(leftStyle?.float, 'left');
```

Елемент з ідентифікатором `right` повинен мати значення властивості `float` рівне `right`.

```js
const rightElement = document.querySelector('#right');
const rightStyle = window.getComputedStyle(rightElement);
assert.equal(rightStyle?.float, 'right');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
    #left {

      width: 50%;
    }
    #right {

      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```

# --solutions--

```html
<head>
  <style>
    #left {
      float: left;
      width: 50%;
    }
    #right {
      float: right;
      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```
