---
id: 587d78a9367417b2b2512aea
title: Зробіть рух більш природним за допомогою кривої Безьє
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7akWUv'
forumTopicId: 301063
dashedName: make-motion-more-natural-using-a-bezier-curve
---

# --description--

У цьому завданні анімується елемент, щоб відтворити рух м’яча, який жонглюють. Попередні завдання охоплювали кубічні криві Безьє `linear` і `ease-out`, проте жодна з них не відображає рух жонглювання точно. Вам потрібно кастомізувати криву Безьє для цього.

`animation-timing-function` автоматично циклічно повторюється на кожному ключовому кадрі, коли `animation-iteration-count` встановлено в infinite. Оскільки в середині тривалості анімації (на `50%`) встановлено правило ключового кадру, це призводить до двох ідентичних прогресій анімації під час підйому і опускання м’яча.

Наступна кубічна крива Безьє імітує рух жонглювання:

```css
cubic-bezier(0.3, 0.4, 0.5, 1.6);
```

Зверніть увагу, що значення y2 більше за 1. Хоча кубічна крива Безьє спроєктована на систему координат 1 на 1 і може приймати x-значення лише від 0 до 1, y-значення можна встановлювати більшими за одиницю. Це створює рух відскоку, який ідеально підходить для імітації жонглювального м’яча.

# --instructions--

Змініть значення `animation-timing-function` елемента з ідентифікатором `green` на функцію `cubic-bezier` з відповідними значеннями x1, y1, x2, y2: 0.311, 0.441, 0.444, 1.649.

# --hints--

Значення властивості `animation-timing-function` для елемента з ідентифікатором `green` має бути функцією `cubic-bezier` з вказаними значеннями x1, y1, x2, y2.

```js
const greenElement = document.querySelector('#green');
 const greenStyle = window.getComputedStyle(greenElement);
assert.equal(greenStyle?.animationTimingFunction, 'cubic-bezier(0.311, 0.441, 0.444, 1.649)');
```

# --seed--

## --seed-contents--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.69, 0.1, 1, 0.1);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.311, 0.441, 0.444, 1.649);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```
