---
id: 587d78a4367417b2b2512ad4
title: Налаштуйте відтінок кольору
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp38TZ'
forumTopicId: 301036
dashedName: adjust-the-hue-of-a-color
---

# --description--

Кольори мають кілька характеристик, зокрема відтінок, насиченість і світлість. CSS3 представив функцію `hsl()` як альтернативний спосіб вибору кольору, безпосередньо вказуючи ці характеристики.

**Відтінок** — це те, що люди зазвичай розуміють під «кольором». Якщо уявити спектр кольорів, що починається з червоного зліва, проходить через зелений посередині і закінчується синім справа, відтінок — це місце кольору на цій лінії. У `hsl()` відтінок використовує концепцію колірного кола замість спектра, де кут кольору на колі задається значенням від 0 до 360.

**Насиченість** — це кількість сірого в кольорі. Повністю насичений колір не містить сірого, а мінімально насичений майже повністю сірий. Це подається у відсотках, де 100% — повна насиченість.

**Світлість** — це кількість білого або чорного в кольорі. Значення подається у відсотках від 0% (чорний) до 100% (білий), де 50% — нормальний колір.

Ось кілька прикладів використання `hsl()` з повністю насиченими кольорами нормальної світлості:

<table><thead><tr><th>Колір</th><th>HSL</th></tr></thead><tbody><tr><td>red</td><td>hsl(0, 100%, 50%)</td></tr><tr><td>yellow</td><td>hsl(60, 100%, 50%)</td></tr><tr><td>green</td><td>hsl(120, 100%, 50%)</td></tr><tr><td>cyan</td><td>hsl(180, 100%, 50%)</td></tr><tr><td>blue</td><td>hsl(240, 100%, 50%)</td></tr><tr><td>magenta</td><td>hsl(300, 100%, 50%)</td></tr></tbody></table>

# --instructions--

Змініть `background-color` кожного `div` елемента на основі назв класів (`green`, `cyan` або `blue`) за допомогою `hsl()`. Усі три мають мати повну насиченість і нормальну світлість.

# --hints--

Ваш код має використовувати функцію `hsl()` для оголошення кольору green.

```js
assert.match(code,/\.green\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Ваш код має використовувати функцію `hsl()` для оголошення кольору cyan.

```js
assert.match(code,/\.cyan\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Ваш код має використовувати функцію `hsl()` для оголошення кольору blue.

```js
assert.match(code,/\.blue\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Елемент `div` з класом `green` має мати `background-color` green.

```js
const greenElement = document.querySelector(".green");
const greenStyle = window.getComputedStyle(greenElement); 
assert.equal(greenStyle?.backgroundColor, 'rgb(0, 255, 0)');
```

Елемент `div` з класом `cyan` має мати `background-color` cyan.

```js
const cyanElement = document.querySelector(".cyan");
const cyanStyle = window.getComputedStyle(cyanElement); 
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

Елемент `div` з класом `blue` має мати `background-color` blue.

```js
const blueElement = document.querySelector(".blue");
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .blue {
    background-color: #000000;
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>

<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: hsl(120, 100%, 50%);
  }

  .cyan {
    background-color:   hsl(180, 100%, 50%);
  }

  .blue {
    background-color: hsl(240, 100%, 50%);
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```
