---
id: 58a7a6ebf9a6318348e2d5aa
title: змініть режим заповнення анімації
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDmcE'
forumTopicId: 301064
dashedName: modify-fill-mode-of-an-animation
---

# --description--

Це чудово, але поки що не працює правильно. Зверніть увагу, як анімація скидається після того, як минув `500ms`, через що кнопка повертається до початкового кольору. Ви хочете, щоб кнопка залишалася підсвіченою.

Цього можна досягти, встановивши властивість `animation-fill-mode` у значення `forwards`. `animation-fill-mode` визначає стиль, який застосовується до елемента після завершення анімації. Ви можете встановити це так:

```css
animation-fill-mode: forwards;
```

# --instructions--

Встановіть властивість `animation-fill-mode` у `button:hover` для `forwards`, щоб кнопка залишалася підсвіченою, коли користувач наводить на неї курсор.

# --hints--

`button:hover` має мати властивість `animation-fill-mode` зі значенням `forwards`.

```js
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-fill-mode\s*?:\s*?forwards\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-name\s*?:\s*?background-color\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-duration\s*?:\s*?500ms\s*?;[\s\S]*}/gi );
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    /* Only change code below this line */

    /* Only change code above this line */
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    animation-fill-mode: forwards;
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
