---
id: 5d822fd413a79914d39e98d3
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Annida quattro elementi `div` nel contenitore `.bb1`. Assegna loro le classi `bb1a`, `bb1b`, `bb1c` e `bb1d` in questo ordine. Questa build avrà quattro sezioni.

# --hints--

Dovresti creare quattro nuovi elementi `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

Dovresti assegnare a uno dei nuovi elementi `div` la classe `bb1a`.

```js
assert.exists(document.querySelector('div.bb1a'));
```

Dovresti assegnare a uno dei nuovi elementi `div` la classe `bb1b`.

```js
assert.exists(document.querySelector('div.bb1b'));
```

Dovresti assegnare a uno dei nuovi elementi `div` la classe `bb1c`.

```js
assert.exists(document.querySelector('div.bb1c'));
```

Dovresti assegnare a uno dei nuovi elementi `div` la classe `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

Dovresti posizionare i nuovi elementi `div` nell’ordine corretto.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

Dovresti posizionare i nuovi elementi `div` all’interno dell’elemento `.bb1`.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
    
```
