---
id: 5d822fd413a79914d39e98cc
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Benvenuto al progetto skyline delle variabili CSS! Inizia aggiungendo un elemento `link` che colleghi il tuo file `styles.css` all’interno dell’elemento `head`.

# --hints--

Il tuo codice dovrebbe avere un elemento `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Non dovresti modificare i tag `head` esistenti. Assicurati di non aver eliminato il tag di chiusura.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Dovresti avere un elemento void `link`.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Il tuo elemento `link` dovrebbe essere all’interno del tuo elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con valore `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con valore `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
