---
id: 5d822fd413a79914d39e98d4
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Give the parts of your building `width` and `height` properties with these values: `70%` and `10%` to `.bb1a`, `80%` and `10%` to `.bb1b`, `90%` and `10%` to `.bb1c`, and `100%` and `70%` to `.bb1d`. Remember that these percentages are relative to the parent and note that the heights will add up to 100% - vertically filling the container.

# --hints--

Dovresti usare un selettore di classe per lo stile di `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Dovresti dare a `.bb1a` una proprietà `width` del `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Dovresti dare `.bb1a` una proprietà `height` del `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Dovresti usare un selettore di classe per lo stile di `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Dovresti dare a `.bb1b` una proprietà `width` dell'`80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Dovresti dare `.bb1b` una proprietà `height` del `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Dovresti usare un selettore di classe per lo stile di `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Dovresti dare a `.bb1c` una proprietà `width` del `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Dovresti dare `.bb1c` una proprietà `height` del `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Dovresti usare un selettore di classe per lo stile di `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Dovresti dare a `.bb1d` una proprietà `width` del `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Dovresti dare `.bb1d` una proprietà `height` del `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

