---
id: 5d822fd413a79914d39e98dd
title: Passaggio 18
challengeType: 0
dashedName: step-18
---

# --description--

Gli edifici sono troppo distanziati. Avvicinali aggiungendo due elementi `div` vuoti in cima all'elemento `.background-buildings`, due altri in fondo a esso e un altro ancora tra `.bb3` e `.bb4`. Questi verranno aggiunti come elementi distribuiti uniformemente nel contenitore, spostando efficacemente gli edifici più vicino al centro.

# --hints--

Dovresti aggiungere due nuovi elementi `div` prima dell'elemento `.bb1`.

```js
const bBuildings = document.querySelector('.background-buildings')?.children;
assert.isTrue(![...bBuildings?.[0]?.classList]?.includes('bb1'));
assert.isTrue(![...bBuildings?.[1]?.classList]?.includes('bb1'));
```

Dovresti aggiungere un nuovo elemento `div` tra gli elementi `.bb3` e `.bb4`.

```js
assert.equal(document.querySelector('.bb3')?.nextElementSibling, document.querySelector('.bb4')?.previousElementSibling);
```

Dovresti aggiungere due nuovi elementi `div` dopo l'elemento `.bb4`.

```js
const bb4 = document.querySelector('.bb4');
assert.exists(bb4?.nextElementSibling);
assert.exists(bb4?.nextElementSibling?.nextElementSibling);
```

Dovresti aggiungere 5 nuovi elementi `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 14);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div class="bb4"></div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
}

.bb3 {
  width: 10%;
  height: 55%;
}

.bb4 {
  width: 11%;
  height: 58%;
}
    
```
