---
id: 68813c1b28ec3fd41838daf6
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將透過建置一個無障礙的筆記應用程式來練習使用 `aria-live` 和 `aria-label` 屬性。 

大部分的 HTML 和所有的 CSS 都已經為你提供。前幾個步驟將涉及添加剩餘的標記。 

如果你查看預覽，會有文字告訴你點擊卡片以進行編輯。但如果你嘗試這麼做，將不會生效。若要使該 `div` 元素可編輯，你可以像這樣使用 `contenteditable` 屬性：

```html
<div contenteditable="true"></div>
```

在開啟的 `div` 頁籤中，為你的 `contenteditable` 屬性添加並將其值設為 `"true"`。現在你應該能夠點擊該元素並編輯文字。

# --hints--

你的開啟 `div` 頁籤應該有 `contenteditable` 屬性。

```js
assert.isTrue(document.getElementById("note").hasAttribute("contenteditable"));
```

你的 `div` 元素應該將 `contenteditable` 屬性設為 `"true"`。

```js
assert.equal(document.getElementById("note").getAttribute("contenteditable"), "true");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    --fcc-editable-region--
    <div id="note" class="note">
    --fcc-editable-region--
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```
