---
id: 68814221fc4752e17b9696bd
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

當使用者點擊筆記、進行編輯，然後點擊筆記外部時，你需要一種方法來取得那個新編輯的版本。你可以透過使用 `blur` 事件來達成。當元素失去焦點時，這個事件會被觸發。以下是一個範例：

```js
element.addEventListener("blur", () => {
  // do something here
});
```

為你的 `noteEl` 變數附加一個 `addEventListener` 方法。第一個引數應該是 `"blur"` 事件，第二個引數應該是一個箭頭函式。在該箭頭函式的函式體內，使用 `const` 創建一個名為 `newContent` 的新變數，並將其指定為 `noteEl.innerHTML` 的值。 

這裡使用 `innerHTML` 來幫助保留註解的間距和格式。 

# --hints--

你應該將 `addEventListener` 方法附加到 `noteEl` 變數。

```js
assert.match(code, /noteEl\.addEventListener/);
```

你的 `addEventListener` 方法的第一個引數應該是 `"blur"` 事件。

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],/);
```

在箭頭函式（程式）本體內，你應該使用 `const` 來創建一個名為 `newContent` 的新變數，並將 `noteEl.innerHTML` 的值指定給它。

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],\s*\(\)\s*=>\s*{[^}]*const\s+newContent\s*=\s*noteEl\.innerHTML[^}]*}\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--
  
--fcc-editable-region--

window.addEventListener("DOMContentLoaded", () => {
  currentContent = noteEl.textContent;
});
```
