---
id: 688141e05e367de0e3ef7635
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

目前，`currentContent` 已初始化為空字串的值。但期望的結果是讓 `currentContent` 保持目前筆記文字的值。

在之前的課程中，你學會了如何這樣處理 `DOMContentLoaded` 事件：

```js
window.addEventListener("DOMContentLoaded", () => {
  // Do stuff here
});
```

當 HTML 文件中的所有內容都已載入並解析後，你將能安全地存取 `noteEl`。

首先，為 `window` 物件附加一個 `addEventListener` 方法。`addEventListener` 方法的第一個引數應該是 `"DOMContentLoaded"` 事件。第二個引數應該是一個箭頭函式。在箭頭函式的主體內，將 `noteEl.textContent` 的值指定給 `currentContent`。

# --hints--

你應該將 `addEventListener` 方法附加到 `window` 物件。

```js
assert.match(code, /window\.addEventListener/);
```

你的事件監聽器應該監聽 `"DOMContentLoaded"` 事件。

```js
assert.match(code, /window\.addEventListener\(['"]DOMContentLoaded['"]\,/);
```

在箭頭函式（arrow function）的主體內，將 `noteEl.textContent` 的值指定給 `currentContent`。

```js
const event = new Event("DOMContentLoaded");
window.dispatchEvent(event);

assert.equal(currentContent, noteEl.textContent);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English.

      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.

      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--

--fcc-editable-region--
```
