---
id: 66bb6a9c2dd58b73cd759034
title: 成績管理アプリを作成する
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `getAverage` という名前の関数があり、テストのスコアの配列をパラメータとして受け取り、平均スコアを返す必要があります。
1. `getGrade` という名前の関数があり、学生のスコアをパラメータとして受け取り、そのスコアに基づいた文字で表す成績を返す必要があります。スコアと対応する成績は以下の通りです：

    | スコア範囲 | 成績   |
    | ----------- | ------- |
    | `100`       | `"A+"`  |
    | `90 - 99`   | `"A"`   |
    | `80 - 89`   | `"B"`   |
    | `70 - 79`   | `"C"`   |
    | `60 - 69`   | `"D"`   |
    | `0 - 59`    | `"F"`   |

1. `hasPassingGrade` という名前の関数があり、スコアをパラメータとして受け取り、そのスコアが合格かどうかに応じて `true` または `false` を返す必要があります。
1. `hasPassingGrade` 関数は `getGrade` 関数を使って成績を取得し、その成績が合格かどうかを判定します。合格の成績は `"F"` 以外のすべてです。
1. `studentMsg` という名前の関数があり、スコアの配列と学生のスコアをパラメータとして受け取ります。この関数は以下の形式の文字列を返す必要があります：

    - 学生がコースに合格した場合は `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`。
    - 学生がコースに不合格だった場合は `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`。

    `average-goes-here` は全スコアの平均に、`grade-goes-here` は学生の成績に置き換えてください。
    平均スコアを取得するには `getAverage` を、学生の成績を取得するには `getGrade` を使ってください。

# --hints--

`getAverage` という名前の関数が必要です。

```js
assert.isFunction(getAverage);
```

`getAverage` 関数は数値を返す必要があります。

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` は `71.7` を返すべきです。

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` は `85.4` を返すべきです。

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` は `92.4` を返すべきです。

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` は `47.25` を返すべきです。

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` は `47.125` を返すべきです。

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

`getAverage` 関数はテストスコアの平均スコアを返す必要があります。

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

`getGrade` という名前の関数が必要です。

```js
assert.isFunction(getGrade);
```

`getGrade` 関数はスコアが `"A+"` の場合に `100` を返す必要があります。

```js
assert.strictEqual(getGrade(100), "A+");
```

`getGrade` 関数はスコアが `"A"` と `90` の間の場合に `99` を返す必要があります。

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

`getGrade` 関数はスコアが `"B"` と `80` の間の場合に `89` を返す必要があります。

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

`getGrade` 関数はスコアが `"C"` と `70` の間の場合に `79` を返す必要があります。

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

`getGrade` 関数はスコアが `"D"` と `60` の間の場合に `69` を返す必要があります。

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

`getGrade` 関数はスコアが `"F"` と `0` の間の場合に `59` を返す必要があります。

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


`hasPassingGrade` という名前の関数が必要です。

```js
assert.isFunction(hasPassingGrade);
```

`hasPassingGrade` 関数はブール値を返す必要があります。

```js
assert.isBoolean(hasPassingGrade(100));
```

`hasPassingGrade` 関数は成績が `true` の場合に `"A"` を返す必要があります。

```js
assert.isTrue(hasPassingGrade(100));
```

`hasPassingGrade` 関数は成績が `false` の場合に `"F"` を返す必要があります。

```js
assert.isFalse(hasPassingGrade(53));
```

`hasPassingGrade` 関数は合格の成績すべてに対して `true` を返す必要があります。

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

`studentMsg` という名前の関数が必要です。

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` は次のメッセージを返すべきです：`"Class average: 71.7. Your grade: F. You failed the course."`。

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` は次のメッセージを返すべきです：`"Class average: 50.8. Your grade: A+. You passed the course."`。

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` は次のメッセージを返すべきです：`"Class average: 48.25. Your grade: B. You passed the course."`。

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` は次のメッセージを返すべきです：`"Class average: 45.625. Your grade: C. You passed the course."`。

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

`studentMsg` 関数は学生のスコアとクラスの平均に基づいて正しいメッセージを返す必要があります。

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
