---
id: 587d781e367417b2b2512acb
title: Ein Element mit absoluter Positionierung am übergeordneten Element fixieren
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJ7c3'
forumTopicId: 301060
dashedName: lock-an-element-to-its-parent-with-absolute-positioning
---

# --description--

Die nächste Option für die CSS-Eigenschaft `position` ist `absolute`, die das Element relativ zu seinem übergeordneten Container fixiert. Im Gegensatz zur `relative`-Position wird das Element dabei aus dem normalen Dokumentfluss entfernt, sodass umliegende Elemente es ignorieren. Die CSS-Versatz-Eigenschaften (oben oder unten sowie links oder rechts) werden verwendet, um die Position anzupassen.

Ein Detail bei absoluter Positionierung ist, dass sie relativ zum nächsten *positionierten* Vorfahren fixiert wird. Wenn Sie vergessen, eine Positionsregel für das übergeordnete Element hinzuzufügen (dies geschieht typischerweise mit `position: relative;`), sucht der Browser weiter nach oben in der Kette und verwendet letztlich den `body`-Tag als Standard.

# --instructions--

Fixieren Sie das `#searchbar`-Element oben rechts an seinem `section`-Elternelement, indem Sie seine `position` auf `absolute` setzen. Geben Sie ihm `top`- und `right`-Versätze von jeweils 50 Pixeln.

# --hints--

Das `#searchbar`-Element sollte eine `position` mit dem Wert `absolute` haben.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.position, 'absolute');
```

Ihr Code sollte den CSS-Versatz `top` von 50 Pixeln am `#searchbar`-Element verwenden.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.top, '50px');
```

Ihr Code sollte den CSS-Versatz `right` von 50 Pixeln am `#searchbar`-Element verwenden.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.right, '50px');
```

# --seed--

## --seed-contents--

```html
<style>
  #searchbar {



  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```

# --solutions--

```html
<style>
  #searchbar {
    position: absolute;
    top: 50px;
    right: 50px;
  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```
