---
id: 587d78a6367417b2b2512adb
title: "Die CSS-Transform-Eigenschaft skewX verwenden, um ein Element entlang der X-Achse zu verzerren"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLP8Sr'
forumTopicId: 301074
dashedName: use-the-css-transform-property-skewx-to-skew-an-element-along-the-x-axis
---

# --description--

Die nächste Funktion der `transform`-Eigenschaft ist `skewX()`, die das ausgewählte Element entlang seiner X- (horizontalen) Achse um einen bestimmten Grad verzerrt.

Der folgende Code verzerrt das Absatz-Element um -32 Grad entlang der X-Achse.

```css
p {
  transform: skewX(-32deg);
}
```

# --instructions--

Verzerren Sie das Element mit der ID `bottom` um 24 Grad entlang der X-Achse, indem Sie die `transform`-Eigenschaft verwenden.

# --hints--

Das Element mit der ID `bottom` sollte um 24 Grad entlang seiner X-Achse verzerrt werden.

```js
assert.match(code, /#bottom\s*?{\s*?.*?\s*?transform:\s*?skewX\(24deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;

  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
