---
id: 587d78a8367417b2b2512ae4
title: Erstellen Sie einen CSS-Herzschlag mit unendlicher Animationsanzahl
challengeType: 0
videoUrl: 'https://scrimba.com/c/cDZpDUr'
forumTopicId: 301062
dashedName: make-a-css-heartbeat-using-an-infinite-animation-count
---

# --description--

Hier ist ein weiteres Beispiel für eine kontinuierliche Animation mit der `animation-iteration-count`-Eigenschaft, die das Herz verwendet, das Sie in einer vorherigen Aufgabe entworfen haben.

Die einsekündige Herzschlag-Animation besteht aus zwei animierten Teilen. Die `heart`-Elemente (einschließlich der `:before`- und `:after`-Teile) werden animiert, um ihre Größe mit der `transform`-Eigenschaft zu ändern, und der Hintergrund-`div` wird animiert, um seine Farbe mit der `background`-Eigenschaft zu ändern.

# --instructions--

Lassen Sie das Herz schlagen, indem Sie die `animation-iteration-count`-Eigenschaft sowohl für die `back`-Klasse als auch für die `heart`-Klasse hinzufügen und den Wert auf `infinite` setzen. Die `heart:before`- und `heart:after`-Selektoren benötigen keine Animations-Eigenschaften.

# --hints--

Die `animation-iteration-count`-Eigenschaft für die `heart`-Klasse sollte den Wert `infinite` haben.

```js
const heartElement = document.querySelector('.heart');
 const heartStyle = window.getComputedStyle(heartElement);
 assert.equal(heartStyle?.animationIterationCount, 'infinite');
```

Die `animation-iteration-count`-Eigenschaft für die `back`-Klasse sollte den Wert `infinite` haben.

```js
const backElement = document.querySelector('.back');
 const backStyle = window.getComputedStyle(backElement);
 assert.equal(backStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>
  .back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: white;
    animation-name: backdiv;
    animation-duration: 1s;

  }

  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
    animation-name: beat;
    animation-duration: 1s;

  }
  .heart:after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart:before {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }

  @keyframes backdiv {
    50% {
      background: #ffe6f2;
    }
  }

  @keyframes beat {
    0% {
      transform: scale(1) rotate(-45deg);
    }
    50% {
      transform: scale(0.6) rotate(-45deg);
    }
  }

</style>
<div class="back"></div>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: white;
    animation-name: backdiv;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
    animation-name: beat;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }
  .heart:after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart:before {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }

  @keyframes backdiv {
    50% {
      background: #ffe6f2;
    }
  }

  @keyframes beat {
    0% {
      transform: scale(1) rotate(-45deg);
    }
    50% {
      transform: scale(0.6) rotate(-45deg);
    }
  }
</style>
<div class="back"></div>
<div class="heart"></div>
```
