---
id: 587d78a3367417b2b2512ad0
title: Ein Element horizontal mit der Eigenschaft margin zentrieren
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJqU4'
forumTopicId: 301043
dashedName: center-an-element-horizontally-using-the-margin-property
---

# --description--

Eine weitere Positionierungstechnik besteht darin, ein Blockelement horizontal zu zentrieren. Eine Möglichkeit, dies zu tun, ist, seine `margin` auf den Wert auto zu setzen.

Diese Methode funktioniert auch für Bilder. Bilder sind standardmäßig Inline-Elemente, können aber zu Blockelementen geändert werden, wenn Sie die `display`-Eigenschaft auf `block` setzen.

# --instructions--

Zentrieren Sie das `div` auf der Seite, indem Sie eine `margin`-Eigenschaft mit dem Wert `auto` hinzufügen.

# --hints--

Das `div` sollte einen `margin` haben, der auf `auto` gesetzt ist.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('div')?.margin ,'auto');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;

  }
</style>
<div></div>
```

# --solutions--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;
    margin: auto;
  }
</style>
<div></div>
```
