---
id: 587d78a5367417b2b2512ad8
title: Textur erstellen durch Hinzufügen eines dezenten Musters als Hintergrundbild
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQdwJC8'
forumTopicId: 301052
dashedName: create-texture-by-adding-a-subtle-pattern-as-a-background-image
---

# --description--

Eine Möglichkeit, einem Hintergrund mehr Textur und Interesse zu verleihen und ihn hervorzuheben, besteht darin, ein dezentes Muster hinzuzufügen. Der Schlüssel ist das Gleichgewicht, da der Hintergrund nicht zu sehr hervorstechen und vom Vordergrund ablenken soll. Die `background`-Eigenschaft unterstützt die `url()`-Funktion, um mit einem Bild der gewählten Textur oder des Musters verknüpft zu werden. Die Linkadresse wird in Anführungszeichen innerhalb der runden Klammern gesetzt.

# --instructions--

Verwenden Sie die URL von `https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png`, um die `background` der gesamten Seite mit dem `body`-Selektor festzulegen.

# --hints--

Ihr `body`-Element sollte eine `background`-Eigenschaft mit einem `url()` und dem angegebenen Link haben.

```js
assert.match(code,/background(-image)?:\s*?url\(\s*("|'|)https:\/\/cdn-media-1\.freecodecamp\.org\/imgr\/MJAkxbh\.png\2\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {

  }
</style>
```

# --solutions--

```html
<style>
  body {
    background: url("https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png");
  }
</style>
```
