---
id: 587d78a8367417b2b2512ae3
title: Elemente kontinuierlich mit einer unendlichen Animationsanzahl animieren
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDVfq'
forumTopicId: 301041
dashedName: animate-elements-continually-using-an-infinite-animation-count
---

# --description--

Die vorherigen Aufgaben behandelten, wie man einige der Animations-Eigenschaften und die `@keyframes`-Regel verwendet. Eine weitere Animations-Eigenschaft ist die `animation-iteration-count`, mit der Sie steuern können, wie oft die Animation wiederholt werden soll. Hier ist ein Beispiel:

```css
animation-iteration-count: 3;
```

In diesem Fall stoppt die Animation nach 3 Durchläufen, aber es ist möglich, die Animation kontinuierlich laufen zu lassen, indem Sie diesen Wert auf `infinite` setzen.

# --instructions--

Um den Ball rechts kontinuierlich springen zu lassen, ändern Sie die `animation-iteration-count`-Eigenschaft in `infinite`.

# --hints--

Die `animation-iteration-count`-Eigenschaft sollte den Wert `infinite` haben.

```js
const ballElement = document.querySelector('#ball');
const ballStyle = window.getComputedStyle(ballElement); 
assert.equal(ballStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: 3;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```
