---
id: 587d781e367417b2b2512acc
title: Ein Element mit fester Positionierung am Browserfenster fixieren
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDNUR'
forumTopicId: 301061
dashedName: lock-an-element-to-the-browser-window-with-fixed-positioning
---

# --description--

Das nächste Layout-Schema, das CSS anbietet, ist die `fixed`-Position, eine Art der absoluten Positionierung, die ein Element relativ zum Browserfenster fixiert. Ähnlich wie bei der absoluten Positionierung wird sie mit den CSS-Versatz-Eigenschaften verwendet und entfernt das Element auch aus dem normalen Fluss des Dokuments. Andere Elemente „merken“ dann nicht mehr, wo es positioniert ist, was an anderer Stelle einige Layout-Anpassungen erforderlich machen kann.

Ein wesentlicher Unterschied zwischen der `fixed`- und der `absolute`-Position ist, dass sich ein Element mit fester Position nicht bewegt, wenn der Benutzer scrollt.

# --instructions--

Die Navigationsleiste im Code ist mit einer id von `navbar` gekennzeichnet. Ändern Sie ihre `position` in `fixed` und setzen Sie den Versatz auf 0 Pixel vom `top` und 0 Pixel vom `left`. Nachdem Sie den Code hinzugefügt haben, scrollen Sie im Vorschaufenster, um zu sehen, wie die Navigation an Ort und Stelle bleibt.

# --hints--

Das `#navbar`-Element sollte eine `position` mit dem Wert `fixed` haben.

```js
const navbarElement = document.querySelector('#navbar');
const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.position, 'fixed');
```

Ihr Code sollte den `top`-CSS-Versatz von 0 Pixeln am `#navbar`-Element verwenden.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
assert.equal(navbarStyle?.top, '0px');
```

Ihr Code sollte den `left`-CSS-Versatz von 0 Pixeln am `#navbar`-Element verwenden.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.left, '0px');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {



    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```

# --solutions--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {
    position: fixed;
    top: 0;
    left: 0;
    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```
