---
id: 587d78a7367417b2b2512ae1
title: Bewegung mit CSS-Animation erzeugen
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7amZfW'
forumTopicId: 301051
dashedName: create-movement-using-css-animation
---

# --description--

Wenn Elemente eine bestimmte `position` haben, wie `fixed` oder `relative`, können die CSS-Versatz-Eigenschaften `right`, `left`, `top` und `bottom` in Animationsregeln verwendet werden, um Bewegung zu erzeugen.

Wie im folgenden Beispiel gezeigt, können Sie das Element nach unten und dann nach oben verschieben, indem Sie die `top`-Eigenschaft des `50%`-Keyframes auf 50px setzen, während sie für den ersten (`0%`) und den letzten (`100%`) Keyframe auf 0px bleibt.

```css
@keyframes rainbow {
  0% {
    background-color: blue;
    top: 0px;
  }
  50% {
    background-color: green;
    top: 50px;
  }
  100% {
    background-color: yellow;
    top: 0px;
  }
}
```

# --instructions--

Fügen Sie der `div`-Animation eine horizontale Bewegung hinzu. Verwenden Sie die `left`-Versatz-Eigenschaft und ergänzen Sie die `@keyframes`-Regel so, dass der Regenbogen bei `0%` bei 0 Pixeln startet, sich bei `50%` auf 25 Pixel bewegt und bei `100%` bei -25 Pixeln endet. Ersetzen Sie die `top`-Eigenschaft im Editor nicht – die Animation soll sowohl vertikale als auch horizontale Bewegung enthalten.

# --hints--

Die `@keyframes`-Regel für `0%` sollte den `left`-Versatz von 0px verwenden.

```js
assert.match(code,/[^50]0%\s*?{[\s\S]*?left:\s*?0px(;[\s\S]*?|\s*?)}/gi);
```

Die `@keyframes`-Regel für `50%` sollte den `left`-Versatz von 25px verwenden.

```js
assert.match(code,/50%\s*?{[\s\S]*?left:\s*?25px(;[\s\S]*?|\s*?)}/gi);
```

Die `@keyframes`-Regel für `100%` sollte den `left`-Versatz von -25px verwenden.

```js
assert.match(code,/100%\s*?{[\s\S]*?left:\s*?-25px(;[\s\S]*?|\s*?)}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;

    }
    50% {
      background-color: green;
      top: 50px;

    }
    100% {
      background-color: yellow;
      top: 0px;

    }
  }
</style>

<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;
      left: 0px;
    }
    50% {
      background-color: green;
      top: 50px;
      left: 25px;
    }
    100% {
      background-color: yellow;
      top: 0px;
      left: -25px;
    }
  }
</style>
<div id="rect"></div>
```
