---
id: 587d78a8367417b2b2512ae5
title: Elemente mit variablen Geschwindigkeiten animieren
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ89WA4'
forumTopicId: 301040
dashedName: animate-elements-at-variable-rates
---

# --description--

Es gibt verschiedene Möglichkeiten, die Animationsgeschwindigkeiten ähnlich animierter Elemente zu verändern. Bisher wurde dies erreicht, indem eine `animation-iteration-count`-Eigenschaft angewendet und `@keyframes`-Regeln festgelegt wurden.

Zur Veranschaulichung besteht die Animation rechts aus zwei Sternen, die jeweils bei 20 % in der `@keyframes`-Regel in Größe und Deckkraft abnehmen, was die Funkelanimation erzeugt. Sie können die `@keyframes`-Regel für eines der Elemente ändern, sodass die Sterne mit unterschiedlichen Geschwindigkeiten funkeln.

# --instructions--

Ändern Sie die Animationsgeschwindigkeit für das Element mit dem Klassennamen `star-1`, indem Sie seine `@keyframes`-Regel auf 50 % setzen.

# --hints--

Die `@keyframes`-Regel für die `star-1`-Klasse sollte 50 % betragen.

```js
assert.match(code,/twinkle-1\s*?{\s*?50%/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-name: twinkle-1;
    animation-duration: 1s;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-name: twinkle-2;
    animation-duration: 1s;
  }

  @keyframes twinkle-1 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  @keyframes twinkle-2 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>

<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
```

# --solutions--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-name: twinkle-1;
    animation-duration: 1s;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-name: twinkle-2;
    animation-duration: 1s;
  }

  @keyframes twinkle-1 {
    50% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  @keyframes twinkle-2 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>
<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
```
