---
id: 587d781d367417b2b2512ac8
title: Den Hover-Zustand eines Anker-Tags anpassen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakRGcm'
forumTopicId: 301035
dashedName: adjust-the-hover-state-of-an-anchor-tag
---

# --description--

Diese Aufgabe behandelt die Verwendung von Pseudoklassen. Eine Pseudoklasse ist ein Schlüsselwort, das zu Selektoren hinzugefügt werden kann, um einen bestimmten Zustand des Elements auszuwählen.

Zum Beispiel kann das Styling eines Anker-Tags für seinen Hover-Zustand mit dem `:hover` Pseudoklassen-Selektor geändert werden. Hier ist das CSS, um die `color` des Anker-Tags während seines Hover-Zustands rot zu färben:

```css
a:hover {
  color: red;
}
```

# --instructions--

Der Code-Editor enthält eine CSS-Regel, die alle `a`-Tags schwarz stylt. Fügen Sie eine Regel hinzu, sodass der `a` blau wird, wenn der Benutzer mit der Maus über das `color`-Tag fährt.

# --hints--

Das Anker-Tag `color` soll schwarz bleiben, fügen Sie nur CSS-Regeln für den `:hover`-Zustand hinzu.

```js
const anchorElement = document.querySelector("a"); 
const anchorStyle = window.getComputedStyle(anchorElement);
assert.equal(anchorStyle?.color, 'rgb(0, 0, 0)');
```

Das Anker-Tag sollte beim Hover eine `color` in Blau haben.

```js
assert.match(code,
    /a:hover\s*?{\s*?color:\s*?(blue|rgba\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?,\s*?1\s*?\)|#00F|rgb\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?\))\s*?;\s*?}/gi
  );
```

# --seed--

## --seed-contents--

```html
<style>
  a {
    color: #000;
  }



</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```

# --solutions--

```html
<style>
  a {
    color: #000;
  }
  a:hover {
    color: rgba(0,0,255,1);
  }
</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```
