---
id: 587d781c367417b2b2512ac2
title: Die Schriftgröße für mehrere Überschriftselemente festlegen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpQNT3'
forumTopicId: 301067
dashedName: set-the-font-size-for-multiple-heading-elements
---

# --description--

Die Eigenschaft `font-size` wird verwendet, um anzugeben, wie groß der Text in einem bestimmten Element ist. Diese Regel kann für mehrere Elemente verwendet werden, um eine visuelle Konsistenz des Texts auf einer Seite zu schaffen. In dieser Aufgabe legen Sie die Werte für alle `h1` bis `h6` Tags fest, um die Größen der Überschriften auszubalancieren.

# --instructions--

Legen Sie in den `style` Tags die `font-size` der:

- `h1` Tags auf 68px fest.
- `h2` Tags auf 52px fest.
- `h3` Tags auf 40px fest.
- `h4` Tags auf 32px fest.
- `h5` Tags auf 21px fest.
- `h6` Tags auf 14px fest.

# --hints--

Ihr Code sollte die Eigenschaft `font-size` für den `h1` Tag auf 68 Pixel setzen.

```js
 const fontSizeOfh1 = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh1 ,'68px');
```

Ihr Code sollte die Eigenschaft `font-size` für den `h2` Tag auf 52 Pixel setzen.

```js
 const fontSizeOfh2 = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh2 ,'52px');
```

Ihr Code sollte die Eigenschaft `font-size` für den `h3` Tag auf 40 Pixel setzen.

```js
 const fontSizeOfh3 = new __helpers.CSSHelp(document).getStyle('h3')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh3 ,'40px');
```

Ihr Code sollte die Eigenschaft `font-size` für den `h4` Tag auf 32 Pixel setzen.

```js
 const fontSizeOfh4 = new __helpers.CSSHelp(document).getStyle('h4')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh4 , '32px');
```

Ihr Code sollte die Eigenschaft `font-size` für den `h5` Tag auf 21 Pixel setzen.

```js
 const fontSizeOfh5 = new __helpers.CSSHelp(document).getStyle('h5')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh5 ,'21px');
```

Ihr Code sollte die Eigenschaft `font-size` für den `h6` Tag auf 14 Pixel setzen.

```js
 const fontSizeOfh6 = new __helpers.CSSHelp(document).getStyle('h6')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh6 , '14px');
```

# --seed--

## --seed-contents--

```html
<style>






</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
  }
  h2 {
    font-size: 52px;
  }
  h3 {
    font-size: 40px;
  }
  h4 {
    font-size: 32px;
  }
  h5 {
    font-size: 21px;
  }
  h6 {
    font-size: 14px;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
