---
id: 587d78a5367417b2b2512ada
title: "Verwenden Sie die CSS-Transform-Eigenschaft scale, um ein Element beim Hover zu skalieren"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLPJuM'
forumTopicId: 301077
dashedName: use-the-css-transform-scale-property-to-scale-an-element-on-hover
---

# --description--

Die `transform`-Eigenschaft hat verschiedene Funktionen, mit denen Sie Ihre Elemente skalieren, verschieben, drehen, schräg stellen usw. können. Wenn sie mit Pseudoklassen wie `:hover` verwendet wird, die einen bestimmten Zustand eines Elements angeben, kann die `transform`-Eigenschaft Ihren Elementen ganz einfach Interaktivität verleihen.

Hier ist ein Beispiel, um die Absatz-Elemente auf das 2,1-fache ihrer ursprünglichen Größe zu skalieren, wenn ein Benutzer mit der Maus darüber fährt:

```css
p:hover {
  transform: scale(2.1);
}
```

**Hinweis:** Das Anwenden einer Transformation auf ein `div`-Element wirkt sich auch auf alle darin enthaltenen Kindelemente im div aus.

# --instructions--

Fügen Sie eine CSS-Regel für den `hover`-Zustand des `div` hinzu und verwenden Sie die `transform`-Eigenschaft, um das `div`-Element beim Überfahren mit der Maus auf das 1,1-fache seiner ursprünglichen Größe zu skalieren.

# --hints--

Die Größe des `div`-Elements sollte sich beim Überfahren mit der Maus auf das 1,1-fache skalieren.

```js
assert.match(code, /div:hover\s*?{\s*?transform:\s*?scale\(1\.1\);/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }



</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }
  div:hover {
    transform: scale(1.1);
  }
</style>
<div></div>
```
