---
id: 587d78a3367417b2b2512ace
title: Elemente mit der Eigenschaft float nach links oder rechts schieben
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDqu2'
forumTopicId: 301066
dashedName: push-elements-left-or-right-with-the-float-property
---

# --description--

Das nächste Positionierungswerkzeug verwendet nicht tatsächlich `position`, sondern setzt die `float`-Eigenschaft eines Elements. Gefloatete Elemente werden aus dem normalen Fluss eines Dokuments entfernt und entweder an den `left` oder `right` ihres übergeordneten Elements geschoben. Es wird häufig zusammen mit der `width`-Eigenschaft verwendet, um anzugeben, wie viel horizontalen Platz das gefloatete Element benötigt.

# --instructions--

Die gegebene Auszeichnung würde sich gut als zweispaltiges Layout eignen, bei dem die `section`- und `aside`-Elemente nebeneinander stehen. Geben Sie dem `#left`-Element eine `float` von `left` und dem `#right`-Element eine `float` von `right`.

# --hints--

Das Element mit der ID `left` sollte einen `float`-Wert von `left` haben.

```js
const leftElement = document.querySelector('#left');
const leftStyle = window.getComputedStyle(leftElement);
assert.equal(leftStyle?.float, 'left');
```

Das Element mit der ID `right` sollte einen `float`-Wert von `right` haben.

```js
const rightElement = document.querySelector('#right');
const rightStyle = window.getComputedStyle(rightElement);
assert.equal(rightStyle?.float, 'right');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
    #left {

      width: 50%;
    }
    #right {

      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```

# --solutions--

```html
<head>
  <style>
    #left {
      float: left;
      width: 50%;
    }
    #right {
      float: right;
      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```
