---
id: 587d78a4367417b2b2512ad4
title: Farbton einer Farbe anpassen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp38TZ'
forumTopicId: 301036
dashedName: adjust-the-hue-of-a-color
---

# --description--

Farben haben mehrere Eigenschaften, darunter Farbton, Sättigung und Helligkeit. CSS3 führte die `hsl()` Funktion als alternative Möglichkeit ein, eine Farbe auszuwählen, indem diese Eigenschaften direkt angegeben werden.

**Farbton** ist das, was die meisten Menschen allgemein als „Farbe“ verstehen. Wenn Sie sich ein Farbspektrum vorstellen, das links mit Rot beginnt, in der Mitte Grün und rechts Blau zeigt, dann beschreibt der Farbton, wo eine Farbe auf dieser Linie liegt. In `hsl()` wird anstelle des Spektrums das Konzept eines Farbrads verwendet, bei dem der Winkel der Farbe auf dem Kreis als Wert zwischen 0 und 360 angegeben wird.

**Sättigung** ist der Anteil von Grau in einer Farbe. Eine vollständig gesättigte Farbe enthält kein Grau, und eine minimal gesättigte Farbe ist fast vollständig grau. Dies wird als Prozentsatz angegeben, wobei 100 % vollständig gesättigt bedeutet.

**Helligkeit** ist der Anteil von Weiß oder Schwarz in einer Farbe. Es wird ein Prozentsatz von 0 % (Schwarz) bis 100 % (Weiß) angegeben, wobei 50 % die normale Farbe darstellt.

Hier sind einige Beispiele für die Verwendung von `hsl()` mit voll gesättigten Farben und normaler Helligkeit:

<table><thead><tr><th>Farbe</th><th>HSL</th></tr></thead><tbody><tr><td>rot</td><td>hsl(0, 100%, 50%)</td></tr><tr><td>gelb</td><td>hsl(60, 100%, 50%)</td></tr><tr><td>grün</td><td>hsl(120, 100%, 50%)</td></tr><tr><td>cyan</td><td>hsl(180, 100%, 50%)</td></tr><tr><td>blau</td><td>hsl(240, 100%, 50%)</td></tr><tr><td>magenta</td><td>hsl(300, 100%, 50%)</td></tr></tbody></table>

# --instructions--

Ändern Sie den `background-color` jedes `div` Elements basierend auf den Klassennamen (`green`, `cyan` oder `blue`) mit `hsl()`. Alle drei sollten volle Sättigung und normale Helligkeit haben.

# --hints--

Ihr Code sollte die `hsl()` Funktion verwenden, um die Farbe grün zu deklarieren.

```js
assert.match(code,/\.green\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Ihr Code sollte die `hsl()` Funktion verwenden, um die Farbe cyan zu deklarieren.

```js
assert.match(code,/\.cyan\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Ihr Code sollte die `hsl()` Funktion verwenden, um die Farbe blau zu deklarieren.

```js
assert.match(code,/\.blue\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Das `div` Element mit der Klasse `green` sollte eine `background-color` von grün haben.

```js
const greenElement = document.querySelector(".green");
const greenStyle = window.getComputedStyle(greenElement); 
assert.equal(greenStyle?.backgroundColor, 'rgb(0, 255, 0)');
```

Das `div` Element mit der Klasse `cyan` sollte eine `background-color` von cyan haben.

```js
const cyanElement = document.querySelector(".cyan");
const cyanStyle = window.getComputedStyle(cyanElement); 
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

Das `div` Element mit der Klasse `blue` sollte eine `background-color` von blau haben.

```js
const blueElement = document.querySelector(".blue");
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .blue {
    background-color: #000000;
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>

<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: hsl(120, 100%, 50%);
  }

  .cyan {
    background-color:   hsl(180, 100%, 50%);
  }

  .blue {
    background-color: hsl(240, 100%, 50%);
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```
