---
id: 587d78a5367417b2b2512ad6
title: Einen allmählichen CSS-Linearverlauf erstellen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4dpt9'
forumTopicId: 301047
dashedName: create-a-gradual-css-linear-gradient
---

# --description--

Das Anwenden einer Farbe auf HTML-Elemente ist nicht auf einen einheitlichen Farbton beschränkt. CSS bietet die Möglichkeit, Farbverläufe, auch bekannt als Gradients, auf Elemente anzuwenden. Dies erfolgt über die `background` Eigenschafts-`linear-gradient()`-Funktion. Hier ist die allgemeine Syntax:

```css
background: linear-gradient(gradient_direction, color 1, color 2, color 3, ...);
```

Das erste Argument gibt die Richtung an, aus der der Farbverlauf startet – es kann als Grad angegeben werden, wobei `90deg` einen horizontalen Verlauf (von links nach rechts) und `45deg` einen diagonalen Verlauf (von unten links nach oben rechts) erzeugt. Die folgenden Argumente geben die Reihenfolge der im Verlauf verwendeten Farben an.

Beispiel:

```css
background: linear-gradient(90deg, red, yellow, rgb(204, 204, 255));
```

# --instructions--

Verwenden Sie einen `linear-gradient()` für die `div`-Eigenschaft des Elements `background` und setzen Sie ihn auf eine Richtung von 35 Grad, um die Farbe von `#CCFFFF` zu `#FFCCCC` zu ändern.

# --hints--

Das `div`-Element sollte eine `linear-gradient` `background` mit der angegebenen Richtung und den Farben haben.

```js
const divElement = document.querySelector('div');
const divStyle = window.getComputedStyle(divElement); 
assert.match(divStyle?.background, /linear-gradient\(35deg, rgb\(204, 255, 255\), rgb\(255, 204, 204\)\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;

  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;
    background: linear-gradient(35deg, #CCFFFF, #FFCCCC);
  }
</style>
<div></div>
```
