---
id: 587d7791367417b2b2512ab4
title: Die Breite eines Elements mit der Eigenschaft width anpassen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cvVLPtN'
forumTopicId: 301039
dashedName: adjust-the-width-of-an-element-using-the-width-property
---

# --description--

Sie können die Breite eines Elements mit der Eigenschaft `width` in CSS festlegen. Werte können in relativen Längeneinheiten (wie `em`), absoluten Längeneinheiten (wie `px`) oder als Prozentsatz des übergeordneten Elements angegeben werden. Hier ist ein Beispiel, das die Breite eines Bildes auf 220px ändert:

```css
img {
  width: 220px;
}
```

# --instructions--

Fügen Sie der gesamten Karte eine Eigenschaft `width` hinzu und setzen Sie sie auf einen absoluten Wert von 245px. Verwenden Sie die Klasse `fullCard`, um das Element auszuwählen.

# --hints--

Ihr Code sollte die Eigenschaft `width` der Karte auf 245 Pixel ändern, indem Sie den Klassenselektor `fullCard` verwenden.

```js
const fullCard = code.match(/\.fullCard\s*{[\s\S]+?[^}]}/g);
const fullCardElement = document.querySelector('.fullCard');
const fullCardStyle = window.getComputedStyle(fullCardElement); 
assert.match(code,/\.fullCard\s*{[\s\S]+?[^}]}/g); 

assert.match(fullCard?.[0],/width\s*:\s*245px\s*(;|})/gi);
assert.equal(fullCardStyle?.maxWidth, 'none');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
  }
  p {
    text-align: justify;
  }
  .links {
    margin-right: 20px;
    text-align: left;
  }
  .fullCard {

    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
  }
  p {
    text-align: justify;
  }
  .links {
    margin-right: 20px;
    text-align: left;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
