---
id: 587d78a6367417b2b2512ade
title: Erstellen Sie eine komplexere Form mit CSS und HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpz4fr'
forumTopicId: 301050
dashedName: create-a-more-complex-shape-using-css-and-html
---

# --description--

Eine der beliebtesten Formen weltweit ist die Herzform, und in dieser Aufgabe erstellen Sie eine solche mit reinem CSS. Zuerst müssen Sie jedoch die `::before` und `::after` Pseudoelemente verstehen. `::before` erzeugt ein Pseudoelement, das das erste Kind des ausgewählten Elements ist; `::after` erzeugt ein Pseudoelement, das das letzte Kind des ausgewählten Elements ist. Im folgenden Beispiel wird ein `::before` Pseudoelement verwendet, um einem Element mit der Klasse `heart` ein Rechteck hinzuzufügen:

```css
.heart::before {
  content: "";
  background-color: yellow;
  border-radius: 25%;
  position: absolute;
  height: 50px;
  width: 70px;
  top: -50px;
  left: 5px;
}
```

Damit die `::before` und `::after` Pseudoelemente richtig funktionieren, müssen sie eine definierte `content` Eigenschaft haben. Diese Eigenschaft wird normalerweise verwendet, um Dinge wie ein Foto oder Text zum ausgewählten Element hinzuzufügen. Wenn die `::before` und `::after` Pseudoelemente verwendet werden, um Formen zu erstellen, ist die `content` Eigenschaft weiterhin erforderlich, wird aber auf einen leeren String gesetzt. Im obigen Beispiel hat das Element mit der Klasse `heart` ein `::before` Pseudoelement, das ein gelbes Rechteck mit einer Höhe und Breite von `50px` bzw. `70px` erzeugt. Dieses Rechteck hat abgerundete Ecken aufgrund seines 25% `border-radius` und ist absolut positioniert `5px` vom linken Rand und `50px` oberhalb der Oberkante des Elements.

# --instructions--

Verwandeln Sie das Element auf dem Bildschirm in ein Herz. Ändern Sie im `.heart::after` Selektor die `background-color` in `pink` und die `border-radius` auf 50%.

Zielen Sie als Nächstes auf das Element mit der Klasse `heart` (nur `heart`) und füllen Sie die `transform` Eigenschaft aus. Verwenden Sie die `rotate()` Funktion mit -45 Grad.

Setzen Sie schließlich im `.heart::before` Selektor dessen `content` Eigenschaft auf einen leeren String.

# --hints--

Die `background-color` Eigenschaft des `.heart::after` Selektors sollte `pink` sein.

```js
const heartAfter = code.match(/\.heart::after\s*{[\s\S]+?[^\}]}/g)?.[0];
assert.match(heartAfter, /({|;)\s*background-color\s*:\s*pink\s*(;|})/g);
```

Die `border-radius` des `.heart::after` Selektors sollte 50% sein.

```js
assert.lengthOf(code.match(/border-radius\s*?:\s*?50%/gi),2);
```

Die `transform` Eigenschaft für die `heart` Klasse sollte eine `rotate()` Funktion mit -45 Grad verwenden.

```js
assert.match(code,/transform\s*?:\s*?rotate\(\s*?-45deg\s*?\)/gi);
```

Die `content` des `.heart::before` Selektors sollte ein leerer String sein.

```js
assert.match(code,/\.heart::before\s*?{\s*?content\s*?:\s*?("|')\1\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: ;
  }
  .heart::after {
    background-color: blue;
    content: "";
    border-radius: 25%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: ;
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
  }
  .heart::after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: "";
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```
