---
id: 587d781e367417b2b2512aca
title: Ein relativ positioniertes Element mit CSS-Versätzen verschieben
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bQEA4'
forumTopicId: 301065
dashedName: move-a-relatively-positioned-element-with-css-offsets
---

# --description--

Die CSS-Versätze von `top` oder `bottom` und `left` oder `right` geben dem Browser an, wie weit ein Element relativ zu seiner normalen Position im Dokumentfluss verschoben werden soll. Sie verschieben ein Element von einem bestimmten Punkt weg, wodurch sich das Element von der referenzierten Seite entfernt (effektiv in die entgegengesetzte Richtung). Wie Sie in der letzten Aufgabe gesehen haben, verschob der `top`-Versatz das `h2` nach unten. Ebenso verschiebt ein `left`-Versatz ein Element nach rechts.

# --instructions--

Verwenden Sie CSS-Versätze, um das `h2` 15 Pixel nach rechts und 10 Pixel nach oben zu verschieben.

# --hints--

Ihr Code sollte einen CSS-Versatz verwenden, um das `h2` relativ 10px nach oben zu positionieren. Anders gesagt, verschieben Sie es 10px weg von der `bottom` seiner normalen Position.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.bottom, '10px');
```

Ihr Code sollte einen CSS-Versatz verwenden, um das `h2` relativ 15px nach rechts zu positionieren. Anders gesagt, verschieben Sie es 15px weg von der `left` seiner normalen Position.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.left , '15px');
```

# --seed--

## --seed-contents--

```html
<head>
<style>
  h2 {
    position: relative;


  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<head>
<style>
  h2 {
    position: relative;
    left: 15px;
    bottom: 10px;
  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
