---
id: 587d781b367417b2b2512abc
title: Die background-color-Eigenschaft von Text anpassen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDqwA6'
forumTopicId: 301032
dashedName: adjust-the-background-color-property-of-text
---

# --description--

Anstatt den gesamten Hintergrund oder die Textfarbe anzupassen, um den Vordergrund gut lesbar zu machen, können Sie dem Element, das den hervorzuhebenden Text enthält, eine `background-color` hinzufügen. Diese Aufgabe verwendet `rgba()` anstelle von `hex` Codes oder normalen `rgb()`.

<blockquote>rgba steht für:<br>  r = rot<br>  g = grün<br>  b = blau<br>  a = alpha/Transparenzgrad</blockquote>

Die RGB-Werte können von 0 bis 255 reichen. Der Alpha-Wert kann von 1, was vollständig undurchsichtig oder eine feste Farbe bedeutet, bis 0, was vollständig transparent oder durchsichtig ist, reichen. `rgba()` ist hier ideal, da es Ihnen erlaubt, die Transparenz anzupassen. Das bedeutet, Sie müssen den Hintergrund nicht komplett blockieren.

Für diese Aufgabe verwenden Sie `background-color: rgba(45, 45, 45, 0.1)`. Es erzeugt eine dunkelgraue Farbe, die aufgrund des niedrigen Transparenzwerts von 0,1 fast durchsichtig ist.

# --instructions--

Um den Text besser hervorzuheben, passen Sie die `background-color` des `h4`-Elements auf den angegebenen `rgba()`-Wert an.

Entfernen Sie außerdem für das `h4` die `height`-Eigenschaft und fügen Sie eine `padding` von 10px hinzu.

# --hints--

Ihr Code sollte dem `background-color`-Element eine `h4`-Eigenschaft mit dem Wert `rgba(45, 45, 45, 0.1)` hinzufügen.

```js
assert(
  /(background-color|background):rgba\(45,45,45,0?\.1\)(;?}|;)/gi.test(
    code.replace(/\s/g, '')
  )
);
```

Ihr Code sollte dem `padding`-Element eine `h4`-Eigenschaft hinzufügen und diese auf 10 Pixel setzen.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.equal(h4style?.paddingTop, "10px");
assert.equal(h4style?.paddingRight, "10px");
assert.equal(h4style?.paddingBottom, "10px");
assert.equal(h4style?.paddingLeft, "10px");
```

Die `height`-Eigenschaft am `h4`-Element sollte entfernt werden.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.notEqual(h4style?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;


  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    padding: 10px;
    background-color: rgba(45, 45, 45, 0.1);
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
