---
id: 587d78a4367417b2b2512ad5
title: Den Ton einer Farbe anpassen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDJvT7'
forumTopicId: 301038
dashedName: adjust-the-tone-of-a-color
---

# --description--

Die `hsl()`-Option in CSS erleichtert es auch, den Ton einer Farbe anzupassen. Das Mischen von Weiß mit einem reinen Farbton erzeugt eine Tönung dieser Farbe, und das Hinzufügen von Schwarz ergibt eine Schattierung. Alternativ entsteht ein Ton durch das Hinzufügen von Grau oder durch sowohl Tönung als auch Schattierung. Denken Sie daran, dass das „s“ und „l“ von `hsl()` für Sättigung und Helligkeit stehen. Der Sättigungsprozentsatz ändert die Menge an Grau, und der Helligkeitsprozentsatz bestimmt, wie viel Weiß oder Schwarz in der Farbe enthalten ist. Dies ist nützlich, wenn Sie einen Basisfarbton haben, den Sie mögen, aber verschiedene Variationen davon benötigen.

# --instructions--

Alle Elemente haben eine Standard-`background-color` von `transparent`. Unser `nav`-Element erscheint derzeit mit einem `cyan`-Hintergrund, weil das dahinterliegende Element eine `background-color` mit `cyan` gesetzt hat. Fügen Sie dem `background-color`-Element eine `nav` hinzu, sodass es denselben `cyan`-Farbton verwendet, aber `80%`-Sättigungs- und `25%`-Helligkeitswerte hat, um seinen Ton und seine Schattierung zu ändern.

# --hints--

Das `nav`-Element sollte eine `background-color` des angepassten Cyan-Tons mit der `hsl()`-Eigenschaft haben.

```js
// Computed style of hsl(180, 80%, 25%) results in rgb(13,115,115)
assert.equal(
  new __helpers.CSSHelp(document).getStyle('nav').getPropVal('background-color', true), 
  'rgb(13,115,115)'
)
```

# --seed--

## --seed-contents--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {

  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>

<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```

# --solutions--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {
    background-color: hsl(180, 80%, 25%);
  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```
