---
id: 587d78a8367417b2b2512ae7
title: Animationszeitpunkt mit Schlüsselwörtern ändern
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKvwCM'
forumTopicId: 301045
dashedName: change-animation-timing-with-keywords
---

# --description--

Bei CSS-Animationen steuert die `animation-timing-function`-Eigenschaft, wie schnell sich ein animiertes Element über die Dauer der Animation verändert. Wenn die Animation ein Auto ist, das sich in einer bestimmten Zeit (Ihrer `animation-duration`) von Punkt A nach Punkt B bewegt, beschreibt die `animation-timing-function`, wie das Auto während der Fahrt beschleunigt und abbremst.

Es gibt eine Reihe vordefinierter Schlüsselwörter für gängige Optionen. Zum Beispiel ist der Standardwert `ease`, der langsam beginnt, in der Mitte schneller wird und am Ende wieder langsamer wird. Weitere Optionen sind `ease-out`, das am Anfang schnell ist und dann langsamer wird, `ease-in`, das am Anfang langsam ist und am Ende schneller wird, oder `linear`, das eine konstante Animationsgeschwindigkeit über die gesamte Dauer anwendet.

# --instructions--

Fügen Sie für die Elemente mit der ID `ball1` und `ball2` jeweils eine `animation-timing-function`-Eigenschaft hinzu und setzen Sie `#ball1` auf `linear` sowie `#ball2` auf `ease-out`. Beachten Sie den Unterschied darin, wie sich die Elemente während der Animation bewegen, aber am Ende zusammenkommen, da sie dieselbe `animation-duration` von 2 Sekunden teilen.

# --hints--

Der Wert der `animation-timing-function`-Eigenschaft für das Element mit der ID `ball1` sollte `linear` sein.

```js
const ballOne =document.querySelector('#ball1'); 
const ballOneStyle = window.getComputedStyle(ballOne); 

const ball1Animation = __helpers.removeWhiteSpace(ballOneStyle?.animationTimingFunction);
assert.isTrue(ball1Animation == 'linear' || ball1Animation == 'cubic-bezier(0,0,1,1)');
```

Der Wert der `animation-timing-function`-Eigenschaft für das Element mit der ID `ball2` sollte `ease-out` sein.

```js
const ballTwo = document.querySelector('#ball2'); 
const ballTwoStyle = window.getComputedStyle(ballTwo); 

const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)');
```

# --seed--

## --seed-contents--

```html
<style>

  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;

  }
  #ball2 {
    left:56%;

  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left:56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
