---
id: 587d78a9367417b2b2512aea
title: Bewegung natürlicher gestalten mit einer Bézierkurve
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7akWUv'
forumTopicId: 301063
dashedName: make-motion-more-natural-using-a-bezier-curve
---

# --description--

Diese Aufgabe animiert ein Element, um die Bewegung eines jonglierten Balls nachzuahmen. Frühere Aufgaben behandelten die `linear` und `ease-out` kubischen Bézierkurven, jedoch stellt keine die Jonglierbewegung genau dar. Sie müssen eine Bézierkurve dafür anpassen.

Die `animation-timing-function` läuft automatisch bei jedem Keyframe in einer Schleife, wenn die `animation-iteration-count` auf unendlich gesetzt ist. Da in der Mitte der Animationsdauer (bei `50%`) eine Keyframe-Regel definiert ist, entstehen zwei identische Animationsverläufe bei der Aufwärts- und Abwärtsbewegung des Balls.

Die folgende kubische Bézierkurve simuliert eine Jonglierbewegung:

```css
cubic-bezier(0.3, 0.4, 0.5, 1.6);
```

Beachten Sie, dass der Wert von y2 größer als 1 ist. Obwohl die kubische Bézierkurve auf ein 1-zu-1-Koordinatensystem abgebildet wird und nur x-Werte von 0 bis 1 akzeptieren kann, kann der y-Wert auf Zahlen größer als eins gesetzt werden. Dies führt zu einer hüpfenden Bewegung, die ideal ist, um den jonglierenden Ball zu simulieren.

# --instructions--

Ändern Sie den Wert der `animation-timing-function` des Elements mit der ID `green` in eine `cubic-bezier`-Funktion mit den Werten für x1, y1, x2, y2 jeweils 0,311, 0,441, 0,444, 1,649.

# --hints--

Der Wert der `animation-timing-function`-Eigenschaft für das Element mit der ID `green` sollte eine `cubic-bezier`-Funktion mit den angegebenen x1-, y1-, x2- und y2-Werten sein.

```js
const greenElement = document.querySelector('#green');
 const greenStyle = window.getComputedStyle(greenElement);
assert.equal(greenStyle?.animationTimingFunction, 'cubic-bezier(0.311, 0.441, 0.444, 1.649)');
```

# --seed--

## --seed-contents--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.69, 0.1, 1, 0.1);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.311, 0.441, 0.444, 1.649);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```
