---
id: 587d78a3367417b2b2512acf
title: Die Position überlappender Elemente mit der z-index-Eigenschaft ändern
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM94aHk'
forumTopicId: 301046
dashedName: change-the-position-of-overlapping-elements-with-the-z-index-property
---

# --description--

Wenn Elemente so positioniert sind, dass sie sich überlappen (also mit `position: absolute | relative | fixed | sticky`), erscheint standardmäßig das Element, das im HTML-Markup später steht, oben auf den anderen Elementen. Die `z-index`-Eigenschaft kann jedoch die Reihenfolge festlegen, in der Elemente übereinander gestapelt werden. Sie muss eine ganze Zahl (also eine Ganzzahl ohne Dezimalstellen) sein, und höhere Werte für die `z-index`-Eigenschaft eines Elements bringen es weiter nach oben im Stapel als Elemente mit niedrigeren Werten.

# --instructions--

Fügen Sie dem Element mit dem Klassennamen `z-index` (dem roten Rechteck) eine `first`-Eigenschaft hinzu und setzen Sie sie auf den Wert 2, damit es das andere Element (blaues Rechteck) überdeckt.

# --hints--

Das Element mit der Klasse `first` sollte einen `z-index`-Wert von 2 haben.

```js
const firstElement = document.querySelector('.first');
const firstStyle = window.getComputedStyle(firstElement); 
assert.equal(firstStyle?.zIndex, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;

  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>
```

# --solutions--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;
    z-index: 2;
  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>
<div class="first"></div>
<div class="second"></div>
```
