---
id: 587d781c367417b2b2512ac3
title: Die Schriftstärke für mehrere Überschriften festlegen
challengeType: 0
videoUrl: 'https://scrimba.com/c/crVWRHq'
forumTopicId: 301069
dashedName: set-the-font-weight-for-multiple-heading-elements
---

# --description--

Sie haben im letzten Abschnitt die `font-size` jedes Überschriften-Tags festgelegt, hier passen Sie die `font-weight` an.

Die `font-weight`-Eigenschaft legt fest, wie dick oder dünn Zeichen in einem Textabschnitt dargestellt werden.

# --instructions--

<ul><li>Setzen Sie die <code>font-weight</code> des <code>h1</code>-Tags auf 800.</li><li>Setzen Sie die <code>font-weight</code> des <code>h2</code>-Tags auf 600.</li><li>Setzen Sie die <code>font-weight</code> des <code>h3</code>-Tags auf 500.</li><li>Setzen Sie die <code>font-weight</code> des <code>h4</code>-Tags auf 400.</li><li>Setzen Sie die <code>font-weight</code> des <code>h5</code>-Tags auf 300.</li><li>Setzen Sie die <code>font-weight</code> des <code>h6</code>-Tags auf 200.</li></ul>

# --hints--

Ihr Code sollte die `font-weight`-Eigenschaft für das `h1`-Tag auf 800 setzen.

```js
const h1Element =document.querySelector('h1')
const h1Style = window.getComputedStyle(h1Element);
assert.equal(h1Style?.fontWeight, '800');
```

Ihr Code sollte die `font-weight`-Eigenschaft für das `h2`-Tag auf 600 setzen.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.fontWeight, '600');
```

Ihr Code sollte die `font-weight`-Eigenschaft für das `h3`-Tag auf 500 setzen.

```js
const h3Element =document.querySelector('h3')
const h3Style = window.getComputedStyle(h3Element);
assert.equal(h3Style?.fontWeight, '500');
```

Ihr Code sollte die `font-weight`-Eigenschaft für das `h4`-Tag auf 400 setzen.

```js
const h4Element =document.querySelector('h4')
const h4Style = window.getComputedStyle(h4Element);
assert.equal(h4Style?.fontWeight, '400');
```

Ihr Code sollte die `font-weight`-Eigenschaft für das `h5`-Tag auf 300 setzen.

```js
const h5Element =document.querySelector('h5')
const h5Style = window.getComputedStyle(h5Element);
assert.equal(h5Style?.fontWeight, '300');
```

Ihr Code sollte die `font-weight`-Eigenschaft für das `h6`-Tag auf 200 setzen.

```js
const h6Element =document.querySelector('h6')
const h6Style = window.getComputedStyle(h6Element);
assert.equal(h6Style?.fontWeight, '200');
```

# --seed--

## --seed-contents--

```html
<style>
  h1 {
    font-size: 68px;

  }
  h2 {
    font-size: 52px;

  }
  h3 {
    font-size: 40px;

  }
  h4 {
    font-size: 32px;

  }
  h5 {
    font-size: 21px;

  }
  h6 {
    font-size: 14px;

  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
    font-weight: 800;
  }
  h2 {
    font-size: 52px;
    font-weight: 600;
  }
  h3 {
    font-size: 40px;
    font-weight: 500;
  }
  h4 {
    font-size: 32px;
    font-weight: 400;
  }
  h5 {
    font-size: 21px;
    font-weight: 300;
  }
  h6 {
    font-size: 14px;
    font-weight: 200;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
