---
id: 587d78a6367417b2b2512adc
title: "Die CSS-Eigenschaft transform skewY verwenden, um ein Element entlang der Y-Achse zu verzerren"
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZ2uB'
forumTopicId: 301075
dashedName: use-the-css-transform-property-skewy-to-skew-an-element-along-the-y-axis
---

# --description--

Da die `skewX()` Funktion das ausgewählte Element entlang der X-Achse um einen bestimmten Grad verzerrt, ist es keine Überraschung, dass die `skewY()` Eigenschaft ein Element entlang der Y- (vertikalen) Achse verzerrt.

# --instructions--

Verzerren Sie das Element mit der ID `top` um -10 Grad entlang der Y-Achse, indem Sie die `transform` Eigenschaft verwenden.

# --hints--

Das Element mit der ID `top` sollte um -10 Grad entlang seiner Y-Achse verzerrt werden.

```js
assert.match(code, /#top\s*?{\s*?.*?\s*?transform:\s*?skewY\(-10deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;

  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;
    transform: skewY(-10deg);
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
