---
id: 587d78a7367417b2b2512ae0
title: "Verwenden Sie CSS-Animation, um den Hover-Zustand eines Buttons zu ändern"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4vZAa'
forumTopicId: 301073
dashedName: use-css-animation-to-change-the-hover-state-of-a-button
---

# --description--

Sie können CSS `@keyframes` verwenden, um die Farbe eines Buttons im Hover-Zustand zu ändern.

Hier ist ein Beispiel, wie die Breite eines Bildes beim Hover geändert wird:

```html
<style>
  img {
    width: 30px;
  }
  img:hover {
    animation-name: width;
    animation-duration: 500ms;
  }

  @keyframes width {
    100% {
      width: 40px;
    }
  }
</style>

<img src="https://cdn.freecodecamp.org/curriculum/applied-visual-design/google-logo.png" alt="Google's Logo" />
```

# --instructions--

Beachten Sie, dass `ms` für Millisekunden steht, wobei 1000ms 1s entsprechen.

Verwenden Sie CSS `@keyframes`, um die `background-color` des `button`-Elements zu ändern, sodass es `#4791d0` wird, wenn ein Benutzer mit der Maus darüber fährt. Die `@keyframes`-Regel sollte nur einen Eintrag für `100%` enthalten.

# --hints--

Die @keyframes-Regel sollte die `animation-name` background-color verwenden.

```js
assert.match(code, /@keyframes\s+?background-color\s*?{/g);
```

Es sollte eine Regel unter `@keyframes` geben, die die `background-color` bei 100 % auf `#4791d0` ändert.

```js
assert.match(code, /100%\s*?{\s*?background-color:\s*?#4791d0;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }


</style>

<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }

  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
