---
id: 587d7791367417b2b2512ab3
title: Visuelles Gleichgewicht mit der Eigenschaft text-align erzeugen
challengeType: 0
videoUrl: 'https://scrimba.com/c/c3b4EAp'
forumTopicId: 301053
dashedName: create-visual-balance-using-the-text-align-property
---

# --description--

Dieser Abschnitt des Curriculums konzentriert sich auf angewandtes visuelles Design. Die erste Gruppe von Aufgaben baut auf dem vorgegebenen Kartenlayout auf, um eine Reihe von Grundsätzen zu zeigen.

Text ist oft ein großer Teil von Webinhalten. CSS bietet mehrere Möglichkeiten, ihn mit der Eigenschaft `text-align` auszurichten.

`text-align: justify;` verteilt den Text so, dass jede Zeile die gleiche Breite hat.

`text-align: center;` zentriert den Text.

`text-align: right;` richtet den Text rechtsbündig aus.

Und `text-align: left;` (die Standardeinstellung) richtet den Text linksbündig aus.

# --instructions--

Richten Sie den Text des `h4`-Tags, der „Google“ sagt, zentriert aus. Dann justieren Sie den Absatz-Tag, der Informationen darüber enthält, wie Google gegründet wurde.

# --hints--

Ihr Code sollte die Eigenschaft text-align am `h4`-Tag verwenden, um ihn auf `center` zu setzen.

```js
const h4Element =document.querySelector('h4')
const h4Style = window.getComputedStyle(h4Element);
assert.equal(h4Style?.textAlign, 'center');
```

Ihr Code sollte die Eigenschaft text-align am `p`-Tag verwenden, um ihn auf `justify` zu setzen.

```js
const pElement =document.querySelector('p')
const pStyle = window.getComputedStyle(pElement);
assert.equal(pStyle?.textAlign, 'justify');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {

  }
  p {

  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
  }
  p {
    text-align: justify;
  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
