---
id: 587d78a9367417b2b2512ae8
title: "Lernen, wie Bézierkurven funktionieren"
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bDrs8'
forumTopicId: 301058
dashedName: learn-how-bezier-curves-work
---

# --description--

Die letzte Aufgabe hat die `animation-timing-function`-Eigenschaft und einige Schlüsselwörter eingeführt, die die Geschwindigkeit einer Animation über ihre Dauer verändern. CSS bietet eine weitere Möglichkeit neben Schlüsselwörtern, die noch feinere Kontrolle darüber ermöglicht, wie die Animation abläuft – durch die Verwendung von Bézierkurven.

In CSS-Animationen werden Bézierkurven mit der `cubic-bezier`-Funktion verwendet. Die Form der Kurve stellt dar, wie die Animation abläuft. Die Kurve liegt in einem 1-mal-1-Koordinatensystem. Die X-Achse dieses Koordinatensystems ist die Dauer der Animation (denken Sie daran als Zeitskala), und die Y-Achse ist die Veränderung in der Animation.

Die `cubic-bezier`-Funktion besteht aus vier Hauptpunkten, die auf diesem 1-mal-1-Raster liegen: `p0`, `p1`, `p2` und `p3`. `p0` und `p3` sind für Sie festgelegt – sie sind die Anfangs- und Endpunkte, die immer jeweils am Ursprung (0, 0) und (1, 1) liegen. Sie setzen die x- und y-Werte für die anderen beiden Punkte, und wo Sie sie im Raster platzieren, bestimmt die Form der Kurve, der die Animation folgt. Dies geschieht in CSS, indem Sie die x- und y-Werte der `p1`- und `p2`-„Anker“-Punkte in der Form `(x1, y1, x2, y2)` angeben. Zusammengenommen hier ein Beispiel für eine Bézierkurve im CSS-Code:

```css
animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
```

Im obigen Beispiel sind die x- und y-Werte für jeden Punkt gleich (x1 = 0,25 = y1 und x2 = 0,75 = y2), was, wenn Sie sich an den Geometrieunterricht erinnern, zu einer Linie führt, die sich vom Ursprung bis zum Punkt (1, 1) erstreckt. Diese Animation ist eine lineare Veränderung eines Elements während der Dauer einer Animation und entspricht der Verwendung des `linear`-Schlüsselworts. Anders gesagt, sie ändert sich mit konstanter Geschwindigkeit.

# --instructions--

Ändern Sie für das Element mit der ID `ball1` den Wert der `animation-timing-function`-Eigenschaft von `linear` auf den entsprechenden Wert der `cubic-bezier`-Funktion. Verwenden Sie die im obigen Beispiel angegebenen Punktwerte.

# --hints--

Der Wert der `animation-timing-function`-Eigenschaft für das Element mit der ID `ball1` sollte der linearen Entsprechung der `cubic-bezier`-Funktion entsprechen.

```js
const ballOne = document.querySelector('#ball1');
const ballOneStyle = window.getComputedStyle(ballOne); 
assert.equal(
  ballOneStyle?.animationTimingFunction,
    'cubic-bezier(0.25, 0.25, 0.75, 0.75)'
);
```

Der Wert der `animation-timing-function`-Eigenschaft für das Element mit der ID `ball2` sollte sich nicht ändern.

```js
const ballTwoElement = document.querySelector('#ball2');
const ballTwoStyle = window.getComputedStyle(ballTwoElement); 
const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(
  ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
