---
id: 587d78a4367417b2b2512ad3
title: Farbe verschiedener Elemente auf komplementäre Farben anpassen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cWmPpud'
forumTopicId: 301033
dashedName: adjust-the-color-of-various-elements-to-complementary-colors
---

# --description--

Die Aufgabe zu komplementären Farben zeigte, dass sich gegenüberliegende Farben auf dem Farbrad gegenseitig lebendiger erscheinen lassen, wenn sie nebeneinander platziert werden. Allerdings kann der starke visuelle Kontrast störend wirken, wenn er auf einer Website zu häufig verwendet wird, und manchmal den Text schwer lesbar machen, wenn er auf einem Hintergrund mit komplementärer Farbe steht. In der Praxis ist eine der Farben meist dominant, und die Komplementärfarbe wird verwendet, um bestimmte Inhalte auf der Seite visuell hervorzuheben.

# --instructions--

Diese Seite verwendet einen Türkiston (`#09A7A1`) als dominante Farbe und dessen orangefarbene (`#FF790E`) Komplementärfarbe, um die Anmeldebuttons optisch hervorzuheben. Ändern Sie die `background-color` von sowohl dem `header` als auch dem `footer` von Schwarz in die türkise Farbe. Dann ändern Sie den `h2` Text `color` ebenfalls in Türkis. Schließlich ändern Sie die `background-color` des `button` in die orange Farbe.

# --hints--

Das `header`-Element sollte eine `background-color` von #09A7A1 haben.

```js
const headerElement = document.querySelector('header'); 
const headerStyle = window.getComputedStyle(headerElement);
assert.equal(headerStyle?.backgroundColor, 'rgb(9, 167, 161)');
```

Das `footer`-Element sollte eine `background-color` von #09A7A1 haben.

```js
const footerElement = document.querySelector('footer'); 
const footerStyle = window.getComputedStyle(footerElement);
assert.equal(footerStyle?.backgroundColor, 'rgb(9, 167, 161)');
```

Das `h2`-Element sollte eine `color` von #09A7A1 haben.

```js
const h2Element = document.querySelector('h2'); 
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.color, 'rgb(9, 167, 161)');
```

Das `button`-Element sollte eine `background-color` von #FF790E haben.

```js
const buttonElement = document.querySelector('button'); 
const buttonStyle = window.getComputedStyle(buttonElement);
assert.equal(buttonStyle?.backgroundColor, 'rgb(255, 121, 14)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: white;
  }
  header {
    background-color: black;
    color: white;
    padding: 0.25em;
  }
  h2 {
    color: black;
  }
  button {
    background-color: white;
  }
  footer {
    background-color: black;
    color: white;
    padding: 0.5em;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
</header>
<main>
  <article>
    <h2>Machine Learning in the Kitchen</h2>
    <p>Join this two day workshop that walks through how to implement cutting-edge snack-getting algorithms with a command line interface. Coding usually involves writing exact instructions, but sometimes you need your computer to execute flexible commands, like <code>fetch Pringles</code>.</p>
    <button>Sign Up</button>
  </article>
  <article>
    <h2>Bisection Vegetable Chopping</h2>
    <p>This week-long retreat will level-up your coding ninja skills to actual ninja skills. No longer is the humble bisection search limited to sorted arrays or coding interview questions, applying its concepts in the kitchen will have you chopping carrots in O(log n) time before you know it.</p>
    <button>Sign Up</button>
  </article>
</main>
<br>
<footer>&copy; 2018 FCC Kitchen</footer>
```

# --solutions--

```html
<style>
  body {
    background-color: white;
  }
  header {
    background-color: #09A7A1;
    color: white;
    padding: 0.25em;
  }
  h2 {
    color: #09A7A1;
  }
  button {
    background-color: #FF790E;
  }
  footer {
    background-color: #09A7A1;
    color: white;
    padding: 0.5em;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
</header>
<main>
  <article>
    <h2>Machine Learning in the Kitchen</h2>
    <p>Join this two day workshop that walks through how to implement cutting-edge snack-getting algorithms with a command line interface. Coding usually involves writing exact instructions, but sometimes you need your computer to execute flexible commands, like <code>fetch Pringles</code>.</p>
    <button>Sign Up</button>
  </article>
  <article>
    <h2>Bisection Vegetable Chopping</h2>
    <p>This week-long retreat will level-up your coding ninja skills to actual ninja skills. No longer is the humble bisection search limited to sorted arrays or coding interview questions, applying its concepts in the kitchen will have you chopping carrots in O(log n) time before you know it.</p>
    <button>Sign Up</button>
  </article>
</main>
<br>
<footer>&copy; 2018 FCC Kitchen</footer>
```
