---
id: 587d78a7367417b2b2512ae2
title: "Visuelle Richtung erzeugen, indem ein Element von links nach rechts verblasst"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJqqAE'
forumTopicId: 301054
dashedName: create-visual-direction-by-fading-an-element-from-left-to-right
---

# --description--

Für diese Aufgabe ändern Sie die `opacity` eines animierten Elements, sodass es allmählich verblasst, wenn es die rechte Seite des Bildschirms erreicht.

In der angezeigten Animation bewegt sich das runde Element mit dem Farbverlauf-Hintergrund gemäß der `@keyframes`-Regel bis zur 50-%-Marke der Animation nach rechts.

# --instructions--

Zielen Sie auf das Element mit der ID `ball` und fügen Sie die `opacity`-Eigenschaft mit dem Wert 0,1 bei `50%` hinzu, damit das Element beim Bewegen nach rechts verblasst.

# --hints--

Die `keyframes`-Regel für das Verblassen sollte die `opacity`-Eigenschaft bei 50 % auf 0,1 setzen.

```js
assert.match(code,/@keyframes fade\s*?{\s*?50%\s*?{\s*?(?:left:\s*?60%;\s*?opacity:\s*?0?\.1;|opacity:\s*?0?\.1;\s*?left:\s*?60%;)/gi
);
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;

    }
  }

</style>

<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;
      opacity: 0.1;
    }
  }
</style>
<div id="ball"></div>
```
