---
id: 587d78a9367417b2b2512ae9
title: "Eine Bézierkurve verwenden, um eine Grafik zu bewegen"
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bnRCK'
forumTopicId: 301071
dashedName: use-a-bezier-curve-to-move-a-graphic
---

# --description--

In einer vorherigen Aufgabe wurde das Schlüsselwort `ease-out` besprochen, das eine Animationsänderung beschreibt, die zuerst beschleunigt und am Ende der Animation langsamer wird. Rechts wird der Unterschied zwischen dem Schlüsselwort `ease-out` (für das blaue Element) und dem Schlüsselwort `linear` (für das rote Element) demonstriert. Ähnliche Animationsverläufe wie mit dem Schlüsselwort `ease-out` können durch die Verwendung einer benutzerdefinierten kubischen Bézierkurvenfunktion erreicht werden.

Im Allgemeinen bewirkt das Ändern der `p1` und `p2` Ankerpunkte die Erstellung verschiedener Bézierkurven, die steuern, wie die Animation im Zeitverlauf abläuft. Hier ist ein Beispiel für eine Bézierkurve mit Werten, die den `ease-out` Stil nachahmen:

```css
animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
```

Denken Sie daran, dass alle `cubic-bezier` Funktionen mit `p0` bei (0, 0) beginnen und mit `p3` bei (1, 1) enden. In diesem Beispiel bewegt sich die Kurve schneller entlang der Y-Achse (beginnt bei 0, geht zum `p1` y-Wert von 0, dann zum `p2` y-Wert von 1) als entlang der X-Achse (0 zum Start, dann 0 für `p1`, bis 0,58 für `p2`). Dadurch verläuft die Änderung im animierten Element schneller als die Zeit der Animation für diesen Abschnitt. Gegen Ende der Kurve kehrt sich das Verhältnis zwischen der Änderung der x- und y-Werte um – der y-Wert bewegt sich von 1 zu 1 (keine Änderung), und die x-Werte bewegen sich von 0,58 zu 1, wodurch die Animationsänderungen im Vergleich zur Animationsdauer langsamer verlaufen.

# --instructions--

Um die Wirkung dieser Bézierkurve in Aktion zu sehen, ändern Sie die `animation-timing-function` des Elements mit der ID `red` in eine `cubic-bezier` Funktion mit x1, y1, x2, y2, die jeweils auf `0, 0, 0.58, 1` als Werte gesetzt sind. Dadurch durchlaufen beide Elemente die Animation ähnlich.

# --hints--

Der Wert der `animation-timing-function` Eigenschaft des Elements mit der ID `red` sollte eine `cubic-bezier` Funktion mit den x1, y1, x2, y2 Werten 0, 0, 0,58, 1 sein.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.equal(
  redStyle?.animationTimingFunction, 'cubic-bezier(0, 0, 0.58, 1)'
);
```

Das Element mit der ID `red` sollte die `animation-timing-function` Eigenschaft von `linear` nicht mehr haben.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.notEqual(redStyle?.animationTimingFunction, 'linear');
```

Der Wert der `animation-timing-function` Eigenschaft für das Element mit der ID `blue` sollte sich nicht ändern.

```js
const blueElement = document.querySelector('#blue');
const blueStyle = window.getComputedStyle( blueElement);
const blueBallAnimation = __helpers.removeWhiteSpace(
  blueStyle?.animationTimingFunction
);
assert.isTrue(
  blueBallAnimation == 'ease-out' ||
    blueBallAnimation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```

# --solutions--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```
