---
id: 58a7a6ebf9a6318348e2d5aa
title: Füllmodus einer Animation ändern
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDmcE'
forumTopicId: 301064
dashedName: modify-fill-mode-of-an-animation
---

# --description--

Das ist schon gut, aber es funktioniert noch nicht richtig. Beachten Sie, wie die Animation zurückgesetzt wird, nachdem `500ms` vorbei ist, wodurch der Button wieder die ursprüngliche Farbe annimmt. Sie möchten, dass der Button hervorgehoben bleibt.

Das erreichen Sie, indem Sie die `animation-fill-mode`-Eigenschaft auf `forwards` setzen. `animation-fill-mode` gibt den Stil an, der auf ein Element angewendet wird, wenn die Animation beendet ist. Sie können es so einstellen:

```css
animation-fill-mode: forwards;
```

# --instructions--

Setzen Sie die `animation-fill-mode`-Eigenschaft von `button:hover` auf `forwards`, damit der Button beim Darüberfahren mit der Maus hervorgehoben bleibt.

# --hints--

`button:hover` sollte eine `animation-fill-mode`-Eigenschaft mit dem Wert `forwards` haben.

```js
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-fill-mode\s*?:\s*?forwards\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-name\s*?:\s*?background-color\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-duration\s*?:\s*?500ms\s*?;[\s\S]*}/gi );
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    /* Only change code below this line */

    /* Only change code above this line */
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    animation-fill-mode: forwards;
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
