---
id: 60a3e3396c7b40068ad69970
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Bevor Sie mit dem Stylen des `div` beginnen können, das Sie hinzugefügt haben, müssen Sie Ihr CSS mit Ihrem HTML verknüpfen.

Fügen Sie ein `link`-Element hinzu, um Ihre `styles.css`-Datei zu verlinken. Setzen Sie das `href` auf `styles.css` und denken Sie daran, das `rel`-Attribut auf `stylesheet` zu setzen.

# --hints--

Ihr Code sollte ein `link`-Element enthalten.

```js
assert(/<link/.test(code))
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
