---
id: 60a3e3396c7b40068ad6996b
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

FĂĽgen Sie innerhalb des `head`-Elements einen `meta`-Tag hinzu, der das `charset` auf `UTF-8` setzt, und ein `title`-Element mit dem Wert `Rothko Painting`.

FĂĽgen Sie innerhalb des `body`-Elements ein `img`-Element mit einem `src` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` hinzu.
# --hints--

Ihr Code sollte einen `meta`-Tag enthalten.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Der `meta`-Tag sollte das `charset` auf `UTF-8` setzen.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Ihr Code sollte ein `title`-Element enthalten.

```js
assert(document.querySelectorAll('title').length === 1);
```

Das `title` sollte `Rothko Painting` sein.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

Ihr Code sollte ein `img`-Element enthalten.

```js
assert(document.querySelectorAll('img').length === 1);
```

Das `img`-Element sollte ein `src` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` haben.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
