---
id: 60a3e3396c7b40068ad69974
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Jedes Gemälde braucht einen Rahmen.

Fassen Sie das `.canvas`-Element in ein weiteres `div`-Element ein. Geben Sie diesem `div` die Klasse `frame`.
# --hints--

Sie sollten ein neues `div`-Element hinzufügen.

```js
assert(document.querySelectorAll('div').length === 2)
```

Ihr `.canvas`-Element sollte im neuen `div`-Element verschachtelt sein.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

Ihr neues `div` sollte eine `class` mit dem Wert `frame` haben.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

Ihr neues `div` sollte sich innerhalb Ihres `body`-Elements befinden.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
