---
id: 60a3e3396c7b40068ad6996c
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Im CSS-Boxmodell wird jedes HTML-Element als eine Box mit vier Bereichen behandelt.

Stellen Sie sich vor, Sie erhalten eine Box von Ihrem Lieblings-Onlinehändler – der Inhalt ist der Artikel in der Box oder in unserem Fall ein Kopfzeilen-, Absatz- oder Bildelement.

Ändern Sie das `src`-Attribut im `<img>` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` zu `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

Das `img`-Element sollte ein `src` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png` haben

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

