---
id: 60a3e3396c7b40068ad6996a
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Sie sollten nun mit den grundlegenden Elementen vertraut sein, die eine HTML-Seite haben sollte.

Richten Sie Ihren Code mit einer `DOCTYPE`-Deklaration, einem `html`-Element mit der Sprache auf Englisch eingestellt, einem `head`-Element und einem `body`-Element ein.

# --hints--

Ihr Code sollte eine `<!DOCTYPE html>`-Deklaration enthalten.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Sie sollten einen öffnenden `<html>`-Tag mit einem `lang`-Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr Code sollte ein `head`-Element innerhalb des `html`-Elements enthalten.

```js
assert(document.querySelectorAll('head').length === 1);
```

Ihr Code sollte ein `body`-Element innerhalb des `html`-Elements enthalten.

```js
assert(document.querySelectorAll('body').length === 1);
```

Ihr `head`-Element sollte vor Ihrem `body`-Element stehen.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
