---
id: 60a3e3396c7b40068ad69980
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Fügen Sie direkt unter Ihrem `div`-Element ein weiteres `class` mit einem `two`-Wert von `one` hinzu. Dies wird Ihr zweites Rechteck sein.

# --hints--

Sie sollten das bestehende `.one`-Element nicht verändern.

```js
assert(document.querySelectorAll('.one').length === 1);
```

Sie sollten ein zweites `div`-Element in Ihrem `.canvas`-Element haben.

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

Ihr zweites `div`-Element sollte einen `class`-Wert von `two` haben.

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

Ihr `.two`-Element sollte nach Ihrem `.one`-Element kommen.

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
