---
id: 60a3e3396c7b40068ad6997a
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Schreiben Sie eine neue Regel, die `.one` anspricht, und setzen Sie dessen `width` auf 425 Pixel.

# --hints--

Sie sollten einen `.one`-Selektor haben.

```js
const hasOne = new __helpers.CSSHelp(document).getStyle('.one');
assert(hasOne);
```

Sie sollten die `width`-Eigenschaft auf `425px` setzen.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '425px');
assert(hasWidth);
```

Ihr `.one`-Element sollte einen `width`-Wert von `425px` haben.

```js
const oneWidth = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('width');
assert(oneWidth === '425px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
