---
id: 60a3e3396c7b40068ad69971
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Zeit für CSS.

Auch wenn Ihr `<div>` keinen Text enthält, wird es trotzdem als ein Kasten mit Inhalt behandelt.
Schreiben Sie eine CSS-Regel, die den `.canvas`-Klassenselektor verwendet und dessen `width` auf 500 Pixel setzt.
Hier ist eine CSS-Regel, die die Breite der Klasse `card` auf 300 Pixel setzt:

```css
.card {
  width: 300px;
}
```

# --hints--

Ihr Code sollte einen `.canvas`-Selektor enthalten.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

Sie sollten die `width`-Eigenschaft auf `500px` setzen.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

Ihr `.canvas`-Element sollte eine `width` von `500px` haben.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
