---
id: 6ac0b35f3b29a27f7e207295
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Katika hatua iliyopita, uliunda kisimbu `sentence`. Sasa chapisha thamani yake kwenye terminal kwa kuandika `print(sentence)` kwenye mstari uliangaziwa.

Kwa hiyo, umemaliza programu yako ya kwanza ya Python.

# --hints--

Unapaswa kuwa na `print(sentence)` katika msimbo wako. Hakikisha unaanza mwanzoni mwa mstari.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
