---
id: 6ac0a4463662d9738c393e21
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Katika hatua iliyopita, uliunda **mfuatano wa herufi**: mfululizo wa herufi uliofungwa katika alama za nukuu moja au alama za nukuu maradufu. Katika programu, mfuatano wa herufi hutumika kuwakilisha maandishi.

**Terminal** ni mahali ambapo unaweza kuona matokeo ya msimbo wako.

Ili uweze kuona matokeo, badilisha mstari wako wa msimbo uliopo kuwa:

```py
print('Hello')
```

Hakikisha unaanza mwanzoni mwa mstari.

# --hints--

Unapaswa kuwa na `print('Hello')` katika msimbo wako. Hakikisha unaanza mwanzoni mwa mstari.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
