---
id: 6ac0aaba914d81b32567c076
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` ni kitendakazi kilichojengwa ndani ambacho unaweza kutumia kutuma matokeo kwenye terminali. Kama unavyoona, maandishi `Hello` sasa yanaonekana kwenye terminali.

Sasa, chapisha maandishi `freeCodeCamp` kwenye terminali kwa kwenda kwenye mstari ulioangaziwa na kuandika:

```py
print('freeCodeCamp')
```

Hakikisha unaanza mwanzoni mwa mstari.

# --hints--

Unapaswa kuwa na `print('freeCodeCamp')` katika msimbo wako. Hakikisha unaanza mwanzoni mwa mstari.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
