---
id: 6ac0ae563f6de70930f47ef7
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Unapoandika msimbo, mara nyingi utataka kutumia tena au kurejea thamani zilizopo katika msimbo wako. Vigezo hukuwezesha kufanya hivyo.

Ili kuweka thamani kwa kigezo, andika jina la kigezo kifuatiwe na opereta ya mgawo (`=`) na thamani unayotaka kuweka kwa kigezo hicho:

```py
sentence = 'I am learning Python'
```

Nenda kwenye mstari ulioangaziwa, tengeneza kigezo kinachoitwa `sentence` na weka thamani yake kuwa mfuatano wa herufi `I am learning Python`.

# --hints--

Unapaswa kuwa na `sentence = 'I am learning Python'` katika msimbo wako. Hakikisha unaanza mwanzoni mwa mstari.

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
