---
id: 6ac09f8420b4b67e075d2cd2
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda programu yako ya kwanza ya Python.

Huu ni mhariri wako wa msimbo, ambapo utaandika msimbo wa Python. Mstari wa 1 umeangaziwa. Nenda hapo na andika:

```py
'Hello'
```

Hakikisha unaanza mwanzoni mwa mstari na usisahau alama za nukuu.

Ukipomaliza, bofya kitufe cha "kagua msimbo wako" kuona kama ni sahihi.

# --hints--

Unapaswa kuandika `'Hello'`. Hakikisha unaanza mwanzoni mwa mstari na usisahau alama za nukuu.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
