---
id: 66a207974c806a19d6607073
title: 创建一个书籍库存应用
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `h1` 元素，其中包含文本 `Book Inventory` 。
1. 你应该有一个 `table` 元素，其中包含列：`Title`、`Author`、`Category`、`Status`和`Rate`。
1. 你的表格至少应有四行，第一行用于列标题，其余行填写书籍相关信息。
1. 表体内的每一行都应该带有以下类之一：`read`、`to-read` 或 `in-progress`。
1. `Status` 列中的 `td` 元素应当包含一个 `span` 元素，该元素的 `class` 属性为 `status` ，分别绕着文本 `Read`、`To Read` 或 `In Progress`，具体取决于该行的类。
1. `Rate` 列中的 `td` 元素应当包含一个 `class` 属性为 `rate` 的 `span` 元素，该元素内包含三个空的 `span` 元素。
1. `.read` 行内的 `.rate` 元素应带有额外的类，其值可为 `one`、`two` 或 `three`，具体取决于个人的评分。 此值应在 `rate` 之后。
1. 你应该创建三个属性选择器，选择具有 `read`、`to-read` 和 `in-progress`类的元素，将他们的 `background-image` 属性设置为使用自定义的 `linear-gradient` 的值。
1. 你应该将每个 `span` 元素的 `display` 属性设置为 `inline-block` 。
1. 你应该使用一个属性选择器来选择那些位于 `to-read` 类的 `tr` 元素内的、具有 `status` 类的`span` 元素，设置他们的 `border` 和 `background-image` 属性。
1. 你应该使用一个属性选择器来选择那些位于 `read` 类的 `tr` 元素内的、具有 `status` 类的`span` 元素，设置他们的 `border` 和 `background-image` 属性。
1. 你应该使用一个属性选择器来选择那些位于 `in-progress` 类的 `tr` 元素内的、具有 `status` 类的`span` 元素，设置他们的 `border` 和 `background-image` 属性。
1. 你应该使用一个属性选择器来选择具有 `status` 类的 `span` 元素，以及类值以 `rate` 开头的 `span` 元素，设置他们的 `height`、`width` 和 `padding` 属性。
1. 你应该使用一个属性选择器来选择那些作为 `span` 元素直接子元素的、且其 `class` 属性值以 `rate` 开头的 `span` 元素，设置他们的 `border`、`border-radius`、`margin`、`height`、`width` 和 `background-color` 等属性。
1. 你应该使用一个属性选择器来选择 `span` 元素的第一个子元素，这些子元素的 `class` 值中要包含单词 `one` ，然后将该元素的 `background-image` 属性设置为使用 `linear-gradient` 的值。
1. 你应该使用一个属性选择器来选择 `span` 元素的前两个子元素，这些子元素的 `class` 属性值中包含单词 `two`，将其 `background-image` 属性设置为使用 `linear-gradient` 的值。
1. 你应该使用一个属性选择器来选择位于具有包含单词 `three` 的 `class` 值的 `span` 元素内的三个 `span` 元素，将他们的 `background-image` 属性设置为使用 `linear-gradient` 的值。

# --hints--

你应该有一个 `h1` 元素，其中包含文本 `Book Inventory` 。

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

你应该只有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你应该有一个 `table` 元素。

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

你应该在 `table` 内有一个 `thead` 元素和一个 `tbody` 元素。

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

在 `thead` 内应有一个包含 5 个 `th` 元素的 `tr`。

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

你的第一列的标题应该是 `Title`。

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

你的第二列的标题应该是 `Author`。

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

你的第三列的标题应该是 `Category` 。

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

你的第四列的标题应该是 `Status`。

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

你的第五列的标题应该是 `Rate` 。

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

你的表格应该至少有四行。

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

每一行应始终有 5 列。

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

除了标题行之外的每一行表格都应该带有类`read`、`to-read` 或 `in-progress`。

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

`Status` 列中的 `td` 元素应当包含一个 `span` 元素。

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

`Status` 列中的每个 `span` 元素都应具有 `status` 类。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

每个 `.status` 元素应包含文本：`Read`、`To Read` 或 `In Progress`，具体取决于其所在行的类。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

`Rate` 列中的 `td` 元素应当包含一个 `span` 元素。

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

每个作为 `Rate` 列中 `td` 元素直接子元素的 `span` 元素，其第一个类应设置为 `rate` 。

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

每个 `.rate` 元素都应包含三个空的 `span` 元素。

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

位于 `.read` 行内的 `.rate` 元素应在 `rate` 类之后添加一个额外的类，该类的值为 `one`、`two` 或 `three`，具体取决于个人评分。

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

你应该有一个属性选择器，以具有 `read` 类的行为目标。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

你应该使用一个属性选择器来选择具有 `read` 类的行，将他们的 `background-image`属性设置为自定义的线性渐变效果。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

你应该有一个属性选择器，选择具有 `to-read` 类的行。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

你应该使用一个属性选择器来选择具有 `to-read` 类的行，将他们的 `background-image`属性设置为自定义的线性渐变效果。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

你应该有一个属性选择器，选择具有 `in-progress` 类的行。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

你应该使用一个属性选择器来选择具有 `in-progress` 类的行，将他们的 `background-image`属性设置为自定义的线性渐变效果。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

你应该将每个 `span` 元素的 `display` 属性设置为 `inline-block` 。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

你应该有一个属性选择器，选择具有 `status` 类且位于类是 `to-read` 的 `tr` 元素内的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

你应该使用一个属性选择器来选择那些位于 `to-read` 类的 `tr` 元素内的、具有 `status` 类的 `span` 元素，设置他们的 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

你应该使用一个属性选择器来选择那些位于 `to-read` 类的 `tr` 元素内的、具有 `status` 类的 `span` 元素，设置他们的 `background-image` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

你应该有一个属性选择器，选择具有 `status` 类且位于类是 `read` 的 `tr` 元素内的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

你应该使用一个属性选择器来选择那些位于 `read` 类的 `tr` 元素内的、具有 `status` 类的 `span` 元素，设置他们 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

你应该使用一个属性选择器来选择那些属于 `read` 类的 `tr` 元素下的、具有 `status` 类的 `span` 元素，设置 `background-image` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

你应该有一个属性选择器，选择具有 `status` 类且位于类是 `in-progress` 的 `tr` 元素内的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

你应该使用一个属性选择器来选择那些位于 `in-progress` 类的 `tr` 元素内的、具有 `status` 类的 `span` 元素，设置他们的 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

你应该使用一个属性选择器来选择那些位于 `in-progress` 类的 `tr` 元素内的、具有 `status` 类的 `span` 元素，设置他们的 `background-image` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

你应该有一个属性选择器，选择具有 `status` 类的 `span` 元素以及类值以 `rate` 开头的 `span` 元素。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

你应该使用一个属性选择器来选择具有 `status` 类的 `span` 元素，以及类值以 `rate` 开头的 `span` 元素，设置他们的 `height` 属性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

你应该使用一个属性选择器来选择具有 `status` 类的 `span` 元素，以及类值以 `rate` 开头的 `span` 元素，设置他们的 `width` 属性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

你应该使用一个属性选择器来选择具有 `status` 类的 `span` 元素，以及类值以 `rate` 开头的 `span` 元素，设置他们的 `padding` 属性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

你应该有一个属性选择器，选择那些直接位于 `class` 值是以 `rate` 开头 `span` 元素内的 `span` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

你应该使用一个属性选择器来选择那些直接位于 `class` 值是以 `rate` 开头 `span` 元素内的 `span` 元素，设置他们的 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

你应该使用一个属性选择器来选择那些直接位于 `class` 值是以 `rate` 开头 `span` 元素内的 `span` 元素，设置他们的 `border-radius` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

你应该使用一个属性选择器来选择那些直接位于 `class` 值是以 `rate` 开头 `span` 元素内的 `span` 元素，设置他们的 `margin` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

你应该使用一个属性选择器来选择那些直接位于 `class` 值是以 `rate` 开头 `span` 元素内的 `span` 元素，设置他们的 `height` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

你应该使用一个属性选择器来选择那些直接位于 `class` 值是以 `rate` 开头 `span` 元素内的 `span` 元素，设置他们的 `width` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

你应该使用一个属性选择器来选择那些直接位于 `class` 值是以 `rate` 开头 `span` 元素内的 `span` 元素，设置他们的 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

你应该有一个属性选择器，选择 `span` 元素的第一个子元素，这些子元素的 `class` 值中包含单词 `one` 。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你应该使用一个属性选择器来选择 `span` 元素的第一个子元素，这些子元素的 `class` 值中要包含单词 `one` ，然后将该元素的 `background-image` 属性设置为使用 `linear-gradient` 的值。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

你应该有一个属性选择器，选择 `span` 元素的前两个子元素，这些子元素的 `class` 值中包含单词 `two` 。

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你应该使用一个属性选择器来选择 `span` 元素的前两个子元素，这些子元素的 `class` 属性值中包含单词 `two`，将其 `background-image` 属性设置为使用 `linear-gradient` 的值。

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

你应该有一个属性选择器，选择那些位于具有 `three` 这个词作为其 `class` 值一部分的 `span` 元素内的 `span` 元素。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你应该使用一个属性选择器来选择位于具有包含单词 `three` 的 `class` 值的 `span` 元素内的`span`元素，将他们的 `background-image` 属性设置为使用 `linear-gradient` 的值。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
