---
id: 67116d7584d0b469b14579c0
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Створіть ще один об’єкт всередині масиву `library` з такими властивостями та значеннями:

| Властивість | Значення   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


Другим елементом вашого масиву `library` має бути об’єкт.

```js
assert.isObject(library[1]);
```

Ваш об’єкт має містити властивість `title`.

```js
assert.property(library[1], "title")
```

Властивість `title` має мати значення `"Atomic Habits"`.

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

Ваш об’єкт має містити властивість `author`.

```js
assert.property(library[1], "author")
```

Властивість `author` має мати значення `"James Clear"`.

```js
assert.propertyVal(library[1], "author", "James Clear");
```

Ваш об’єкт має містити властивість `about`.

```js
assert.property(library[1], "about")
```

Властивість `about` має мати значення `"A practical book about discarding bad habits and building good ones"`.

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

Ваш об’єкт має містити властивість `pages`.

```js
assert.property(library[1], "pages")
```

Значення властивості `pages` має бути числом.

```js
assert.isNumber(library[1]?.pages)
```

Властивість `pages` має бути встановлена в `320`.

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
