---
id: 67116d7584d0b469b14579bf
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Всередині масиву `library` створіть об’єкт із такими властивостями та значеннями:

| Властивість | Значення   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

У вас має бути об’єкт всередині масиву `library`.

```js
assert.isObject(library[0])
```

Ваш об’єкт має мати властивість `title`.

```js
assert.property(library[0], "title")
```

Значення властивості `title` має бути `"Your Next Five Moves: Master the Art of Business Strategy"`.

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

Ваш об’єкт має мати властивість `author`.

```js
assert.property(library[0], "author")
```

Властивість `author` має мати значення `"Patrick Bet-David and Greg Dinkin"`.

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

Ваш об’єкт має мати властивість `about`.

```js
assert.property(library[0], "about")
```

Властивість `about` має мати значення `"A book on how to plan ahead"`.

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

Ваш об’єкт має мати властивість `pages`.

```js
assert.property(library[0], "pages")
```

Значення властивості `pages` має бути числом.

```js
assert.isNumber(library[0]?.pages)
```

Властивість `pages` має бути встановлена в `320`.

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
