---
id: 67116c175a77cb64c5c10d49
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У попередніх уроках ви навчились працювати з функціями вищого порядку, такими як `map`, `filter` і `reduce`.

У цьому практичному занятті ви створите застосунок для керування бібліотекою, що дасть вам можливість попрактикуватися у роботі з різними функціями вищого порядку.

Щоб почати, створіть змінну з назвою `library` і призначте їй порожній масив.

# --hints--

Ви повинні створити масив `library`.

```js
assert.isArray(library)
```

Ваш масив `library` має бути порожнім.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
