---
id: 669e81368e52b3a5c35a2dc5
title: Costruisci una pagina di compilazione video
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `main` come unico figlio dell'elemento `body`.
1. Dovresti avere un elemento `h1` con l'argomento della tua pagina.
1. Dovresti avere un paragrafo che introduce l'argomento della tua pagina sotto il tuo elemento `h1`.
1. Dovresti avere tre elementi `section` sotto il tuo primo paragrafo.
1. Ogni sezione dovrebbe contenere un elemento `h2`, un paragrafo e un elemento `iframe`, in questo ordine.
1. I tre elementi `iframe` dovrebbero avere un attributo `src` impostato su un video valido.
1. Ogni elemento `iframe` dovrebbe avere anche un attributo `title` per descrivere il contenuto incorporato, e un attributo `height` e un attributo `width` per impostare l'elemento a una dimensione adeguata.

# --hints--

Dovresti avere un elemento `main` dentro il tuo elemento `body`.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

Il tuo elemento `main` dovrebbe essere l'unico figlio dell'elemento `body`.

```js
assert.equal(document.querySelector('body').children.length, 1);
```

Dovresti avere un elemento `h1` con l'argomento della tua pagina dentro l'elemento `main`.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Dovresti avere un paragrafo che introduce l'argomento della tua pagina sotto il tuo elemento `h1`.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

Dovresti avere tre elementi `section` sotto il tuo primo elemento `p`.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

Ogni elemento `section` dovrebbe iniziare con un elemento `h2` che funge da titolo per quella sezione.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

Ogni elemento `section` dovrebbe contenere un elemento `p` per introdurre il contenuto video come suo secondo figlio.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

Ogni elemento `section` dovrebbe contenere un elemento `iframe` come suo terzo figlio.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

Il tuo primo elemento `iframe` dovrebbe avere un attributo `src` impostato su un video valido.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

Il tuo primo elemento `iframe` dovrebbe avere un attributo `title` per descrivere il contenuto incorporato.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

Il tuo primo elemento `iframe` dovrebbe avere un attributo `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

Il tuo primo elemento `iframe` dovrebbe avere un attributo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

Il tuo secondo elemento `iframe` dovrebbe avere un attributo `src` impostato su un video valido.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

Il tuo secondo elemento `iframe` dovrebbe avere un attributo `title` per descrivere il contenuto incorporato.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

Il tuo secondo elemento `iframe` dovrebbe avere un attributo `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

Il tuo secondo elemento `iframe` dovrebbe avere un attributo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

Il tuo terzo elemento `iframe` dovrebbe avere un attributo `src` impostato su un video valido.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

Il tuo terzo elemento `iframe` dovrebbe avere un attributo `title` per descrivere il contenuto incorporato.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

Il tuo terzo elemento `iframe` dovrebbe avere un attributo `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

Il tuo terzo elemento `iframe` dovrebbe avere un attributo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front-End Web Development</h1>
        <p>Front-End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front-end
            developers ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
