---
id: 694470b9a65b278f981ec62e
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

最後に、`countup` 文の中で引数 `5` を使って `console.log()` 関数を呼び出し、動作を確認してください。

出力は `[1, 2, 3, 4, 5]` になるはずです。

# --hints--

`countup` 文の中で引数 `5` を使って `console.log()` 関数を呼び出してください。

```js
assert.deepStrictEqual(countup(5), [1, 2, 3, 4, 5]);
assert.match(code, /console\.log\s*\(\s*countup\s*\(\s*5\s*\)\s*\)/)
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
console.log(countup(5));
```
