---
id: 69446f18fdcb11aa516f7c5c
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、再帰を使って1から指定された数までの数字を含む配列を作成する`countup`関数を作成します。

まず、1つのパラメータを受け取る空の`countup`関数を作成してください。

# --hints--

`countup`という名前の関数が必要です。

```js
assert.isFunction(countup);
```

`countup`関数は1つのパラメータを受け取る必要があります。

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
