---
id: 69447054012a22f38c944ff0
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

`if` 関数の中に `countup` 文を追加してください。

`number` が `1` より小さいかどうかを確認し、空の配列を返してください。

これは再帰の基本ケースとして機能します。

# --hints--

`number` が `1` より小さい場合、関数は空の配列を返す必要があります。

```js
assert.deepStrictEqual(countup(0), []);
```

`number` が `1` より小さくない場合、関数は何も返してはいけません。

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
