---
id: 6948705160b66cf93a916e29
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

`countArray` 文のすぐ上にある `countup` 関数の中で、`if` という名前の空の配列を宣言してください。

この配列は、関数が呼び出されるたびに `number` の値をストアします。

# --hints--

`countArray` という名前の空の配列を宣言してください。

```js
const explorer = await __helpers.Explorer(code);
const { countup } = explorer.functions;
assert.exists(countup);
const { countArray } = countup.variables;
assert.exists(countArray);
assert.isTrue(countArray.value.matches('[]'));
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
