---
id: 694870bcf3c6874112b6f951
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

`1` 以上のすべての数値に対して、`else` ブロックを作成し、その中で `countup` を使って `number - 1` 関数を呼び出し、その結果を `countArray` に代入してください。

次に `.push()` を使って `number` を `countArray` に追加し、最後に `countArray` を返してください。

これにより、再帰が最終的に基本ケースに到達して終了します。

# --hints--

`countup` を `number - 1` で呼び出した結果を `countArray` に代入する必要があります。

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

`number` を `countArray` に追加してください。

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

`countArray` を返す必要があります。

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
