---
id: 67ec2f7924ade2450c06bcd8
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在之前的课程中，你学习了如何使用 `useEffect` Hook 处理副作用，比如数据获取。在本次研讨会中，你将构建一个应用，该应用根据用户输入从 API 获取水果数据并动态显示结果。

首先，让你的 `FruitsSearch` 组件返回一个 `id` 为 `search-container` 的 `div` 元素。

# --hints--

你应该有一个 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll("div"), 2);
```

你的 `div` 元素应有一个名为 `search-container` 的 `id`。

```js
const el = document.getElementById("search-container");
assert.exists(el);
assert.strictEqual(el.tagName, "DIV");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

--fcc-editable-region--
export function FruitsSearch() {
  
}
--fcc-editable-region--
```
