---
id: 587d7db2367417b2b2512b8b
title: Verstehen der sofort ausgeführten Funktionsausdrucks (IIFE)
challengeType: 1
forumTopicId: 301328
dashedName: understand-the-immediately-invoked-function-expression-iife
---

# --description--

Ein häufiges Muster in JavaScript ist es, eine Funktion sofort nach ihrer Deklaration auszuführen:

```js
(function () {
  console.log("Chirp, chirp!");
})();
```

Dies ist ein anonymer Funktionsausdruck, der sofort ausgeführt wird und `Chirp, chirp!` unmittelbar ausgibt.

Beachten Sie, dass die Funktion keinen Namen hat und nicht in einer Variablen gespeichert wird. Die zwei runden Klammern () am Ende des Funktionsausdrucks bewirken, dass sie sofort ausgeführt oder aufgerufen wird. Dieses Muster ist als <dfn>sofort ausgeführter Funktionsausdruck</dfn> oder <dfn>IIFE</dfn> bekannt.

# --instructions--

Schreiben Sie die Funktion `makeNest` so um, dass sie anonym ist und entfernen Sie den Aufruf, sodass sie stattdessen ein anonymes sofort ausgeführtes Funktionsausdruck (IIFE) ist.

# --hints--

Die Funktion sollte anonym sein.

```js
assert(/\((function|\(\))(=>|\(\)){?/.test(__helpers.removeJSComments(code).replace(/\s/g, '')));
```

Ihre Funktion sollte am Ende des Ausdrucks runde Klammern haben, um sie sofort aufzurufen.

```js
assert(/\(.*(\)\(|\}\(\))\)/.test(__helpers.removeJSComments(code).replace(/[\s;]/g, '')));
```

# --seed--

## --seed-contents--

```js
function makeNest() {
  console.log("A cozy nest is ready");
}

makeNest();
```

# --solutions--

```js
(function () {
  console.log("A cozy nest is ready");
})();
```

---

```js
(function () {
  console.log("A cozy nest is ready");
}());
```

---

```js
(() => {
  console.log("A cozy nest is ready");
})();
```

---

```js
(() =>
  console.log("A cozy nest is ready")
)();
```
