---
id: 587d7db1367417b2b2512b85
title: Setzen Sie das Prototype des Kindes auf eine Instanz des Elternteils
challengeType: 1
forumTopicId: 301325
dashedName: set-the-childs-prototype-to-an-instance-of-the-parent
---

# --description--

In der vorherigen Aufgabe haben Sie den ersten Schritt gesehen, um Verhalten vom Supertyp (oder Elternteil) `Animal` zu erben: eine neue Instanz von `Animal` zu erstellen.

Diese Aufgabe behandelt den nächsten Schritt: Setzen Sie das `prototype` des Subtyps (oder Kindes) – in diesem Fall `Bird` – auf eine Instanz von `Animal`.

```js
Bird.prototype = Object.create(Animal.prototype);
```

Denken Sie daran, dass das `prototype` wie das „Rezept“ zur Erstellung eines Objekts ist. In gewisser Weise enthält das Rezept für `Bird` jetzt alle wichtigen „Zutaten“ von `Animal`.

```js
let duck = new Bird("Donald");
duck.eat();
```

`duck` erbt alle Eigenschaften von `Animal`, einschließlich der `eat`-Methode.

# --instructions--

Ändern Sie den Code so, dass Instanzen von `Dog` von `Animal` erben.

# --hints--

`Dog.prototype` sollte eine Instanz von `Animal` sein.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }

// Only change code below this line


let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);

let beagle = new Dog();
beagle.eat();
```
