---
id: 587d7db0367417b2b2512b82
title: Die Prototypkette verstehen
challengeType: 1
forumTopicId: 301329
dashedName: understand-the-prototype-chain
---

# --description--

Alle Objekte in JavaScript (mit wenigen Ausnahmen) haben eine `prototype`. Außerdem ist das `prototype` eines Objekts selbst ein Objekt.

```js
function Bird(name) {
  this.name = name;
}

typeof Bird.prototype;
```

Da ein `prototype` ein Objekt ist, kann ein `prototype` seine eigene `prototype` haben! In diesem Fall ist das `prototype` von `Bird.prototype` `Object.prototype`:

```js
Object.prototype.isPrototypeOf(Bird.prototype);
```

Wie ist das nützlich? Sie erinnern sich vielleicht an die `hasOwnProperty`-Methode aus einer vorherigen Aufgabe:

```js
let duck = new Bird("Donald");
duck.hasOwnProperty("name");
```

Die `hasOwnProperty`-Methode ist in `Object.prototype` definiert, auf die über `Bird.prototype` zugegriffen werden kann, und von dort aus über `duck`. Dies ist ein Beispiel für die `prototype`-Kette. In dieser `prototype`-Kette ist `Bird` das `supertype` für `duck`, während `duck` das `subtype` ist. `Object` ist ein `supertype` für sowohl `Bird` als auch `duck`. `Object` ist ein `supertype` für alle Objekte in JavaScript. Daher kann jedes Objekt die `hasOwnProperty`-Methode verwenden.

# --instructions--

Ändern Sie den Code so, dass die korrekte Prototypkette angezeigt wird.

# --hints--

Ihr Code sollte zeigen, dass `Object.prototype` das Prototyp von `Dog.prototype` ist

```js
assert(/Object\.prototype\.isPrototypeOf/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

Dog.prototype.isPrototypeOf(beagle);  // yields true

// Fix the code below so that it evaluates to true
???.isPrototypeOf(Dog.prototype);
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
Object.prototype.isPrototypeOf(Dog.prototype);
```
