---
id: 587d7daf367417b2b2512b7f
title: Ändern Sie das Prototype auf ein neues Objekt
challengeType: 1
forumTopicId: 301316
dashedName: change-the-prototype-to-a-new-object
---

# --description--

Bis jetzt haben Sie Eigenschaften einzeln zum `prototype` hinzugefügt:

```js
Bird.prototype.numLegs = 2;
```

Das wird nach einigen Eigenschaften mühsam.

```js
Bird.prototype.eat = function() {
  console.log("nom nom nom");
}

Bird.prototype.describe = function() {
  console.log("My name is " + this.name);
}
```

Eine effizientere Methode ist, das `prototype` auf ein neues Objekt zu setzen, das bereits die Eigenschaften enthält. So werden die Eigenschaften auf einmal hinzugefügt:

```js
Bird.prototype = {
  numLegs: 2, 
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --instructions--

Fügen Sie die Eigenschaft `numLegs` und die beiden Methoden `eat()` und `describe()` zum `prototype` von `Dog` hinzu, indem Sie das `prototype` auf ein neues Objekt setzen.

# --hints--

`Dog.prototype` sollte auf ein neues Objekt gesetzt werden.

```js
assert(/Dog\.prototype\s*?=\s*?{/.test(__helpers.removeJSComments(code)));
```

`Dog.prototype` sollte die Eigenschaft `numLegs` haben.

```js
assert(Dog.prototype.numLegs !== undefined);
```

`Dog.prototype` sollte die Methode `eat()` haben.

```js
assert(typeof Dog.prototype.eat === 'function');
```

`Dog.prototype` sollte die Methode `describe()` haben.

```js
assert(typeof Dog.prototype.describe === 'function');
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype = {
  // Only change code below this line

};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
numLegs: 4,
  eat () {
    console.log('nom nom nom');
  },
  describe () {
    console.log('My name is ' + this.name);
  }
};
```
