---
id: 587d7dad367417b2b2512b75
title: Eine Methode in einem Objekt erstellen
challengeType: 1
forumTopicId: 301318
dashedName: create-a-method-on-an-object
---

# --description--

Objekte können eine besondere Art von Eigenschaft haben, die <dfn>Methode</dfn> genannt wird.

Methoden sind Eigenschaften, die Funktionen sind. Das fügt einem Objekt ein anderes Verhalten hinzu. Hier ist das `duck`-Beispiel mit einer Methode:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + duck.name + ".";}
};
duck.sayName();
```

Das Beispiel fügt die `sayName`-Methode hinzu, die eine Funktion ist, die einen Satz zurückgibt, der den Namen des `duck` angibt. Beachten Sie, dass die Methode in der return-Anweisung auf die `name`-Eigenschaft zugegriffen hat, und zwar mit `duck.name`. Die nächste Aufgabe behandelt eine andere Möglichkeit, dies zu tun.

# --instructions--

Verwenden Sie das `dog`-Objekt und geben Sie ihm eine Methode namens `sayLegs`. Die Methode soll den Satz `This dog has 4 legs.` zurückgeben.

# --hints--

`dog.sayLegs()` sollte eine Funktion sein.

```js
assert(typeof dog.sayLegs === 'function');
```

`dog.sayLegs()` sollte den angegebenen String zurückgeben – beachten Sie, dass Satzzeichen und Leerzeichen wichtig sind.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,

};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
