---
id: 587d7db1367417b2b2512b86
title: Zurücksetzen einer geerbten Konstruktor-Eigenschaft
challengeType: 1
forumTopicId: 301324
dashedName: reset-an-inherited-constructor-property
---

# --description--

Wenn ein Objekt seine `prototype` von einem anderen Objekt erbt, erbt es auch die Konstruktor-Eigenschaft des Supertyps.

Hier ist ein Beispiel:

```js
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
let duck = new Bird();
duck.constructor
```

Aber `duck` und alle Instanzen von `Bird` sollten zeigen, dass sie vom `Bird` und nicht vom `Animal` konstruiert wurden. Dazu können Sie die Konstruktor-Eigenschaft von `Bird` manuell auf das `Bird`-Objekt setzen:

```js
Bird.prototype.constructor = Bird;
duck.constructor
```

# --instructions--

Beheben Sie den Code so, dass `duck.constructor` und `beagle.constructor` ihre jeweiligen Konstruktoren zurückgeben.

# --hints--

`Bird.prototype` sollte eine Instanz von `Animal` sein.

```js
assert(Animal.prototype.isPrototypeOf(Bird.prototype));
```

`duck.constructor` sollte `Bird` zurückgeben.

```js
assert(duck.constructor === Bird);
```

`Dog.prototype` sollte eine Instanz von `Animal` sein.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

`beagle.constructor` sollte `Dog` zurückgeben.

```js
assert(beagle.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Animal() { }
function Bird() { }
function Dog() { }

Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);

// Only change code below this line



let duck = new Bird();
let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
function Bird() { }
function Dog() { }
Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Bird.prototype.constructor = Bird;
let duck = new Bird();
let beagle = new Dog();
```
