---
id: 587d7db0367417b2b2512b84
title: Vererben von Verhalten von einem Supertyp
challengeType: 1
forumTopicId: 301319
dashedName: inherit-behaviors-from-a-supertype
---

# --description--

In der vorherigen Aufgabe haben Sie eine `supertype` namens `Animal` erstellt, die Verhalten definiert, die alle Tiere gemeinsam haben:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
```

Diese und die nächste Aufgabe behandeln, wie Sie die Methoden von `Animal` innerhalb von `Bird` und `Dog` wiederverwenden können, ohne sie erneut zu definieren. Dabei wird eine Technik namens Vererbung verwendet. Diese Aufgabe behandelt den ersten Schritt: Erstellen Sie eine Instanz des `supertype` (oder Elternteils). Sie kennen bereits eine Möglichkeit, eine Instanz von `Animal` mit dem `new`-Operator zu erstellen:

```js
let animal = new Animal();
```

Es gibt einige Nachteile bei der Verwendung dieser Syntax für Vererbung, die für den Umfang dieser Aufgabe zu komplex sind. Stattdessen hier ein alternativer Ansatz ohne diese Nachteile:

```js
let animal = Object.create(Animal.prototype);
```

`Object.create(obj)` erstellt ein neues Objekt und setzt `obj` als den `prototype` des neuen Objekts. Denken Sie daran, dass der `prototype` wie das „Rezept“ zur Erstellung eines Objekts ist. Indem Sie den `prototype` von `animal` so setzen, dass er der `prototype` von `Animal` ist, geben Sie der `animal`-Instanz effektiv dasselbe „Rezept“ wie jeder anderen Instanz von `Animal`.

```js
animal.eat();
animal instanceof Animal;
```

Die `instanceof`-Methode hier würde `true` zurückgeben.

# --instructions--

Verwenden Sie `Object.create`, um zwei Instanzen von `Animal` mit den Namen `duck` und `beagle` zu erstellen.

# --hints--

Die `duck`-Variable sollte definiert sein.

```js
assert(typeof duck !== 'undefined');
```

Die `beagle`-Variable sollte definiert sein.

```js
assert(typeof beagle !== 'undefined');
```

Die `duck`-Variable sollte mit `Object.create` initialisiert werden.

```js
assert(
  /(let|const|var)\s{1,}duck\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

Die `beagle`-Variable sollte mit `Object.create` initialisiert werden.

```js
assert(
  /(let|const|var)\s{1,}beagle\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

`duck` sollte einen `prototype` von `Animal` haben.

```js
assert(duck instanceof Animal);
```

`beagle` sollte einen `prototype` von `Animal` haben.

```js
assert(beagle instanceof Animal);
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

// Only change code below this line

let duck; // Change this line
let beagle; // Change this line
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
let duck = Object.create(Animal.prototype);
let beagle = Object.create(Animal.prototype);

duck.eat();
beagle.eat();
```
