---
id: 587d7dae367417b2b2512b7a
title: Überprüfen Sie den Konstruktor eines Objekts mit instanceof
challengeType: 1
forumTopicId: 301337
dashedName: verify-an-objects-constructor-with-instanceof
---

# --description--

Jedes Mal, wenn eine Konstruktorfunktion ein neues Objekt erstellt, sagt man, dass dieses Objekt eine <dfn>Instanz</dfn> seines Konstruktors ist. JavaScript bietet eine praktische Möglichkeit, dies mit dem `instanceof`-Operator zu überprüfen. `instanceof` ermöglicht es Ihnen, ein Objekt mit einem Konstruktor zu vergleichen und gibt `true` oder `false` zurück, je nachdem, ob das Objekt mit diesem Konstruktor erstellt wurde oder nicht. Hier ist ein Beispiel:

```js
let Bird = function(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}

let crow = new Bird("Alexis", "black");

crow instanceof Bird;
```

Diese `instanceof`-Methode würde `true` zurückgeben.

Wenn ein Objekt ohne Verwendung eines Konstruktors erstellt wird, wird `instanceof` bestätigen, dass es keine Instanz dieses Konstruktors ist:

```js
let canary = {
  name: "Mildred",
  color: "Yellow",
  numLegs: 2
};

canary instanceof Bird;
```

Diese `instanceof`-Methode würde `false` zurückgeben.

# --instructions--

Erstellen Sie eine neue Instanz des `House`-Konstruktors, nennen Sie sie `myHouse` und übergeben Sie eine Anzahl von Schlafzimmern. Verwenden Sie dann `instanceof`, um zu überprüfen, ob es eine Instanz von `House` ist.

# --hints--

`myHouse` sollte ein `numBedrooms`-Attribut haben, das auf eine Zahl gesetzt ist.

```js
assert(typeof myHouse.numBedrooms === 'number');
```

Sie sollten überprüfen, ob `myHouse` eine Instanz von `House` ist, indem Sie den `instanceof`-Operator verwenden.

```js
assert(/myHouse\s*instanceof\s*House/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}

// Only change code below this line
```

# --solutions--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}
const myHouse = new House(4);
console.log(myHouse instanceof House);
```
