---
id: 587d7db1367417b2b2512b87
title: Methoden nach Vererbung hinzufügen
challengeType: 1
forumTopicId: 301315
dashedName: add-methods-after-inheritance
---

# --description--

Eine Konstruktorfunktion, die ihr `prototype` Objekt von einer Supertyp-Konstruktorfunktion erbt, kann zusätzlich zu den geerbten Methoden auch eigene Methoden haben.

Zum Beispiel ist `Bird` eine Konstruktorfunktion, die ihr `prototype` von `Animal` erbt:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
Bird.prototype.constructor = Bird;
```

Zusätzlich zu dem, was von `Animal` geerbt wird, möchten Sie Verhalten hinzufügen, das für `Bird` Objekte einzigartig ist. Hier erhält `Bird` eine `fly()` Funktion. Funktionen werden dem `Bird's` `prototype` auf dieselbe Weise hinzugefügt wie jeder Konstruktorfunktion:

```js
Bird.prototype.fly = function() {
  console.log("I'm flying!");
};
```

Jetzt haben Instanzen von `Bird` sowohl `eat()` als auch `fly()` Methoden:

```js
let duck = new Bird();
duck.eat();
duck.fly();
```

`duck.eat()` würde den String `nom nom nom` in der Konsole ausgeben, und `duck.fly()` würde den String `I'm flying!` ausgeben.

# --instructions--

Fügen Sie allen notwendigen Code hinzu, damit das `Dog` Objekt von `Animal` erbt und der `Dog` `prototype` Konstruktor auf `Dog` gesetzt wird. Fügen Sie dann eine `bark()` Methode zum `Dog` Objekt hinzu, sodass `beagle` sowohl `eat()` als auch `bark()` ausführen kann. Die `bark()` Methode soll `Woof!` in der Konsole ausgeben.

# --hints--

`Animal` sollte nicht auf die `bark()` Methode reagieren.

```js
assert(typeof Animal.prototype.bark == 'undefined');
```

`Dog` sollte die `eat()` Methode von `Animal` erben.

```js
assert(typeof Dog.prototype.eat == 'function');
```

Das `Dog` Prototyp sollte eine `bark()` Methode haben.

```js
assert('bark' in Dog.prototype);
```

`beagle` sollte ein `instanceof` `Animal` sein.

```js
assert(beagle instanceof Animal);
```

Der Konstruktor für `beagle` sollte auf `Dog` gesetzt werden.

```js
assert(beagle.constructor === Dog);
```

`beagle.eat()` sollte den String `nom nom nom` ausgeben.

```js
const spy = __helpers.spyOn(console, 'log');
beagle.eat();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'nom nom nom');
```

`beagle.bark()` sollte den String `Woof!` ausgeben.

```js
const spy = __helpers.spyOn(console, 'log');
beagle.bark();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'Woof!');
```

# --seed--

## --seed-contents--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }

// Only change code below this line




// Only change code above this line

let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Dog.prototype.bark = function () {
  console.log('Woof!');
};
let beagle = new Dog();

beagle.eat();
beagle.bark();
```
