---
id: 587d7daf367417b2b2512b7d
title: Alle Eigenschaften durchlaufen
challengeType: 1
forumTopicId: 301320
dashedName: iterate-over-all-properties
---

# --description--

Sie haben jetzt zwei Arten von Eigenschaften gesehen: <dfn>eigene Eigenschaften</dfn> und `prototype` Eigenschaften. Eigene Eigenschaften sind direkt auf der Objektinstanz selbst definiert. Und Prototyp-Eigenschaften sind auf dem `prototype` definiert.

```js
function Bird(name) {
  this.name = name;  //own property
}

Bird.prototype.numLegs = 2; // prototype property

let duck = new Bird("Donald");
```

So fügen Sie die eigenen Eigenschaften von `duck` zum Array `ownProps` und die `prototype` Eigenschaften zum Array `prototypeProps` hinzu:

```js
let ownProps = [];
let prototypeProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  } else {
    prototypeProps.push(property);
  }
}

console.log(ownProps);
console.log(prototypeProps);
```

`console.log(ownProps)` würde `["name"]` in der Konsole anzeigen, und `console.log(prototypeProps)` würde `["numLegs"]` anzeigen.

# --instructions--

Fügen Sie alle eigenen Eigenschaften von `beagle` zum Array `ownProps` hinzu. Fügen Sie alle `prototype` Eigenschaften von `Dog` zum Array `prototypeProps` hinzu.

# --hints--

Das `ownProps` Array sollte nur `name` enthalten.

```js
assert.deepEqual(ownProps, ['name']);
```

Das `prototypeProps` Array sollte nur `numLegs` enthalten.

```js
assert.deepEqual(prototypeProps, ['numLegs']);
```

Sie sollten diese Aufgabe lösen, ohne die eingebaute Methode `Object.keys()` zu verwenden.

```js
assert(!/\Object.keys/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];
for (let prop in beagle) {
  if (beagle.hasOwnProperty(prop)) {
    ownProps.push(prop);
  } else {
    prototypeProps.push(prop);
  }
}
```
