---
id: 587d7dad367417b2b2512b78
title: "Einen Konstruktor verwenden, um Objekte zu erstellen"
challengeType: 1
forumTopicId: 18233
dashedName: use-a-constructor-to-create-objects
---

# --description--

Hier ist der `Bird` Konstruktor aus der vorherigen Aufgabe:

```js
function Bird() {
  this.name = "Albert";
  this.color  = "blue";
  this.numLegs = 2;
}

let blueBird = new Bird();
```

**HINWEIS:** `this` innerhalb des Konstruktors bezieht sich immer auf das zu erstellende Objekt.

Beachten Sie, dass der `new` Operator beim Aufrufen eines Konstruktors verwendet wird. Dies weist JavaScript an, eine neue Instanz von `Bird` mit dem Namen `blueBird` zu erstellen. Ohne den `new` Operator würde `this` innerhalb des Konstruktors nicht auf das neu erstellte Objekt zeigen, was unerwartete Ergebnisse zur Folge hätte. Nun hat `blueBird` alle Eigenschaften, die im `Bird` Konstruktor definiert sind:

```js
blueBird.name;
blueBird.color;
blueBird.numLegs;
```

Wie bei jedem anderen Objekt können seine Eigenschaften aufgerufen und geändert werden:

```js
blueBird.name = 'Elvira';
blueBird.name;
```

# --instructions--

Verwenden Sie den `Dog` Konstruktor aus der letzten Lektion, um eine neue Instanz von `Dog` zu erstellen und weisen Sie sie einer Variablen `hound` zu.

# --hints--

`hound` sollte mit dem `Dog` Konstruktor erstellt werden.

```js
assert(hound instanceof Dog);
```

Ihr Code sollte den `new` Operator verwenden, um eine Instanz von `Dog` zu erstellen.

```js
assert(__helpers.removeJSComments(code).match(/new/g));
```

# --seed--

## --seed-contents--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
// Only change code below this line
```

# --solutions--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
const hound = new Dog();
```
