---
id: 587d7daf367417b2b2512b80
title: "Denken Sie daran, die Konstruktor-Eigenschaft beim Ändern des Prototyps zu setzen"
challengeType: 1
forumTopicId: 301323
dashedName: remember-to-set-the-constructor-property-when-changing-the-prototype
---

# --description--

Es gibt einen entscheidenden Nebeneffekt, wenn Sie den Prototyp manuell auf ein neues Objekt setzen. Dabei wird die `constructor`-Eigenschaft gelöscht! Diese Eigenschaft kann verwendet werden, um zu überprüfen, welche Konstruktorfunktion die Instanz erstellt hat. Da die Eigenschaft jedoch überschrieben wurde, liefert sie jetzt falsche Ergebnisse:

```js
duck.constructor === Bird;
duck.constructor === Object;
duck instanceof Bird;
```

Diese Ausdrücke würden der Reihe nach zu `false`, `true` und `true` ausgewertet werden.

Um das zu beheben, denken Sie daran, die `constructor`-Eigenschaft zu definieren, wann immer ein Prototyp manuell auf ein neues Objekt gesetzt wird:

```js
Bird.prototype = {
  constructor: Bird,
  numLegs: 2,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name); 
  }
};
```

# --instructions--

Definieren Sie die `constructor`-Eigenschaft auf dem `Dog`-`prototype`.

# --hints--

`Dog.prototype` sollte die `constructor`-Eigenschaft setzen.

```js
assert(Dog.prototype.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
Dog.prototype = {

  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
  constructor: Dog,
  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```
