---
id: 587d7db0367417b2b2512b81
title: "Verstehen, woher das Prototype eines Objekts stammt"
challengeType: 1
forumTopicId: 301330
dashedName: understand-where-an-objects-prototype-comes-from
---

# --description--

So wie Menschen Gene von ihren Eltern erben, erbt ein Objekt sein `prototype` direkt von der Konstruktorfunktion, die es erstellt hat. Zum Beispiel erstellt hier der `Bird`-Konstruktor das `duck`-Objekt:

```js
function Bird(name) {
  this.name = name;
}

let duck = new Bird("Donald");
```

`duck` erbt sein `prototype` von der `Bird`-Konstruktorfunktion. Sie können diese Beziehung mit der `isPrototypeOf`-Methode anzeigen:

```js
Bird.prototype.isPrototypeOf(duck);
```

Dies würde `true` zurückgeben.

# --instructions--

Verwenden Sie `isPrototypeOf`, um das `prototype` von `beagle` zu überprüfen.

# --hints--

Sie sollten zeigen, dass `Dog.prototype` das `prototype` von `beagle` ist.

```js
assert(/Dog\.prototype\.isPrototypeOf\(beagle\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
```
