---
id: 587d7db1367417b2b2512b88
title: Überschreiben geerbter Methoden
challengeType: 1
forumTopicId: 301322
dashedName: override-inherited-methods
---

# --description--

In vorherigen Lektionen haben Sie gelernt, dass ein Objekt sein Verhalten (Methoden) von einem anderen Objekt erben kann, indem es auf sein `prototype` Objekt verweist:

```js
ChildObject.prototype = Object.create(ParentObject.prototype);
```

Dann erhielt das `ChildObject` eigene Methoden, indem diese an sein `prototype` angehängt wurden:

```js
ChildObject.prototype.methodName = function() {...};
```

Es ist möglich, eine geerbte Methode zu überschreiben. Das geschieht auf dieselbe Weise – indem man eine Methode zum `ChildObject.prototype` hinzufügt, die denselben Methodennamen wie die zu überschreibende Methode verwendet. Hier ist ein Beispiel, wie `Bird` die von `eat()` geerbte `Animal` Methode überschreibt:

```js
function Animal() { }
Animal.prototype.eat = function() {
  return "nom nom nom";
};
function Bird() { }

Bird.prototype = Object.create(Animal.prototype);

Bird.prototype.eat = function() {
  return "peck peck peck";
};
```

Wenn Sie eine Instanz `let duck = new Bird();` haben und `duck.eat()` aufrufen, sucht JavaScript die Methode in der `prototype` Kette von `duck` folgendermaßen:

1.  `duck` => Ist `eat()` hier definiert? Nein.
2.  `Bird` => Ist `eat()` hier definiert? => Ja. Führen Sie sie aus und beenden Sie die Suche.
3.  `Animal` => `eat()` ist ebenfalls definiert, aber JavaScript hat die Suche beendet, bevor es diese Ebene erreicht hat.
4.  Object => JavaScript hat die Suche beendet, bevor es diese Ebene erreicht hat.

# --instructions--

Überschreiben Sie die `fly()` Methode für `Penguin`, sodass sie den String `Alas, this is a flightless bird.` zurückgibt.

# --hints--

`penguin.fly()` sollte den String `Alas, this is a flightless bird.` zurückgeben.

```js
assert(penguin.fly() === 'Alas, this is a flightless bird.');
```

Die `bird.fly()` Methode sollte den String `I am flying!` zurückgeben.

```js
assert(new Bird().fly() === 'I am flying!');
```

# --seed--

## --seed-contents--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;

// Only change code below this line



// Only change code above this line

let penguin = new Penguin();
console.log(penguin.fly());
```

# --solutions--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;
Penguin.prototype.fly = () => 'Alas, this is a flightless bird.';
let penguin = new Penguin();
console.log(penguin.fly());
```
