---
id: 587d7db2367417b2b2512b89
title: "Verwenden Sie ein Mixin, um gemeinsames Verhalten zwischen nicht verwandten Objekten hinzuzufügen"
challengeType: 1
forumTopicId: 301331
dashedName: use-a-mixin-to-add-common-behavior-between-unrelated-objects
---

# --description--

Wie Sie gesehen haben, wird Verhalten durch Vererbung geteilt. Es gibt jedoch Fälle, in denen Vererbung nicht die beste Lösung ist. Vererbung funktioniert nicht gut für nicht verwandte Objekte wie `Bird` und `Airplane`. Beide können fliegen, aber ein `Bird` ist kein Typ von `Airplane` und umgekehrt.

Für nicht verwandte Objekte ist es besser, <dfn>Mixins</dfn> zu verwenden. Ein Mixin ermöglicht es anderen Objekten, eine Sammlung von Funktionen zu nutzen.

```js
let flyMixin = function(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  }
};
```

Das `flyMixin` nimmt ein beliebiges Objekt und gibt ihm die `fly`-Methode.

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let plane = {
  model: "777",
  numPassengers: 524
};

flyMixin(bird);
flyMixin(plane);
```

Hier werden `bird` und `plane` an `flyMixin` übergeben, das dann die `fly`-Funktion jedem Objekt zuweist. Nun können sowohl `bird` als auch `plane` fliegen:

```js
bird.fly();
plane.fly();
```

Die Konsole würde den String `Flying, wooosh!` zweimal anzeigen, einmal für jeden `.fly()`-Aufruf.

Beachten Sie, wie das Mixin erlaubt, dieselbe `fly`-Methode von den nicht verwandten Objekten `bird` und `plane` wiederzuverwenden.

# --instructions--

Erstellen Sie ein Mixin namens `glideMixin`, das eine Methode namens `glide` definiert. Verwenden Sie dann das `glideMixin`, um sowohl `bird` als auch `boat` die Fähigkeit zum Gleiten zuzuweisen.

# --hints--

Ihr Code sollte eine `glideMixin`-Variable deklarieren, die eine Funktion ist.

```js
assert(typeof glideMixin === 'function');
```

Ihr Code sollte das `glideMixin` am `bird`-Objekt verwenden, um ihm die `glide`-Methode zu geben.

```js
assert(typeof bird.glide === 'function');
```

Ihr Code sollte das `glideMixin` am `boat`-Objekt verwenden, um ihm die `glide`-Methode zu geben.

```js
assert(typeof boat.glide === 'function');
```

# --seed--

## --seed-contents--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};

// Only change code below this line
```

# --solutions--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};
function glideMixin (obj) {
  obj.glide = () => 'Gliding!';
}

glideMixin(bird);
glideMixin(boat);
```
