---
id: 587d7db0367417b2b2512b83
title: "Vererbung verwenden, damit Sie sich nicht wiederholen"
challengeType: 1
forumTopicId: 301334
dashedName: use-inheritance-so-you-dont-repeat-yourself
---

# --description--

Es gibt ein Prinzip in der Programmierung, das <dfn>Don't Repeat Yourself (DRY)</dfn> heißt. Der Grund, warum sich wiederholender Code ein Problem darstellt, ist, dass jede Änderung erfordert, den Code an mehreren Stellen zu korrigieren. Das bedeutet in der Regel mehr Arbeit für Programmierende und mehr Raum für Fehler.

Beachten Sie im folgenden Beispiel, dass die `describe`-Methode von `Bird` und `Dog` gemeinsam genutzt wird:

```js
Bird.prototype = {
  constructor: Bird,
  describe: function() {
    console.log("My name is " + this.name);
  }
};

Dog.prototype = {
  constructor: Dog,
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

Die `describe`-Methode wird an zwei Stellen wiederholt. Der Code kann so bearbeitet werden, dass er dem DRY-Prinzip folgt, indem eine `supertype` (oder Elternklasse) namens `Animal` erstellt wird:

```js
function Animal() { };

Animal.prototype = {
  constructor: Animal, 
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

Da `Animal` die `describe`-Methode enthält, können Sie sie aus `Bird` und `Dog` entfernen:

```js
Bird.prototype = {
  constructor: Bird
};

Dog.prototype = {
  constructor: Dog
};
```

# --instructions--

Die `eat`-Methode wird sowohl in `Cat` als auch in `Bear` wiederholt. Bearbeiten Sie den Code im Sinne von DRY, indem Sie die `eat`-Methode in die `Animal`-`supertype` verschieben.

# --hints--

`Animal.prototype` sollte die `eat`-Eigenschaft haben.

```js
assert(Animal.prototype.hasOwnProperty('eat'));
```

`Bear.prototype` sollte nicht die `eat`-Eigenschaft haben.

```js
assert(!Bear.prototype.hasOwnProperty('eat'));
```

`Cat.prototype` sollte nicht die `eat`-Eigenschaft haben.

```js
assert(!Cat.prototype.hasOwnProperty('eat'));
```

# --seed--

## --seed-contents--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,

};
```

# --solutions--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
```
