---
id: 587d7dac367417b2b2512b74
title: "Punktnotation verwenden, um auf die Eigenschaften eines Objekts zuzugreifen"
challengeType: 1
forumTopicId: 301333
dashedName: use-dot-notation-to-access-the-properties-of-an-object
---

# --description--

Die letzte Aufgabe hat ein Objekt mit verschiedenen Eigenschaften erstellt. Jetzt sehen Sie, wie Sie auf die Werte dieser Eigenschaften zugreifen. Hier ist ein Beispiel:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
console.log(duck.name);
```

Die Punktnotation wird am Objektnamen `duck` verwendet, gefolgt vom Namen der Eigenschaft `name`, um auf den Wert von `Aflac` zuzugreifen.

# --instructions--

Geben Sie beide Eigenschaften des Objekts `dog` in Ihrer Konsole aus.

# --hints--

Ihr Code sollte `console.log` verwenden, um den Wert der Eigenschaft `name` des Objekts `dog` auszugeben.

```js
assert(/console.log\(.*dog\.name.*\)/g.test(__helpers.removeJSComments(code)));
```

Ihr Code sollte `console.log` verwenden, um den Wert der Eigenschaft `numLegs` des Objekts `dog` auszugeben.

```js
assert(/console.log\(.*dog\.numLegs.*\)/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
// Only change code below this line
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
console.log(dog.name);
console.log(dog.numLegs);
```
