---
id: 587d7dae367417b2b2512b7c
title: "Prototyp-Eigenschaften verwenden, um doppelten Code zu reduzieren"
challengeType: 1
forumTopicId: 301336
dashedName: use-prototype-properties-to-reduce-duplicate-code
---

# --description--

Da `numLegs` wahrscheinlich für alle Instanzen von `Bird` denselben Wert hat, besitzen Sie im Grunde eine doppelte Variable `numLegs` innerhalb jeder `Bird`-Instanz.

Das ist vielleicht kein Problem, wenn es nur zwei Instanzen gibt, aber stellen Sie sich vor, es wären Millionen von Instanzen. Das wären sehr viele doppelte Variablen.

Eine bessere Methode ist, die `prototype` von `Bird` zu verwenden. Eigenschaften im `prototype` werden unter ALLEN Instanzen von `Bird` geteilt. So fügen Sie `numLegs` zum `Bird prototype` hinzu:

```js
Bird.prototype.numLegs = 2;
```

Jetzt haben alle Instanzen von `Bird` die `numLegs`-Eigenschaft.

```js
console.log(duck.numLegs);
console.log(canary.numLegs);
```

Da alle Instanzen automatisch die Eigenschaften auf dem `prototype` haben, können Sie sich einen `prototype` als „Rezept“ zur Erstellung von Objekten vorstellen. Beachten Sie, dass der `prototype` für `duck` und `canary` Teil des `Bird`-Konstruktors als `Bird.prototype` ist.

# --instructions--

Fügen Sie dem `numLegs` von `prototype` eine `Dog`-Eigenschaft hinzu

# --hints--

`beagle` sollte eine `numLegs`-Eigenschaft haben.

```js
assert(beagle.numLegs !== undefined);
```

`beagle.numLegs` sollte eine Zahl sein.

```js
assert(typeof beagle.numLegs === 'number');
```

`numLegs` sollte eine `prototype`-Eigenschaft sein, keine eigene Eigenschaft.

```js
assert(beagle.hasOwnProperty('numLegs') === false);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}



// Only change code above this line
let beagle = new Dog("Snoopy");
```

# --solutions--

```js
function Dog (name) {
  this.name = name;
}
Dog.prototype.numLegs = 4;
let beagle = new Dog("Snoopy");
```
