---
id: 587d7db2367417b2b2512b8c
title: "Verwenden Sie eine IIFE, um ein Modul zu erstellen"
challengeType: 1
forumTopicId: 301332
dashedName: use-an-iife-to-create-a-module
---

# --description--

Ein sofort ausgeführter Funktionsausdruck (IIFE) wird häufig verwendet, um zusammengehörige Funktionalitäten in einem einzigen Objekt oder <dfn>Modul</dfn> zu bündeln. Zum Beispiel definierte eine frühere Aufgabe zwei Mixins:

```js
function glideMixin(obj) {
  obj.glide = function() {
    console.log("Gliding on the water");
  };
}
function flyMixin(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  };
}
```

Wir können diese Mixins wie folgt in einem Modul zusammenfassen:

```js
let motionModule = (function () {
  return {
    glideMixin: function(obj) {
      obj.glide = function() {
        console.log("Gliding on the water");
      };
    },
    flyMixin: function(obj) {
      obj.fly = function() {
        console.log("Flying, wooosh!");
      };
    }
  }
})();
```

Beachten Sie, dass Sie einen sofort ausgeführten Funktionsausdruck (IIFE) haben, der ein Objekt zurückgibt `motionModule`. Dieses zurückgegebene Objekt enthält alle Mixin-Verhalten als Eigenschaften des Objekts. Der Vorteil des Modul-Musters besteht darin, dass alle Bewegungsverhalten in einem einzigen Objekt verpackt werden können, das dann von anderen Teilen Ihres Codes verwendet werden kann. Hier ist ein Beispiel für die Verwendung:

```js
motionModule.glideMixin(duck);
duck.glide();
```

# --instructions--

Erstellen Sie ein Modul namens `funModule`, um die beiden Mixins `isCuteMixin` und `singMixin` zu umschließen. `funModule` sollte ein Objekt zurückgeben.

# --hints--

`funModule` sollte definiert sein und ein Objekt zurückgeben.

```js
assert(typeof funModule === 'object');
```

`funModule.isCuteMixin` sollte auf eine Funktion zugreifen.

```js
assert(typeof funModule.isCuteMixin === 'function');
```

`funModule.singMixin` sollte auf eine Funktion zugreifen.

```js
assert(typeof funModule.singMixin === 'function');
```

# --seed--

## --seed-contents--

```js
let isCuteMixin = function(obj) {
  obj.isCute = function() {
    return true;
  };
};
let singMixin = function(obj) {
  obj.sing = function() {
    console.log("Singing to an awesome tune");
  };
};
```

# --solutions--

```js
const funModule = (function () {
  return {
    isCuteMixin: obj => {
      obj.isCute = () => true;
    },
    singMixin: obj => {
      obj.sing = () => console.log("Singing to an awesome tune");
    }
  };
})();
```
