---
id: 587d7fb2367417b2b2512bf5
title: Отримати параметр маршруту від клієнта
challengeType: 2
forumTopicId: 301513
dashedName: get-route-parameter-input-from-the-client
---

# --description--

Під час створення API ми маємо дозволити користувачам повідомляти нам, що саме вони хочуть отримати від нашої служби. Наприклад, якщо клієнт запитує інформацію про користувача, збереженого в базі даних, йому потрібен спосіб повідомити, про якого користувача йде мова. Один із можливих способів досягти цього — використання параметрів маршруту. Параметри маршруту — це іменовані сегменти URL, розділені косими рисками (/). Кожен сегмент захоплює значення частини URL, яка відповідає його позиції. Захоплені значення можна знайти в об’єкті `req.params`.

<blockquote>route_path: '/user/:userId/book/:bookId'<br>actual_request_URL: '/user/546/book/6754' <br>req.params: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

Побудуйте сервер-ехо, змонтований на маршруті `GET /:word/echo`. Відповідайте JSON-об’єктом зі структурою `{echo: word}`. Слово, яке потрібно повторити, можна знайти в `req.params.word`. Ви можете протестувати свій маршрут у адресному рядку браузера, відвідавши відповідні маршрути, наприклад `your-app-rootpath/freecodecamp/echo`.

# --hints--

Test 1 : Ваш сервер-ехо має правильно повторювати слова

```js
  const response = await fetch(code + '/eChOtEsT/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'eChOtEsT',
    'Test 1: the echo server is not working as expected'
  );
```

Test 2 : Ваш сервер-ехо має правильно повторювати слова

```js
  const response = await fetch(code + '/ech0-t3st/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'ech0-t3st',
    'Test 2: the echo server is not working as expected'
  );
```
