---
id: 587d7fb2367417b2b2512bf6
title: Отримання вхідних даних параметра запиту від клієнта
challengeType: 2
forumTopicId: 301512
dashedName: get-query-parameter-input-from-the-client
---

# --description--

Ще один поширений спосіб отримати вхідні дані від клієнта — це кодувати їх після шляху маршруту, використовуючи рядок запиту. Рядок запиту відокремлюється знаком питання (?), і містить пари поле=значення. Кожна пара розділяється амперсандом (&). Express може аналізувати дані з рядка запиту і заповнювати об’єкт `req.query`. Деякі символи, як-от відсоток (%), не можуть бути у URL і мають бути закодовані в іншому форматі перед відправленням. Якщо ви використовуєте API з JavaScript, ви можете застосовувати спеціальні методи для кодування/декодування цих символів.

<blockquote>route_path: '/library'<br>actual_request_URL: '/library?userId=546&#x26;bookId=6754' <br>req.query: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

Створіть кінцеву точку API, змонтовану на `GET /name`. Відповідайте JSON-документом, що має структуру `{ name: 'firstname lastname'}`. Параметри імені та прізвища мають бути закодовані в рядку запиту, наприклад `?first=firstname&last=lastname`.

**Примітка:** У наступному завданні ви отримаєте дані з POST-запиту на тому ж шляху маршруту `/name`. Якщо хочете, можете використати метод `app.route(path).get(handler).post(handler)`. Цей синтаксис дозволяє ланцюжити різні обробники HTTP-методів на одному маршруті. Це допоможе трохи зекономити набір коду і зробити його чистішим.

# --hints--

Test 1 : Ваша кінцева точка API має відповідати `{ "name": "Mick Jagger" }`, коли викликається кінцева точка `/name` з `?first=Mick&last=Jagger`

```js
  const response = await fetch(code + '/name?first=Mick&last=Jagger');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "GET /name" route does not behave as expected'
  );
```

Test 2 : Ваша кінцева точка API має відповідати `{ "name": "Keith Richards" }`, коли викликається кінцева точка `/name` з `?first=Keith&last=Richards`

```js
  const response = await fetch(code + '/name?last=Richards&first=Keith');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "GET /name" route does not behave as expected'
  );
```
