---
id: 587d7fb2367417b2b2512bf8
title: Отримання даних із POST-запитів
challengeType: 2
forumTopicId: 301511
dashedName: get-data-from-post-requests
---

# --description--

Створіть обробник POST на шляху `/name`. Це той самий шлях, що й раніше. Ми підготували форму на html головній сторінці. Вона відправить ті самі дані, що й у завданні 10 (Query string). Якщо body-parser налаштований правильно, ви повинні знайти параметри в об’єкті `req.body`. Ось типовий приклад із бібліотеки:

<blockquote>маршрут: POST '/library'<br>urlencoded_body: userId=546&#x26;bookId=6754 <br>req.body: {userId: '546', bookId: '6754'}</blockquote>

Відповідайте тим самим JSON-об’єктом, що й раніше: `{name: 'firstname lastname'}`. Перевірте, чи працює ваша кінцева точка, використовуючи html-форму, яку ми надали на головній сторінці застосунку.

Підказка: Існують кілька інших http-методів, окрім GET і POST. За конвенцією існує відповідність між http-дієсловом і операцією, яку ви виконуєте на сервері. Звичайне відображення таке:

POST (іноді PUT) — створити новий ресурс, використовуючи інформацію, надіслану з запитом,

GET — прочитати існуючий ресурс без змін,

PUT або PATCH (іноді POST) — оновити ресурс, використовуючи надіслані дані,

DELETE — видалити ресурс.

Також є кілька інших методів, які використовуються для узгодження з’єднання із сервером. Окрім GET, усі інші методи, перелічені вище, можуть містити корисне навантаження (тобто дані в тілі запиту). Підпрограмне забезпечення body-parser працює з цими методами також.

# --hints--

Test 1 : Ваша кінцева точка API має відповідати з правильним іменем

```js
  const response = await fetch(code + '/name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/x-www-form-urlencoded'
    },
    body: new URLSearchParams({ first: 'Mick', last: 'Jagger' })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "POST /name" route does not behave as expected'
  );
```

Test 2 : Ваша кінцева точка API має відповідати з правильним іменем

```js
  const response = await fetch(code + '/name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/x-www-form-urlencoded'
    },
    body: new URLSearchParams({
      first: 'Keith',
      last: 'Richards'
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "POST /name" route does not behave as expected'
  );
```
