---
id: 587d7fb0367417b2b2512bef
title: Віддавати HTML-файл
challengeType: 2
forumTopicId: 301516
dashedName: serve-an-html-file
---

# --description--

Ви можете відповідати на запити файлом, використовуючи метод `res.sendFile(path)`. Його можна розмістити всередині обробника маршруту `app.get('/', ...)`. Під капотом цей метод встановить відповідні заголовки, щоб браузер знав, як обробляти файл, який ви хочете надіслати, залежно від його типу. Потім він прочитає і надішле файл. Цей метод потребує абсолютного шляху до файлу. Рекомендуємо використовувати глобальну змінну Node `__dirname` для обчислення шляху ось так:

```js
absolutePath = __dirname + '/relativePath/file.ext'
```

# --instructions--

Надішліть файл `/views/index.html` у відповідь на GET-запити за шляхом `/`. Якщо ви переглянете свій живий застосунок, має з’явитися великий HTML-заголовок (і форма, яку ми використаємо пізніше…), без застосованих стилів.

**Примітка:** Ви можете відредагувати рішення попереднього завдання або створити нове. Якщо створюєте нове рішення, пам’ятайте, що Express оцінює маршрути зверху вниз і виконує обробник для першого збігу. Вам потрібно перетворити на коментарі попереднє рішення, інакше сервер продовжить відповідати рядком.

# --hints--

Ваш застосунок має віддавати файл views/index.html

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /<h1>.*<\/h1>/,
    'Your app does not serve the expected HTML'
  );
```
