---
id: 66c057041df6394ca796bf33
title: Crie um construtor de frases
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

Neste laboratório, você criará duas histórias diferentes usando um modelo de frase. Você usará variáveis para armazenar diferentes partes da história e, em seguida, exibirá as histórias no console.

Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve declarar as seguintes variáveis usando `let`:

- `adjective`
- `noun`
- `verb`
- `place`
- `adjective2`
- `noun2`

1. Você deve atribuir às variáveis acima alguns valores de string à sua escolha.

2. Você deve declarar uma variável `firstStory` usando `const` para armazenar a primeira história.

3. Você deve usar o seguinte modelo de história para criar a primeira história e atribuí-la à variável `firstStory`:

`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`;

1. Você deve exibir a primeira história no console usando a mensagem `"First story: [firstStory]"`.

1. Você deve atribuir novos valores às variáveis `adjective`, `noun`, `verb`, `place`, `adjective2` e `noun2`.

1. Crie outra história usando o mesmo modelo e atribua-a à variável `secondStory`.

1. Você deve exibir a segunda história no console usando a mensagem `"Second story: [secondStory]"`.


# --hints--

Você deve declarar uma variável `adjective`.

```js
assert.isNotNull(adjective);
```

Você deve declarar uma variável `noun`.

```js
assert.isNotNull(noun);
```

Você deve declarar uma variável `verb`.

```js
assert.isNotNull(verb);
```

Você deve declarar uma variável `place`.

```js
assert.isNotNull(place);
```

Você deve declarar uma variável `adjective2`.

```js
assert.isNotNull(adjective2);
```

Você deve declarar uma variável `noun2`.

```js
assert.isNotNull(noun2);
```

Você deve atribuir um valor de string à variável `adjective`.

```js
assert.isString(adjective);
```

Você deve atribuir um valor de string à variável `noun`.

```js
assert.isString(noun);
```

Você deve atribuir um valor de string à variável `verb`.

```js
assert.isString(verb);
```

Você deve atribuir um valor de string à variável `place`.

```js
assert.isString(place);
```

Você deve atribuir um valor de string à variável `adjective2`.

```js
assert.isString(adjective2);
```

Você deve atribuir um valor de string à variável `noun2`.

```js
assert.isString(noun2);
```

Você deve declarar uma variável `firstStory`.

```js
assert.isNotNull(firstStory);
```

You should use the correct story format for the first story: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Preste atenção nos espaços.

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*)\1;?\s*`));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

Você deve registrar a primeira história usando a mensagem `"First story: [firstStory]"`.

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

Você deve declarar uma variável `secondStory`.

```js
assert.isNotNull(secondStory);
```

Você deve reatribuir a variável `adjective` para a segunda história.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

Você deve reatribuir a variável `noun` para a segunda história.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

Você deve reatribuir a variável `verb` para a segunda história.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

Você deve reatribuir a variável `place` para a segunda história.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

Você deve reatribuir a variável `adjective2` para a segunda história.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

Você deve reatribuir a variável `noun2` para a segunda história.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

You should use the correct story format for the second story: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Pay attention to spaces.

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

Você deve registrar a segunda história usando o formato `"Second story: [secondStory]"`.

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

A `firstStory` não deve ser igual à `secondStory`.

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```

