---
id: 6720b3191bbafd269bae0ae8
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Створіть елемент `div`, який має `id` зі значенням `greeting-card` та `class` зі значенням `card`.

Всередині елемента `div` додайте `h1` з текстом `Happy Birthday!`. Потім додайте абзац з `class`, який має назву `message`, і текстом `Wishing you all the happiness and joy on your special day!`.

# --hints--

У вас має бути елемент `div`.

```js
assert.exists(document.querySelector("div"));
```

Елемент `div` повинен мати `id` зі значенням `greeting-card`.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

Елемент `div` повинен мати `class` зі значенням `card`.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

У вас має бути елемент `h1` всередині елемента `div`.

```js
assert.exists(document.querySelector("div > h1"));
```

Елемент `h1` повинен містити текст `Happy Birthday!`.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

У вас має бути елемент `p` всередині `div`.

```js
assert.exists(document.querySelector("div > p"));
```

Елемент `p` повинен мати `class` зі значенням `message`.

```js
assert.exists(document.querySelector("div > p.message"));
```

Елемент `p` повинен містити текст `Wishing you all the happiness and joy on your special day!`.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
