---
id: 6720b1f2680ec706ed534acf
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з різними псевдокласами та псевдоелементами, створюючи листівку-привітання. Для вас уже підготовлено шаблонний код HTML.

Почніть практичне заняття з підключення файлу `styles.css`.

# --hints--

Ваш елемент `link` має бути всередині елемента `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Ваш елемент `link` має мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Ваш елемент `link` має мати атрибут `href` зі значенням `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
