---
id: 6720b7f3155a9278d6b6a8b8
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Тепер настав час стилізувати вашу листівку.

Додайте селектор для елемента `body`, потім:

- змініть `font-family` на `Arial`, а потім додайте загальний `sans-serif`,
- задайте відступ з усіх боків у розмірі `40px`,
- встановіть властивість `text-align` зі значенням `center`.

# --hints--

У вас має бути селектор `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Ви маєте встановити `font-family` у значення `"Arial", sans-serif` у селекторі `body`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

Ви маєте встановити `padding` у значення `40px` у селекторі `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

Ви маєте встановити `text-align` у значення `center` у селекторі `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
