---
id: 672b4434dd6d6a1edc405419
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Додайте два елементи `a` всередину елемента `.card-links`.

Перший має містити текст `Send Card`, атрибут `class` зі значенням `send-link` та атрибут `href` зі значенням `#send`.

Другий має містити текст `Share on Social Media`, атрибут `class` зі значенням `share-link` та атрибут `href` зі значенням `#share`.

# --hints--

Перший елемент `a` всередині `.card-links` має містити текст `Send Card`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

Перший елемент `a` має мати атрибут `class` зі значенням `send-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

Перший елемент `a` має мати атрибут `href` зі значенням `#send`.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

Другий елемент `a` всередині `.card-links` має містити текст `Share on Social Media`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

Другий елемент `a` має мати атрибут `class` зі значенням `share-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

Другий елемент `a` має мати атрибут `href` зі значенням `#share`.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
