---
id: 67332db0f793bf08f2a18528
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Час заповнити другий `section`!

Додайте `h2` елемент до `#share` елемента, який містить текст `Sharing your card...`, а потім додайте `p` елемент з текстом `Your card was shared on social media!`.

# --hints--

`#share` елемент повинен містити `h2` елемент.

```js
assert.exists(document.querySelector("#share > h2"));
```

`h2` елемент повинен містити текст `Sharing your card...`.

```js
assert.strictEqual(document.querySelector("#share > h2").innerText.trim(), "Sharing your card...");
```

Другий елемент всередині `#share` має бути `p` елементом.

```js
assert.strictEqual(document.querySelector("#share").children[1].tagName, "P");
```

`p` елемент повинен мати текст `Your card was shared on social media!`.

```js
assert.strictEqual(document.querySelector("#share").children[1].innerText.trim(), "Your card was shared on social media!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
