---
id: 686daa7ed79ceacd0b264e7d
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Unapaswa kupanga ndani kipengele kimoja cha `path` ndani ya kipengele chako cha `svg` ili kutoa umbo la picha.

Unda kipengele cha `path`.

# --hints--

Unapaswa kuwa na kipengele cha `path` kilichopangwa ndani ya kipengele chako cha `svg`.

```js
const path = document.querySelector('svg path');
assert.exists(path);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    <svg>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </svg>
  </body>
</html>
```
