---
id: 686d749d6f32f708fdf2b7ad
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika funzo lililopita, ulijifunza kuhusu vipengele vya `svg` na jinsi vinavyotumika mara nyingi kama alama katika miradi. Katika misimbo halisi ya dunia, kawaida unategemea maktaba za alama, hivyo huna haja ya kuunda vipengele vya `svg` kutoka mwanzo. Hata hivyo, katika warsha hii, utajenga alama ya moyo ili ujifunze kuhusu sifa za msingi zinazotumika ndani ya kipengele cha `svg`.

Anza kwa kuunda kipengele cha `svg` kwenye ukurasa.

# --hints--

Unapaswa kuwa na kipengele cha `svg` kwenye ukurasa.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
