---
id: 686daa7ed79ceacd0b264e80
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Unakaribia zaidi. Sasa, angalia mfano huu:

```html
<svg viewBox="0 0 50 50">
</svg>
```

Sifa ya kipengele `viewBox` inadhibiti sehemu gani ya picha inaonekana ndani ya SVG.

- Nambari mbili za kwanza (`0 0`) zinaweka nafasi ya kuanzia ya `viewBox` — kona ya juu-kushoto (x na y).
- Nambari mbili zinazofuata (`50 50`) zinafafanua upana na urefu wa `viewBox`.

Weka sifa ya kipengele `viewBox` kuwa `0 0 24 24`.

# --hints--

Unapaswa kuwa na sifa ya kipengele `viewBox`.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
assert.exists(viewBox);
```

Unapaswa kuweka nafasi ya x ya `viewBox` kuwa 0.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const x = viewBox?.trim().split(' ')[0];
assert.strictEqual(x, '0');
```

Unapaswa kuweka nafasi ya y ya `viewBox` kuwa 0.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const y = viewBox?.trim().split(' ')[1];
assert.strictEqual(y, '0');
```

Unapaswa kuweka upana wa `viewBox` kuwa 24.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const width = viewBox?.trim().split(' ')[2];
assert.strictEqual(width, '24');
```

Unapaswa kuweka urefu wa `viewBox` kuwa 24.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const height = viewBox?.trim().split(' ')[3];
assert.strictEqual(height, '24');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24">
    --fcc-editable-region--
      <path
        d="M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z"
      ></path>
    </svg>
  </body>
</html>
```
