---
id: 69446f18fdcb11aa516f7c5c
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 재귀를 사용해 1부터 주어진 숫자까지의 숫자를 포함하는 배열을 생성하는 `countup` 함수를 만듭니다.

먼저 하나의 매개변수를 받는 빈 `countup` 함수를 만드세요.

# --hints--

`countup`라는 함수가 있어야 합니다.

```js
assert.isFunction(countup);
```

`countup` 함수는 하나의 매개변수를 받아야 합니다.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
