---
id: 69447054012a22f38c944ff0
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

`if` 함수 안에 `countup` 문장을 추가하세요.

`number`가 `1`보다 작은지 확인하고 빈 배열을 반환하세요.

이것이 재귀의 기본 사례가 될 것입니다.

# --hints--

`number`가 `1`보다 작을 때 함수는 빈 배열을 반환해야 합니다.

```js
assert.deepStrictEqual(countup(0), []);
```

`number`가 `1`보다 **작지 않을 때** 함수는 **아무것도 반환하지 않아야** 합니다.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
