---
id: 694470b9a65b278f981ec62e
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

마지막으로, `countup` 문장 안에서 인수 `5`을 사용해 `console.log()` 함수를 호출하여 테스트하세요.

출력은 `[1, 2, 3, 4, 5]`이어야 합니다.

# --hints--

`countup` 문장 안에서 인수 `5`를 사용해 함수 `console.log()`를 호출해야 합니다.

```js
assert.deepStrictEqual(countup(5), [1, 2, 3, 4, 5]);
assert.match(code, /console\.log\s*\(\s*countup\s*\(\s*5\s*\)\s*\)/)
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
console.log(countup(5));
```
