---
id: 694870bcf3c6874112b6f951
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

`1`보다 크거나 같은 모든 숫자에 대해 `else` 블록을 만들고 그 안에서 `countup`를 사용해 `number - 1` 함수를 호출한 결과를 `countArray`에 할당하세요.

그다음 `.push()`를 사용해 `number`을 `countArray`에 추가하고, 마지막으로 `countArray`을 반환하세요.

이렇게 하면 재귀가 결국 기본 사례에 도달해 종료됩니다.

# --hints--

`countup`을 사용해 `number - 1`를 호출한 결과를 `countArray`에 할당해야 합니다.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

`number`를 `countArray`에 추가해야 합니다.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

`countArray`을 반환해야 합니다.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
