---
id: 6948705160b66cf93a916e29
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

`countArray` 함수 안에 `countup` 문장 바로 위에 `if`라는 빈 배열을 선언하세요.

이 배열은 각 함수 호출 시 `number` 값을 저장할 것입니다.

# --hints--

`countArray`라는 빈 배열을 선언해야 합니다.

```js
assert.match(countup.toString(), /(let|const|var)\s+countArray\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
