---
id: a6e40f1041b06c996f7b2406
title: 构建第一个元素查找器
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

在本实验中，你将创建一个函数，该函数遍历一个数组并返回第一个通过测试函数（也称为“truth test”）的元素。 

该函数将迭代数组并使用提供的测试函数测试每个元素。最后，它将返回测试函数返回 `true` 的第一个元素。

例如：

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个名为 `findElement` 的函数，它接受一个数组和一个函数作为参数。
2. 该函数应返回数组中第一个通过真值测试的项。这意味着，调用传入的函数 `func`，给定一个元素 `x`，如果 `func(x)` 为 `true`，则通过真值测试。
3. 如果没有元素通过测试，函数应返回 `undefined`。

# --hints--

你应该有一个 `findElement` 函数。

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` 应返回 `8`。

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` 应返回 `undefined`。

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` 应该返回 `3`。

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` 应该返回 `"javascript"`。

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` 应该返回 `undefined`。

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` 应该返回 `2`。

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` 应该返回 `undefined`。

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
