---
id: 61695c4aad56f95497c19583
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Remember that the `title` element gives search engines extra information about the page. It also displays the content of that `title` element in two more ways:

* in the title bar when the page is open
* na aba do navegador para a página quando você passar o mouse sobre ela. Mesmo se essa aba não estiver ativa, assim que passar o mouse na aba, o texto de `title` será exibido.

Dentro do elemento `head`, aninhe um elemento `title` com o texto `Colored Markers`.

# --hints--

Você deve acrescentar uma tag de abertura`title`.

```js
assert(code.match(/<title\s*>/i));
```

Você deve acrescentar uma tag de fechamento`title`.

```js
assert(code.match(/<\/title\s*>/i));
```

O projeto deve ter o título `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Lembre-se, maiúsculas, minúsculas e a ortografia são importantes para o título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
