---
id: 61695ab9f6ffe951c16d03dd
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Nest a `head` element within the `html` element. Just after the `head` element, add a `body` element.

# --hints--

Você deve acrescentar uma tag de abertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Você deve ter uma tag de fechamento para `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Você deve ter uma tag de abertura para `body`.

```js
assert(code.match(/<body\s*>/i));
```

Você deve ter uma tag de fechamento para `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Deve haver apenas uma tag de abertura do elemento `head`.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

Deve haver apenas uma tag de fechamento do elemento `head`.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

Deve haver apenas uma tag de abertura do elemento `body`.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

Deve haver apenas uma tag de fechamento do elemento `body`.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

O elemento `head` deve estar vazio.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

O elemento `body` deve estar vazio.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

O elemento `body` deve ser inserido depois do elemento `head`.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

O elemento `head` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

O elemento `body` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
