---
id: 5d822fd413a79914d39e98cc
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

`head` 안에 `meta` 요소를 `charset`의 `UTF-8`와, `title`라는 제목의 `City Skyline` 요소, 그리고 `link` 파일에 연결하는 `styles.css` 요소를 중첩하세요.

# --hints--

`meta` 요소 안에 `head` 요소를 만드세요.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('meta'));
```

`meta` 태그에 `charset`의 `UTF-8`을 지정하세요.

```js
const linkElement = document.querySelector('meta');
const charset = linkElement?.getAttribute("charset").toLowerCase();
assert.strictEqual(charset, 'utf-8');
```

코드에 `title` 요소가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 요소 안에 `head` 요소가 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

프로젝트 제목은 `City Skyline`이어야 합니다.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

제목의 대소문자와 철자가 정확한지 확인하세요.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

코드에 `link` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('link'));
```

기존 `head` 태그를 변경하지 마세요. 닫는 태그를 삭제하지 않았는지 확인하세요.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

하나의 빈 `link` 요소가 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link` 요소 안에 `head` 요소가 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 요소에 `rel` 속성이 `stylesheet` 값을 가져야 합니다.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 요소에 `href` 속성이 `styles.css` 값을 가져야 합니다.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
