---
id: 5d822fd413a79914d39e98cb
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

다음으로, `head` 요소 안에 여는 `body` 태그와 닫는 `html` 태그를 추가하세요.

# --hints--

여는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<head\s*>/i));
```

닫는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<\/head\s*>/i));
```

여는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<body\s*>/i));
```

닫는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 요소와 `body` 요소는 형제 관계여야 합니다.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

`head` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

`body` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
  
</html>
--fcc-editable-region--

```
