---
id: 5d822fd413a79914d39e98ca
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`DOCTYPE` 아래에 여는 `html` 태그와 닫는 `html` 태그를 추가해서 코드를 작성할 공간을 만드세요. 언어는 영어로 설정하세요.

# --hints--

`DOCTYPE` 선언은 HTML 맨 앞에 있어야 합니다.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 요소는 여는 태그와 `lang` 속성이 `en` 값을 가져야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/html\s*>/));
```

`html` 태그는 올바른 순서로 있어야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

`html` 요소는 하나만 있어야 합니다.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
