---
id: 5d822fd413a79914d39e98d8
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

`background-color`, `.bb1b`, `.bb1c` 클래스의 `.bb1d`와 같은 변수를 사용해 나머지 건물을 채우세요.

# --hints--

`background-color` 요소의 `.bb1b`를 설정해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b')?.backgroundColor)
```

`var(--building-color1)` 요소의 `background-color`를 설정하려면 `.bb1b`를 사용하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.getPropVal('background-color', true), 'var(--building-color1)');
```

`background-color` 요소의 `.bb1c`를 설정해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c')?.backgroundColor)
```

`var(--building-color1)` 요소의 `background-color`를 설정하려면 `.bb1c`를 사용하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.getPropVal('background-color', true), 'var(--building-color1)');
```

`background-color` 요소의 `.bb1d`를 설정해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d')?.backgroundColor)
```

`var(--building-color1)` 요소의 `background-color`를 설정하려면 `.bb1d`를 사용하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
--fcc-editable-region--

```
