---
id: 5d822fd413a79914d39e98ce
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

모든 요소에 `box-sizing` `border-box`를 추가하세요. 이렇게 하면 추가한 테두리가 요소 크기에 영향을 주지 않습니다.

# --hints--

`box-sizing` 속성을 사용하세요.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

기존 `*` 선택자를 활용하세요.

```js
// Two selectors create two CSSStyleRule objects
assert.equal(new __helpers.CSSHelp(document).getStyleDeclarations('*').length, 1);
```

모든 요소에 `box-sizing` `border-box`를 적용하세요.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```
