---
id: 5d822fd413a79914d39e98db
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

새 건물에 `width`와 `height` 프로프를 각각 `10%`와 `50%`로, `.bb2`, `10%`와 `55%` 프로프를 `.bb3`로, 그리고 `11%`와 `58%` 프로프를 `.bb4`로 지정하세요. 이 프로젝트에서는 거의 모든 단위를 백분율 기반으로 사용하고 일부는 flexbox를 활용하므로 모든 것이 완전히 반응형이 될 것입니다.

# --hints--

`.bb2`에 `width`를 `10%`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb2')?.width, '10%');
```

`.bb2`에 `height`를 `50%`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb2')?.height, '50%');
```

`.bb3`에 `width`를 `10%`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb3')?.width, '10%');
```

`.bb3`에 `height`를 `55%`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb3')?.height, '55%');
```

`.bb4`에 `width`를 `11%`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb4')?.width, '11%');
```

`.bb4`에 `height`를 `58%`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb4')?.height, '58%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div class="bb4"></div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

--fcc-editable-region--

--fcc-editable-region--

```
