---
id: 5d822fd413a79914d39e98d3
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

`div` 컨테이너 안에 `.bb1` 요소 네 개를 중첩하세요. 이 요소들에 차례대로 `bb1a`, `bb1b`, `bb1c`, `bb1d` 클래스를 부여하세요. 이 건물은 네 개의 섹션을 갖게 됩니다.

# --hints--

새로운 `div` 요소 네 개를 만들어야 합니다.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

새로운 `div` 요소 중 하나에 `bb1a` 클래스를 부여해야 합니다.

```js
assert.exists(document.querySelector('div.bb1a'));
```

새로운 `div` 요소 중 하나에 `bb1b` 클래스를 부여해야 합니다.

```js
assert.exists(document.querySelector('div.bb1b'));
```

새로운 `div` 요소 중 하나에 `bb1c` 클래스를 부여해야 합니다.

```js
assert.exists(document.querySelector('div.bb1c'));
```

새로운 `div` 요소 중 하나에 `bb1d` 클래스를 부여해야 합니다.

```js
assert.exists(document.querySelector('div.bb1d'));
```

새로운 `div` 요소들을 올바른 순서로 배치해야 합니다.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

새로운 `div` 요소들을 `.bb1` 요소 안에 배치해야 합니다.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
    
```
