---
id: 5d822fd413a79914d39e98d2
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

배경 건물 컨테이너 안에 `div`를 `bb1` 클래스로 중첩하세요. `styles.css` 파일을 열고 `.bb1`에 `width`의 `10%`와 `height`의 `70%`을 주세요. "bb"는 "background building"의 약자로, 이것이 첫 번째 건물이 됩니다.

# --hints--

새로운 `div` 요소를 만들어야 합니다.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

새로운 `div`에 `bb1` 클래스를 지정해야 합니다.

```js
assert.exists(document.querySelector('div.bb1'));
```

요소를 스타일링할 때 `.bb1` 클래스 선택자를 사용해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

`.bb1` 요소에 `width`의 `10%`을 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

`.bb1` 요소에 `height`의 `70%`을 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
    
```
