---
id: 5d822fd413a79914d39e98e4
title: 28단계
challengeType: 0
dashedName: step-28
---

# --description--

배경 건물들이 점점 멋져 보이기 시작합니다. `div` 요소 아래에 새로운 `.background-buildings`를 만들고 `foreground-buildings` 클래스를 지정하세요. 이것은 더 많은 건물을 위한 또 다른 컨테이너가 될 것입니다.

# --hints--

새로운 `div` 요소를 만들어야 합니다.

```js
assert.equal(document.querySelectorAll('div')?.length, 15);
```

새로운 `div`는 `div.background-buildings` 요소 다음에 와야 합니다.

```js
assert.exists(document.querySelector('div.background-buildings + div'));
```

새로운 `div`에는 `foreground-buildings` 클래스를 지정해야 합니다.

```js
assert.exists(document.querySelector('div.foreground-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

```
