---
id: 5d822fd413a79914d39e98cf
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`body`를 볼 수 있습니다(페이지에서 가장 안쪽 상자입니다); 그 주위의 상자는 `html` 요소입니다. `body`가 뷰포트를 가득 채우도록 `height`를 `100vh`로 지정하세요. `margin`의 기본 `body`를 `margin`를 `0`로 설정해 제거하세요. 마지막으로, 뷰포트를 벗어나는 내용이 있을 때 나타나는 스크롤바를 숨기려면 `overflow` 속성을 `hidden`로 설정하세요.

# --hints--

`body` 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body`에 `height`를 `100vh`로 지정하세요.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

`body`에 `margin`를 `0`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

`body`에 `overflow` 속성을 `hidden`로 설정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```
