---
id: 5d822fd413a79914d39e98e1
title: 25단계
challengeType: 0
dashedName: step-25
---

# --description--

이 방법은 작동하지 않았습니다. `.bb1`에서 선언한 변수들이 `.bb2`와 `.bb3` 형제 요소로 전달되지 않기 때문입니다. CSS가 작동하는 방식이 바로 그렇습니다. 이 때문에 변수들은 종종 `:root` 선택자에서 선언됩니다. 이 선택자는 CSS에서 가장 상위 수준 선택자이며, 여기에 변수를 선언하면 어디서든 사용할 수 있습니다. 스타일시트 상단에 `:root` 선택자를 추가하고 모든 변수 선언을 그곳으로 옮기세요.

# --hints--

스타일시트 상단에 `:root` 선택자를 선언해야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(':root'));
```

`--building-color1` 선택자에서 `#aa80ff`를 `:root` 값으로 정의해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color1')?.trim(), '#aa80ff');
```

`--building-color2` 선택자에서 `#66cc99`를 `:root` 값으로 정의해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color2')?.trim(), '#66cc99');
```

`--building-color3` 선택자에서 `#cc6699`를 `:root` 값으로 정의해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color3')?.trim(), '#cc6699');
```

`.bb1`에서 사용자 정의 속성 변수를 제거해야 합니다.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color2'));
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color3'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2, green);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3, pink);
}

.bb4 {
  width: 11%;
  height: 58%;
}
    
```
