---
id: 5d822fd413a79914d39e98d0
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`div` 안에 `body` 요소를 만들고 `background-buildings` 클래스를 지정하세요. 이 요소는 건물 그룹을 담는 컨테이너가 됩니다.

# --hints--

`div` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('div'));
```

`div` 요소는 `body` 안에 있어야 합니다.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

`div` 요소에는 `background-buildings` 클래스가 있어야 합니다.

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```
