---
id: 5d822fd413a79914d39e98e9
title: 33단계
challengeType: 0
dashedName: step-33
---

# --description--

코드를 최적화하려면 먼저 `position`에서 `top`로 `.foreground-buildings`와 `.background-buildings` 속성과 값을 옮기세요. 그런 다음 기존 `.background-buildings` 선택자를 변경해 `.foreground-buildings`도 대상으로 지정하여 두 곳 모두에 스타일이 적용되도록 하세요. 선택자를 구분할 때는 쉼표(`,`)를 이렇게 사용할 수 있습니다: `selector1, selector2`.

# --hints--

`.foreground-buildings` 선언은 삭제하지 마세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.foreground-buildings'));
```

`position`에서 `.foreground-buildings` 속성은 삭제하세요.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.foreground-buildings')?.position);
```

`top`에서 `.foreground-buildings` 속성은 삭제하세요.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.foreground-buildings')?.top);
```

`position`에 `absolute`의 `.background-buildings, .foreground-buildings` 속성을 추가하세요.

```js
function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle('.background-buildings, .foreground-buildings') ?? a.getStyle('.foreground-buildings, .background-buildings');
}
assert.equal(eitherOr()?.position, 'absolute');
```

`top`에 `0`의 `.background-buildings, .foreground-buildings` 속성을 추가하세요.

```js
function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle('.background-buildings, .foreground-buildings') ?? a.getStyle('.foreground-buildings, .background-buildings');
}
assert.equal(eitherOr()?.top, '0px');
```

같은 스타일 선언에서 `.foreground-buildings`와 `.background-buildings` 선택자를 모두 사용하려면 쉼표를 사용하세요.

```js
function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle('.background-buildings, .foreground-buildings') ?? a.getStyle('.foreground-buildings, .background-buildings');
}
assert.exists(eitherOr());
```

별도의 `.background-buildings` 선택자는 없어야 합니다.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--
.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}
--fcc-editable-region--
.fb1 {
  width: 10%;
  height: 60%;
}

.fb2 {
  width: 10%;
  height: 40%;
}

.fb3 {
  width: 10%;
  height: 35%;
}

.fb4 {
  width: 8%;
  height: 45%;
}

.fb5 {
  width: 10%;
  height: 33%;
}

.fb6 {
  width: 9%;
  height: 38%;
}
    
```
