---
id: 5d822fd413a79914d39e98d4
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

건물 `width`와 `height`의 프로퍼티에 다음 값을 지정하세요: `70%`와 `10%`는 `.bb1a`, `80%`와 `10%`는 `.bb1b`, `90%`와 `10%`는 `.bb1c`, `100%`와 `70%`는 `.bb1d`로 설정하세요. 이 백분율은 부모 요소를 기준으로 하며, 높이 값들이 합쳐서 100%가 되어 컨테이너를 수직으로 채운다는 점을 기억하세요.

# --hints--

`.bb1a`에 스타일을 적용하려면 클래스 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

`.bb1a`에 `width` 값을 `70%` 프로퍼티로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

`.bb1a`에 `height` 값을 `10%` 프로퍼티로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

`.bb1b`에 스타일을 적용하려면 클래스 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

`.bb1b`에 `width` 값을 `80%` 프로퍼티로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

`.bb1b`에 `height` 값을 `10%` 프로퍼티로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

`.bb1c`에 스타일을 적용하려면 클래스 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

`.bb1c`에 `width` 값을 `90%` 프로퍼티로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

`.bb1c`에 `height` 값을 `10%` 프로퍼티로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

`.bb1d`에 스타일을 적용하려면 클래스 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

`.bb1d`에 `width` 값을 `100%` 프로퍼티로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

`.bb1d`에 `height` 값을 `70%` 프로퍼티로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```
