---
id: 5d822fd413a79914d39e98dd
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

건물들이 너무 떨어져 있습니다. `div` 요소 상단에 빈 `.background-buildings` 요소 두 개를 추가하고, 하단에 두 개 더, 그리고 `.bb3`와 `.bb4` 사이에 하나를 추가해 건물들을 서로 더 가깝게 모으세요. 이렇게 하면 컨테이너 전체에 고르게 배치되어 건물들이 중앙 쪽으로 더 가까워집니다.

# --hints--

`div` 요소 앞에 `.bb1` 요소 두 개를 추가하세요.

```js
const bBuildings = document.querySelector('.background-buildings')?.children;
assert(![...bBuildings?.[0]?.classList]?.includes('bb1'));
assert(![...bBuildings?.[1]?.classList]?.includes('bb1'));
```

`div`과 `.bb3` 요소 사이에 `.bb4` 요소 하나를 추가하세요.

```js
assert(document.querySelector('.bb3')?.nextElementSibling === document.querySelector('.bb4')?.previousElementSibling);
```

`div` 요소 뒤에 `.bb4` 요소 두 개를 추가하세요.

```js
const bb4 = document.querySelector('.bb4');
assert.exists(bb4?.nextElementSibling);
assert.exists(bb4?.nextElementSibling?.nextElementSibling);
```

`div` 요소 다섯 개를 추가하세요.

```js
assert.equal(document.querySelectorAll('div')?.length, 14);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div class="bb4"></div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
}

.bb3 {
  width: 10%;
  height: 55%;
}

.bb4 {
  width: 11%;
  height: 58%;
}
    
```
