---
id: 69baafd6d683dd9921b75db3
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви використаєте JavaScript, щоб змоделювати потік відвідувачів на музичному фестивалі.

Ви працюватимете з двома наборами даних, які описують ворота фестивалю під час ранкової та нічної змін:

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

Кожен об’єкт воріт містить такі властивості:

- `id`: рядок, який ідентифікує ворота.
- `capacity`: число відвідувачів, яких ворота можуть пропустити за один такт.
- `queue`: масив чисел, що представляє, скільки відвідувачів прибуває до воріт під час конкретного такту.

Для початку створіть порожню функцію з назвою `initializeThroughput`, яка приймає параметр з назвою `gates`.

# --hints--

Ви повинні створити функцію з назвою `initializeThroughput`.

```js
assert.isFunction(initializeThroughput);
```

Ваша функція `initializeThroughput` має мати параметр `gates`.

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
