---
id: 69babdf3f896711c11516373
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Під час кожного такту симуляції:

- Певна кількість відвідувачів прибуває до воріт (з їхнього `queue`).
- Ворота обробляють відвідувачів на основі свого `capacity`.
- Якщо прибуває більше відвідувачів, ніж ворота можуть обробити, деякі залишаться (переповнення).

Тепер ви створите функцію, яка реалізує цю логіку для одних воріт за один такт.

Створіть порожню функцію з назвою `processGateFlow` з такими параметрами:

- `gate`: об’єкт воріт, який обробляється.
- `tickIndex`: поточна позиція тіку в масиві `queue`.

# --hints--

Вам слід створити функцію з назвою `processGateFlow`.

```js
assert.isFunction(processGateFlow);
```

Ваша функція `processGateFlow` має мати параметри `gate` і `tickIndex`.

```js
const explorer = await __helpers.Explorer(code);
const { processGateFlow } = explorer.allFunctions;
const parameters = processGateFlow?.parameters;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), "gate");
assert.equal(parameters[1].toString(), "tickIndex");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
