---
id: 69bac067cd15c4685d07906d
title: Крок 15
challengeType: 1
dashedName: step-15
---

# --description--

Тепер, коли у вас є функції для обробки окремих поділок і керування переповненням, настав час створити функцію-контролер для однієї брами. Ця функція буде виконувати всі дії для однієї брами під час однієї поділки симуляції.

Створіть порожню функцію з назвою `handleGateAtTick` з такими параметрами:

- `gates`: Повний масив об’єктів брами.
- `gate`: Поточна брама, яку обробляють.
- `tickIndex`: Поточна позиція поділки в масиві `queue`.
- `throughputSummary`: Об’єкт, що відстежує загальну кількість оброблених відвідувачів для кожної брами.

# --hints--

Ви повинні створити функцію з назвою `handleGateAtTick`.

```js
assert.isFunction(handleGateAtTick);
```

Ваша функція `handleGateAtTick` має мати параметри `gates`, `gate`, `tickIndex` і `throughputSummary`.

```js
const explorer = await __helpers.Explorer(code);
const { handleGateAtTick } = explorer.allFunctions;
const parameters = handleGateAtTick?.parameters;
assert.lengthOf(parameters, 4);
assert.equal(parameters[0].toString(), "gates");
assert.equal(parameters[1].toString(), "gate");
assert.equal(parameters[2].toString(), "tickIndex");
assert.equal(parameters[3].toString(), "throughputSummary");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
