---
id: 69bab8292750d84cac5b52a0
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

У цьому практичному занятті пропускна здатність означає загальну кількість відвідувачів, оброблених кожними воротами за всі поділки в одній симуляції (ранок або ніч). Ваша функція `initializeThroughput` підготує об’єкт для відстеження, скільки відвідувачів обробляють кожні ворота.

Спочатку створіть порожній об’єкт з іменем `summary`.

# --hints--

Ваша `initializeThroughput` функція має створити порожній об’єкт і призначити його змінній з ім’ям `summary`.

```js
const explorer = await __helpers.Explorer(code);
const { initializeThroughput } = explorer.allFunctions;
const { summary } = initializeThroughput?.variables ?? {};
const validExpressions = ["{}", "new Object()"];
assert.isTrue(
  validExpressions.some((expression) => summary?.value.matches(expression))
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
