---
id: 69d048937c7bb4b4c22ebb18
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Тепер настав час пропустити учасників через ворота. Усередині циклу `while`:

- Зменшіть `currentTickQueue` на 1, щоб показати, що один учасник пройшов. Використайте оператор декременту (`--`).
- Збільшіть `processed` на 1, щоб відстежувати, скільки учасників було оброблено. Використайте оператор інкременту (`++`).

# --hints--

Ваш цикл `while` має зменшувати `currentTickQueue` за допомогою оператора декременту `--` і збільшувати `processed` за допомогою оператора інкременту `++`.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const viableSolutions = [
  "while(currentTickQueue>0&&processed<gate.capacity){processed++;currentTickQueue--;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){processed++currentTickQueue--;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){processed++;currentTickQueue--}}",
  "while(currentTickQueue>0&&processed<gate.capacity){processed++currentTickQueue--}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--;processed++;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--processed++;}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--;processed++}}",
  "while(currentTickQueue>0&&processed<gate.capacity){currentTickQueue--processed++}}"
];
const endsCorrectly = viableSolutions.some(sol => cleaned.endsWith(sol));
assert(
  endsCorrectly,
  "Your `while` loop should contain both the `processed++` increment and the `currentTickQueue--` decrement, and nothing else."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
