---
id: 655773b0591c5f9f4045883e
title: Hatua 24
challengeType: 20
dashedName: step-24
---

# --description--

Algoriti itaanza katika nodi iliyotajwa. Kisha itachunguza grafu ili kupata njia fupi zaidi kati ya nodi ya kuanzia, au *chanzo*, na nodi zote nyingine.

Kwa hiyo kitendakazi chako kinahitaji vigezo viwili: `graph`, na `start`. Viungeze katika tamko la kitendakazi chako.

# --hints--

Kitendakazi chako kinapaswa kuchukua `graph` na `start` kama vigezo, kwa mpangilio huu.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
