---
id: 6557709b0aee699a6a00528c
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Kwa mfano, grafu inaweza kutumika kuwakilisha pointi mbili katika nafasi, `A` na `B`, zilizounganishwa na njia. Grafu kama hii itaundwa na nodi mbili zilizounganishwa na kingo.

Badilisha ufunguo uliopo wa `'species'` na mfuatano wa herufi `'A'`. Kisha, badilisha thamani inayolingana na mfuatano wa herufi `'B'` ili kuwakilisha muunganisho kati ya nodi za `'A'` na `'B'`.

# --hints--

Kamusi yako inapaswa kuwa na ufunguo wa `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Ufunguo wako wa `'A'` unapaswa kuwa na `'B'` kama thamani.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Kamusi yako inapaswa kuwa na funguo mbili.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
