---
id: 65576ff7888f9e96f52a4be1
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Kamusi huhifadhi data kwa njia ya jozi za *funguo*-*thamani*. Funguo hutenganishwa na thamani inayolingana nayo kwa alama ya nukta mbili. Na kila jozi ya funguo-thamani hutenganishwa na jozi inayofuata kwa koma:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Ongeza jozi mpya ya funguo-thamani kwenye kamusi yako. Tumia mfuatano wa herufi `'species'` kama funguo, na mfuatano wa herufi `'guinea pig'` kama thamani.

# --hints--

Unapaswa kuwa na funguo `'species'` yenye thamani `'guinea pig'` ndani ya kamusi yako ya `copper`.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

Kamusi yako ya `copper` inapaswa kuwa na jozi moja tu ya funguo-thamani.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
