---
id: 65578d0f6c78a0b868a43b9c
title: Hatua 40
challengeType: 20
dashedName: step-40
---

# --description--

Ndani ya mzunguko wa `while`, jambo la kwanza kufanya ni kufafanua nodi ya sasa ya kutembelea. Kwa hilo unaweza kutumia kitendakazi cha `min()`. Hurejesha kipengele kidogo zaidi kutoka kwa inayoweza kurudiwa iliyotumwa kama hoja.

Futa `pass`, kisha tengeneza kigezo kinachoitwa `current` na weka thamani ya `min(unvisited)`.

# --hints--

Unapaswa kutengeneza kigezo cha `current` ndani ya mzunguko wako wa `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

Unapaswa kuweka thamani ya `min(unvisited)` kwa kigezo chako cha `current`. Kumbuka kufuta `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
