---
id: 65578d4fc3afc3b8f554c882
title: Hatua 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` pia huchukua hoja ya maneno pekee. Kwa kupitisha kitendakazi kama hoja ya ziada kwa `min()`, unaweza kubadilisha jinsi vitu vya orodha vinavyolinganishwa.

Matokeo ya mstari uliouandika katika hatua iliyopita ni nodi inayofuata kwanza kwa mpangilio wa alfabeti. Badala yake unataka kuchagua nodi isiyotembelewa yenye umbali mdogo zaidi kutoka kwa nodi ya kuanzia.

Pitia `key=distances.get` kama hoja ya pili kwa wito wako wa `min()`. Kwa njia hii, ulinganisho utafanyika kulingana na thamani ambayo kila kipengee cha orodha ya `unvisited` kina ndani ya kamusi ya `distances`.

# --hints--

Unapaswa kupitisha `key=distances.get` kama hoja ya pili kwa wito wako wa `min()`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
