---
id: 655771d889132f9ccd341060
title: Hatua 20
challengeType: 20
dashedName: step-20
---

# --description--

Grafu huitwa grafu *enye uzito* wakati makali yake yanahusishwa na uzito, unaowakilisha umbali, muda au thamani nyingine ya kihesabu.

Katika kesi yako, uzito huu utakuwa ni umbali kati ya kila nodi, au sehemu katika nafasi. Ili kuwakilisha grafu yenye uzito unaweza kubadilisha kamusi yako, ukitumia orodha ya tuple kwa kila thamani.

Kipengele cha kwanza katika tuple kitakuwa nodi inayounganishwa, na kipengele cha pili kitakuwa nambari kamili inayobainisha umbali.

Badilisha `my_graph['A']` kuwa orodha ya tuple, ukizingatia kwamba umbali wa `A-B` ni `3` na umbali wa `A-D` ni `1`.

# --hints--

`my_graph["A"]` inapaswa kuwa orodha inayojumuisha tuple `('B', 3)` na `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` inapaswa kuwa na funguo 4 zenye majina `'A'`, `'B'`, `'C'`, na `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
