---
id: 6557743527cb92a06417ea97
title: Hatua 26
challengeType: 20
dashedName: step-26
---

# --description--

Unda mzunguko wa `for` kurudia juu ya grafu yako, na tumia njia ya `.append()` kuongeza nodi kila moja mwishoni mwa orodha ya `unvisited`.

# --hints--

Unapaswa kuunda mzunguko wa `for` kurudia juu ya `graph` ndani ya kitendakazi cha `shortest_path`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "shortest_path");    
    assert(function_body.match(/^\s*for\s+\w+\s+in\s+graph\s*:/m));
  }
})
```

Unapaswa kuongeza kila nodi kwenye `unvisited` ndani ya mzunguko wako wa `for`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const block_regex = /for\s+(\w+)\s+in\s+graph\s*/;
    const {block_body} = __helpers.python.getBlock(commentless_code, block_regex);
    const loop_condition = commentless_code.match(block_regex);
    const regex = new RegExp(`^\\s+unvisited\\.append\\s*\\(\\s*${loop_condition[1]}\\s*\\)`, "m");
    assert(block_body.match(regex));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
--fcc-editable-region--
```
