---
id: 65579228c669fcbebffd01d5
title: Hatua 52
challengeType: 20
dashedName: step-52
---

# --description--

Hitilafu nyingine ni nyembamba. Wakati umbali mfupi zaidi unapopatikana kwa nodi jirani, `paths[current]` hupewa thamani ya njia ya nodi jirani, `paths[node]`.

Hii ina maana kwamba vigezo vyote viwili vinaelekeza kwenye orodha ile ile. Kwa kuwa orodha ni zinazobadilika, unapoongeza nodi jirani kwenye njia yake, `paths[node]` na `paths[current]` zote hubadilika kwa sababu ni orodha ile ile. Hii husababisha njia zisizo sahihi, ingawa umbali ni sahihi.

Unaweza kurekebisha hitilafu hiyo kwa kuweka nakala ya `paths[current]` kwenye njia ya nodi jirani. Kwa hiyo unaweza kutumia sintaksia ya kukata (slice):

```py
my_list[:]
```

Ambapo `my_list` ni orodha unayotaka kunakili. Badilisha mgawo wa `paths[node] = paths[current]` uliopo ndani ya blokhi yako ya `if` kwa kukata `paths[current]`.

# --hints--

Unapaswa kuweka thamani ya `paths[current][:]` kwenye njia ya nodi jirani.

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
