---
id: 655773f8b8b5db9fc6d0ae76
title: Hatua 25
challengeType: 20
dashedName: step-25
---

# --description--

Ili kufuatilia nodi zilizotembelewa, unahitaji orodha ya nodi zote katika grafu. Mara nodi inapotekelezwa, itafutwa kutoka kwenye orodha hiyo.

Sasa, badilisha neno kuu la `pass` na liweke thamani ya orodha tupu yenye jina la `unvisited`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `unvisited` ndani ya kitendakazi cha `shortest_path`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

Unapaswa kuweka thamani ya orodha tupu kwa kigezo chako cha `unvisited`. Kumbuka kufuta `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
