---
id: 6557746aad2844a0cd864e12
title: Hatua 27
challengeType: 20
dashedName: step-27
---

# --description--

Wakati algoriti inapotafiti grafu, inapaswa kufuatilia umbali mfupi unaojulikana kwa sasa kati ya nodi ya kuanzia na nodi nyingine.

Kabla ya mzunguko wako wa `for`, tengeneza kigezo kipya kinachoitwa `distances` na kitoe kamusi tupu.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

Kigezo chako cha `distances` kinapaswa kuwa kamusi tupu.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
