---
id: 655774d01daeeaa1978b99d5
title: Hatua 29
challengeType: 20
dashedName: step-29
---

# --description--

Mwanzo, nodi zote nyingine katika grafu zinachukuliwa kuwa ziko umbali usio na kikomo kutoka kwa nodi chanzo, kwa sababu umbali bado haujafafanuliwa.

Tengeneza kifungu cha `else` na weka thamani isiyo na kikomo kwa nodi katika kamusi ya `distances`. Kwa ajili ya hilo, tumia kitendakazi cha `float()` na mfuatano wa herufi `'inf'` kama hoja ili kuzalisha nambari ya desimali inayowakilisha ukomo chanya.

# --hints--

Unapaswa kuwa na kifungu cha `else`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

Unapaswa kuweka `float('inf')` kwa `distances[node]` ndani ya kifungu chako kipya cha `else`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
