---
id: 6557712d77ce2d9bd7e63afd
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Ongeza nodi nyingine inayounganishwa na `B` kwenye grafu yako na uitaje `C`.

Badilisha kamusi yako iliyopo ili wakilishe mpangilio huu: ongeza ufunguo mwingine `'C'` hadi `my_graph` na uipe thamani ya mfuatano wa herufi `'B'`.

Pia, badilisha thamani ya ufunguo uliopo `'B'` kuwa orodha `['A', 'C']` ili wakilishe muunganisho mwingi wa nodi yako `'B'`.

# --hints--

Kamusi yako inapaswa kuwa na funguo 3 — `'A'`, `'B'`, na `'C'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

Thamani ya `my_graph['A']` inapaswa kuwa mfuatano wa herufi `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` inapaswa kuwa orodha.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

Thamani ya `my_graph['B']` inapaswa kuwa orodha inayojumuisha `'A'` na `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

Thamani ya `my_graph['C']` inapaswa kuwa mfuatano wa herufi `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
