---
id: 65578cb031cd93b77a285db2
title: Hatua 38
challengeType: 20
dashedName: step-38
---

# --description--

Kitendakazi chako kitachunguza nodi zote zinazounganishwa na nodi ya kuanzia. Kitahesabu njia fupi zaidi kwa zote. Kisha, kitatoa nodi ya kuanzia kutoka kwa nodi ambazo hazijatembelewa.

Ifuatayo, nodi jirani iliyo karibu zaidi itatembelewa na mchakato utaendelea kurudiwa hadi nodi zote zimetembelewa.

Kuanzia sasa, utafanya kazi kwenye mzunguko mkuu unaochunguza nodi katika grafu. Ili kuepuka matatizo ya kuendesha mzunguko usio na mwisho wakati wa ukuzaji wa algoriti, badilisha wito wa kitendakazi chako kuwa maoni.

# --hints--

Unapaswa kubadilisha wito wa kitendakazi chako kuwa maoni.

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
