---
id: 6557924d47c325bf27afbe51
title: Hatua 53
challengeType: 20
dashedName: step-53
---

# --description--

Algoriti imekamilika lakini unaweza kuboresha matokeo. Pia, unaweza kutoa kitendakazi na hoja ya ziada ili kurudisha njia kati ya nodi mbili tu.

Ongeza `target` kama kigezo cha tatu katika tamko la kitendakazi chako na umpe thamani ya msingi ya mfuatano wa herufi tupu.

# --hints--

Kitendakazi chako kinapaswa kuchukua vigezo vitatu: `graph`, `start`, na `target`. Mpangilio ni muhimu.

```js
({ test: () => assert.match(code, /^def\s+shortest_path\s*\(\s*graph\s*,\s*start\s*,\s*target\s*=?\s*.*\s*\)\s*:/m) })
```

Kigezo cha `target` kinapaswa kuwa na thamani ya msingi ya mfuatano wa herufi tupu.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == "(graph, start, target='')"
  `))
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
--fcc-editable-region--
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
