---
id: 65577236b056379d5dbc7000
title: Hatua 23
challengeType: 20
dashedName: step-23
---

# --description--

Sasa utaanza kuendeleza algoriti ya kuhesabu njia fupi zaidi kati ya kila nodi katika grafu yako mpya.

Tangaza kitendakazi tupu kinachoitwa `shortest_path`. Tumia neno kuu la `pass` kujaza mwili wa kitendakazi.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `shortest_path`.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
