---
id: 655774955b097ea14897db12
title: Hatua 28
challengeType: 20
dashedName: step-28
---

# --description--

Umbali kutoka nodi ya kuanzia ni sifuri, kwa sababu algoriti inaanza tathmini yake moja kwa moja kutoka hapo.

Baada ya kuongeza `node` kwa `unvisited` katika mzunguko wako, tengeneza usemi wa `if` unaochochea ikiwa nodi ni sawa na nodi ya kuanzia. Kisha weka thamani ya `0` kwa nodi hiyo ndani ya kamusi ya `distances`.

# --hints--

Unapaswa kutengeneza usemi wa `if` unaotekelezwa wakati `node` ni sawa na `start`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

Ndani ya usemi wako mpya wa `if` unapaswa kuweka thamani ya `0` kwa nodi katika kamusi ya `distances`.

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
