---
id: 6566195b0a021bb660b2b4b1
title: Hatua 21
challengeType: 20
dashedName: step-21
---

# --description--

Sasa badilisha `my_graph['B']` kuwa orodha ya tuple, ambapo kipengele cha kwanza katika tuple ni nodi iliyounganishwa, na kipengele cha pili ni umbali. Umbali wa `B-C` ni `4`.

# --hints--

`my_graph['B']` inapaswa kuwa orodha inayojumuisha tuple za `('A', 3)` na `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph` inapaswa kuwa na funguo 4 zinazoitwa `'A'`, `'B'`, `'C'`, na `'D'`.

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
