---
id: 65661b72d6745ebec6a96923
title: Hatua 22
challengeType: 20
dashedName: step-22
---

# --description--

Kwa njia ile ile, badilisha orodha mbili zilizobaki ukizingatia kuwa umbali wa `C-D` ni `7`.

# --hints--

`my_graph['C']` inapaswa kuwa orodha inayojumuisha tuple za `('B', 4)` na `('D', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

`my_graph['D']` inapaswa kuwa orodha inayojumuisha tuple za `('A', 1)` na `('C', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

`my_graph` inapaswa kuwa na funguo 4 zenye majina `'A'`, `'B'`, `'C'`, na `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` inapaswa kuwa orodha inayojumuisha tuple za `('B', 3)` na `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph['B']` inapaswa kuwa orodha inayojumuisha tuple za `('A', 3)` na `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
