---
id: 6578b13757611e2825beb8a5
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Funguo lazima ziwe za kipekee ndani ya kamusi na zinaweza kuwa aina za data zisizobadilika tu. Hii inamaanisha huwezi kutumia orodha au kamusi nyingine kama funguo.

Ongeza ufunguo mwingine `'age'` kwenye kamusi yako na umpe nambari kamili `2` kama thamani.

# --hints--

Unapaswa kuwa na ufunguo `'age'` wenye thamani `2` mara tu baada ya `'species': 'guinea pig'`, ndani ya kamusi yako ya `copper`. Usisahau koma.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Unapaswa kuwa na jozi mbili za ufunguo-thamani katika kamusi yako ya `copper`.

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
