---
id: 657891ab9c1903f4e55433ba
title: Hatua 44
challengeType: 20
dashedName: step-44
---

# --description--

Wakati hali ya `if` yako mpya ni kweli, njia fupi zaidi kuelekea nodi jirani imepatikana.

Ndani ya blokki yako mpya ya `if`, futa `pass` na weka thamani upya umbali wa nodi jirani kuwa jumla ya umbali wa nodi jirani pamoja na umbali wa `current`.

# --hints--

Unapaswa kuweka thamani ya `distance + distances[current]` kwa `distances[node]` ndani ya `if` yako mpya.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+distance\s*\+\s*distances\s*\[\s*current\s*\]\s*<\s*distances\s*\[\s*node\s*\]\s*/);
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*distance\s*\+\s*distances\s*\[\s*current\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)

    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
--fcc-editable-region--
            if distance + distances[current] < distances[node]:
                pass
--fcc-editable-region--
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
