---
id: 65789506b30453080f77470c
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Hadi sasa, tayari umejifunza aina tofauti za data:

- Aina za data zisizobadilika, kama nambari kamili, mfuatano wa herufi, tuple, na thamani ya kweli au siyo kweli.
- Aina za data zinazobadilika, kama orodha na kamusi.

Kamusi hutambulika kwa jozi ya mabano ya mviringo, `{}`.

Anza kwa kuunda kigezo kinachoitwa `copper` na ukaweke thamani ya kamusi tupu kwa kutumia jozi ya mabano ya mviringo, kwa njia ile ile unavyounda orodha tupu kwa jozi ya mabano ya mraba.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

Kigezo chako cha `copper` kinapaswa kuwa na thamani ya kamusi tupu. Tumia jozi ya mabano ya mviringo kwa ajili ya hiyo.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
