---
id: 6578b57361f2f132a02e2a18
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Unaweza kupata data iliyohifadhiwa katika kamusi kupitia funguo zake:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

Baada ya kamusi yako, fuata mfano ulio hapo juu kupata thamani ya funguo `'species'` za `copper` na chapisha matokeo.

# --hints--

Haupaswi kubadilisha mgawo wa thamani wa kamusi yako.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Unapaswa kutumia `copper['species']` kupata thamani ya funguo `'species'`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

Unapaswa kuita `print()` ukipitisha `copper['species']` kama hoja.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
