---
id: 6579cbab9825b8170974c69a
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Unaweza kufuta jozi ya ufunguo-na-thamani kutoka kwenye kamusi kwa kutumia neno kuu la `del`:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

Mara moja kabla ya mzunguko wako wa `for`, tumia neno kuu la `del` kufuta ufunguo wa `'age'` pamoja na thamani yake kutoka `copper`.

# --hints--

Unapaswa kutumia neno kuu la `del` kufuta `copper['age']` kabla ya mzunguko wa `for`.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
