---
id: 6579ca0923cfa7162089d2f0
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Sintaksia ile ile inaweza kutumika kubadilisha thamani ya ufunguo uliopo.

Mara tu kabla ya wito la `print()`, pata upatikanaji wa ufunguo wa `'species'` na weka thamani yake kuwa `'Cavia porcellus'` tena.

# --hints--

Unapaswa kuweka thamani ya `copper['species']` kuwa mfuatano wa herufi `'Cavia porcellus'` kabla ya wito la `print()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
