---
id: 6579717f0920131304286804
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Ili kuongeza jozi mpya ya ufunguo-na-thamani baada ya kutamka kamusi, unaweza kuonyesha ufunguo kwa njia ile ile unavyopata ufunguo uliopo, na weka thamani ya ufunguo mpya kwa kutumia opereta ya mgawo:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

Futa wito lako la `print()`. Kisha, baada ya tamko la `copper`, ongeza ufunguo `'food'` kwenye kamusi yako na weka thamani yake kuwa `'hay'`.

# --hints--

Haupaswi kuwa na `print(copper['age'])` katika msimbo wako.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

Unapaswa kuongeza ufunguo `'food'` kwenye `copper` baada ya kutamka kamusi.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

Unapaswa weka `copper['food']` kuwa `'hay'` baada ya kutamka kamusi.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
