---
id: 6579cd5f6dd62c189e53ddbb
title: Hatua 14
challengeType: 20
dashedName: step-14
---

# --description--

Sasa baada ya kukagua vipengele vya msingi vya kamusi, unaweza kuendelea kujenga algoriti ya njia fupi zaidi.

Futa kila mstari wa msimbo baada ya tamko la kamusi ya `copper`.

# --hints--

Unapaswa kufuta mistari yote ya msimbo baada ya tamko la kamusi yako.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
