---
id: 65796fac81f983127558f3f4
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Sasa, badilisha wito wako wa `print()` uliopo ili uchapishe thamani ya ufunguo wa `age`.


# --hints--

Unapaswa kutumia `copper['age']` kupata upatikanaji wa thamani ya ufunguo wa `age`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

Unapaswa kuita `print()` ukimpa `copper['age']` kama hoja.

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
