---
id: 6579dd49fa8a8e1fd06b85a9
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Michoro ni miundo ya data inayowakilisha uhusiano kati ya jozi za vipengele.
Vipengele hivi, vinavyoitwa *nodi*, vinaweza kuwa vitu halisi vya maisha, vyombo, pointi katika nafasi au vinginevyo. Muunganiko kati ya nodi huitwa *mipaka*.

Hapa kuna uwakilishi wa picha wa mchoro:

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


Badilisha jina la kamusi `copper` kuwa `my_graph`. Hii itawakilisha mchoro ili kujaribu algoriti yako.

# --hints--

Unapaswa kubadilisha jina la kamusi `copper` kuwa `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Kigezo chako cha `my_graph` kinapaswa kuwa kamusi.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
