---
id: 65affe602a2a828b0382a058
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Badilisha funguo la `'age'` kwa mfuatano wa herufi `'B'` na weka thamani yake kuwa mfuatano wa herufi `'A'` ili wakilisha muunganisho kati ya nodi zote mbili kwa mwelekeo zote mbili.

# --hints--

Kamusi yako inapaswa kuwa na funguo la `'B'`.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

Funguo lako la `'B'` linapaswa kuwa na `'A'` kama thamani.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

Kamusi yako inapaswa kuwa na funguo la `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Funguo lako la `'A'` linapaswa kuwa na `'B'` kama thamani.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Kamusi yako inapaswa kuwa na funguo mbili.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
