---
id: 65b7cf140d34058d7ea8935f
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Ikiwa unataka kurudia juu ya thamani za funguo za kamusi, njia moja ni kutumia njia ya `.values()`.

Badilisha mzunguko wako wa `for` kurudia juu ya `copper.values()` badala ya `copper` na angalia matokeo.

# --hints--

Unapaswa kubadilisha mzunguko wako wa `for` kurudia juu ya `copper.values()` badala ya `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
