---
id: 66faac4139dbbd5dd632c860
title: Crea un modulo di candidatura per un lavoro
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `div` con la classe `container`.
1. All'interno dell'elemento `div` dovresti avere un elemento `form`.
1. Il modulo dovrebbe contenere un elemento `input` con il tipo `text` e l'id `name` per inserire il nome completo dell'utente.
1. Dovresti avere un altro elemento `input` con il tipo `email` e l'id `email` per inserire l'indirizzo email dell'utente.
1. Il modulo dovrebbe includere un elemento `select` con l'id `position` che permette agli utenti di selezionare una posizione lavorativa.
1. Dovresti avere un elemento `fieldset` con la classe `radio-group`.
1. All'interno di `.radio-group` dovresti avere un insieme di elementi `input` con il tipo `radio` e le etichette appropriate per selezionare le opzioni di disponibilità (ad esempio, full-time, part-time). Il gruppo `name` dovrebbe essere `availability`.
1. Dovresti avere un elemento `textarea` con l'id `message` per inserire un messaggio.
1. Dovresti associare ogni elemento `input` a un elemento `label`.
1. Dovresti avere un elemento `button` con il tipo `submit` per inviare il modulo.
1. Aggiungi una pseudo-classe `:focus` agli elementi `input` e `textarea` per cambiare il colore del bordo quando sono in focus.
1. Gli elementi `input`, `select` e `textarea` dovrebbero avere una pseudo-classe `:invalid` che cambia il colore del bordo in rosso quando viene rilevato un input non valido.
1. Gli elementi `input`, `select` e `textarea` dovrebbero avere una pseudo-classe `:valid` che cambia il colore del bordo in verde quando viene inserito un input valido.
1. L'elemento `button` dovrebbe avere una pseudo-classe `:hover` che cambia il colore di sfondo quando ci si passa sopra con il mouse.
1. Usa la pseudo-classe `:checked` su `.radio-group input[type="radio"]` per aggiungere un colore al bordo, un colore di sfondo e un'ombra al riquadro quando il pulsante di opzione è selezionato.
1. Usa la pseudo-classe `:checked` sui pulsanti di opzione per cambiare il colore del testo dell'elemento `label` associato quando l'opzione è selezionata.
1. Aggiungi una pseudo-classe `:first-of-type` all'elemento `input` per definire uno stile diverso al primo campo di input (ad esempio, angoli arrotondati).

**Nota:** Assicurati di collegare il tuo foglio di stile nel tuo HTML e di applicare il CSS.

# --hints--

Dovresti avere un elemento `div` con la classe `container`.

```js
assert.exists(document.querySelector("div.container"));
```

All'interno dell'elemento `div` dovresti avere un elemento `form`.

```js
assert.exists(document.querySelector("div.container > form"));
```

Il modulo dovrebbe contenere un elemento `input` con il tipo `text` e l'id `name` per inserire il nome completo dell'utente.

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

Dovresti avere un altro elemento `input` con il tipo `email` e l'id `email` per inserire l'indirizzo email dell'utente.

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

Il modulo dovrebbe includere un elemento `select` con l'id `position` con alcuni elementi `option`.

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

Dovresti avere un elemento `fieldset` o `section` con la classe `.radio-group`.

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

All'interno di `.radio-group` dovresti avere un gruppo di elementi `input` con il tipo `radio` per selezionare le opzioni di disponibilità. Il gruppo `name` dovrebbe essere `availability`.

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

Dovresti avere un elemento `textarea` con l'id `message` per inserire un messaggio.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

Dovresti associare ogni elemento `input` a un elemento `label`.

```js
const els = document.querySelectorAll('input');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

Dovresti avere un elemento `button` con il tipo `submit` per inviare il modulo.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

Dovresti aggiungere una pseudo-classe `:focus` agli elementi `input` e `textarea` per cambiare il colore del bordo quando sono in focus. Usa un selettore di lista nell'ordine dato.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'))
```

Gli elementi `input`, `select` e `textarea` dovrebbero avere una pseudo-classe `:invalid` che cambia il colore del bordo in rosso quando viene rilevato un input non valido. Usa un selettore di lista nell'ordine dato.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

Gli elementi `input`, `select` e `textarea` dovrebbero avere una pseudo-classe `:valid` che cambia il colore del bordo in verde quando viene inserito un input valido. Usa un selettore di lista nell'ordine dato.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

L'elemento `button` dovrebbe avere una pseudo-classe `:hover` che cambia il colore di sfondo quando ci si passa sopra con il mouse.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

Dovresti usare la pseudo-classe `:checked` su `.radio-group input[type="radio"]` per aggiungere un colore al bordo quando il pulsante di opzione è selezionato.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

Dovresti usare la pseudo-classe `:checked` su `.radio-group input[type="radio"]` per aggiungere un colore di sfondo quando il pulsante di opzione è selezionato.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

Dovresti usare la pseudo-classe `:checked` su `.radio-group input[type="radio"]` per aggiungere un'ombra al riquadro quando il pulsante di opzione è selezionato.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

Dovresti usare la pseudo-classe `:checked` sui pulsanti di opzione per cambiare il colore del testo dell'elemento `label` associato quando l'opzione è selezionata.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

Aggiungi una pseudo-classe `:first-of-type` agli elementi `input` per definire uno stile diverso al primo campo di input.

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
