---
id: 68f890b73bbfc9efa7379e7b
title: Depurar erros de coerção de tipo em um aplicativo com bugs
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Você acabou de entrar em uma empresa local de desenvolvimento web e sua primeira tarefa é limpar um código com bugs deixado pelos desenvolvedores anteriores.

Eles estavam tentando operações aritméticas simples, mas algo deu errado, os resultados não fazem sentido. Sua tarefa é revisar, depurar e corrigir o código para que ele funcione conforme o esperado e seja mais fácil de ler.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma variável chamada `firstResult` que adiciona corretamente os números `5` e `10` para produzir o valor `15`.
2. Você deve ter uma variável chamada `secondResult` que subtrai corretamente `5` de `8` para produzir o valor `3`.
3. Você deve ter uma variável chamada `thirdResult` que produz o valor `6` somando dois números. Substitua o booleano atualmente usado na expressão por um número.
4. Você deve ter uma variável chamada `fourthResult` que produz o valor `8` somando dois números. Substitua o booleano atualmente usado na expressão por um número.
5. Você deve ter uma variável chamada `fifthResult` que multiplica corretamente dois números para obter o produto `20`. Substitua a string atualmente usada na expressão por um número.
6. Você deve ter uma variável chamada `sixthResult` que adiciona corretamente dois números para produzir o valor `22`. Substitua o `null` atualmente usado na expressão por um número.

# --hints--

Você deve ter uma variável chamada `firstResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.firstResult);
```

`firstResult` deve ter o valor numérico `15`.

```js
assert.strictEqual(firstResult, 15);
```

Você deve ter uma variável chamada `secondResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.secondResult);
```

`secondResult` deve ter o valor numérico `3`.

```js
assert.strictEqual(secondResult, 3);
```

Você deve ter uma variável chamada `thirdResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.thirdResult);
```

`thirdResult` não deve incluir um valor booleano (`true` ou `false`).

```js
const explorer = await __helpers.Explorer(code);
const { thirdResult } = explorer.variables;
assert.notInclude(thirdResult?.value?.toString(), "true")
assert.notInclude(thirdResult?.value?.toString(), "false")
```

`thirdResult` deve ter o valor numérico `6`.

```js
assert.strictEqual(thirdResult, 6);
```

Você deve ter uma variável chamada `fourthResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.fourthResult);
```

`fourthResult` não deve incluir um valor booleano (`true` ou `false`).

```js
const explorer = await __helpers.Explorer(code);
const { fourthResult } = explorer.variables;
assert.notInclude(fourthResult?.value?.toString(), "true")
assert.notInclude(fourthResult?.value?.toString(), "false")
```

`fourthResult` deve ter o valor numérico `8`.

```js
assert.strictEqual(fourthResult, 8);
```

Você deve ter uma variável chamada `fifthResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.fifthResult);
```

`fifthResult` não deve incluir uma multiplicação de string (por exemplo, "10" * 2).

```js
const explorer = await __helpers.Explorer(code);
const { fifthResult } = explorer.variables;
const expression = fifthResult?.value?.toString() || '';
assert.notMatch(expression, /(["'`]).*\1\s*\*/);
```

`fifthResult` deve ter o valor numérico `20`.

```js
assert.strictEqual(fifthResult, 20);
```

Você deve ter uma variável chamada `sixthResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.sixthResult);
```

`sixthResult` não deve incluir `null` em sua expressão.

```js
const explorer = await __helpers.Explorer(code);
const { sixthResult } = explorer.variables;
assert.notInclude(sixthResult?.value?.toString(), "null");
```

`sixthResult` deve ter o valor numérico `22`.

```js
assert.strictEqual(sixthResult, 22);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
