---
id: 66ea7fc048a6e26e33c598fe
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Ahora, es hora de generar un número aleatorio usando el método `Math.random()`.

Crea una variable llamada `randomNum` y asígnale el valor del resultado de llamar al método `Math.random()`.

Luego, registra la variable `randomNum` en la consola.

Intenta agregar un espacio en el código para reiniciar el bot y ver diferentes números aleatorios que se generan.

# --hints--

Deberías tener una variable llamada `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomNum);
```

Deberías asignar el resultado de llamar al método `Math.random()` a la variable `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
const randomNum = explorer.variables.randomNum;
assert.equal(randomNum?.value?.toString(), "Math.random()");
```

Deberías registrar la variable `randomNum` en la consola.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
