---
id: 66ea8239b9eeb76f12299030
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

La siguiente parte del taller es revisar cómo generar un número aleatorio entre dos valores.

Empiece añadiendo otro `console.log()` que registre el mensaje `"Now, generate a random number between two values."` en la consola.

Luego, cree una variable llamada `min` y asígnele el valor de `1` y cree una variable llamada `max` y asígnele el valor de `100`.

En el siguiente paso, generará un número aleatorio entre estos dos valores.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debería tener un `console.log()` que registre el mensaje `"Now, generate a random number between two values."`.

```js
assert.equal(getLogs()[3], "Now, generate a random number between two values.");
```

Debería tener una variable llamada `min`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.min);
```

La variable `min` debe ser un número.

```js
assert.isNumber(min);
```

Debe asignar el número `1` a su variable `min`.

```js
assert.equal(min, 1);
```

Debería tener una variable llamada `max`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.max);
```

La variable `max` debe ser un número.

```js
assert.isNumber(max);
```

Debe asignar el número `100` a su variable `max`.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
