---
id: 66ea7e7399244c6d4b6052a1
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

La próxima parte del taller revisará cómo funciona el método `Math.random()`.

Comienza agregando otro `console.log()` que registre la cadena `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` en la consola.
 
# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes tener un `console.log()` que registre la cadena `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."`.

```js
assert.equal(
  getLogs()[1],
  "The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."
);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
