---
id: 66ea7adae8053065a64f9002
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, revisarás cómo trabajar con los diferentes métodos del `Math` al construir un Mathbot. Este Mathbot registrará algunas operaciones matemáticas y mensajes en la consola en un esfuerzo por enseñarte sobre el objeto `Math` de JavaScript.

Para este primer paso, comienza creando una variable llamada `botName` y asígnale el valor de cadena `"MathBot"`.

Luego, crea otra variable llamada <code>greeting` y asígnale la frase `"¡Hola!" Me llamo [aquí va el nombre del bot] y estoy aquí para enseñarte sobre el objeto Math!"</code>. En lugar del marcador de posición `[aquí va el nombre del bot]`, utiliza la variable `botName`. 

Eres libre de usar literales de plantilla o concatenación de cadenas con el operador `+` para lograrlo.

Finalmente, registra la variable `greeting` en la consola.

# --hints--

Debe tener una variable llamada `botName`.

```js
assert.isNotNull(botName);
```

Su variable `botName` debe ser una cadena.

```js
assert.isString(botName);
```

Debes asignar el valor de cadena `"MathBot"` a la variable `botName`.

```js
assert.strictEqual(botName, "MathBot");
```

Debes tener una variable llamada `greeting`.

```js
assert.isNotNull(greeting);
```

La variable `greeting` debería ser una cadena.

```js
assert.isString(greeting);
```

You should assign the following sentence to the `greeting` variable: `"Hi there! Me llamo [MathBot va aquí] y estoy aquí para enseñarte sobre el objeto Math!"</code>. Asegúrate de reemplazar `[MathBot va aquí]` con la variable `botName`.

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

Deberías registrar la variable `greeting` en la consola.

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
