---
id: 66ea8239b9eeb76f12299030
title: Pasos 4
challengeType: 1
dashedName: step-4
---

# --description--

La siguiente parte del taller es revisar cómo generar un número aleatorio entre dos valores.

Empiece añadiendo otro `console.log()` que registre el mensaje `"Ahora, genera un número aleatorio entre dos valores."` en la consola.

Luego, cree una variable llamada `min` y asígnele el valor de `1` y cree una variable llamada `max` y asígnele el valor de `100`.

En el siguiente paso, generará un número aleatorio entre estos dos valores.

# --hints--

Debería tener un `console.log()` que registre el mensaje `"Ahora, genera un número aleatorio entre dos valores."`.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

Debería tener una variable llamada `min`.

```js
assert.isNotNull(min);
```

La variable `min` debe ser un número.

```js
assert.isNumber(min);
```

Debe asignar el número `1` a su variable `min`.

```js
assert.equal(min, 1);
```

Debería tener una variable llamada `max`.

```js
assert.isNotNull(max);
```

La variable `max` debe ser un número.

```js
assert.isNumber(max);
```

Debe asignar el número `100` a su variable `max`.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
