---
id: 66ea845f54195c7020cc521b
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

La fórmula para generar un número aleatorio entre dos valores es la siguiente:

```js
Math.random() * (maximum - minimum) + minimum;
```

Esto producirá un resultado que es un número de punto flotante entre dos valores.

Crea una variable llamada `randomNum2` y asígnale el resultado de generar un valor entre los valores `min` y `max`.

Luego, registra la variable `randomNum2` para ver el resultado. Intenta actualizar la página para ver diferentes resultados.

# --hints--

Debe tener una variable llamada `randomNum2`.

```js
assert.isNotNull(randomNum2);
```

Su variable `randomNum2` debe ser un número.

```js
assert.isNumber(randomNum2);
```

Debe asignar el resultado de generar un valor entre los valores `min` y `max`. Asegúrate de no establecer un valor fijo.  You should use the formula provided above.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

Debe registrar la variable `randomNum2` en la consola.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
