---
id: 66eb488c0129e192542f0865
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

In previous lessons, you learned how to work with the `Math.ceil()` method like this:

```js
const price = 10.01;
Math.ceil(price); // 11
```

El método `Math.ceil()` redondea el valor hacia arriba al número entero más cercano.

Crea una variable llamada `numRoundedUp` y asígnale el resultado de redondear el número de punto flotante `3.2` hacia arriba al número entero más cercano.

Luego, registra en el console la variable `numRoundedUp` para ver el resultado.

# --hints--

Deberías tener una variable `numRoundedUp`.

```js
assert.isNotNull(numRoundedUp);
```

La variable `numRoundedUp` debería ser un número.

```js
assert.isNumber(numRoundedUp);
```

Tu variable `numRoundedUp` debería ser igual al resultado de redondear el número de punto flotante `3.2` hacia abajo al número entero más cercano.

```js
assert.equal(numRoundedUp, 4);
```

No debes codificar de manera fija el valor `4` para la variable `numRoundedUp`.

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

Deberías registrar en el console la variable `numRoundedUp`.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
