---
id: 69218e46f4490b33f2481a1a
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

La fórmula para generar un entero aleatorio entre dos valores es la siguiente:

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

Esto producirá un resultado que es un entero entre dos valores.

Crea una variable llamada `randomInt` y asígnale el resultado de generar un valor entre los valores `min` y `max`.

Luego, registra la variable `randomInt` para ver el resultado. Intenta actualizar la página para ver resultados diferentes.

# --hints--

Debes tener una variable llamada `randomInt`.

```js
assert.isNotNull(randomInt);
```

Tu variable `randomInt` debe ser un número.

```js
assert.isNumber(randomInt);
```

Debe asignar el resultado de generar un valor entre los valores `min` y `max`. Asegúrate de no establecer un valor fijo.  You should use the formula provided above.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

Debes registrar la variable `randomInt` en la consola.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
