---
id: 68e4c81577243e4db1676491
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, practicarás trabajando con HTML semántico construyendo una página web que incluye algunos de los consejos de Quincy Larson para conseguir un trabajo como desarrollador. La estructura básica de HTML ya ha sido preparada para ti.

Comienza creando un elemento `h1` con el texto `Quincy's Tips for Getting a Developer Job`.

# --hints--

Debería tener un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe tener el texto `Quincy's Tips for Getting a Developer Job`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
   
--fcc-editable-region--
  </body>
</html>
```
