---
id: 68e65ec5348e40a739903e6d
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Ahora, crea un elemento de párrafo debajo del elemento `h1`. Dentro de este párrafo, agrega el texto `Learning to code is hard, but as Quincy Larson says, You can become a developer.`.

# --hints--

Debes tener un elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Tu párrafo debe tener el texto `Learning to code is hard, but as Quincy Larson says, You can become a developer.`. Revisa bien tu ortografía.

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

Tu párrafo debe estar debajo del elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
   
--fcc-editable-region--
  </body>
</html>
```
