---
id: 68e964b84f954c322376e592
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

If the source of a quote is a website, you can cite it with the `cite` attribute. El valor de este atributo debe ser una URL válida. Aunque este atributo no cambia la presentación de la cita en bloque, es muy útil para dar a los lectores de pantalla y motores de búsqueda más información sobre la cita.

Aquí tienes un ejemplo de un elemento de cita en línea con un atributo `cite`:

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

Agrega el atributo `cite` al elemento de cita en línea con esta URL:

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

Tu elemento `q` debe tener un atributo `cite`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

El atributo `cite` de tu elemento `q` debe tener el valor `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
