---
id: 68e90335bd53a901c41dfc13
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

En una lección anterior aprendiste que el elemento `q` se usa para distinguir el texto citado del contenido circundante.

Aquí hay un ejemplo:

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

La mayoría de los navegadores modernos añadirán comillas alrededor de una cita en línea automáticamente cuando uses el elemento `q`.

Dentro del párrafo envuelve `You can become a developer.` en un elemento de cita en línea, manteniendo el resto del párrafo sin cambios.

# --hints--

Debes tener un elemento `q`.

```js
assert.exists(document.querySelector('q'));
```

Tu elemento `q` debe estar anidado dentro del párrafo que está debajo del elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

Tu elemento `q` debe tener el texto `You can become a developer.`.

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

Dentro de tu elemento `q`, el texto no debe tener espacios en blanco al inicio ni al final (sin espacios, tabulaciones ni saltos de línea antes del primer carácter o después del último).

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

Solo debes agregar etiquetas `q` alrededor de la parte apropiada de la oración sin cambiar el texto original. El elemento `p` debe seguir leyendo: `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
