---
id: 68eaa9ef8a0867071f195f59
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Hasta ahora has estado usando el atributo `cite` para atribuir la fuente de la cita, pero el atributo realmente no muestra la fuente al usuario.

Si quieres atribuir la fuente visualmente, puedes agregar un elemento de cita, `cite`, fuera del elemento de cita en bloque. El elemento de cita es un elemento HTML que puedes usar para marcar el título de una obra creativa referenciada, como un libro, artículo, canción, película, sitio web o artículo de investigación.

Aquí tienes un ejemplo:

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>—Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

Dentro del elemento `p` debajo del elemento de cita en bloque, envuelve `How to Learn to Code and Get a Developer Job [Full Book]` en un elemento `cite`.

# --hints--

Debes tener un elemento `cite` dentro del elemento `p` debajo del elemento `blockquote`.

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

Tu elemento `cite` debe tener el texto `How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

Ignorando el elemento `cite` que acabas de agregar, el texto del elemento `p` debería seguir siendo `—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
--fcc-editable-region--
        <p>
          —Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
        </p>
--fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
