---
id: 68ea7e7b916e8270483689b6
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Debajo del elemento `h2`, agregarás otra cita de Quincy. Esta vez, la cita no formará parte de un párrafo más grande. En cambio, todo el párrafo será una cita. Para distinguir el texto citado de esta manera, debes usar el elemento de cita en bloque: `blockquote`. En el navegador, verás que el texto está ligeramente indentado.

Aquí tienes un ejemplo de un elemento de cita en bloque con texto citado:

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Ahora, dentro de la primera sección, agrega un elemento de cita en bloque debajo del elemento `h2` con el texto `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

# --hints--

Dentro del primer elemento `section`, debes tener un elemento `blockquote` debajo del elemento `h2`.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

Tu elemento `blockquote` debe contener el texto `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
--fcc-editable-region--
      <section>
        <h2>Envisioning Success</h2>

      </section>
--fcc-editable-region--
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
