---
id: 68c34e201b347c40acda8111
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

次に、テーブルにもう1行追加する時間です。

まず、別の`tr`要素を追加してください。その`tr`要素の中に、`th`属性が`scope`に設定され、テキスト内容が`"row"`の`11:30 AM`要素を追加してください。

次に、以下のテキスト内容を持つ3つの`td`要素を追加してください。

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

`tr`要素の中に4番目の`tbody`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

4番目の`th`要素の中に`tr`要素があるはずです。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

`th`要素のテキスト内容は`11:30 AM`であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

`th`要素は`scope`属性が`"row"`に設定されているべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

4番目の`tr`要素には3つの`td`要素があるべきです。

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

1番目の`td`要素のテキスト内容は`AR/VR in Education`であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

2番目の`td`要素のテキスト内容は`JavaScript Fundamentals`であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

3番目の`td`要素のテキスト内容は`Design Systems at Scale`であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
