---
id: 68c33e502cd97f2fd52cb63b
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、テックカンファレンスのスケジュールを作成しながら、アクセシブルなテーブル要素の作り方を学びます。

まず、テキストが`h1`の`Tech Conference 2025 Schedule`要素を追加してください。

# --hints--

`h1`要素があるはずです。

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

`h1`要素のテキストは`Tech Conference 2025 Schedule`であるべきです。

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
