---
id: 68c34d57ed874e4004f7b5d4
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

現在、テキスト内容が `td` の `Break` 要素は1つの列にしかまたがっていません。しかし、3つの列すべてにまたがると便利です。

以前のワークショップやレッスンで学んだように、テーブルのセルを複数の列にまたがらせるには `colspan` 属性を使えます。

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

`colspan` 要素に `td` 属性を追加し、その値を `3` に設定してください。

# --hints--

テキスト内容が `td` の `Break` 要素には、`colspan` 属性が `3` に設定されている必要があります。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
