---
id: 68c346d03709343bff132707
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

次のいくつかのステップでは、テーブルの本文を作成していきます。

まず、`tbody` 要素の下に `thead` 要素を追加してください。次に、`tbody` の中に `tr` 要素を追加します。

# --hints--

`tbody` 要素の下に `thead` 要素があるはずです。

```js
assert.exists(document.querySelector("tbody"));
```

`tr` 要素の中に `tbody` 要素があるはずです。

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
