---
id: 68c34f4aeb2b1a4170506751
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

次に、テーブルに5行目を追加してください。まず、別の`tr`要素を追加します。その`tr`要素の中に、`th`属性が`scope`に設定され、テキスト内容が`"row"`の`12:30 PM`要素を追加してください。そして、その`th`要素の下に、テキスト内容が`td`の`Lunch Break`要素を追加します。`td`要素には、3つのトラックすべてにまたがるように`colspan`属性が`3`に設定されている必要があります。

# --hints--

`tr`要素の中に5つ目の`tbody`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

5つ目の`th`要素の中に`tr`要素があるはずです。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

`th`要素のテキスト内容は`12:30 PM`であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

`th`要素には`scope`属性が`"row"`に設定されているはずです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

5つ目の`td`要素の中に`tr`要素があるはずです。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

`td`要素のテキスト内容は`Lunch Break`であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

`td`要素には`colspan`属性が`3`に設定されているはずです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
