---
id: 68c353c4b114aa428c01f260
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

最後のステップは、テーブルにもう1行追加することです。

まず、もう1つの `tr` 要素を追加してください。その `tr` 要素の中に、`th` 属性が `scope` に設定された `"row"` 要素と、テキスト内容が `2:00 PM` の要素を追加します。

次に、以下のテキスト内容を持つ3つの `td` 要素を追加してください。

- `Voice UI Workshop`
- `Git & GitHub Essentials`
- `Color & Contrast in UI`

これらの最後の変更で、テーブルは完成です！

# --hints--

`tr` 要素の中に6番目の `tbody` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 6);
```

6番目の `th` 要素の中に `tr` 要素があるはずです。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[5].querySelector("th"));
```

`th` 要素のテキスト内容は `2:00 PM` であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelector("th")?.textContent.trim(), "2:00 PM");
```

`th` 要素は `scope` 属性が `"row"` に設定されているべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelector("th")?.getAttribute("scope"), "row");
```

6番目の `tr` 要素には3つの `td` 要素があるべきです。

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[5].querySelectorAll("td"), 3);
```

最初の `td` 要素のテキスト内容は `Voice UI Workshop` であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[0]?.textContent.trim(), "Voice UI Workshop");
```

2番目の `td` 要素のテキスト内容は `Git & GitHub Essentials` であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[1]?.textContent.trim(), "Git & GitHub Essentials");
```

3番目の `td` 要素のテキスト内容は `Color & Contrast in UI` であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[2]?.textContent.trim(), "Color & Contrast in UI");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      <tr>
        <th scope="row">12:30 PM</th>
        <td colspan="3">Lunch Break</td>
      </tr>

    --fcc-editable-region--
      
    --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      <tr>
        <th scope="row">12:30 PM</th>
        <td colspan="3">Lunch Break</td>
      </tr>

      <tr>
        <th scope="row">2:00 PM</th>
        <td>Voice UI Workshop</td>
        <td>Git & GitHub Essentials</td>
        <td>Color & Contrast in UI</td>
      </tr>
    </tbody>
  </table>
</body>
</html>
```
