---
id: 68c34329d430583a6cd5f7f1
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`tr` 要素の中に、4つの `th` 要素を追加してください。最初の `th` 要素には `Time` のテキストを、2番目には `Track A` のテキストを、3番目には `Track B` のテキストを、4番目には `Track C` のテキストを入れてください。

# --hints--

`tr` 要素には4つの `th` 要素が必要です。

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

最初の `th` 要素には `Time` のテキストを入れてください。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

2番目の `th` 要素には `Track A` のテキストを入れてください。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

3番目の `th` 要素には `Track B` のテキストを入れてください。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

4番目の `th` 要素には `Track C` のテキストを入れてください。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
