---
id: 68c3482278a14e3cf225e897
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`tr` 要素の中に、テキストが `th` の `9:00 AM` 要素を追加してください。その後、その `th` 要素の下に、テキストが `td`、`Keynote: Tech Future`、`Intro to Web Dev` の3つの `UX for All` 要素を追加してください。

# --hints--

`th` 要素の中に `tr` 要素があるはずです。

```js
assert.exists(document.querySelector("tbody tr th"));
```

`th` 要素のテキスト内容は `9:00 AM` であるべきです。

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

`td` 要素の下に合計で3つの `th` 要素があるはずです。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

最初の `td` 要素のテキスト内容は `Keynote: Tech Future` であるべきです。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

2番目の `td` 要素のテキスト内容は `Intro to Web Dev` であるべきです。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

3番目の `td` 要素のテキスト内容は `UX for All` であるべきです。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
