---
id: 68c341332d98fe38fae733af
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

いよいよテーブルの作成を始めましょう。

まず、`table` 要素の下に `h1` 要素を追加してください。`table` 要素の中に、テキスト `caption` を持つ `Schedule by Track and Time` 要素を追加しましょう。

# --hints--

`table` 要素があるはずです。

```js
assert.exists(document.querySelector("table"));
```

`table` 要素は `caption` 要素を持つべきです。

```js
assert.exists(document.querySelector("table caption"));
```

`caption` 要素にはテキスト `Schedule by Track and Time` が含まれているはずです。

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
