---
id: 68c3420cb0176439a95af1b8
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

次のステップはテーブルヘッドを追加することです。

`thead` 要素の下に `caption` 要素を追加してください。`thead` の中に `tr` 要素を追加します。

# --hints--

`thead` 要素があるはずです。

```js
assert.exists(document.querySelector("table thead"));
```

`tr` の中に `thead` 要素があるはずです。

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
