---
id: 68c34c608cf9e53f5d36b1d5
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

次に、テーブルに3行目を追加してください。まず、別の`tr`要素を追加します。その`tr`要素の中に、`th`属性が`scope`に設定された`"row"`要素を追加し、テキスト内容は`11:00 AM`にしてください。その後、その`th`要素の下に、テキスト内容が`td`の`Break`要素を追加してください。

# --hints--

`tr`要素の中に3つ目の`tbody`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

3つ目の`th`要素の中に`tr`要素があるはずです。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

`th`要素のテキスト内容は`11:00 AM`であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

`th`要素は`scope`属性が`"row"`に設定されているべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

3つ目の`td`要素の中に`tr`要素があるはずです。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

`td`要素のテキスト内容は`Break`であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
