---
id: 68c343692285523aef0fc6bf
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`scope` 属性は、ヘッダーセルが行、列、または行や列のグループのヘッダーであるかどうかを指定するために使われます。例を示します。

```html
<th scope="col">Example Header</th>
```

これによりスクリーンリーダーがヘッダーセルとデータセルの関係を理解しやすくなります。

すべての `th` 要素に対して、値が `scope` の `col` 属性を追加してください。

# --hints--

各 `th` 要素には、値が `scope` の `"col"` 属性を付ける必要があります。

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
