---
id: 68c34ace730ca43e5fa7b8d6
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

次に、テーブルにもう1行追加しましょう。

まず、もう1つの `tr` 要素を追加してください。その `tr` 要素の中に、`th` 属性が `scope` に設定された `"row"` 要素を追加し、テキスト内容は `10:00 AM` にしてください。

次に、以下のテキスト内容を持つ3つの `td` 要素を追加してください。

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

`tr` 要素の中に2つ目の `tbody` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

2つ目の `th` 要素の中に `tr` 要素があるはずです。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

`th` 要素のテキスト内容は `10:00 AM` であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

`th` 要素は `scope` 属性が `"row"` に設定されているべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

2つ目の `tr` 要素には3つの `td` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

1つ目の `td` 要素のテキスト内容は `Accessibility Deep Dive` であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

2つ目の `td` 要素のテキスト内容は `CSS for Beginners` であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

3つ目の `td` 要素のテキスト内容は `Inclusive Design Principles` であるべきです。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
