---
id: 5efc4f528d6a74d05e68af74
title: Passo 56
challengeType: 0
dashedName: step-56
---

# --description--

Existe outra maneira de associar o texto de um elemento `input` com o elemento em si. Você pode colocar o texto dentro de um elemento `label` e adicionar um atributo `for` com o mesmo valor que o atributo `id` do elemento `input`.

Tendo um elemento `input` como este abaixo:

```html
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

Um exemplo de um elemento `label` que está associado a este elemento `input` é:

```html
<label for="breakfast">Breakfast</label>
```

Associe o texto `Loving` com a caixa de seleção inserindo apenas o texto `Loving` em um elemento `label` e dando a ele um atributo `for` apropriado.

# --hints--

Você deve se certificar de que o elemento de caixa de seleção ainda está presente e que você não fez nenhuma alteração nele. Ele deve estar como no código inicial.

```js
const checkboxElementRegex = /\s*<input\s+id\s*=\s*"loving"\s*type\s*=\s*"checkbox"\s*>/;
assert.match(code, checkboxElementRegex);
```

A caixa de seleção deve ter o atributo `id` com o valor `loving`.  Eram esperados `--fcc-expected--` elementos, mas foram encontrados `--fcc-actual--`.

```js
const checkboxElementId = document.querySelector('input[type="checkbox"]')?.getAttribute('id');
assert.strictEqual(checkboxElementId, 'loving');
```

Você deve adicionar um novo elemento `label` após o elemento de caixa de seleção. Certifique-se de que os dois tenham uma tag de abertura e fechamento.

```js
assert.lengthOf(document.querySelectorAll('label'), 3);
assert.lengthOf(code.match(/<\/label\>/g), 3);
```

O elemento `label` deve estar localizado após a caixa de seleção.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElement = checkboxElement?.nextElementSibling?.nodeName;
assert.strictEqual(checkboxSiblingElement, 'LABEL');
```

O elemento `label` deve ter o texto `Loving`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementText = checkboxElement?.nextElementSibling?.innerText;

assert.strictEqual(checkboxSiblingElementText, 'Loving');
```

O texto do label `Loving` deve estar dentro do novo elemento `label` apenas. Certifique-se de que você não duplicou o texto fora do elemento `label`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingNode = checkboxElement?.nextSibling?.nodeValue ?? '';
const labelElement = checkboxElement?.nextElementSibling;
const labelSiblingNode = labelElement?.nextSibling?.nodeValue ?? '';

assert.notInclude(checkboxSiblingNode, 'Loving');
assert.notInclude(labelSiblingNode, 'Loving');
```

O elemento `label` não deve ter um atributo `for`. Certifique-se de que há um espaço entre a tag de abertura de `label` e o atributo `for`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.isNotNull(checkboxSiblingElementAttr);
```

O elemento `label` deve ter o atributo `for` com o valor de `loving`.  Eram esperados `--fcc-expected--` elementos, mas foram encontrados `--fcc-actual--`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.strictEqual(checkboxSiblingElementAttr, 'loving');
```

Deve haver um espaço entre a caixa de seleção e o novo elemento `label`.

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

