---
id: 66fbcf750a62784cf11f562d
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Erstellen Sie ein Objekt mit dem Namen `recipe1`. Innerhalb des `recipe1`-Objekts erstellen Sie eine `name`-Eigenschaft mit dem Wert `"Spaghetti Carbonara"`.

Erstellen Sie außerdem innerhalb des `recipe1`-Objekts eine `ingredients`-Eigenschaft mit einem Array als Wert. Das Array soll `spaghetti`, `Parmesan cheese`, `pancetta` und `black pepper` enthalten.

# --hints--

Sie sollten ein Objekt mit dem Namen `recipe1` erstellen.

```js
assert.isObject(recipe1);
```

Ihr `recipe1`-Objekt sollte eine `name`-Eigenschaft haben.

```js
assert.property(recipe1, 'name');
```

Ihre `name`-Eigenschaft sollte auf `Spaghetti Carbonara` gesetzt sein.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

Ihr `recipe1`-Objekt sollte eine `ingredients`-Eigenschaft haben.

```js
assert.property(recipe1, 'ingredients');
```

Ihre `ingredients`-Eigenschaft sollte einen Array-Wert haben.

```js
assert.isArray(recipe1.ingredients);
```

Der Array-Wert Ihrer `ingredients`-Eigenschaft sollte `spaghetti`, `Parmesan cheese`, `pancetta` und `black pepper` enthalten.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
