---
id: 66fbcdfed6848e48947dbbc9
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop erstellen Sie einen Rezept-Tracker mit JavaScript-Objekten.

Beginnen Sie damit, ein leeres Array mit dem Namen `recipes` zu erstellen. In dieses Array werden Sie später die Rezeptobjekte einfügen.

# --hints--

Sie sollten ein Array mit dem Namen `recipes` erstellen.

```js
assert.isArray(recipes);
```

Ihr `recipes`-Array sollte leer sein.

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
