---
id: 66fbcf750a62784cf11f5632
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Erstellen Sie eine `getTotalIngredients`-Funktion, die ein einzelnes Argument entgegennimmt, das ein Array mit Zutaten darstellt, und die Anzahl der Zutaten aus dem an die Funktion übergebenen Array zurückgibt.

# --hints--

Sie sollten eine `getTotalIngredients`-Funktion erstellen.

```js
assert.isFunction(getTotalIngredients);
```

Ihre `getTotalIngredients`-Funktion sollte einen einzelnen Parameter haben.

```js
assert.lengthOf(getTotalIngredients, 1);
```

Ihre `getTotalIngredients`-Funktion sollte eine Zahl zurückgeben.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)` sollte `4` zurückgeben.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` sollte `5` zurückgeben.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` sollte `3` zurückgeben.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
