---
id: 66fbcf750a62784cf11f5634
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Es ist Zeit, jede der Funktionen zu testen. Sie können dafür eines der Rezepte verwenden, aber dieses Tutorial beginnt mit `recipe1`.

Erstellen Sie zwei neue Variablen: `recipe1TotalIngredients` und `recipe1DifficultyLevel`. Weisen Sie ihnen die Werte zu, indem Sie die entsprechende Funktion für jede Variable aufrufen und die passende `recipe1`-Eigenschaft übergeben.

Protokollieren Sie schließlich jede Variable in der Konsole, um die Ergebnisse zu sehen.

# --hints--

Sie sollten eine `recipe1TotalIngredients`-Variable erstellen.

```js
assert.isNotNull(recipe1TotalIngredients);
```

Ihre `recipe1TotalIngredients`-Variable sollte auf `getTotalIngredients(recipe1.ingredients)` gesetzt werden.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

Sie sollten `recipe1TotalIngredients` in der Konsole protokollieren.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

Sie sollten eine `recipe1DifficultyLevel`-Variable erstellen.

```js
assert.isNotNull(recipe1DifficultyLevel);
```

Ihre `recipe1DifficultyLevel`-Variable sollte auf `getDifficultyLevel(recipe1.cookingTime)` gesetzt werden.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

Sie sollten `recipe1DifficultyLevel` in der Konsole protokollieren.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
