---
id: 66fbcf750a62784cf11f562e
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Erstellen Sie ein `recipe2`-Objekt mit den folgenden Eigenschaften und Werten:

| Schlüssel | Wert   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

Sie sollten ein Objekt mit dem Namen `recipe2` erstellen.

```js
assert.isObject(recipe2);
```

Ihr `recipe2`-Objekt sollte eine `name`-Eigenschaft haben.

```js
assert.property(recipe2, "name");
```

Ihre `name`-Eigenschaft sollte auf `Chicken Curry` gesetzt sein.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

Ihr `recipe2`-Objekt sollte eine `ingredients`-Eigenschaft haben.

```js
assert.property(recipe2, "ingredients");
```

Ihre `ingredients`-Eigenschaft sollte einen Array-Wert haben.

```js
assert.isArray(recipe2.ingredients);
```

Der Array-Wert Ihrer `ingredients`-Eigenschaft sollte `chicken breast`, `coconut milk`, `curry powder`, `onion` und `garlic` enthalten.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

Ihr `recipe2`-Objekt sollte eine `cookingTime`-Eigenschaft haben.

```js
assert.property(recipe2, "cookingTime");
```

Der Wert Ihrer `cookingTime`-Eigenschaft sollte eine Zahl sein.

```js
assert.isNumber(recipe2.cookingTime);
```

Ihre `cookingTime`-Eigenschaft sollte auf `42` gesetzt sein.

```js
assert.equal(recipe2.cookingTime, 42);
```

Ihr `recipe2`-Objekt sollte eine `totalIngredients`-Eigenschaft haben.

```js
assert.property(recipe2, "totalIngredients");
```

Ihre `totalIngredients`-Eigenschaft sollte auf `null` gesetzt sein.

```js
assert.isNull(recipe2.totalIngredients);
```

Ihr `recipe2`-Objekt sollte eine `difficultyLevel`-Eigenschaft haben.

```js
assert.property(recipe2, "difficultyLevel");
```

Ihre `difficultyLevel`-Eigenschaft sollte auf einen leeren String gesetzt sein.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
