---
id: 66fbcf750a62784cf11f5630
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Sie sollten nun die drei Objekte in das `recipes` Array pushen. Dazu können Sie die `push()` Methode verwenden.

Verwenden Sie die `push()` Methode, um alle Rezeptobjekte in das `recipes` Array zu pushen. Achten Sie darauf, `recipe1`, `recipe2` und `recipe3` in dieser Reihenfolge zu pushen.

Löschen Sie außerdem die Variablen `recipe1Name`, `recipe2CookingTime` und `recipe3Ingredients` sowie die `console.log` Anweisungen, die diese Variablen ausgeben.

# --hints--

Sie sollten alle Rezeptobjekte in das `recipes` Array pushen. Achten Sie darauf, sie in der Reihenfolge zu pushen, in der sie deklariert sind.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

Sie sollten die Variable `recipe1Name` löschen.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

Sie sollten die Anweisung `console.log(recipe1Name);` löschen.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

Sie sollten die Variable `recipe2CookingTime` löschen.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

Sie sollten die Anweisung `console.log(recipe2CookingTime);` löschen.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

Sie sollten die Variable `recipe3Ingredients` löschen.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

Sie sollten die Anweisung `console.log(recipe3Ingredients);` löschen.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
