---
id: 66fbcf750a62784cf11f5633
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Erstellen Sie eine `getDifficultyLevel`-Funktion, die eine Zahl als Parameter entgegennimmt, welche die Kochzeit angibt.

Wenn die Kochzeit kleiner oder gleich `30` ist, soll die Funktion `"easy"` zurückgeben. Wenn sie kleiner oder gleich `60` ist, soll die Funktion `"medium"` zurückgeben. Andernfalls soll die Funktion `"hard"` zurückgeben.

# --hints--

Sie sollten eine `getDifficultyLevel`-Funktion erstellen.

```js
assert.isFunction(getDifficultyLevel);
```

Ihre `getDifficultyLevel`-Funktion sollte einen einzelnen Parameter haben.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

Ihre `getDifficultyLevel`-Funktion sollte `"easy"` zurückgeben, wenn die Kochzeit kleiner oder gleich `30` ist.

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

Ihre `getDifficultyLevel`-Funktion sollte `"medium"` zurückgeben, wenn die Kochzeit größer als `31` und kleiner oder gleich `60` ist.

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

Ihre `getDifficultyLevel`-Funktion sollte `"hard"` zurückgeben, wenn die Kochzeit größer als `60` ist.

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
