---
id: 66fbcf750a62784cf11f5635
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Sie können jetzt jedes Element des `recipes`-Arrays mit Werten für die `totalIngredients`- und `difficultyLevel`-Eigenschaften füllen.

Greifen Sie vorerst auf die `totalIngredients`- und `difficultyLevel`-Eigenschaften von `recipe1` zu und setzen Sie diese auf die entsprechenden Ergebnisse von Funktionsaufrufen und Argumenten.

# --hints--

Sie sollten auf die `totalIngredients`-Eigenschaft von `recipe1` zugreifen.

```js
assert.isNotNull(recipe1.totalIngredients);
```

Sie sollten das Ergebnis des Aufrufs von `getTotalIngredients` mit `recipe1.ingredients` der `totalIngredients`-Eigenschaft von `recipe1` zuweisen.

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

Sie sollten auf die `difficultyLevel`-Eigenschaft von `recipe1` zugreifen.

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

Sie sollten das Ergebnis des Aufrufs von `getDifficultyLevel` mit `recipe1.cookingTime` der Eigenschaft `difficultyLevel` von `recipe1` zuweisen.

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
