---
id: 66ff02792b3042a2b811f475
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Fügen Sie dem `recipe1`-Objekt die folgenden Eigenschaften hinzu:

| Schlüssel | Wert   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

Die Eigenschaften mit `null` und leeren String-Werten werden später aktualisiert, nachdem Sie sie berechnet haben.

# --hints--

Ihr `recipe1`-Objekt sollte eine `cookingTime`-Eigenschaft haben.

```js
assert.property(recipe1, "cookingTime");
```

Ihre `cookingTime`-Eigenschaft sollte auf `22` gesetzt werden.

```js
assert.equal(recipe1.cookingTime, 22);
```

Ihr `recipe1`-Objekt sollte eine `totalIngredients`-Eigenschaft haben.

```js
assert.property(recipe1, "totalIngredients");
```

Ihre `totalIngredients`-Eigenschaft sollte auf `null` gesetzt werden.

```js
assert.isNull(recipe1.totalIngredients);
```

Ihr `recipe1`-Objekt sollte eine `difficultyLevel`-Eigenschaft haben.

```js
assert.property(recipe1, "difficultyLevel");
```

Ihre `difficultyLevel`-Eigenschaft sollte auf einen leeren String gesetzt werden.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
