---
id: 670e4f45f7116c0f216a5177
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Bevor Sie weitermachen, sollten Sie üben, wie man auf Eigenschaften eines Objekts zugreift.

Sie können dafür entweder die Punkt- (`.`) oder die Klammernotation (`[]`) verwenden. Hier ist ein Beispiel:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Greifen Sie auf die Eigenschaft `name` von `recipe1` zu und weisen Sie sie der Variablen `recipe1Name` zu.

Greifen Sie als Nächstes auf die Eigenschaft `cookingTime` von `recipe2` zu und weisen Sie sie der Variablen `recipe2CookingTime` zu.

Greifen Sie schließlich auf die Eigenschaft `ingredients` von `recipe3` zu und weisen Sie sie der Variablen `recipe3Ingredients` zu.

Stellen Sie sicher, dass alle von Ihnen erstellten Variablen in der Konsole ausgegeben werden.

# --hints--

Sie sollten eine `recipe1Name`-Variable erstellen.

```js
assert.isDefined(recipe1Name);
```

Sie sollten den Wert der Eigenschaft `name` von `recipe1` Ihrer `recipe1name`-Variable zuweisen.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Sie sollten `recipe1Name` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Sie sollten eine `recipe2CookingTime`-Variable erstellen.

```js
assert.isDefined(recipe2CookingTime);
```

Sie sollten den Wert der Eigenschaft `cookingTime` von `recipe2` Ihrer `recipe2CookingTime`-Variable zuweisen.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Sie sollten `recipe2CookingTime` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

Sie sollten eine `recipe3Ingredients`-Variable erstellen.

```js
assert.isDefined(recipe3Ingredients);
```

Sie sollten den Wert der Eigenschaft `ingredients` von `recipe3` Ihrer `recipe3Ingredients`-Variable zuweisen.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

Sie sollten `recipe3Ingredients` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
