---
id: 69161f4c0c9c8c48b3c843c5
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Erstellen Sie ein `recipe3`-Objekt mit den folgenden Eigenschaften und Werten:

|Schlüssel|Wert|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

Sie sollten ein Objekt mit dem Namen `recipe3` erstellen.

```js
assert.isObject(recipe3);
```

Ihr `recipe3`-Objekt sollte eine `name`-Eigenschaft haben.

```js
assert.property(recipe3, "name");
```

Ihre `name`-Eigenschaft sollte auf `Vegetable Stir Fry` gesetzt sein.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

Ihr `recipe3`-Objekt sollte eine `ingredients`-Eigenschaft haben.

```js
assert.property(recipe3, "ingredients");
```

Ihre `ingredients`-Eigenschaft sollte einen Array-Wert haben.

```js
assert.isArray(recipe3.ingredients);
```

Der Array-Wert Ihrer `ingredients`-Eigenschaft sollte `broccoli`, `carrot` und `bell pepper` enthalten.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

Ihr `recipe3`-Objekt sollte eine `cookingTime`-Eigenschaft haben.

```js
assert.property(recipe3, "cookingTime");
```

Der Wert Ihrer `cookingTime`-Eigenschaft sollte eine Zahl sein.

```js
assert.isNumber(recipe3.cookingTime);
```

Ihre `cookingTime`-Eigenschaft sollte auf `15` gesetzt sein.

```js
assert.equal(recipe3.cookingTime, 15);
```

Ihr `recipe3`-Objekt sollte eine `totalIngredients`-Eigenschaft haben.

```js
assert.property(recipe3, "totalIngredients");
```

Ihre `totalIngredients`-Eigenschaft sollte auf `null` gesetzt sein.

```js
assert.isNull(recipe3.totalIngredients);
```

Ihr `recipe3`-Objekt sollte eine `difficultyLevel`-Eigenschaft haben.

```js
assert.property(recipe3, "difficultyLevel");
```

Ihre `difficultyLevel`-Eigenschaft sollte auf einen leeren String gesetzt sein.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
