---
id: 68cabc534d2e741f33a3fe1d
title: 게임 캐릭터 능력치 추적기 만들기
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

이 실습에서는 게임 캐릭터 능력치 추적기를 만듭니다. 이 프로그램은 특정 속성을 가진 캐릭터를 생성하고, 그 속성을 업데이트하며, 현재 캐릭터의 능력치를 조회할 수 있게 합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 게임 캐릭터를 나타내고 캐릭터 능력치를 관리하는 `GameCharacter`라는 클래스를 만드세요.
2. 인스턴스화할 때, 새 `GameCharacter` 객체는 다음 속성을 가져야 합니다:
    - `_name`는 인스턴스화 시점에 주어진 문자열로 설정됩니다.
    - `_health`는 `100`로 설정됩니다.
    - `_mana`는 `50`로 설정됩니다.
    - `_level`는 `1`로 설정됩니다.
3. 캐릭터 이름에 대한 읽기 전용 접근을 제공하는 `name` 속성을 만드세요.
4. `health` 속성에 대해:
     - 현재 체력을 반환하는 getter를 정의하세요.
     - 체력이 `0` 미만으로 설정되는 것을 방지하고, `100`를 초과하지 않도록 제한하는 setter를 정의하세요.
5. `mana` 속성에 대해:
     - 현재 마나를 반환하는 getter를 정의하세요.
     - 마나가 `0` 미만으로 설정되는 것을 방지하고, `50`를 초과하지 않도록 제한하는 setter를 정의하세요.
6. 캐릭터의 현재 레벨을 반환하는 `level` getter를 만드세요.
7. 다음을 수행하는 `level_up` 메서드를 정의하세요:
    - 캐릭터 레벨을 1만큼 증가시킵니다.
    - 체력은 `100`로, 마나는 `50`로 각각 해당 속성 setter를 사용해 초기화합니다.
    - `<name> leveled up to <level>!` 형식의 메시지를 출력하는데, `<name>`과 `<level>`는 각각 캐릭터 이름과 새 레벨로 교체되어야 합니다.
8. 다음을 포함하는 형식화된 문자열을 반환하는 `__str__` 메서드를 정의하세요:
    - 캐릭터 이름
    - 캐릭터 레벨
    - 캐릭터 현재 체력
    - 캐릭터 현재 마나
   예를 들어, `Kratos`라는 이름의 캐릭터는 인스턴스화 직후 다음과 같이 나타내야 합니다:

   ```md
   Name: Kratos
   Level: 1
   Health: 100
   Mana: 50
   ```

## 사용 예시

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

`GameCharacter` 클래스를 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

`__init__` 클래스에 `GameCharacter` 메서드가 있어야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

`__init__` 메서드는 두 개의 매개변수를 가져야 하며, 첫 번째는 `self`여야 합니다.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

인스턴스화 시점에 주어진 문자열을 `self._name`에 `__init__` 메서드 안에서 할당해야 합니다.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

`100` 메서드 안에서 `self._health`에 `__init__`를 할당해야 합니다.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

`50` 메서드 안에서 `self._mana`에 `__init__`를 할당해야 합니다.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

`1` 메서드 안에서 `self._level`에 `__init__`를 할당해야 합니다.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

`name` 클래스에 `GameCharacter` 메서드가 있어야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

`name` 메서드는 `self`라는 하나의 매개변수만 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

`self._name` 메서드에서 `name`를 반환해야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

`name` 메서드는 `@property` 데코레이터를 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

`health` 클래스에 `GameCharacter` 메서드가 있어야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

`health` 메서드는 `self`라는 하나의 매개변수만 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

`self._health` 메서드에서 `health`를 반환해야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

`health` 메서드는 `@property` 데코레이터를 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

`health` setter 메서드를 만들어야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

`health` setter 메서드는 `@health.setter` 데코레이터를 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

`health` setter 메서드는 두 개의 매개변수를 가져야 하며, 첫 번째는 `self`여야 합니다.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

`self._health` setter에 주어진 값이 `0`보다 작으면 `health`에 `0`를 설정해야 합니다.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

`self._health` setter에 주어진 값이 `100`보다 크면 `health`에 `100`를 설정해야 합니다.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

`self._health` setter에 주어진 값이 `health`와 `0` 사이일 때 `100`에 그 값을 설정해야 합니다.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
    char.health = 0
    assert char.health == 0
    char.health = 100
    assert char.health == 100
`)})
```

`mana` 클래스에 `GameCharacter` 메서드가 있어야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

`mana` 메서드는 `self`라는 하나의 매개변수만 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

`self._mana` 메서드에서 `mana`를 반환해야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

`mana` 메서드는 `@property` 데코레이터를 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

`mana` setter 메서드를 만들어야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

`mana` setter 메서드는 `@mana.setter` 데코레이터를 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

`mana` setter 메서드는 두 개의 매개변수를 가져야 하며, 첫 번째는 `self`여야 합니다.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

`self._mana` setter에 주어진 값이 `0`보다 작으면 `mana`에 `0`를 설정해야 합니다.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

`self._mana` setter에 주어진 값이 `50`보다 크면 `mana`에 `50`를 설정해야 합니다.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

`self._mana` setter에 주어진 값이 `mana`와 `0` 사이일 때 `50`에 그 값을 설정해야 합니다.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

`level` 클래스에 `GameCharacter` 메서드가 있어야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

`level` 메서드는 `self`라는 하나의 매개변수만 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

`self._level` 메서드에서 `level`를 반환해야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

`level` 메서드는 `@property` 데코레이터를 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

`level_up` 클래스에 `GameCharacter` 메서드가 있어야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

`level_up` 메서드는 `self`라는 하나의 매개변수만 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

`self._level` 메서드에서 `1`을 `level_up`만큼 증가시켜야 합니다.

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

`self.health` 메서드에서 `100`에 `level_up`를 설정해야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

`self.mana` 메서드에서 `50`에 `level_up`를 설정해야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

`level_up` 메서드는 `<name> leveled up to <level>!`를 출력해야 하며, `<name>`와 `<level>`는 각각 캐릭터 이름과 새 레벨로 교체되어야 합니다.

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

`__str__` 클래스에 `GameCharacter` 메서드가 있어야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

`__str__` 메서드는 `self`라는 하나의 매개변수만 가져야 합니다.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

`__str__` 메서드는 제공된 형식으로 캐릭터 능력치를 포함하는 문자열을 반환해야 합니다.

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
