---
id: 684cd027744c2a0520ba4d58
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS 使用直接套用到標記的工具類別。這些類別包含字型、顏色、Flexbox 版面配置、Grid 等的樣式。 

要在 Tailwind CSS 中套用背景顏色，你可以使用 `bg` 前置式，接著是顏色名稱和數字。該數字表述該顏色的色階。

這裡是一個將深藍色背景色套用到段落元素的範例。

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

在你的開啟 `div` 頁籤中，為你的類別添加一個實用類別 `bg-red-100`，該類別表現為淺紅色背景。

# --hints--

你的開啟 `div` 頁籤應該有一個類別屬性。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

你的類別屬性應該包含工具類別 `bg-red-100`。

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
