---
id: 684cd8bf5445de089af9c837
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

現在是時候套用邊框顏色。若要使用 Tailwind CSS 設定 `border-color` 屬性，你可以使用具有以下樣式的工具類別：`border-color-number`，其中 `number` 表示邊框顏色的色階。 

以下是將紫色邊框顏色套用到元素的範例：

```html
<p class="border-purple-400">This is border-color.</p>
```

在你的開啟 `div` 頁籤中，套用紅色邊框顏色，色階為 `300`。 

# --hints--

你的開啟 `div` 頁籤應該有 `border-red-300` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
