---
id: 684ccbdc1ef89e042b744e09
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

工作坊的下一個零件是為你的 dismiss 按鈕添加功能。 

在你的 `p` 元素下方，為你的 `button` 元素添加一個 `aria-label` 屬性，設定為 `"Dismiss"`。在該 `button` 元素內，添加一個帶有 `aria-hidden` 屬性的 `span` 元素。該 `span` 元素的文字應為 `&times;`。 

*注意*：由於本工作坊將只專注於使用 Tailwind CSS 樣式，`button` 不會被添加任何互動功能。 

# --hints--

你應該有一個 `button` 元素。

```js
assert.exists(document.querySelector("button"));
```
 
你的 `button` 元素應該有一個 `aria-label` 屬性。

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

你的 `button` 元素應該設定 `aria-label` 屬性為 `"Dismiss"`。

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

你的 `button` 元素應該在內部巢狀一個 `span` 元素。

```js
assert.exists(document.querySelector("button span"));
```

你的 `span` 元素應該有一個 `aria-hidden` 屬性。

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

你的 `span` 元素應該有文字 `&times;`。

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
