---
id: 684cff0bf0a4ca0ad7e42990
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

下一步是將邊框的角落設為圓角。若要使用 Tailwind CSS 設定 `border-radius` 屬性，你可以使用符合以下樣式的工具類別：`rounded-size`，其中 `size` 表示半徑。 

以下是將超大邊框半徑套用到元素的範例：

```html
<p class="rounded-xl">This is border radius.</p>
```

在你的開啟 `div` 頁籤中，套用大小為 `md` 的邊框圓角。 

# --hints--

你的開啟 `div` 頁籤應該有 `rounded-md` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
