---
id: 684cc59eed2706fa50e8933c
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次工作坊中，你將學習如何透過建置錯誤訊息組件來開始使用 Tailwind CSS 的基礎。 

在最初的幾個步驟中，你將專注於建置初始的 HTML 標記。 

在 `body` 元素內，為你的 `div` 元素添加一個。然後在該 `div` 元素內，為你的 `p` 元素添加一個。

# --hints--

你應該有一個 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

你應該有一個 `p` 元素。

```js
assert.exists(document.querySelector("p"));
```

你的 `p` 元素應該在 `div` 元素內。

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
