---
id: 684d009994a3840b98c9aa3a
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

目前，`div` 元素的寬度橫跨整個頁面的寬度。要在 Tailwind CSS 中套用寬度，你有幾個選項，可使用 `w-number` 樣式來設定固定寬度，或使用 `w-fraction` 樣式來設定百分比寬度。  

以下是將寬度設定為其父容器 `66%` 的範例：

```html
<p class="w-2/3">This is a percentage width.</p>
```

在你的開啟 `div` 頁籤內，將寬度設定為其父容器的 `50%`。

# --hints--

你的開啟 `div` 頁籤應該有 `w-1/2` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
