---
id: 684d0d369e2b3b0f288d2c20
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

將 `div` 元素水平置中會很不錯。若要同時對兩側套用 `margin`，你可以使用以下樣式：`mx-number` 或 `my-number`。 

`mx` 設定左右的邊距，而 `my` 設定上下的邊距。

以下是將上方和下方邊距套用到元素的範例：

```html
<p class="my-8">This is margin.</p>
```

在你的開啟 `div` 頁籤中，使用正確的工具類別將左右邊距設為 `auto`。

# --hints--

你的開啟 `div` 頁籤應該有 `mx-auto` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
