---
id: 684d0eeb9c81e5109e52fe8b
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

`div` 元素開始看起來比較好。但現在關閉 `button` 有點太靠近段落元素。

要在彈性容器的項目之間創建間距，你可以使用以下樣式：`gap-number`，其中每個 `number` 單位代表在彈性佈局的列和行之間固定的 0.25rem（或 4px）間距。

以下是將 `gap` 應用於元素的範例：

```html
<p class="flex gap-8">This is gap.</p>
```

在你的開啟 `div` 標籤內，使用正確的工具類別將 `gap` 設定為 `16px`。

# --hints--

你的開啟 `div` 頁籤應該有 `gap-4` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
