---
id: 684d17a3b350bb13d3c19c0f
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

現在是時候為 `button` 元素套用一些樣式。 

首先，為你的 `button` 元素添加類別屬性。接著，套用正確的工具類別以設定紅色文字顏色，色調設為 `700`。

然後，在你的開啟 `button` 頁籤內，使用正確的工具類別將 `font-size` 設定為 `2xl`，該設定表現為 `1.5rem` 或（`24px`）。

# --hints--

你的 `button` 元素應該有一個類別屬性。

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

你的開啟 `button` 頁籤應該有 `text-red-700` 實用類別。

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

你的開啟 `button` 頁籤應該有 `text-2xl` 實用類別。

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
