---
id: 684d105e488dbd1166a717ae
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

如果你增加預覽視窗的大小，你會注意到所有的項目仍然在容器中靠左對齊。如果彈性項目能置中會更好。 

要使用 Tailwind CSS 將彈性項目沿著容器的主軸定位，你可以使用以下樣式：`justify-value`。`value` 可以是 `start`、`end`、`center`、`between` 等。任何在原生 CSS 中 `justify-content` 可用的有效值，都可以在 Tailwind CSS 中使用。 

在你的開啟 `div` 頁籤中，使用正確的工具類別來設定 `justify-content` 屬性為 `center`。

# --hints--

你的開啟 `div` 頁籤應該有 `justify-center` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
