---
id: 684d09cad4ec900e41db2e74
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

目前 `div` 元素被放置在預覽螢幕的左上角。但在元素頂部添加一些空間會更好。要使用 Tailwind CSS 為元素套用 `margin`，你可以使用以下樣式：`m-number`，其中每個 `number` 單位代表 0.25rem（或 4 像素），用於套用到四個邊。

如果你只想為一側套用邊距，可以使用以下任一樣式：`mt-number`（上邊距）、`mb-number`（下邊距）、`mr-number`（右邊距）和 `ml-number`（左邊距）。

以下是為元素套用 32px `margin-bottom` 的範例：

```html
<p class="mb-8">This is margin.</p>
```

在你的開啟 `div` 標籤內，使用正確的工具類別來套用 `16px` 的 `margin-top`。

# --hints--

你的開啟 `div` 頁籤應該有 `mt-4` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
