---
id: 684d11ffd15a531225d081e9
title: 第 17 步
challengeType: 0
dashedName: step-17
---

# --description--

既然樣式已經套用到 `div` 元素，現在是時候為 `p` 元素套用一些樣式。 

目前，段落元素的文字顏色是黑色。但如果顏色是深紅色會更好。

在先前的步驟中，你學會了如何使用像 `bg-red-100` 這樣的背景顏色。你可以透過遵循這個樣式：`text-color-shade`，對文字顏色做類似的操作。 

以下是如何將藍色文字顏色套用到元素的範例：

```html
<p class="text-blue-500">This has blue text.</p>
```

首先，為你的 `p` 元素添加類別屬性。然後，套用正確的工具類別以設定紅色文字顏色，色階設為 `700`。

# --hints--

你的 `p` 元素應該有一個類別屬性。

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

你的開啟 `p` 頁籤應該有 `text-red-700` 實用類別。

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
