---
id: 684d160a2370b01310ae0fda
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

新的紅色看起來不錯，但如果文字較大會更好。若要使用 Tailwind CSS 為元素設定 `font-size`，你可以使用以下樣式：`text-value`。 

以下是將元素的 `font-size` 設定為 `0.875rem` 或（`14px`）的範例：

```html
<p class="text-sm">This is for font-size.</p>
```

在你開啟的 `p` 元素內，使用正確的工具類別將 `font-size` 設定為 `xl`，該值表述為 `1.25rem` 或（`20px`）。

# --hints--

你的開啟 `p` 頁籤應該有 `text-xl` 實用類別。

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
