---
id: 684d0e11094c3e0fe1fb98f5
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

現在，段落元素堆疊在關閉 `button` 元素的上方。但如果 `button` 元素能與段落位於同一列會更好。這就是 Flexbox 派上用場的地方。

要將 `div` 元素設為 Tailwind CSS 中的 Flexbox 佈局，你可以使用 `flex` 工具類別。

# --hints--

你的開啟 `div` 頁籤應該有 `flex` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("flex"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
