---
id: 684d1942a1904a14a3709c6c
title: 第 20 步
challengeType: 0
dashedName: step-20
---

# --description--

你的錯誤訊息組件的樣式幾乎完成，但還有一件事需要修正。由於你為段落和按鈕文字設定了不同的字型大小，這些 flex 項目之間的垂直對齊出現偏差。 

要使用 Tailwind CSS 沿著容器的交叉軸定位彈性項目，你可以使用以下樣式：`items-value`。`value` 可以是 `start`、`end`、`center`、`baseline` 等等。 

在你開啟的 `div` 頁籤中，使用正確的工具類別來設定 `align-items` 屬性為 `center`。

隨著最後一次更改，你的錯誤訊息組件已完成！

# --hints--

你的開啟 `div` 頁籤應該有 `items-center` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
