---
id: 684d03df2a0ddf0c735865f0
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

現在你看到元素寬度是 `body` 的一半大小。這對較大型的裝置效果良好，但對較小型的裝置則不太適用。

Tailwind CSS 採用行動優先的設計方法。所以如果你想為其他裝置創建響應式設計，你可以像這樣使用 Tailwind 的響應式變體，例如 `md` 和 `lg`：

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

以下範例會將寬度設定為父元素的 `66%`，適用於裝置尺寸為 `1024px` 及以上。 

在你的開啟 `div` 頁籤中，修改你的 `w-1/2`，使其僅套用於 `768px` 及以上的裝置尺寸。你可以為這個類別添加 `md` 前置式。 

# --hints--

你的開啟 `div` 頁籤應該有 `md:w-1/2` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

你不應該再在開頭的 `div` 頁籤中使用 `w-1/2` 類別。

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
