---
id: 684d07cbfa43a20d6a002ed7
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

`div` 元素的樣式開始成形了。但如果文字周圍有一些間距會更好。 

要使用 Tailwind CSS 為元素套用 `padding`，你可以使用以下樣式：`p-number`，其中 `number` 表述你想在四個邊上套用的像素數。 

這裡是一個將 `8px` 的 `padding` 應用到元素的範例：

```html
<p class="p-8">This is padding.</p>
```

在你的開啟 `div` 頁籤中，使用正確的工具類別來套用 `4px` 的 `padding`。

# --hints--

你的開啟 `div` 頁籤應該有 `p-4` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
