---
id: 68501d32b69f9d7756334ffc
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

如果你放大預覽視窗，你應該會看到 `div` 元素在頁面上水平置中。但對於較小的裝置，`div` 元素的寬度會延伸至整個頁面寬度。  

更新 `mx-auto` 類別，使其僅套用於 `md` 裝置及以上。

# --hints--

你的開啟 `div` 頁籤應該有 `md:mx-auto` 實用類別。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

你不應該再在你的開啟 `div` 頁籤中使用 `mx-auto` 類別。

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
