---
id: 62a115879a6d51422652cbfc
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

あなたのロールプレイングゲームでは、ユーザーがステータスを追跡し、武器を購入し、モンスターと戦うことができます。ゲームのインタラクティブなJavaScript部分を進める前に、ゲーム情報を表示するHTML要素を最初に作成する必要があります。

`div` 要素の中に4つの `#game` 要素を作成してください。それぞれに、順番に `id`、`stats`、`controls`、`monsterStats` の `text` 値を付けてください。

# --hints--

4つの新しい `div` 要素を作成してください。

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

新しい `div` 要素のうち1つに、`id` の `stats` を付けてください。

```js
assert.exists(document.querySelector('div#stats'));
```

新しい `div` 要素のうち1つに、`id` の `controls` を付けてください。

```js
assert.exists(document.querySelector('div#controls'));
```

新しい `div` 要素のうち1つに、`id` の `monsterStats` を付けてください。

```js
assert.exists(document.querySelector('div#monsterStats'));
```

新しい `div` 要素のうち1つに、`id` の `text` を付けてください。

```js
assert.exists(document.querySelector('div#text'));
```

新しい `div` 要素を正しい順番に配置してください。

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

新しい `div` 要素を `#game` 要素の中に配置してください。

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
