---
id: 62a1166ed9a56d439c0770e7
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`span` 要素の中に `#stats` 要素を3つ作成してください。それぞれにクラス `stat` を付けます。次に、1つ目にはテキスト `XP: 0`、2つ目にはテキスト `Health: 100`、3つ目にはテキスト `Gold: 50` を設定してください。

# --hints--

`span` 要素の中に `#stats` 要素が3つあるはずです。

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

新しく作成した3つの `span` 要素にクラス `stat` を付けてください。

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

1つ目の `.stat` 要素には指定されたテキスト `XP: 0` を設定してください。

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

2つ目の `.stat` 要素には指定されたテキスト `Health: 100` を設定してください。

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

3つ目の `.stat` 要素には指定されたテキスト `Gold: 50` を設定してください。

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
