---
id: 62a23cb9bc467a147516b500
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

`#controls` 要素の中に、3つの `button` 要素を作成してください。1つ目は `id` が `button1` に設定されていて、テキストは `Go to store` です。2つ目は `id` が `button2` に設定されていて、テキストは `Go to cave` です。3つ目は `id` が `button3` に設定されていて、テキストは `Fight dragon` です。

# --hints--

`button` 要素に3つの `#controls` 要素を追加してください。

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

最初のボタンは `id` を `button1` に設定してください。

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

最初のボタンのテキストは `Go to store` にしてください。

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

2つ目のボタンは `id` を `button2` に設定してください。

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

2つ目のボタンのテキストは `Go to cave` にしてください。

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

3つ目のボタンは `id` を `button3` に設定してください。

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

3つ目のボタンのテキストは `Fight dragon` にしてください。

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
--fcc-editable-region--
      <div id="controls">

      </div>
--fcc-editable-region--
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
