---
id: 62a23d1c5f1c93161f3582ae
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

`#stats` 要素と同様に、`#monsterStats` 要素には `span` 要素が2つ必要です。これらにクラス `stat` を付け、最初の要素にはテキスト `Monster Name: ` を、2番目の要素にはテキスト `Health: ` を設定してください。それぞれのテキストの後に、空のネストされた `strong` 要素を持つ `span` 要素を追加します。最初の内側の `span` 要素には `id` を `monsterName` に、2番目の内側の `span` 要素には `id` を `monsterHealth` に設定してください。

# --hints--

`#monsterStats` 要素には `span` 要素が2つ必要です。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

新しい `span` 要素には両方とも `class` の値を `stat` にしてください。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

最初の `span` 要素には、空のネストされた `strong` 要素を持つ `span` 要素を含めてください。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

2番目の `span` 要素には、空のネストされた `strong` 要素を持つ `span` 要素を含めてください。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

最初の `span` 要素にはテキスト `Monster Name: ` を入れてください。スペースが正しいことを確認してください。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+/);
```

2番目の `span` 要素にはテキスト `Health: ` を入れてください。スペースが正しいことを確認してください。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[1].innerHTML, /Health:\s+/);
```

`strong` と `span` 要素はテキストの後に配置してください。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+<strong>/);
assert.match(spans[1].innerHTML, /Health:\s+<strong>/);
```

最初の内側の `span` 要素には `id` を `monsterName` に設定してください。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[0]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterName');
```

2番目の内側の `span` 要素には `id` を `monsterHealth` に設定してください。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[1]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterHealth');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
--fcc-editable-region--
      <div id="monsterStats">

      </div>
--fcc-editable-region--
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
