---
id: 62a2509ba163e020bb9d84ea
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

これでJavaScriptの記述を始められます。まず`script`要素を作成してください。この要素はJavaScriptをHTMLファイルに読み込むために使います。

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

`script`要素があるはずです。

```js
assert.isAtLeast(document.querySelectorAll('script').length, 1);
```

`script`要素には開始タグが必要です。

```js
assert.match(code, /<script\s*>/i);
```

`script`要素には終了タグが必要です。

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
