---
id: 62a39f5a5790eb27c1e5d4bf
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

表示が新しい`script.js`ファイルに切り替わりました。上のタブを使ってファイル間を切り替えられることを覚えておいてください。

このファイルに`console.log("Hello World");`行を追加し、コンソールに表示されるのを確認してください。

# --hints--

コードに`console.log("Hello World");`行があるはずです。

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
