---
id: 62a3b79d520a7f3d0e25afd6
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript は <dfn>Document Object Model</dfn>（DOM）を使って HTML とやり取りします。DOM は HTML を表すオブジェクトのツリーです。HTML 全体を表す `document` オブジェクトを使って HTML にアクセスできます。

HTML 内の特定の要素を見つける方法の一つが `querySelector()` メソッドを使うことです。`querySelector()` メソッドは CSS セレクターを引数に取り、そのセレクターに一致する最初の要素を返します。たとえば、HTML 内の `<h1>` 要素を見つけるには、次のように書きます。

```js
let h1 = document.querySelector("h1");
```

`h1` は文字列で、使う CSS セレクターに一致することに注意してください。

`button1` という変数を作成し、`querySelector()` を使って `id` が `button1` の要素を代入してください。CSS `id` セレクターには `#` の接頭辞が付くことを覚えておいてください。

# --hints--

`let` を使って `button1` 変数を宣言してください。

```js
assert.match(code, /let\s+button1/);
```

`document.querySelector()` を使ってください。

```js
assert.match(code, /document\.querySelector/);
```

`"#button1"` セレクターを使ってください。

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
