---
id: 62a3a0a3c0a4b32915d26a6e
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`console.log("Hello World");` の行を削除して、プロジェクトのコードを書き始めてください。

前のプロジェクトでは、変数と `let` キーワードの使い方を学びました。

```js
let camperbot;
```

また、変数に値を初期化する方法も学びました。

```js
let age = 32;
```

`let` キーワードを使って、`xp` という変数を宣言し、数値の `0` を代入してください。

# --hints--

コードに `console.log("Hello World");` の行があってはいけません。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

変数を宣言するには `let` キーワードを使ってください。

```js
assert.match(code, /let/);
```

`xp` という名前の変数を宣言してください。

```js
assert.match(code, /let\s+xp/);
```

`xp` には `0` の値を代入してください。

```js
assert.equal(xp, 0);
```

`xp` 変数を `0` に初期化してください。

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
