---
id: 62a3b41c9494f937560640ab
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`monsterHealth` と `inventory` という名前の変数をさらに2つ宣言してください。

`inventory` 変数には、文字列 `"stick"` を含む配列の値を代入してください。

前のプロジェクトで配列を扱ったのを思い出してください：

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

`let` を使って `monsterHealth` 変数を宣言してください。

```js
assert.match(code, /let\s+monsterHealth/i);
```

`monsterHealth` 変数の名前はキャメルケースを使ってください。

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` には値を持たせないでください。

```js
assert.isUndefined(monsterHealth);
```

`monsterHealth` 変数には値を代入しないでください。

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

`let` を使って `inventory` 変数を宣言してください。

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` 変数は引き続き配列である必要があります。

```js
assert.isArray(inventory);
```

`inventory` 変数は値を1つだけ持つべきです。

```js
assert.lengthOf(inventory, 1);
```

`inventory` 変数には文字列 `"stick"` を含めてください。

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
