---
id: 62a3a75d8466a12e009eff76
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`health` という名前の変数を `100` の値で初期化し、`gold` という名前の変数を `50` の値で初期化してください。

# --hints--

`let` を使って `health` という名前の変数を宣言してください。

```js
assert.match(code, /let\s+health/);
```

`health` という名前の変数を `100` の値で初期化してください。

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

`let` を使って `gold` という名前の変数を宣言してください。

```js
assert.match(code, /let\s+gold/);
```

`gold` という名前の変数を `50` の値で初期化してください。

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` は `100` の値を持つべきです。

```js
assert.equal(health, 100);
```

`gold` は `50` の値を持つべきです。

```js
assert.equal(gold, 50);
```

`xp` はまだ `0` の値を持つべきです。

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
