---
id: 62a3a7e4f1060e2fc5ffb34b
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`currentWeaponIndex` という名前の別の変数を作成し、`0` に設定してください。

# --hints--

`let` を使って `currentWeaponIndex` という変数を宣言してください。

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

変数名はキャメルケースを使うことをお勧めします。

```js
assert.match(code, /currentWeaponIndex/);
```

`currentWeaponIndex` 変数は `0` に設定してください。

```js
assert.equal(currentWeaponIndex, 0);
```

変数は `0` で初期化してください。

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
