---
id: 62a3bec30ea7f941412512dc
title: ステップ 29
challengeType: 0
dashedName: step-29
---

# --description--

ボタンと同じように、次の`id`の変数を作成し、`querySelector()`を使ってそれらに要素を値として割り当ててください。

`text`、`xpText`、`healthText`、`goldText`、`monsterStats`、および`monsterName`です。

これらは`const`キーワードで宣言し、変数名は`id`に合わせてください。

# --hints--

`text`変数は`const`で宣言してください。

```js
assert.match(code, /const\s+text/);
```

`text`変数には`#text`要素の値を持たせてください。

```js
assert.match(code, /const\s+text\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#text\1\s*\)\s*/);
```

`xpText`変数は`const`で宣言してください。

```js
assert.match(code, /const\s+xpText/);
```

`xpText`変数には`#xpText`要素の値を持たせてください。

```js
assert.match(code, /const\s+xpText\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#xpText\1\s*\)\s*/);
```

`healthText`変数は`const`で宣言してください。

```js
assert.match(code, /const\s+healthText/);
```

`healthText`変数には`#healthText`要素の値を持たせてください。

```js
assert.match(code, /const\s+healthText\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#healthText\1\s*\)\s*/);
```

`goldText`変数は`const`で宣言してください。

```js
assert.match(code, /const\s+goldText/);
```

`goldText`変数には`#goldText`要素の値を持たせてください。

```js
assert.match(code, /const\s+goldText\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#goldText\1\s*\)\s*/);
```

`monsterStats`変数は`const`で宣言してください。

```js
assert.match(code, /const\s+monsterStats/);
```

`monsterStats`変数には`#monsterStats`要素の値を持たせてください。

```js
assert.match(code, /const\s+monsterStats\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#monsterStats\1\s*\)\s*/);
```

`monsterName`変数は`const`で宣言してください。

```js
assert.match(code, /const\s+monsterName/);
```

`monsterName`変数には`#monsterName`要素の値を持たせてください。

```js
assert.match(code, /const\s+monsterName\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#monsterName\1\s*\)\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
