---
id: 62a3b3eab50e193608c19fc6
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`fighting` という名前の変数を宣言してください。ただし、値は初期化しないでください。

# --hints--

`let` を使って `fighting` という変数を宣言してください。

```js
assert.match(code, /let\s+fighting/);
```

`fighting` 変数には値を設定しないでください。

```js
assert.isUndefined(fighting);
```

`fighting` 変数に値を代入しないでください。

```js
assert.match(code, /let\s+fighting\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
--fcc-editable-region--
```
