---
id: 62a8c2bbbd8aa82052f47c53
title: ステップ 71
challengeType: 0
dashedName: step-71
---

# --description--

`locations` 配列に3つ目のオブジェクトを追加してください。ほかの2つのオブジェクトと同じプロパティを持たせてください。

`name` を `cave` に設定してください。`button text` を `"Fight slime"`、`"Fight fanged beast"`、`"Go to town square"` の文字列を含む配列に設定してください。`"button functions"` を `fightSlime`、`fightBeast`、`goTown` の変数を含む配列に設定してください。`text` プロパティを `"You enter the cave. You see some monsters."` に設定してください。

# --hints--

`locations` 配列には3つの値があるはずです。

```js
assert.lengthOf(locations, 3);
```

3つ目の `locations` の値はオブジェクトであるべきです。

```js
assert.isObject(locations[2]);
```

3つ目の `locations` オブジェクトは `name` プロパティを持ち、その値は `"cave"` であるべきです。

```js
assert.equal(locations[2].name, "cave");
```

3つ目の `locations` オブジェクトは配列である `"button text"` プロパティを持つべきです。

```js
assert.isArray(locations[2]["button text"]);
```

`"button text"` プロパティは `"Fight slime"`、`"Fight fanged beast"`、`"Go to town square"` の文字列の値を持つべきです。

```js
assert.equal(locations[2]["button text"][0], "Fight slime");
assert.equal(locations[2]["button text"][1], "Fight fanged beast");
assert.equal(locations[2]["button text"][2], "Go to town square");
```

3つ目の `locations` オブジェクトは配列である `"button functions"` プロパティを持つべきです。

```js
assert.isArray(locations[2]["button functions"]);
```

`"button functions"` プロパティは `fightSlime`、`fightBeast`、`goTown` の関数の値を持つべきです。

```js
assert.equal(locations[2]["button functions"][0], fightSlime);
assert.equal(locations[2]["button functions"][1], fightBeast);
assert.equal(locations[2]["button functions"][2], goTown);
```

3つ目の `locations` オブジェクトは文字列である `text` プロパティを持つべきです。

```js
assert.isString(locations[2].text);
```

3つ目の `locations` オブジェクトは `text` プロパティを持ち、その値は `"You enter the cave. You see some monsters."` であるべきです。

```js
assert.equal(locations[2].text, "You enter the cave. You see some monsters.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
--fcc-editable-region--

--fcc-editable-region--
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}

function buyHealth() {

}

function buyWeapon() {

}

function fightSlime() {

}

function fightBeast() {

}
```
