---
id: 62b46e3a8d4be31be5af793d
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

少し問題が発生しました。ページ内のボタン要素をクエリしようとしていますが、`script` タグが HTML の `head` 内にあります。これは、コードがブラウザーが HTML の読み込みを終える前に実行されていることを意味し、`document.querySelector()` はボタンを認識できません。なぜならブラウザーがまだ処理していないからです。

これを修正するには、`script` 要素を `head` 要素の外に移動し、`body` 要素の最後（閉じる `</body>` タグの直前）に配置してください。

# --hints--

`script` 要素は `head` 要素の中にあってはいけません。

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

`script` 要素は `body` 要素の最後に来るべきです。

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
