---
id: 6577ea4f23e3ba1217007bc7
title: ステップ 28
challengeType: 0
dashedName: step-28
---

# --description--

最後に、ボタンのスタイルをいくつか追加する必要があります。まず、`cursor` プロパティを `pointer` に設定してください。次に、テキストの `color` を `#0a0a23` に、`background-color` を `#feac32` に設定します。

それから、`background-image` プロパティを `linear-gradient(#fecc4c, #ffac33)` に設定してください。最後に、`border` を `3px solid #feac32` に設定します。

# --hints--

`button` セレクターがあるはずです。

```js
const btn = new __helpers.CSSHelp(document).getStyle('button');
assert.exists(btn);
```

`button` セレクターには `cursor` プロパティが `pointer` に設定されているはずです。

```js
const cursor = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('cursor');
assert.equal(cursor, 'pointer');
```

`button` セレクターには `color` プロパティが `#0a0a23` に設定されているはずです。

```js
const color = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('color');
assert.equal(color, 'rgb(10, 10, 35)');
```

`button` セレクターには `background-color` プロパティが `#feac32` に設定されているはずです。

```js
const backgroundColor = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('background-color');
assert.equal(backgroundColor, 'rgb(254, 172, 50)');
```

`button` セレクターには `background-image` プロパティが `linear-gradient(#fecc4c, #ffac33)` に設定されているはずです。

```js
const backgroundImage = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('background-image');
assert.equal(backgroundImage, 'linear-gradient(rgb(254, 204, 76), rgb(255, 172, 51))');
```

`button` セレクターには `border` プロパティが `3px solid #feac32` に設定されているはずです。

```js
const border = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('border');
assert.equal(border, '3px solid rgb(254, 172, 50)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
