---
id: bad87fee1348bd9aedf08827
title: 创建一个无序列表
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cDKVPuv'
forumTopicId: 16814
dashedName: create-a-bulleted-unordered-list
---

# --description--

HTML 有一个特定的元素用于创建<dfn>无序列表</dfn>。

无序列表以一个开头的 `<ul>` 元素开始，后面跟任意数量的 `<li>` 元素。最后，无序列表以 `</ul>` 关闭。

例如：

```html
<ul>
  <li>milk</li>
  <li>cheese</li>
</ul>
```

会创建一个要点列表，包括 `milk` 和 `cheese`。

# --instructions--

请删除页面底部的两个 `p` 元素，然后在底部创建一个无序列表，其中包含你认为猫咪最喜欢的三件东西。

# --hints--

应存在一个 `ul` 无序列表。

```js
assert.isNotEmpty(document.querySelectorAll('ul'));
```

`ul` 无序列表中应包含 3 个 `li` 元素。

```js
assert.lengthOf(document.querySelectorAll('ul li'),3);
```

你的 `ul` 元素应该有一个结束标签。

```js
assert.match(code,/<\/ul>/gi);
assert.match(code,/<ul/gi);
assert.strictEqual(code.match(/<\/ul>/gi).length,code.match(/<ul/gi).length);
```

你的 `li` 元素应该有一个结束标签。

```js
assert.match(code,/<\/li>/gi);
assert.match(code,/<li[\s>]/gi);
assert.strictEqual(code.match(/<\/li>/gi).length,code.match(/<li[\s>]/gi).length);
```

每个 `li` 元素不应只包含空字符串或只包含空格。

```js
assert.isEmpty([...document.querySelectorAll('ul li')].filter((item) => item.textContent.trim() === ""));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <ul>
    <li>milk</li>
    <li>mice</li>
    <li>catnip</li>
  </ul>
</main>
```
