---
id: bad87fee1348bd9aedf08801
title: 用 p 元素代表段落
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ceZ7DtN'
forumTopicId: 18202
dashedName: inform-with-the-paragraph-element
---

# --description--

`p` 元素是网站段落文本的首选元素。`p` 是 "paragraph" 的缩写。

你可以像这样创建一个段落：

```html
<p>I'm a p tag!</p>
```

# --instructions--

在 `h2` 元素下方添加一个 `p` 元素，其内容是 `Hello Paragraph`。

**注意：** 作为惯例，所有超文本标记语言标签均以小写字母书写，例如 `<p></p>` 而不是 `<P></P>`。

# --hints--

应包含一个 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll('p'),1);
```

`p` 元素的内容文本应为 `Hello Paragraph`。

```js
assert.match(document.querySelector('p').textContent,/hello(\s)+paragraph/gi);
```

你的 `p` 元素应该有一个结束标签。

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
```

# --solutions--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
<p>Hello Paragraph</p>
```
