---
id: 673b3d6b7ef7318eef926d5a
title: بناء لوحة مزاجية
challengeType: 25
dashedName: build-a-mood-board
demoType: onClick
---

# --description--

اللوحة المزاجية هي مجموعة من الصور والنصوص التي تنقل فكرة عامة أو هدفًا أو شعورًا حول موضوع معين.

في هذا المختبر، ستنشئ لوحة مزاجية باستخدام مكوِّنات React قابلة لإعادة الاستخدام. تم توفير CSS لك بالفعل.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب عليك إنشاء وتصدير مكوِّن `MoodBoardItem` يقبل ثلاث خصائص: `color`، و`image`، و`description`.
2. يجب أن يُرجع مكوِّن `MoodBoardItem` عنصر `div` مع الصنف `mood-board-item` كعنصر المستوى الأعلى.
3. يجب تعيين لون الخلفية لعنصر `.mood-board-item` إلى قيمة الخاصية `color` باستخدام أنماط مضمَّنة.
4. يجب عرض عنصر `img`، مع صنف `mood-board-image` و`src` الخاص به مضبوط على قيمة الخاصية `image`، داخل عنصر `.mood-board-item`.
5. يجب عرض عنصر `h3`، مع صنف `mood-board-text` ونصه هو قيمة الخاصية `description`، داخل عنصر `.mood-board-item`.
6. يجب إنشاء وتصدير مكوِّن `MoodBoard`.
7. يجب أن يُرجع مكوِّن `MoodBoard` عنصر `div` كعنصر المستوى الأعلى.
8. يجب أن يعرض مكوِّن `MoodBoard` عنصر `h1` مع صنف `mood-board-heading` والنص `Destination Mood Board`.
9. يجب أن يعرض مكوِّن `MoodBoard` عنصر `div` مع صنف `mood-board`.
10. يجب أن يعرض مكوِّن `MoodBoard` ثلاثة مكوِّنات `MoodBoardItem` على الأقل داخل عنصر `.mood-board`، ويجب أن يمرر كل منها الخصائص `color` و`image` و`description` بقيم صحيحة.

يمكنك استخدام الصور التالية في لوحتك المزاجية إذا أردت:

- `https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/shore.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/grass.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/ship.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg`

# --hints--

يجب تصدير مكوِّن `MoodBoardItem`.

```js
assert.isFunction(window.index.MoodBoardItem, 1);
```

يجب أن يُرجع مكوِّن `MoodBoardItem` عنصر `div` مع صنف `mood-board-item` كعنصر المستوى الأعلى.

```js
const item = document.getElementsByClassName('mood-board-item')[0];
assert.exists(item);
assert.equal(item.tagName, 'DIV');
```

يجب تعيين لون الخلفية لعنصر `.mood-board-item` إلى قيمة الخاصية `color` باستخدام أنماط مضمَّنة.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { color: "red" });
  const moodBoardItem = container.querySelector(".mood-board-item");
  assert.equal(moodBoardItem.style.backgroundColor, "red");
```

يجب أن يعرض مكوِّن `MoodBoardItem` عنصر `img` مع صنف `mood-board-image` و`src` مضبوط على قيمة الخاصية `image`.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { image: "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg" });
  const img = container.querySelector(".mood-board-image");
  assert.exists(img);
  assert.equal(img.tagName, "IMG");
  assert.equal(img.src, "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg");
```

يجب أن يعرض مكوِّن `MoodBoardItem` عنصر `h3` مع صنف `mood-board-text` ونصه مضبوط على قيمة الخاصية `description`.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { description: "Caribbean" });
  const text = container.querySelector(".mood-board-text");
  assert.exists(text);
  assert.equal(text.tagName, "H3");
  assert.equal(text.textContent, "Caribbean");
```

يجب تصدير مكوِّن `MoodBoard`.

```js
assert.isFunction(window.index.MoodBoard, 1);
```

يجب أن يُرجع مكوِّن `MoodBoard` عنصر `div` كعنصر المستوى الأعلى.

```js
const root = document.getElementById('root');
assert.equal(root.children[0].tagName, 'DIV');
```

يجب أن يعرض مكوِّن `MoodBoard` عنصر `h1` مع صنف `mood-board-heading` والنص `Destination Mood Board`.

```js
const heading = document.getElementsByClassName('mood-board-heading')[0];
assert.exists(heading);
assert.equal(heading.tagName, 'H1');
assert.equal(heading.textContent, 'Destination Mood Board');
```

يجب أن يعرض مكوِّن `MoodBoard` ثلاثة مكوِّنات `MoodBoardItem` على الأقل، ويمرر كل منها الخصائص `color` و`image` و`description` بقيم صحيحة.

```js
const moodboard = document.getElementsByClassName('mood-board')[0];
const items = moodboard.querySelectorAll(".mood-board-item");
assert.isAtLeast(items.length, 3);     

items.forEach((item) => {
  const color = item.style.backgroundColor;
  const img = item.querySelector("img");
  const text = item.querySelector("h3");
  
  assert.isAtLeast(color.length, 1);
  assert.isAtLeast(img.src.length, 1);
  assert.isAtLeast(text.textContent.length, 1);
});
```

يجب عرض مكوِّن `MoodBoard` في عنصر الصفحة `#root`.

```js
const root = document.getElementById('root');
const moodBoard = root.getElementsByClassName('mood-board')[0];
assert.exists(moodBoard);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx

```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx
export function MoodBoardItem(props) {
  return (
    <div className='mood-board-item' style={{ backgroundColor: props.color }}>
      <img src={props.image} alt='Mood' className='mood-board-image' />
      <h3 className='mood-board-text'>{props.description}</h3>
    </div>
  );
}

export function MoodBoard() {
  return (
    <div>
      <h1 className='mood-board-heading'>Destination Mood Board</h1>
      <div className='mood-board'>
        <MoodBoardItem
          color='#2da64f'
          image='https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg'
          description='Caribbean'
        />
        <MoodBoardItem
          color='#8e44ad'
          image='https://cdn.freecodecamp.org/curriculum/labs/shore.jpg'
          description='Gawadar Beach'
        />
        <MoodBoardItem
          color='#3498db'
          image='https://cdn.freecodecamp.org/curriculum/labs/grass.jpg'
          description='Cape Town'
        />
        <MoodBoardItem
          color='#bf3d7e'
          image='https://cdn.freecodecamp.org/curriculum/labs/ship.jpg'
          description='Suez Canal'
        />
        <MoodBoardItem
          color='#e74c3c'
          image='https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg'
          description='Santorini'
        />
        <MoodBoardItem
          color='#95a5a6'
          image='https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg'
          description='Istanbul'
        />
      </div>
    </div>
  );
}
```
