---
id: 66e45c8140f9fda5c105ae26
title: Erstellen Sie eine Veranstaltungsflyer-Seite
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `header`-Element innerhalb des Bodys haben.
1. Ihr `header`-Element sollte ein Bild für Ihre Veranstaltung enthalten und ein `h1` mit dem Titel Ihrer Veranstaltung, in dieser Reihenfolge. Sie können dieses Bild verwenden, wenn Sie möchten: `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. Sie sollten ein `main`-Element innerhalb des Bodys haben.
1. Ihr `main`-Element sollte mindestens zwei `section`-Elemente enthalten, die die Funktionen der Veranstaltung zeigen.
1. Ihre `section`-Elemente sollten jeweils ein `h2` enthalten.
1. Ihr Body sollte einen Innenabstand (padding) oben und unten von `50px` haben.
1. Ihr Body sollte einen Außenabstand (margin) oben und unten von `0` haben und einen links und rechts, der sich zentriert.
1. Ihr Body sollte eine Breite relativ zur Breite des Viewports haben.
1. Ihr Body sollte die `calc`-Funktion verwenden, um seine `min-height`-Eigenschaft auf 100 % der Höhe des Viewports minus aller oben und unten angewendeten Innenabstände zu setzen.
1. Sie sollten mindestens ein `hr` in Ihrem Flyer haben.
1. Sie sollten die `width` aller `hr`- und `section`-Elemente auf einen Prozentwert relativ zu ihrem übergeordneten Element setzen.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verlinken und Ihr CSS anwenden.

# --hints--

Sie sollten ein `header`-Element innerhalb des `body`-Elements haben.

```js
assert.exists(document.querySelector('body > header'));
```

Sie sollten ein `img`-Element als erstes Kind im `header` haben.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

Sie sollten ein `h1`-Element als zweites Kind im `header` haben.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

Sie sollten nur ein `h1`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Ihr `h1` sollte den Titel Ihrer Veranstaltung enthalten.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

Sie sollten ein `main`-Element innerhalb des `body` haben.

```js
assert.exists(document.querySelector('body > main'));
```

Sie sollten mindestens zwei `section`-Elemente innerhalb Ihres `main`-Elements haben.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

Ihre `section`-Elemente sollten jeweils ein `h2` enthalten.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

Ihre `h2`-Elemente sollten nicht leer sein.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

Ihr `body`-Element sollte einen Innenabstand (padding) oben und unten von `50px` haben.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

Ihr `body`-Element sollte einen Außenabstand (margin) oben und unten von `0` haben.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

Ihr `body`-Element sollte einen Außenabstand (margin) links und rechts von `auto` haben.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

Ihr `body`-Element sollte seine Breite mit `vw` setzen.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

Ihr `body` sollte `calc` verwenden, um seine `min-height` auf `100vh - 100px` zu setzen.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

Sie sollten mindestens ein `hr`-Element haben.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

Die Breite Ihrer `hr`-Elemente sollte mit einem Prozentwert gesetzt werden.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

Die Breite Ihrer `section`-Elemente sollte mit einem Prozentwert gesetzt werden.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
