---
id: 66faac4139dbbd5dd632c860
title: 구직 애플리케이션 양식 만들기
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 도전 과제를 완료하세요.

**사용자 스토리:**

1. `div` 요소에 `container` 클래스를 지정해야 합니다.
1. `div` 요소 안에 `form` 요소가 있어야 합니다.
1. 양식에는 사용자의 전체 이름을 입력할 수 있도록 타입이 `input`이고 id가 `text`인 `name` 요소가 포함되어야 합니다.
1. 사용자의 이메일 주소를 입력할 수 있도록 타입이 `input`이고 id가 `email`인 또 다른 `email` 요소가 있어야 합니다.
1. 사용자가 직무를 선택할 수 있도록 id가 `select`인 `position` 요소가 양식에 포함되어야 합니다.
1. `fieldset` 요소에 `radio-group` 클래스가 있어야 합니다.
1. `.radio-group` 안에는 타입이 `input`이고 관련 레이블이 있는 `radio` 요소 세트가 있어야 하며, (예: 풀타임, 파트타임) 그룹 `name`은 `availability`이어야 합니다.
1. 메시지를 입력할 수 있도록 id가 `textarea`인 `message` 요소가 있어야 합니다.
1. 모든 `input` 요소는 `label` 요소와 연결되어야 합니다.
1. 양식을 제출할 수 있도록 타입이 `button`인 `submit` 요소가 있어야 합니다.
1. `:focus`과 `input` 요소에 `textarea` 가상 클래스를 추가하여 포커스 시 테두리 색상이 변경되도록 하세요.
1. `input`, `select`, `textarea` 요소에는 잘못된 입력이 감지되면 테두리 색상이 빨간색으로 바뀌는 `:invalid` 가상 클래스가 있어야 합니다.
1. `input`, `select`, `textarea` 요소에는 올바른 입력이 들어오면 테두리 색상이 초록색으로 바뀌는 `:valid` 가상 클래스가 있어야 합니다.
1. `button` 요소에는 마우스를 올렸을 때 배경색이 바뀌는 `:hover` 가상 클래스가 있어야 합니다.
1. 라디오 버튼이 선택되었을 때 테두리 색상, 배경색, 박스 그림자가 추가되도록 `:checked`에 `.radio-group input[type="radio"]` 가상 클래스를 사용하세요.
1. 라디오 버튼이 선택되었을 때 관련 `:checked`의 텍스트 색상이 바뀌도록 라디오 버튼에 `label` 가상 클래스를 사용하세요.
1. 첫 번째 입력 필드를 다르게 스타일링하려면 `:first-of-type` 요소에 `input` 가상 클래스를 추가하세요. (예: 둥근 모서리)

**참고:** HTML에 스타일시트를 연결하고 CSS를 적용하는 것을 잊지 마세요.

# --hints--

`div` 요소에 `container` 클래스가 있어야 합니다.

```js
assert.exists(document.querySelector("div.container"));
```

`div` 요소 안에 `form` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("div.container > form"));
```

양식에는 사용자의 전체 이름을 입력할 수 있도록 타입이 `input`이고 id가 `text`인 `name` 요소가 포함되어야 합니다.

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

사용자의 이메일 주소를 입력할 수 있도록 타입이 `input`이고 id가 `email`인 또 다른 `email` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

양식에는 일부 `select` 요소가 포함된 id가 `position`인 `option` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

`fieldset` 또는 `section` 요소에 `.radio-group` 클래스가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

`.radio-group` 안에는 타입이 `input`인 `radio` 요소 그룹이 있어야 하며, 선택 그룹 `name`은 `availability`이어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

메시지를 입력할 수 있도록 id가 `textarea`인 `message` 요소가 있어야 합니다.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

모든 `input` 요소는 `label` 요소와 연결되어야 합니다.

```js
const els = document.querySelectorAll('input');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

양식을 제출할 수 있도록 타입이 `button`인 `submit` 요소가 있어야 합니다.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

포커스 시 테두리 색상이 변경되도록 `:focus`와 `input` 요소에 `textarea` 가상 클래스를 추가하세요. 주어진 순서대로 리스트 선택자를 사용하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'))
```

잘못된 입력이 감지되면 테두리 색상이 빨간색으로 바뀌도록 `input`, `select`, `textarea` 요소에 `:invalid` 가상 클래스가 있어야 합니다. 주어진 순서대로 리스트 선택자를 사용하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

올바른 입력이 들어오면 테두리 색상이 초록색으로 바뀌도록 `input`, `select`, `textarea` 요소에 `:valid` 가상 클래스가 있어야 합니다. 주어진 순서대로 리스트 선택자를 사용하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

마우스를 올렸을 때 배경색이 바뀌도록 `button` 요소에 `:hover` 가상 클래스가 있어야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

라디오 버튼이 선택되었을 때 테두리 색상을 추가하려면 `:checked`에 `.radio-group input[type="radio"]` 가상 클래스를 사용하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

라디오 버튼이 선택되었을 때 배경색을 추가하려면 `:checked`에 `.radio-group input[type="radio"]` 가상 클래스를 사용하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

라디오 버튼이 선택되었을 때 박스 그림자를 추가하려면 `:checked`에 `.radio-group input[type="radio"]` 가상 클래스를 사용하세요.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

라디오 버튼이 선택되었을 때 관련 `:checked`의 텍스트 색상을 변경하려면 라디오 버튼에 `label` 가상 클래스를 사용하세요.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

첫 번째 입력 필드를 다르게 스타일링하려면 `:first-of-type` 요소에 `input` 가상 클래스를 추가하세요.

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
