---
id: 695774002591bbc5f8cf3e53
title: 构建一个折扣计算器
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

在本实验中，你将编写一个函数，用于计算应用百分比折扣后某项的最终价格。

例如，如果一项的价格是 `50` 并且应用了 `20` 的折扣，折扣金额是 `10`，最终价格是 `40`。

**目标：** 完成以下用户故事以完成实验。

**用户故事：**

1. 你应该定义一个名为 `apply_discount` 的函数。

2. `apply_discount` 函数应当正好接受两个参数：`price` 和 `discount`。

3. 如果 `price` 不是数字（`int` 或 `float`），函数应返回字串 `The price should be a number`。

4. 如果 `discount` 不是数字（`int` 或 `float`），该函数应返回字串 `The discount should be a number`。

5. 如果 `price` 小于或相等于 `0`，函数应返回字串 `The price should be greater than 0`。

6. 如果 `discount` 小于 `0` 或大于 `100`，函数应返回字串 `The discount should be between 0 and 100`。

7. 如果两个 `input` 都有效，`函数` 应该计算折扣作为价格的百分比。

8. 该函数应返回应用折扣后的最终价格。

# --hints--

你应该有一个名为 `apply_discount` 的函数。

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

你的 `apply_discount` 函数应接受两个参数：`price` 和 `discount`。

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

当调用 `apply_discount` 并传入的 `price`（第一个参数）不是数字（`int` 或 `float`）时，它应返回 `The price should be a number`。

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

当调用 `apply_discount` 并传入一个不是数字（`int` 或 `float`）的 `discount`（第二个参数）时，它应返回 `The discount should be a number`。

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

当 `apply_discount` 被调用且 `price` 小于或相等于 `0` 时，它应返回 `The price should be greater than 0`。

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

当 `apply_discount` 被调用且 `discount` 小于 `0` 或大于 `100` 时，它应返回 `The discount should be between 0 and 100`。

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` 应该返回 `80`。

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` 应该返回 `100`。

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` 应该返回 `50`。

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

当调用 `apply_discount` 并传入折扣为 `100` 时，它应返回 `0`。

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` 应该返回 `59.6`。

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
