---
id: 684ccbdc1ef89e042b744e09
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Der nächste Teil des Workshops besteht darin, die Schaltfläche zum Schließen hinzuzufügen.

Fügen Sie unter Ihrem `p`-Element ein `button`-Element mit einem `aria-label`-Attribut hinzu, das auf `"Dismiss"` gesetzt ist. Innerhalb dieses `button`-Elements fügen Sie ein `span`-Element mit einem `aria-hidden`-Attribut hinzu. Der Text des `span`-Elements sollte `&times;` sein.

*Hinweis*: Da sich dieser Workshop ausschließlich auf die Arbeit mit Tailwind-CSS-Stilen konzentriert, wird die `button` keine Art von Interaktivität erhalten.

# --hints--

Sie sollten ein `button`-Element haben.

```js
assert.exists(document.querySelector("button"));
```
 
Ihr `button`-Element sollte ein `aria-label`-Attribut besitzen.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

Ihr `button`-Element sollte ein `aria-label`-Attribut haben, das auf `"Dismiss"` gesetzt ist.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

Ihr `button`-Element sollte ein verschachteltes `span`-Element enthalten.

```js
assert.exists(document.querySelector("button span"));
```

Ihr `span`-Element sollte ein `aria-hidden`-Attribut besitzen.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

Ihr `span`-Element sollte den Text `&times;` enthalten.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </div>
  </body>
</html>
```
