---
id: 684cca09aac0fc035db4099f
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie innerhalb Ihres `p`-Elements den Text `Error! Something went wrong. Please try again.` hinzu, wobei `Error!` in ein `strong`-Element eingeschlossen sein soll.

# --hints--

Sie sollten ein `strong`-Element innerhalb Ihres `p`-Elements haben.

```js
assert.exists(document.querySelector("p strong"));
```

Ihr `strong`-Element sollte den Text `Error!` enthalten.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

Ihr `p`-Element sollte den korrekten Text haben, der mit `Something went wrong. Please try again.` endet.

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </p>
    </div>
  </body>
</html>
```
