---
id: 684cd8bf5445de089af9c837
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Jetzt ist es Zeit, eine Rahmenfarbe anzuwenden. Um die `border-color` Eigenschaft mit Tailwind CSS zu setzen, können Sie eine Utility-Klasse mit folgendem Muster verwenden: `border-color-number`, wobei `number` den Farbton für die Rahmenfarbe darstellt.

Hier ist ein Beispiel, wie Sie einem Element eine lila Rahmenfarbe zuweisen:

```html
<p class="border-purple-400">This is border-color.</p>
```

Fügen Sie in Ihrem öffnenden `div` Tag eine rote Rahmenfarbe mit dem Farbton `300` hinzu.

# --hints--

Ihr öffnender `div` Tag sollte die `border-red-300` Utility-Klasse enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
