---
id: 684cc59eed2706fa50e8933c
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop lernen Sie, wie Sie mit den Grundlagen von Tailwind CSS beginnen, indem Sie eine Fehlernachrichten-Komponente erstellen.

In den ersten Schritten liegt Ihr Fokus darauf, die anfängliche HTML-Auszeichnung zu erstellen.

Fügen Sie innerhalb des `body`-Elements ein `div`-Element hinzu. Dann fügen Sie innerhalb dieses `div`-Elements ein `p`-Element hinzu.

# --hints--

Sie sollten ein `div`-Element haben.

```js
assert.exists(document.querySelector("div"));
```

Sie sollten ein `p`-Element haben.

```js
assert.exists(document.querySelector("p"));
```

Ihr `p`-Element sollte sich innerhalb des `div`-Elements befinden.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
