---
id: 684cff0bf0a4ca0ad7e42990
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Der nächste Schritt wäre, die Ecken des Rahmens abzurunden. Um die `border-radius`-Eigenschaft mit Tailwind CSS zu setzen, können Sie eine Utility-Klasse mit folgendem Muster verwenden: `rounded-size`, wobei `size` den Radius darstellt.

Hier ist ein Beispiel, wie Sie einem Element einen extra großen Rahmenradius zuweisen:

```html
<p class="rounded-xl">This is border radius.</p>
```

Fügen Sie in Ihrem öffnenden `div`-Tag einen Rahmenradius mit der Größe `md` hinzu.

# --hints--

Ihr öffnender `div`-Tag sollte die Utility-Klasse `rounded-md` enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
