---
id: 684cd027744c2a0520ba4d58
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS verwendet Utility-Klassen, die direkt auf die Auszeichnung angewendet werden. Diese Klassen umfassen Styling für Schriftarten, Farben, Flexbox-Layouts, Grid und mehr.

Um in Tailwind CSS eine Hintergrundfarbe anzuwenden, können Sie das `bg`-Präfix verwenden, gefolgt vom Farbnamen und einer Zahl. Die Zahl steht für den Farbton.

Hier ist ein Beispiel, wie Sie einem Absatz-Element eine dunkelblaue Hintergrundfarbe zuweisen.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

Fügen Sie in Ihrem öffnenden `div`-Tag eine Klasse mit der Utility-Klasse `bg-red-100` hinzu, die eine hellrote Hintergrundfarbe darstellt.

# --hints--

Ihr öffnender `div`-Tag sollte ein Klassenattribut haben.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

Ihr Klassenattribut sollte die Utility-Klasse `bg-red-100` enthalten.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
