---
id: 684cd59ec94b5e063bf776f0
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Es wäre schön, einen Rahmen um diesen hellroten Hintergrund zu haben.

Um den `border-width` in Tailwind CSS zu setzen, können Sie die `border-number` Utility-Klasse verwenden, wobei `number` die Helligkeit oder Dicke Ihres Rahmens darstellt.

Hier ist ein Beispiel, wie Sie einen `border-width` von `4px` auf allen Seiten anwenden:

```html
<p class="border-4">This is how to set the border-width.</p>
```

Fügen Sie in Ihrem öffnenden `div`-Tag die korrekte Utility-Klasse für einen `border-width` hinzu, der auf allen Seiten auf `2px` gesetzt ist.

# --hints--

Ihr öffnender `div`-Tag sollte eine `border-2`-Klasse haben.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
