---
id: 684d0d369e2b3b0f288d2c20
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Es wäre schön, das `div`-Element horizontal zu zentrieren. Um `margin` gleichzeitig auf zwei Seiten anzuwenden, können Sie die folgenden Muster verwenden: `mx-number` oder `my-number`.

`mx` legt den Außenabstand für links und rechts fest, während `my` den Außenabstand für oben und unten bestimmt.

Hier ist ein Beispiel, wie Sie den Außenabstand oben und unten auf ein Element anwenden:

```html
<p class="my-8">This is margin.</p>
```

Verwenden Sie in Ihrem öffnenden `div`-Tag die richtige Utility-Klasse, um den Außenabstand links und rechts auf `auto` zu setzen.

# --hints--

Ihr öffnender `div`-Tag sollte die Utility-Klasse `mx-auto` enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
