---
id: 684d17a3b350bb13d3c19c0f
title: Schritt 19
challengeType: 0
dashedName: step-19
---

# --description--

Jetzt ist es Zeit, einige Styles auf das `button`-Element anzuwenden.

Beginnen Sie damit, Ihrem `button`-Element ein Klassenattribut hinzuzufügen. Wenden Sie anschließend die richtige Utility-Klasse an, um eine rote Textfarbe mit dem Farbton `700` einzustellen.

Verwenden Sie dann innerhalb Ihres öffnenden `button`-Tags die korrekte Utility-Klasse, um das `font-size` auf `2xl` zu setzen, was `1.5rem` oder (`24px`) darstellt.

# --hints--

Ihr `button`-Element sollte ein Klassenattribut haben.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

Ihr öffnendes `button`-Tag sollte die Utility-Klasse `text-red-700` haben.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

Ihr öffnendes `button`-Tag sollte die Utility-Klasse `text-2xl` haben.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
