---
id: 684d03df2a0ddf0c735865f0
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Jetzt sehen Sie, dass die Elementbreite halb so groß ist wie die von `body`. Das funktioniert gut für größere Geräte, aber nicht so gut für kleinere Geräte.

Tailwind CSS folgt einem Mobile-First-Ansatz im Design. Wenn Sie also responsive Designs für andere Geräte erstellen möchten, können Sie Tailwinds responsive Varianten wie `md` und `lg` folgendermaßen verwenden:

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

Das folgende Beispiel setzt die Breite auf `66%` des Elternelements für Gerätegrößen von `1024px` und größer.

Ändern Sie in Ihrem öffnenden `div`-Tag Ihre `w-1/2` so, dass sie nur für Gerätegrößen von `768px` und größer angewendet wird. Sie können diese Klasse mit der `md`-Variante als Präfix versehen.

# --hints--

Ihr öffnender `div`-Tag sollte die Utility-Klasse `md:w-1/2` enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

Die Klasse `w-1/2` sollte nicht mehr in Ihrem öffnenden `div`-Tag vorhanden sein.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
