---
id: 684d09cad4ec900e41db2e74
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Im Moment befindet sich das `div`-Element in der oberen linken Ecke des Vorschaubildschirms. Es wäre jedoch schön, oben am Element etwas Platz hinzuzufügen. Um mit Tailwind CSS `margin` auf ein Element anzuwenden, können Sie folgendes Muster verwenden: `m-number`, wobei jede `number`-Einheit 0,25rem (oder 4 Pixel) darstellt, die auf alle vier Seiten angewendet werden.

Wenn Sie den Außenabstand nur auf einer Seite anwenden möchten, können Sie eines der folgenden Muster verwenden: `mt-number` (Außenabstand oben), `mb-number` (Außenabstand unten), `mr-number` (Außenabstand rechts) und `ml-number` (Außenabstand links).

Hier ist ein Beispiel, wie Sie `margin-bottom` von 32px auf ein Element anwenden:

```html
<p class="mb-8">This is margin.</p>
```

Fügen Sie in Ihrem öffnenden `div`-Tag die richtige Utility-Klasse ein, um `margin-top` von `16px` anzuwenden.

# --hints--

Ihr öffnendes `div`-Tag sollte die Utility-Klasse `mt-4` enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
