---
id: 684d1942a1904a14a3709c6c
title: Schritt 20
challengeType: 0
dashedName: step-20
---

# --description--

Das Styling für Ihre Fehlernachrichten-Komponente ist fast fertig, aber es gibt noch eine Sache zu beheben. Da Sie unterschiedliche Schriftgrößen für den Absatz und den Button-Text festgelegt haben, stimmt die vertikale Ausrichtung zwischen diesen Flex-Elementen nicht.

Um Flex-Elemente entlang der Querachse eines Containers mit Tailwind CSS zu positionieren, können Sie folgendes Muster verwenden: `items-value`. `value` kann alles Mögliche sein, zum Beispiel `start`, `end`, `center`, `baseline` und mehr.

Fügen Sie in Ihrem öffnenden `div`-Tag die korrekte Utility-Klasse hinzu, um die `align-items`-Eigenschaft auf `center` zu setzen.

Und mit dieser letzten Änderung ist Ihre Fehlernachrichten-Komponente fertig!

# --hints--

Ihr öffnender `div`-Tag sollte die Utility-Klasse `items-center` enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
