---
id: 684d0e11094c3e0fe1fb98f5
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Im Moment liegt das Absatz-Element übereinander mit dem dismiss `button` Element. Es wäre jedoch schön, wenn das `button` Element in derselben Zeile wie der Absatz wäre. Hier kommt Flexbox ins Spiel.

Um das `div` Element in Tailwind CSS auf ein Flexbox-Layout einzustellen, können Sie die `flex` Utility-Klasse verwenden.

# --hints--

Ihr öffnendes `div` Tag sollte die `flex` Utility-Klasse haben.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("flex"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
