---
id: 684d160a2370b01310ae0fda
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Die neue rote Farbe sieht schön aus, aber es wäre besser, wenn der Text größer wäre. Um die `font-size` für ein Element mit Tailwind CSS festzulegen, können Sie folgendes Muster verwenden: `text-value`. 

Hier ist ein Beispiel, wie Sie die `font-size` für ein Element auf `0.875rem` oder (`14px`) setzen.

```html
<p class="text-sm">This is for font-size.</p>
```

Innerhalb Ihres öffnenden `p`-Elements verwenden Sie die richtige Utility-Klasse, um die `font-size` auf `xl` zu setzen, was `1.25rem` oder (`20px`) darstellt.

# --hints--

Ihr öffnender `p`-Tag sollte die Utility-Klasse `text-xl` haben.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
