---
id: 684d11ffd15a531225d081e9
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Da die Styles nun auf das `div`-Element angewendet wurden, ist es Zeit, einige Styles auf das `p`-Element anzuwenden.

Derzeit ist die Textfarbe für das Absatz-Element schwarz. Es wäre jedoch schön, wenn die Farbe stattdessen ein dunkles Rot wäre.

In einem früheren Schritt haben Sie gelernt, wie man mit Hintergrundfarben arbeitet, zum Beispiel so: `bg-red-100`. Etwas Ähnliches können Sie mit der Textfarbe machen, indem Sie diesem Muster folgen: `text-color-shade`.

Hier ist ein Beispiel, wie Sie einem Element eine blaue Textfarbe zuweisen:

```html
<p class="text-blue-500">This has blue text.</p>
```

Beginnen Sie damit, Ihrem `p`-Element ein Klassenattribut hinzuzufügen. Wenden Sie dann die richtige Utility-Klasse an, um eine rote Textfarbe mit dem Farbton `700` einzustellen.

# --hints--

Ihr `p`-Element sollte ein Klassenattribut haben.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

Ihr öffnendes `p`-Tag sollte die Utility-Klasse `text-red-700` enthalten.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
