---
id: 684d009994a3840b98c9aa3a
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Im Moment erstreckt sich die Breite des `div`-Elements über die gesamte Breite der Seite. Um Breiten in Tailwind CSS anzuwenden, haben Sie verschiedene Möglichkeiten, zum Beispiel das `w-number`-Muster für feste Breiten oder das `w-fraction`-Muster für Prozentangaben.

Hier ist ein Beispiel, wie Sie eine Breite anwenden, die auf `66%` des übergeordneten Containers gesetzt ist:

```html
<p class="w-2/3">This is a percentage width.</p>
```

Fügen Sie in Ihrem öffnenden `div`-Tag eine Breite hinzu, die auf `50%` des übergeordneten Containers gesetzt ist.

# --hints--

Ihr öffnender `div`-Tag sollte die Utility-Klasse `w-1/2` enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
