---
id: 684d105e488dbd1166a717ae
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Wenn Sie die Größe des Vorschaufensters vergrößern, werden Sie feststellen, dass alle Elemente im Container weiterhin linksbündig ausgerichtet sind. Es wäre schön, wenn die Flex-Elemente zentriert wären.

Um Flex-Elemente entlang der Hauptachse eines Containers mit Tailwind CSS zu positionieren, können Sie folgendes Muster verwenden: `justify-value`. `value` kann `start`, `end`, `center`, `between` und mehr sein. Jeder gültige Wert, der in Vanilla CSS mit `justify-content` verwendet wird, kann hier mit Tailwind CSS genutzt werden.

Fügen Sie innerhalb Ihres öffnenden `div`-Tags die korrekte Utility-Klasse hinzu, um die `justify-content`-Eigenschaft auf `center` zu setzen.

# --hints--

Ihr öffnender `div`-Tag sollte die `justify-center`-Utility-Klasse enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
