---
id: 684d0eeb9c81e5109e52fe8b
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Das `div`-Element sieht schon besser aus. Aber jetzt ist das dismiss `button` ein wenig zu nah am Absatz-Element.

Um zwischen den Flex-Elementen Abstand zu schaffen, können Sie das folgende Muster verwenden: `gap-number`. `number` steht für den festen Abstand zwischen den Reihen und Spalten in einem Flex-Layout.

Hier ist ein Beispiel, wie Sie einem Element ein `gap` zuweisen:

```html
<p class="flex gap-8">This is gap.</p>
```

Fügen Sie in Ihrem öffnenden `div`-Tag die richtige Utility-Klasse ein, um den `gap` auf `4` zu setzen.

# --hints--

Ihr öffnender `div`-Tag sollte die Utility-Klasse `gap-4` enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
