---
id: 684d07cbfa43a20d6a002ed7
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Die Styles für das `div`-Element nehmen langsam Gestalt an. Aber es wäre schön, wenn um den Text herum etwas Platz wäre.

Um `padding` mit Tailwind CSS auf ein Element anzuwenden, können Sie folgendes Muster verwenden: `p-number`, wobei `number` die Pixelzahl darstellt, die Sie auf allen vier Seiten anwenden möchten.

Hier ist ein Beispiel, wie man `padding` von `8px` auf ein Element anwendet:

```html
<p class="p-8">This is padding.</p>
```

Fügen Sie in Ihrem öffnenden `div`-Tag die korrekte Utility-Klasse ein, um `padding` von `4px` anzuwenden.

# --hints--

Ihr öffnender `div`-Tag sollte die Utility-Klasse `p-4` enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
