---
id: 68501d32b69f9d7756334ffc
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Wenn Sie das Vorschaufenster vergrößern, sollten Sie das `div`-Element horizontal zentriert auf der Seite sehen. Für kleinere Geräte erstreckt sich die Breite des `div`-Elements jedoch über die gesamte Seitenbreite.

Aktualisieren Sie die `mx-auto`-Klasse so, dass sie nur für `md`-Geräte und größere angewendet wird.

# --hints--

Ihr öffnendes `div`-Tag sollte die `md:mx-auto`-Utility-Klasse enthalten.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

Sie sollten die `mx-auto`-Klasse nicht mehr in Ihrem öffnenden `div`-Tag haben.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
