---
id: 66bba6fff611169359d9d36a
title: بناء صفحة بريد إلكتروني سرية
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يحتوي على عنصر `main` مع `id` يساوي `email`.
1. يجب أن يحتوي عنصر `#email` على `padding` بقيمة `50px`، وهامش علوي بقيمة `50px`، و`width` بقيمة `500px`، و`border` بقيمة `2px`.
1. يجب أن يكون العرض الكلي لعنصر `#email`، بما في ذلك الحشوات والحدود، هو `500px`.
1. يجب أن يكون لديك عنصران من نوع `div`، أحدهما يحمل المعرف `confidential` والآخر يحمل المعرف `top-secret`، داخل عنصر `#email`.
1. يجب أن يكون لعناصر `#confidential` و`#top-secret` خاصية `display` بقيمة `inline-block`.
1. يجب أن تحتوي عناصر `#confidential` و`#top-secret` على `padding`، وهامش أيسر، و`border`.
1. يجب أن يحتوي عنصر `#confidential` على النص `CONFIDENTIAL`.
1. يجب أن يحتوي عنصر `#top-secret` على النص `TOP SECRET`.
1. يجب تدوير عنصري `#confidential` و`#top-secret` باستخدام خاصية التحويل في CSS.
1. يجب أن يكون لديك ثلاثة عناصر فقرة على الأقل داخل عنصر `#email`.
1. يجب أن يكون لديك ثلاثة عناصر `span` على الأقل بصنف `blurred`، داخل عناصر الفقرة.
1. يجب أن يكون لديك محدِّد `blurred` يطبق تأثير التمويه بقيمة `3px` باستخدام مرشح CSS.

# --hints--

يجب أن يكون لديك عنصر `main` يحمل `id` بقيمة `email`.

```js
assert.exists(document.querySelector('main#email'));
```

يجب أن يكون لديك محدد `#email` يضبط `padding` لعناصره من جميع الجهات إلى `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

يجب أن يكون لديك محدد `#email` يضبط `margin-top` لعناصره إلى `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

يجب أن يحتوي محدِّد `#email` على خاصية `width` بقيمة `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

يجب أن يحتوي العنصر `#email` على حدود بعرض `2px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

يجب أن يحتوي عنصر `#email` على خاصية `box-sizing` بقيمة `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

يجب أن تحتوي على عنصر `div` يحمل `id` باسم `confidential` داخل عنصر `#email`.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

يجب أن يكون لديك عنصر `div` يحمل `id` باسم `top-secret` داخل عنصر `#email`.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

يجب أن يكون `display` لعنصر `#confidential` مضبوطًا على `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

يجب أن يكون لدى العنصر `#top-secret` الخاص بك خاصية `display` مضبوطة على `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

يجب أن يحتوي عنصر `#confidential` على `padding` من جميع الجهات.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

يجب أن يحتوي عنصر `#top-secret` على `padding` من جميع الجهات.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

يجب أن يحتوي العنصر `#confidential` على هامش أيسر.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

يجب أن يحتوي العنصر `#top-secret` على هامش أيسر.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

يجب أن يحتوي عنصر `#confidential` على `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

يجب أن يحتوي عنصر `#top-secret` على `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

يجب أن يكون نص عنصر `#confidential` هو `CONFIDENTIAL`. إما أنك حذفت النص أو يوجد خطأ مطبعي.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

يجب أن يكون نص عنصر `#top-secret` هو `TOP SECRET`. إما أنك حذفت النص أو يوجد خطأ مطبعي.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

يجب أن يستخدم العنصر `#confidential` خاصية CSS `transform` لتدوير العنصر.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

يجب أن يستخدم العنصر `#top-secret` خاصية CSS `transform` لتدوير العنصر.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

يجب أن يحتوي عنصر `#email` على ثلاثة عناصر فقرة على الأقل.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

يجب أن يحتوي على الأقل ثلاثة عناصر `span` بصنف `blurred` داخل فقراتك.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

يجب ألا تكون عناصر `.blurred` فارغة.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

يجب أن يكون لديك محدِّد `.blurred` يضبط خاصية `filter` لعناصره إلى `blur(3px)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
