---
id: a3bfc1673c0526e06d3ac698
title: Побудуйте калькулятор суми простих чисел
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

У цій лабораторній роботі ви створите калькулятор, який приймає число і повертає суму всіх простих чисел, що менші або рівні цьому числу.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати `sumPrimes` функцію, яка приймає число як аргумент.
1. `sumPrimes` функція має повертати суму всіх простих чисел, менших або рівних заданому числу.
1. Якщо вхідне число менше за `2`, функція має повертати `0`.

# --hints--

Ви повинні мати `sumPrimes` функцію.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` має повертати `17`.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` має повертати `10`.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` має повертати `2`.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` має повертати `0`.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` має повертати `73156`.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
