---
id: af7588ade1100bde429baf20
title: 빠진 글자 감지기 만들기
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `fearNotLetter`라는 이름의 함수가 있어야 합니다.
1. `fearNotLetter` 함수는 알파벳 순서로 된 글자 범위를 나타내는 문자열 하나를 인수로 받아야 하며, 이 범위에는 한 글자가 빠져 있을 수 있습니다.
1. 함수는 전달된 글자 범위에서 빠진 글자를 찾아 반환해야 합니다.
1. 범위 내 모든 글자가 있으면 함수는 `undefined`를 반환해야 합니다.

# --hints--

`fearNotLetter` 함수가 있어야 합니다.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")`는 문자열 `d`를 반환해야 합니다.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")`는 문자열 `i`를 반환해야 합니다.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")`는 문자열 `u`를 반환해야 합니다.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")`는 문자열 `e`를 반환해야 합니다.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")`는 `undefined`를 반환해야 합니다.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
