---
id: 68e51d8f163d3f8b70f9c4e5
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie den Umgang mit den Methoden `includes()` und `slice()`.

Beginnen Sie damit, eine Variable namens `fccSentence` zu erstellen und ihr den String `"freeCodeCamp is a great place to learn web development."` zuzuweisen.

# --hints--

Sie sollten eine Variable namens `fccSentence` deklarieren.

```js
assert.exists(fccSentence);
```

Ihre Variable `fccSentence` sollte ein String sein.

```js
assert.isString(fccSentence);
```

Ihrer Variable `fccSentence` sollte der String `"freeCodeCamp is a great place to learn web development."` zugewiesen werden.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
