---
id: 68e6172975008fea742a63cd
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Jetzt erkunden Sie die `slice()`-Methode, die einen Teil eines Strings extrahiert.

Erstellen Sie eine Variable namens `message` und weisen Sie ihr den String `"Welcome to freeCodeCamp!"` zu.

Verwenden Sie dann `console.log()`, um `"Here are some examples of the slice() method:"` in der Konsole auszugeben.

# --hints--

Sie sollten eine Variable namens `message` deklarieren.

```js
assert.exists(message);
```

Ihre `message`-Variable sollte ein String sein.

```js
assert.isString(message);
```

Ihre `message`-Variable sollte `"Welcome to freeCodeCamp!"` sein.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

Sie sollten `"Here are some examples of the slice() method:"` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
