---
id: 68e6172875008fea742a63c9
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Wie Sie sich aus den vorherigen Lektionen erinnern, prüft die `includes()`-Methode, ob ein String eine bestimmte Teilzeichenfolge enthält, und gibt `true` oder `false` zurück.

Hier ist ein Beispiel, das die `includes()`-Methode verwendet:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Erstellen Sie eine Variable mit dem Namen `hasFreeCodeCamp`. Weisen Sie ihr dann das Ergebnis der Verwendung der `includes()`-Methode auf `fccSentence` zu, um zu prüfen, ob sie `"freeCodeCamp"` enthält.

# --hints--

Sie sollten eine Variable mit dem Namen `hasFreeCodeCamp` deklarieren.

```js
assert.exists(hasFreeCodeCamp);
```

Sie sollten `fccSentence.includes("freeCodeCamp")` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Sie sollten das Ergebnis von `fccSentence.includes("freeCodeCamp")` Ihrer `hasFreeCodeCamp`-Variable zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Das Ergebnis von `fccSentence.includes("freeCodeCamp")` sollte `true` sein.

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
