---
id: 68e6172975008fea742a63d1
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Die `slice()`-Methode kann negative Indizes verwenden, um vom Ende des Strings zu zählen. `-1` bezieht sich auf das letzte Zeichen.

Hier ist ein Beispiel:

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

Erstellen Sie eine Variable namens `endPunctuation`. Weisen Sie ihr dann das Ergebnis zu, das durch die Verwendung von `message.slice(-1)` zum Extrahieren des letzten Zeichens entsteht.

Verwenden Sie anschließend entweder ein Template Literal oder String-Konkatenation, um die Nachricht `The ending punctuation mark is a "<endPunctuation>"` in der Konsole auszugeben. Ersetzen Sie `<endPunctuation>` durch den tatsächlichen Wert der Variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Sie sollten eine Variable namens `endPunctuation` deklarieren.

```js
assert.exists(endPunctuation);
```

Ihre `endPunctuation`-Variable sollte ein String sein.

```js
assert.isString(endPunctuation);
```

Sie sollten `message.slice(-1)` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

Sie sollten das Ergebnis von `message.slice(-1)` Ihrer `endPunctuation`-Variable zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

Ihre `endPunctuation`-Variable sollte gleich `"!"` sein.

```js
assert.strictEqual(endPunctuation, "!");
```

Sie sollten die Nachricht `The ending punctuation mark is a "<endPunctuation>"` ausgeben, wobei `<endPunctuation>` durch den tatsächlichen Wert der Variable ersetzt wird.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
