---
id: 68e6172975008fea742a63cc
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Die `includes()`-Methode unterscheidet zwischen Groß- und Kleinbuchstaben, das heißt, sie ist case-sensitiv.

Erstellen Sie eine Variable namens `hasLowercaseFCC`. Weisen Sie ihr dann das Ergebnis der Verwendung der `includes()`-Methode zu, um zu prüfen, ob `fccSentence` `"freecodecamp"` (alles Kleinbuchstaben) enthält.

Verwenden Sie anschließend entweder ein Template Literal oder String-Konkatenation, um die Nachricht `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` in der Konsole auszugeben. Ersetzen Sie `<hasLowercaseFCC>` durch den tatsächlichen Wert der Variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Sie sollten eine Variable namens `hasLowercaseFCC` deklarieren.

```js
assert.exists(hasLowercaseFCC);
```

Sie sollten `fccSentence.includes("freecodecamp")` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Sie sollten das Ergebnis von `fccSentence.includes("freecodecamp")` Ihrer `hasLowercaseFCC`-Variable zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Das Ergebnis von `fccSentence.includes("freecodecamp")` sollte `false` sein.

```js
assert.isFalse(hasLowercaseFCC);
```

Sie sollten die Nachricht ausgeben, die `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` enthält, wobei `<hasLowercaseFCC>` durch den tatsächlichen Wert der Variable ersetzt werden sollte.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
