---
id: 68e6172975008fea742a63ce
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Die `slice()`-Methode extrahiert einen Abschnitt aus einem String. Sie nimmt zwei Argumente: den Startindex (einschließlich) und den Endindex (ausschließlich).

Hier ist ein Beispiel:

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

String-Indizes beginnen bei 0. Im `message`-String ist das Zeichen am Index 11 `f` (der Anfang von „freeCodeCamp“).

Erstellen Sie eine Variable namens `platform`. Weisen Sie ihr dann das Ergebnis zu, das entsteht, wenn Sie `message.slice(11, 23)` verwenden, um `"freeCodeCamp"` zu extrahieren.

# --hints--

Sie sollten eine Variable namens `platform` deklarieren.

```js
assert.exists(platform);
```

Ihre `platform`-Variable sollte ein String sein.

```js
assert.isString(platform);
```

Sie sollten `message.slice(11, 23)` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Sie sollten das Ergebnis von `message.slice(11, 23)` Ihrer `platform`-Variable zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Ihre `platform`-Variable sollte gleich `"freeCodeCamp"` sein.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
