---
id: 68e6172975008fea742a63ca
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Verwenden Sie nun entweder ein Template Literal oder String-Konkatenation, um die Nachricht `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` in der Konsole auszugeben. Ersetzen Sie `<hasFreeCodeCamp>` durch den tatsächlichen Wert der Variablen.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Sie sollten die Nachricht `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` ausgeben, wobei `<hasFreeCodeCamp>` durch den tatsächlichen Wert der Variablen ersetzt werden sollte.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
