---
id: 68e6172975008fea742a63cb
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Jetzt ist es an der Zeit zu sehen, was passiert, wenn die Teilzeichenkette NICHT gefunden wird.

Erstellen Sie eine Variable namens `hasJavaScript`. Weisen Sie ihr dann das Ergebnis der Verwendung der `includes()`-Methode zu, um zu prüfen, ob `fccSentence` `"JavaScript"` enthält.

Verwenden Sie anschließend entweder ein Template Literal oder String-Konkatenation, um die Nachricht `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` in der Konsole auszugeben. Ersetzen Sie `<hasJavaScript>` durch den tatsächlichen Wert der Variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Sie sollten eine Variable namens `hasJavaScript` deklarieren.

```js
assert.exists(hasJavaScript);
```

Sie sollten `fccSentence.includes("JavaScript")` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Sie sollten das Ergebnis von `fccSentence.includes("JavaScript")` Ihrer `hasJavaScript`-Variable zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Das Ergebnis von `fccSentence.includes("JavaScript")` sollte `false` sein.

```js
assert.isFalse(hasJavaScript);
```

Sie sollten die Nachricht `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` ausgeben, wobei `<hasJavaScript>` durch den tatsächlichen Wert der Variable ersetzt wird.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
