---
id: 68e6172975008fea742a63d0
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Sie können `slice()` verwenden, um vom Anfang eines Strings zu extrahieren.

Erstellen Sie eine Variable namens `greetingWord`. Weisen Sie ihr dann das Ergebnis zu, das Sie mit `message.slice(0, 7)` erhalten, um das erste Wort `"Welcome"` zu extrahieren.

Verwenden Sie anschließend entweder ein Template Literal oder String-Konkatenation, um die Nachricht `The first word is "<greetingWord>".` in der Konsole auszugeben. Ersetzen Sie `<greetingWord>` durch den tatsächlichen Wert der Variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Sie sollten eine Variable namens `greetingWord` deklarieren.

```js
assert.exists(greetingWord);
```

Ihre `greetingWord`-Variable sollte ein String sein.

```js
assert.isString(greetingWord);
```

Sie sollten `message.slice(0, 7)` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Sie sollten das Ergebnis von `message.slice(0, 7)` Ihrer `greetingWord`-Variable zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Ihre `greetingWord`-Variable sollte gleich `"Welcome"` sein.

```js
assert.strictEqual(greetingWord, "Welcome");
```

Sie sollten die Nachricht `The first word is "<greetingWord>".` ausgeben, wobei `<greetingWord>` durch den tatsächlichen Wert der Variable ersetzt wird.

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
