---
id: 68ef6e42fa6e8a9a18a401fa
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Verwenden Sie `console.log()`, um `"Here are some examples of the includes() method:"` in der Konsole auszugeben.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Sie sollten `"Here are some examples of the includes() method:"` in der Konsole ausgeben.

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
