---
id: 5a24c314108439a4d403617c
title: Tumia njia ya maisha componentWillMount
challengeType: 6
forumTopicId: 301423
dashedName: use-the-lifecycle-method-componentwillmount
---

# --description--

Sehemu za React zina njia maalum kadhaa zinazotoa fursa za kufanya vitendo katika nyakati maalum za mzunguko wa maisha wa sehemu. Hizi huitwa njia za maisha, au hooks za maisha, na zinakuwezesha kushika sehemu katika nyakati fulani. Hii inaweza kuwa kabla ya kuonyeshwa, kabla ya kusasishwa, kabla ya kupokea vigezo vinavyopitishwa, kabla ya kuondolewa, na kadhalika. Hapa kuna orodha ya baadhi ya njia kuu za maisha: `componentWillMount()` `componentDidMount()` `shouldComponentUpdate()` `componentDidUpdate()` `componentWillUnmount()` Mafunzo kadhaa yajayo yataelezea baadhi ya matumizi ya msingi ya njia hizi za maisha.

**Kumbuka:** Njia ya maisha `componentWillMount` itafutwa katika toleo lijalo la 16.X na kuondolewa kabisa katika toleo la 17. Jifunze zaidi katika <a href="https://www.freecodecamp.org/news/how-to-safely-use-reacts-life-cycles-with-fiber-s-async-rendering-fd4469ebbd8f/" target="_blank" rel="noopener noreferrer nofollow">makala hii</a>

# --instructions--

Njia `componentWillMount()` inaitwa kabla ya njia `render()` wakati sehemu inapoingizwa kwenye DOM. Andika kitu kwenye konsoli ndani ya `componentWillMount()` - unaweza kufungua konsoli ya kivinjari chako kuona matokeo.

# --hints--

`MyComponent` inapaswa kuonyesha kipengele cha `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').length === 1;
  })()
);
```

`console.log` inapaswa kuitwa ndani ya `componentWillMount`.

```js
assert(
  (function () {
    const lifecycle = React.createElement(MyComponent)
      .type.prototype.componentWillMount.toString()
      .replace(/ /g, '');
    return lifecycle.includes('console.log(');
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line
    console.log('Component is mounting...');
    // Change code above this line
  }
  render() {
    return <div />
  }
};
```
