---
id: 5a24c314108439a4d4036167
title: Onyesha sehemu ya darasa la vitu kwenye DOM
challengeType: 6
forumTopicId: 301404
dashedName: render-a-class-component-to-the-dom
---

# --description--

Huenda ukumbuke kutumia API ya ReactDOM katika changamoto ya awali kuonyesha vipengele vya JSX kwenye DOM. Mchakato wa kuonyesha sehemu za React utaonekana kama huo. Changamoto za hivi karibuni zilizingatia sehemu na muundo, kwa hivyo kuonyesha kulifanywa kwa niaba yako nyuma ya pazia. Hata hivyo, msimbo wowote wa React utakaoandika hautaonyesha kwenye DOM bila kufanya wito kwa API ya ReactDOM.

Hapa kuna ukumbusho wa sintaksia: `ReactDOM.render(componentToRender, targetNode)`. Hoja ya kwanza ni sehemu ya React unayotaka kuonyesha. Hoja ya pili ni nodi ya DOM unayotaka kuonyesha sehemu hiyo ndani yake.

Sehemu za React hupitishwa ndani ya `ReactDOM.render()` kwa njia tofauti kidogo kuliko vipengele vya JSX. Kwa vipengele vya JSX, unapitisha jina la kipengele unachotaka kuonyesha. Hata hivyo, kwa sehemu za React, unahitaji kutumia sintaksia ile ile kama unavyokuwa unaonyesha sehemu zilizopangwa ndani, kwa mfano `ReactDOM.render(<ComponentToRender />, targetNode)`. Unatumia sintaksia hii kwa sehemu za darasa la vitu la ES6 na sehemu za kitendakazi.

# --instructions--

Sehemu zote za `Fruits` na `Vegetables` zimetengenezwa kwa niaba yako nyuma ya pazia. Onyesha sehemu zote kama watoto wa sehemu ya `TypesOfFood`, kisha onyesha `TypesOfFood` kwenye DOM. Kuna `div` yenye `id='challenge-node'` inayopatikana kwa ajili yako kutumia.

# --hints--

Sehemu ya `TypesOfFood` inapaswa kurudisha kipengele kimoja cha `div`.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Sehemu ya `TypesOfFood` inapaswa kuonyesha sehemu ya `Fruits` baada ya kipengele cha `h1`.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

Sehemu ya `TypesOfFood` inapaswa kuonyesha sehemu ya `Vegetables` baada ya `Fruits`.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

Sehemu ya `TypesOfFood` inapaswa kuonyesha kwenye DOM ndani ya `div` yenye kitambulisho cha `challenge-node`.

```js
assert(
  (function () {
    const hasRenderCall = /ReactDOM\.render\(<TypesOfFood\/>,document\.getElementById\((['"`])challenge-node\1\)\);?/.test(
      __helpers.removeWhiteSpace(code)
    );
    if (!hasRenderCall) {
      return false;
    }

    const e = React.createElement;
    globalThis.Fruits = function Fruits() {
      return e(
        'div',
        null,
        e('h2', null, 'Fruits:'),
        e('h4', null, 'Non-Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Apples'),
          e('li', null, 'Blueberries'),
          e('li', null, 'Strawberries'),
          e('li', null, 'Bananas')
        ),
        e('h4', null, 'Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Lemon'),
          e('li', null, 'Lime'),
          e('li', null, 'Orange'),
          e('li', null, 'Grapefruit')
        )
      );
    };
    globalThis.Vegetables = function Vegetables() {
      return e(
        'div',
        null,
        e('h2', null, 'Vegetables:'),
        e(
          'ul',
          null,
          e('li', null, 'Brussel Sprouts'),
          e('li', null, 'Broccoli'),
          e('li', null, 'Squash')
        )
      );
    };

    ReactDOM.render(
      React.createElement(TypesOfFood),
      document.getElementById('challenge-node')
    );

    const root = document.getElementById('challenge-node').childNodes[0];
    if (!root) {
      return false;
    }
    const html = root.innerHTML;
    return (
      html.includes(
        '<div><h2>Fruits:</h2><h4>Non-Citrus:</h4><ul><li>Apples</li><li>Blueberries</li><li>Strawberries</li><li>Bananas</li></ul><h4>Citrus:</h4><ul><li>Lemon</li><li>Lime</li><li>Orange</li><li>Grapefruit</li></ul></div>'
      ) &&
      html.includes(
        '<div><h2>Vegetables:</h2><ul><li>Brussel Sprouts</li><li>Broccoli</li><li>Squash</li></ul></div>'
      )
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
```

# --solutions--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}
          <Fruits />
           <Vegetables />
         {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
ReactDOM.render(<TypesOfFood />, document.getElementById('challenge-node'));
```
