---
id: 5a24c314108439a4d403617e
title: Ongeza wasikilizaji wa tukio
challengeType: 6
forumTopicId: 301377
dashedName: add-event-listeners
---

# --description--

Njia ya `componentDidMount()` pia ni mahali pazuri zaidi pa kuambatisha wasikilizaji wowote wa tukio unayohitaji kuongeza kwa ajili ya utendaji maalum. React hutoa mfumo wa tukio la synthetic unaozunguka mfumo wa tukio la asili uliopo katika vivinjari. Hii inamaanisha kwamba mfumo wa tukio la synthetic hufanya kazi sawa kabisa bila kujali kivinjari cha mtumizi - hata kama matukio ya asili yanaweza kutofautiana kati ya vivinjari tofauti.

Umeanza tayari kutumia baadhi ya wasughulikiaji wa tukio la synthetic kama `onClick()`. Mfumo wa tukio la synthetic wa React ni mzuri kutumia kwa mwingiliano mwingi utakaoendesha kwenye vipengele vya DOM. Hata hivyo, ikiwa unataka kuambatisha mshughulikiaji wa tukio kwa vitu vya document au window, lazima ufanye hivyo moja kwa moja.

# --instructions--

Ambatisha msikilizaji wa tukio katika njia ya `componentDidMount()` kwa matukio ya `keydown` na yafanye matukio haya kusababisha wito wa kurudisha `handleKeyPress()`. Unaweza kutumia `document.addEventListener()` ambayo inachukua tukio (katika nukuu) kama hoja ya kwanza na wito wa kurudisha kama hoja ya pili.

Kisha, katika `componentWillUnmount()`, ondoa msikilizaji huyu wa tukio uleule. Unaweza kupitisha hoja zile zile kwa `document.removeEventListener()`. Ni desturi nzuri kutumia njia hii ya mzunguko wa maisha kufanya usafi wowote kwenye sehemu za React kabla hazijaondolewa na kuharibiwa. Kuondoa wasikilizaji wa tukio ni mfano wa tendo la usafi kama hilo.

# --hints--

`MyComponent` inapaswa onyesha kipengele cha `div` kinachozunguka lebo ya `h1`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').children().find('h1').length === 1;
  })()
);
```

Msikilizaji wa `keydown` unapaswa kuambatishwa kwa document katika `componentDidMount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const didMountString = mockedComponent
      .instance()
      .componentDidMount.toString();
    return new RegExp(
      'document.addEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(didMountString);
  })()
);
```

Msikilizaji wa `keydown` unapaswa kuondolewa kutoka document katika `componentWillUnmount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const willUnmountString = mockedComponent
      .instance()
      .componentWillUnmount.toString();
    return new RegExp(
      'document.removeEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(willUnmountString);
  })()
);
```

Mara sehemu imewekwa, kubonyeza `enter` kunapaswa kusahisha hali yake na lebo ya `h1` iliyoonyeshwa.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const beforeState = mockedComponent.state('message');
  const beforeText = mockedComponent.find('h1').text();
  const pressEnterKey = () => {
    mockedComponent.instance().handleKeyPress({ keyCode: 13 });
    return waitForIt(() => {
      mockedComponent.update();
      return {
        state: mockedComponent.state('message'),
        text: mockedComponent.find('h1').text()
      };
    });
  };
  const afterKeyPress = await pressEnterKey();
  assert(
    beforeState !== afterKeyPress.state && beforeText !== afterKeyPress.text
  );
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleEnter = this.handleEnter.bind(this);
    this.handleKeyPress = this.handleKeyPress.bind(this);
  }
  // Change code below this line
  componentDidMount() {

  }
  componentWillUnmount() {

  }
  // Change code above this line
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleKeyPress = this.handleKeyPress.bind(this);
    this.handleEnter = this.handleEnter.bind(this);  }
  componentDidMount() {
    // Change code below this line
    document.addEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  componentWillUnmount() {
    // Change code below this line
    document.removeEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```
