---
id: 5a24c314108439a4d4036163
title: Unda sehemu ya React
challengeType: 6
forumTopicId: 301386
dashedName: create-a-react-component
---

# --description--

Njia nyingine ya kufafanua sehemu ya React ni kwa kutumia sintaksia ya ES6 `class`. Katika mfano ufuatao, `Kitten` inaongeza `React.Component`:

```jsx
class Kitten extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <h1>Hi</h1>
    );
  }
}
```

Hii inaunda darasa la ES6 `Kitten` ambalo linaongeza darasa la `React.Component`. Hivyo darasa la `Kitten` sasa lina upatikanaji wa vipengele vingi muhimu vya React, kama vile hali ya ndani na hooks za mzunguko wa maisha. Usijali kama bado hujafahamiana na maneno haya, yatajadiliwa kwa kina zaidi katika changamoto zijazo. Pia angalia darasa la `Kitten` lina `constructor` iliyofafanuliwa ndani yake ambayo inaita `super()`. Inatumia `super()` kuita kijenzi cha darasa la mzazi, katika kesi hii `React.Component`. Kijenzi ni njia maalum inayotumika wakati wa uanzishaji wa vitu vinavyoundwa kwa kutumia neno kuu la `class`. Ni mbinu bora kuita `constructor` ya sehemu na `super`, na kupitisha `props` kwa zote mbili. Hii inahakikisha sehemu inaanzishwa ipasavyo. Kwa sasa, fahamu kuwa ni kawaida msimbo huu kujumuishwa. Hivi karibuni utaona matumizi mengine ya kijenzi pamoja na `props`.

# --instructions--

`MyComponent` imefafanuliwa katika mhariri wa msimbo kwa kutumia sintaksia ya darasa. Malizia kuandika njia ya `render` ili irudishe kipengele cha `div` chenye `h1` yenye maandishi `Hello React!`.

# --hints--

Sehemu ya React inapaswa kurudisha kipengele cha `div`.

```js
assert(Enzyme.shallow(React.createElement(MyComponent)).type() === 'div');
```

`div` iliyorejeshwa inapaswa kuonyesha kichwa cha sehemu cha `h1` ndani yake.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.shallow(React.createElement(MyComponent)).html()
  )
);
```

Kichwa cha sehemu cha `h1` kinapaswa kuwa na mfuatano wa herufi `Hello React!`.

```js
assert(
  Enzyme.shallow(React.createElement(MyComponent)).html() ===
    '<div><h1>Hello React!</h1></div>'
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line



    // Change code above this line
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line
    return (
      <div>
        <h1>Hello React!</h1>
      </div>
    );
    // Change code above this line
  }
};
```
