---
id: 5a24c314108439a4d4036180
title: Boresha upyaonyeshaji kwa shouldComponentUpdate
challengeType: 6
forumTopicId: 301398
dashedName: optimize-re-renders-with-shouldcomponentupdate
---

# --description--

Hadi sasa, ikiwa sehemu yoyote inapokea `state` mpya au `props` mpya, hujionyesha upya yenyewe na watoto wake wote. Hii kwa kawaida ni sawa. Lakini React hutoa njia ya maisha ya sehemu (lifecycle method) unayoweza kuita wakati sehemu za watoto zinapopokea `state` mpya au `props` mpya, na kutangaza hasa kama sehemu hizo zinapaswa kusasishwa au la. Njia hiyo ni `shouldComponentUpdate()`, na inachukua `nextProps` na `nextState` kama vigezo.

Njia hii ni njia muhimu ya kuboresha utendaji. Kwa mfano, tabia ya kawaida ni kwamba sehemu yako hujionyesha upya inapopokea `props` mpya, hata kama `props` hazijabadilika. Unaweza kutumia `shouldComponentUpdate()` kuzuia hili kwa kulinganisha `props`. Njia hii lazima irudishe thamani ya `boolean` inayomwambia React kama sehemu inapaswa kusasishwa au la. Unaweza kulinganisha vigezo vinavyopitishwa vya sasa (`this.props`) na vigezo vinavyopitishwa vijavyo (`nextProps`) ili kuamua kama unahitaji kusasisha au la, na kurudisha `true` au `false` ipasavyo.

# --instructions--

Njia ya `shouldComponentUpdate()` imeongezwa katika sehemu inayoitwa `OnlyEvens`. Kwa sasa, njia hii inarudisha `true` hivyo `OnlyEvens` hujionyesha upya kila mara inapopokea `props` mpya. Badilisha njia hiyo ili `OnlyEvens` isasishwe tu ikiwa `value` ya vigezo vinavyopitishwa vipya ni nambari kamili chanya ya hata. Bofya kitufe cha `Add` na tazama mfuatano wa matukio katika konsoli ya kivinjari chako wakati vifungo vya maisha ya sehemu vinapoanzishwa.

# --hints--

Sehemu ya `Controller` inapaswa kuonyesha sehemu ya `OnlyEvens` kama mtoto.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Controller));
    return (
      mockedComponent.find('Controller').length === 1 &&
      mockedComponent.find('OnlyEvens').length === 1
    );
  })()
);
```

Njia ya `shouldComponentUpdate` inapaswa kutangazwa kwenye sehemu ya `OnlyEvens`.

```js
assert(
  (() => {
    const child = React.createElement(OnlyEvens)
      .type.prototype.shouldComponentUpdate.toString()
      .replace(/s/g, '');
    return child !== 'undefined';
  })()
);
```

Sehemu ya `OnlyEvens` inapaswa kurudisha lebo ya `h1` ambayo inaonyesha thamani ya `this.props.value`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 1000 });
    return mockedComponent.find('h1').html();
  };
  const second = () => {
    mockedComponent.setState({ value: 10 });
    return mockedComponent.find('h1').html();
  };
  const firstValue = first();
  const secondValue = second();
  assert(firstValue === '<h1>1000</h1>' && secondValue === '<h1>10</h1>');
})();
```

`OnlyEvens` inapaswa kujionyesha upya tu wakati `nextProps.value` ni nambari kamili chanya ya hata.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 8 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ value: 7 });
    return mockedComponent.find('h1').text();
  };
  const third = () => {
    mockedComponent.setState({ value: 42 });
    return mockedComponent.find('h1').text();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(firstValue === '8' && secondValue === '8' && thirdValue === '42');
})();
```

# --seed--

## --seed-contents--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return true;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```

# --solutions--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return nextProps.value % 2 === 0;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```
