---
id: 5a24c314108439a4d4036160
title: Eleza darasa la HTML katika JSX
challengeType: 6
forumTopicId: 301393
dashedName: define-an-html-class-in-jsx
---

# --description--

Sasa unapoanza kufahamika kuandika JSX, huenda unajiuliza tofauti zake na HTML ni zipi.

Hadi sasa, inaonekana kama HTML na JSX ni sawa kabisa.

Tofauti kuu katika JSX ni kwamba huwezi tena kutumia neno `class` kuunda madarasa ya HTML. Hii ni kwa sababu `class` ni neno lililohifadhiwa katika JavaScript. Badala yake, JSX hutumia `className`.

Kwa kweli, mkataba wa majina kwa sifa zote za kipengele cha HTML na rejea za matukio katika JSX huwa camelCase. Kwa mfano, tukio la bofya katika JSX ni `onClick`, badala ya `onclick`. Vivyo hivyo, `onchange` hubadilika kuwa `onChange`. Ingawa hii ni tofauti ndogo, ni muhimu kuizingatia unapokwenda mbele.

# --instructions--

Tumia darasa la `myDiv` kwa `div` iliyotolewa katika msimbo wa JSX.

# --hints--

Kipengele cha `JSX` kinapaswa kurudisha kipengele cha `div`.

```js
assert.strictEqual(JSX.type, 'div');
```

`div` inapaswa kuwa na darasa la `myDiv`.

```js
assert.strictEqual(JSX.props.className, 'myDiv');
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h1>Add a class to this div</h1>
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div className = 'myDiv'>
  <h1>Add a class to this div</h1>
</div>);
```
