---
id: 5a24c314108439a4d403618c
title: Tumia Array.filter() kuchuja safu ya data kwa njia ya mabadiliko
challengeType: 6
forumTopicId: 301416
dashedName: use-array-filter-to-dynamically-filter-an-array
---

# --description--

Njia ya `map` ya safu ya data ni chombo chenye nguvu utakachotumia mara nyingi unapoandika programu na React. Njia nyingine inayohusiana na `map` ni `filter`, ambayo huchuja yaliyomo katika safu ya data kulingana na sharti, kisha hurudisha safu mpya ya data. Kwa mfano, ikiwa una safu ya watumizi wote wenye sifa ya `online` ambayo inaweza kuwekwa kuwa `true` au `false`, unaweza kuchuja watumizi waliopo mtandaoni kwa kuandika:

```js
let onlineUsers = users.filter(user => user.online);
```

# --instructions--

Katika mhariri wa msimbo, `MyComponent` anaanzisha `state` na safu ya watumizi. Baadhi ya watumizi wako mtandaoni na wengine hawako. Chuja safu ya data ili uone watumizi waliopo mtandaoni tu. Ili kufanya hivyo, kwanza tumia `filter` kurudisha safu mpya ya watumizi ambao sifa yao ya `online` ni `true`. Kisha, katika kigezo cha `renderOnline`, pitia safu iliyochujwa, na rudisha kipengele cha `li` kwa kila mtumizi chenye maandishi ya `username` yao. Hakikisha pia unajumuisha sifa ya `key` ya kipekee, kama ilivyokuwa katika changamoto za mwisho.

# --hints--

`MyComponent` inapaswa kuwepo na ionyeshwe kwenye ukurasa.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('MyComponent').length,
  1
);
```

Hali ya `MyComponent` inapaswa kuanzishwa na safu ya watumizi sita.

```js
assert(
  Array.isArray(
    Enzyme.mount(React.createElement(MyComponent)).state('users')
  ) === true &&
    Enzyme.mount(React.createElement(MyComponent)).state('users').length === 6
);
```

`MyComponent` inapaswa kurudisha `div`, `h1`, na kisha orodha isiyopangwa yenye vipengele vya `li` kwa kila mtumizi ambaye hali yake ya mtandaoni imewekwa kuwa `true`.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MyComponent));
  const users = (bool) => ({
    users: [
      { username: 'Jeff', online: bool },
      { username: 'Alan', online: bool },
      { username: 'Mary', online: bool },
      { username: 'Jim', online: bool },
      { username: 'Laura', online: bool }
    ]
  });
  const result = () => comp.find('li').length;
  const _1 = result();
  const _2 = () => {
    comp.setState(users(true));
    return result();
  };
  const _3 = () => {
    comp.setState(users(false));
    return result();
  };
  const _4 = () => {
    comp.setState({ users: [] });
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  assert(
    comp.find('div').length === 1 &&
      comp.find('h1').length === 1 &&
      comp.find('ul').length === 1 &&
      _1 === 4 &&
      _2_val === 5 &&
      _3_val === 0 &&
      _4_val === 0
  );
})();
```

`MyComponent` inapaswa kuonyesha vipengele vya `li` vinavyoonyesha `username` ya kila mtumizi mtandaoni.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MyComponent));
  const users = (bool) => ({
    users: [
      { username: 'Jeff', online: bool },
      { username: 'Alan', online: bool },
      { username: 'Mary', online: bool },
      { username: 'Jim', online: bool },
      { username: 'Laura', online: bool }
    ]
  });
  const ul = () => {
    comp.setState(users(true));
    return comp.find('ul').html();
  };
  const html = ul();
  assert(
    html ===
      '<ul><li>Jeff</li><li>Alan</li><li>Mary</li><li>Jim</li><li>Laura</li></ul>'
  );
})();
```

Kila kipengele cha orodha kinapaswa kuwa na sifa ya `key` ya kipekee.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(MyComponent)).find('ul');
    console.log(ul.debug());
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key()
    ]);
    return keys.size === 4;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      users: [
        {
          username: 'Jeff',
          online: true
        },
        {
          username: 'Alan',
          online: false
        },
        {
          username: 'Mary',
          online: true
        },
        {
          username: 'Jim',
          online: false
        },
        {
          username: 'Sara',
          online: true
        },
        {
          username: 'Laura',
          online: true
        }
      ]
    };
  }
  render() {
    const usersOnline = null; // Change this line
    const renderOnline = null; // Change this line
    return (
      <div>
        <h1>Current Online Users:</h1>
        <ul>{renderOnline}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      users: [
        {
          username: 'Jeff',
          online: true
        },
        {
          username: 'Alan',
          online: false
        },
        {
          username: 'Mary',
          online: true
        },
        {
          username: 'Jim',
          online: false
        },
        {
          username: 'Sara',
          online: true
        },
        {
          username: 'Laura',
          online: true
        }
      ]
    };
  }
  render() {
    const usersOnline = this.state.users.filter(user => {
      return user.online;
    });
    const renderOnline = usersOnline.map(user => {
      return <li key={user.username}>{user.username}</li>;
    });
    return (
      <div>
        <h1>Current Online Users:</h1>
        <ul>{renderOnline}</ul>
      </div>
    );
  }
}
```
