---
id: 5a24c314108439a4d4036179
title: Unda fomu inayodhibitiwa
challengeType: 6
forumTopicId: 301384
dashedName: create-a-controlled-form
---

# --description--

Changamoto ya mwisho ilionyesha kuwa React inaweza kudhibiti hali ya ndani kwa vipengele fulani kama `input` na `textarea`, ambayo inawafanya kuwa sehemu zinazodhibitiwa. Hii inahusiana pia na vipengele vingine vya fomu, ikiwa ni pamoja na kipengele cha kawaida cha HTML `form`.

# --instructions--

Sehemu ya `MyForm` imeandaliwa na `form` tupu yenye mshughulikiaji wa wasilisho. Mshughulikiaji wa wasilisho utaitwa wakati fomu itakapowasilishwa.

Tumeongeza kitufe kinachowasilisha fomu. Unaweza kuona kina `type` iliyowekwa kuwa `submit` ikionyesha kuwa ni kitufe kinachodhibiti fomu. Ongeza kipengele cha `input` ndani ya `form` na weka sifa zake za `value` na `onChange()` kama changamoto ya mwisho. Kisha kamilisha njia ya `handleSubmit` ili iweke sifa ya hali ya sehemu `submit` kwa thamani ya ingizo ya sasa katika `state` ya ndani.

**Kumbuka:** Pia lazima uitaje `event.preventDefault()` katika mshughulikiaji wa wasilisho, ili kuzuia tabia ya kawaida ya kuwasilisha fomu ambayo itaonyesha upya ukurasa wa mtandao. Kwa urahisi wa mtumiaji, tabia ya kawaida imezuiwa hapa ili kuzuia kuonyesha upya kunakosababisha msimbo wa changamoto kurudi awali.

Mwishowe, unda lebo ya `h1` baada ya `form` ambayo inaonyesha thamani ya `submit` kutoka kwa `state` ya sehemu. Kisha unaweza kuandika kwenye fomu na kubofya kitufe (au kubonyeza enter), na unapaswa kuona ingizo lako linaonyeshwa kwenye ukurasa.

# --hints--

`MyForm` inapaswa kurudisha kipengele cha `div` kinachojumuisha lebo ya `form` na `h1`. Fomu inapaswa kujumuisha `input` na `button`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyForm));
    return (
      mockedComponent.find('div').children().find('form').length === 1 &&
      mockedComponent.find('div').children().find('h1').length === 1 &&
      mockedComponent.find('form').children().find('input').length === 1 &&
      mockedComponent.find('form').children().find('button').length === 1
    );
  })()
);
```

Hali ya `MyForm` inapaswa kuanzishwa na sifa za `input` na `submit`, zote zikiwa mfuatano wa herufi tupu.

```js
assert(
  Enzyme.mount(React.createElement(MyForm)).state('input') === '' &&
    Enzyme.mount(React.createElement(MyForm)).state('submit') === ''
);
```

Kuingiza kwenye kipengele cha `input` kunapaswa kusasisha sifa ya `input` ya hali ya sehemu.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return mockedComponent.state('input');
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return {
      state: mockedComponent.state('input'),
      inputVal: mockedComponent.find('input').props().value
    };
  };
  const before = _1();
  const after = _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
})();
```

Kuwasilisha fomu kunapaswa kuendesha `handleSubmit` ambayo inapaswa kuweka sifa ya `submit` katika hali sawa na ingizo la sasa.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'SubmitInput' } });
    return mockedComponent.state('submit');
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.state('submit');
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'SubmitInput');
})();
```

`handleSubmit` inapaswa kuita `event.preventDefault`

```js
assert(
  __helpers.isCalledWithNoArgs(
    'event.preventDefault',
    MyForm.prototype.handleSubmit.toString()
  )
);
```

Kichwa cha sehemu cha `h1` kinapaswa kuonyesha thamani ya uwanja wa `submit` kutoka kwa hali ya sehemu.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return mockedComponent.find('h1').text();
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.find('h1').text();
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'TestInput');
})();
```

# --seed--

## --seed-contents--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          {/* Change code below this line */}

          {/* Change code above this line */}
          <button type='submit'>Submit!</button>
        </form>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    event.preventDefault();
    this.setState(state => ({
      submit: state.input
    }));
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          <input value={this.state.input} onChange={this.handleChange} />
          <button type='submit'>Submit!</button>
        </form>
        <h1>{this.state.submit}</h1>
      </div>
    );
  }
}
```
