---
id: 5a24c314108439a4d4036189
title: Badilisha CSS ndani ya mstari kwa masharti kulingana na hali ya sehemu
challengeType: 6
forumTopicId: 301380
dashedName: change-inline-css-conditionally-based-on-component-state
---

# --description--

Hadi sasa, umeona matumizi kadhaa ya onyesho la masharti na matumizi ya mitindo ndani ya mstari. Hapa kuna mfano mmoja zaidi unaochanganya mada hizi mbili. Pia unaweza kuonyesha CSS kwa masharti kulingana na hali ya sehemu ya React. Ili kufanya hivyo, unakagua hali fulani, na ikiwa hali hiyo itatimizwa, unabadilisha kitu cha mitindo kilichopewa sehemu za JSX katika njia ya onyesho.

Mfumo huu ni muhimu kuelewa kwa sababu ni mabadiliko makubwa kutoka kwa njia ya jadi ya kutumia mitindo kwa kubadilisha vipengele vya DOM moja kwa moja (ambayo ni ya kawaida sana kwa jQuery, kwa mfano). Katika njia hiyo, lazima ufuatilie wakati vipengele vinabadilika na pia ushughulikie mabadiliko hayo moja kwa moja. Inaweza kuwa vigumu kufuatilia mabadiliko, na hivyo kufanya UI yako isiyotabirika. Unapoweka kitu cha mitindo kulingana na hali, unaeleza jinsi UI inavyopaswa kuonekana kama kazi ya hali ya programu. Kuna mtiririko wazi wa taarifa unaosogea kwa mwelekeo mmoja tu. Hii ndiyo njia inayopendekezwa wakati wa kuandika programu kwa React.

# --instructions--

Mhariri wa msimbo una sehemu rahisi ya ingizo iliyodhibitiwa yenye mpaka uliopambwa. Unataka kupamba mpaka huu kwa rangi nyekundu ikiwa mtumizi ataandika zaidi ya herufi 15 katika kisanduku cha ingizo. Ongeza sharti la kukagua hili na, ikiwa sharti ni halali, weka mtindo wa mpaka wa ingizo kuwa `3px solid red`. Unaweza kujaribu kwa kuingiza maandishi kwenye ingizo.

# --hints--

Sehemu ya `GateKeeper` inapaswa onyesha kipengele cha `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.find('div').length === 1;
  })()
);
```

Sehemu ya `GateKeeper` inapaswa kuanzishwa na funguo ya hali `input` iliyowekwa kuwa mfuatano wa herufi tupu.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.state().input === '';
  })()
);
```

Sehemu ya `GateKeeper` inapaswa onyesha lebo ya `h3` na lebo ya `input`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('h3').length === 1 &&
      mockedComponent.find('input').length === 1
    );
  })()
);
```

Leboa ya `input` inapaswa kuwa na mtindo wa `1px solid black` kwa sifa ya `border` mwanzoni.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('input').props().style.border === '1px solid black'
    );
  })()
);
```

Leboa ya `input` inapaswa kupambwa na mpaka wa `3px solid red` ikiwa thamani ya ingizo katika hali ni ndefu zaidi ya herufi 15.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  let initialStyle = mockedComponent.find('input').props().style.border;
  const state_1 = () => {
    mockedComponent.setState({ input: 'this is 15 char' });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const state_2 = () => {
    mockedComponent.setState({
      input: 'A very long string longer than 15 characters.'
    });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const style_1 = await state_1();
  const style_2 = await state_2();
  assert(
    initialStyle === '1px solid black' &&
      style_1 === '1px solid black' &&
      style_2 === '3px solid red'
  );
```

# --seed--

## --seed-contents--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    // Change code below this line

    // Change code above this line
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    if (this.state.input.length > 15) {
      inputStyle.border = '3px solid red';
    };
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```
