---
id: 5a24c314108439a4d4036177
title: Andika kielekezi rahisi
challengeType: 6
forumTopicId: 301425
dashedName: write-a-simple-counter
---

# --description--

Unaweza kubuni sehemu yenye hali ngumu zaidi kwa kuunganisha dhana zilizojadiliwa hadi sasa. Hizi ni pamoja na kuanzisha `state`, kuandika njia ambazo zinaweka `state`, na kuweka wasimamizi wa bofya ili kuanzisha njia hizi.

# --instructions--

Sehemu ya `Counter` inahifadhi thamani ya `count` katika `state`. Kuna vifungo viwili vinavyofanya wito kwa njia `increment()` na `decrement()`. Andika njia hizi ili thamani ya kielekezi iongezwe au ipunguzwe kwa 1 wakati kifungo kinachofaa kinapobofyanwa. Pia, tengeneza njia ya `reset()` ili wakati kifungo cha kuweka upya kinapobofyanwa, hesabu iwe 0.

**Kumbuka:** Hakikisha haubadili sifa za `className` za vifungo. Pia, kumbuka kuongeza vifungo vinavyohitajika kwa njia mpya zilizoundwa katika kijenzi.

# --hints--

`Counter` inapaswa kurudisha kipengele cha `div` ambacho kina vifungo vitatu vyenye maandishi kwa mpangilio huu `Increment!`, `Decrement!`, `Reset`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Counter));
    return (
      mockedComponent.find('.inc').text() === 'Increment!' &&
      mockedComponent.find('.dec').text() === 'Decrement!' &&
      mockedComponent.find('.reset').text() === 'Reset'
    );
  })()
);
```

Hali ya `Counter` inapaswa kuanzishwa na sifa ya `count` iliyowekwa kuwa `0`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
assert(mockedComponent.find('h1').text() === 'Current Count: 0');
```

Kubofya kifungo cha kuongeza kunapaswa kuongeza hesabu kwa `1`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.inc').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: 1');
```

Kubofya kifungo cha kupunguza kunapaswa kupunguza hesabu kwa `1`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.dec').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: -1');
```

Kubofya kifungo cha kuweka upya kunapaswa kuweka hesabu kuwa `0`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.setState({ count: 5 });
const currentCountElement = mockedComponent.find('h1');
assert(currentCountElement.text() === 'Current Count: 5');
mockedComponent.find('.reset').simulate('click');
assert(currentCountElement.text() === 'Current Count: 0');
```

# --seed--

## --seed-contents--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
  this.increment = this.increment.bind(this);
 this.decrement = this.decrement.bind(this);
 this.reset = this.reset.bind(this);
 }
  reset() {
    this.setState({
      count: 0
    });
  }
  increment() {
    this.setState(state => ({
      count: state.count + 1
    }));
  }
  decrement() {
    this.setState(state => ({
      count: state.count - 1
    }));
  }
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```
