---
id: 5a24c314108439a4d4036183
title: Tumia JavaScript ya Juu katika Njia ya React Onyesha
challengeType: 6
forumTopicId: 301415
dashedName: use-advanced-javascript-in-react-render-method
---

# --description--

Katika changamoto zilizopita, ulijifunza jinsi ya kuingiza msimbo wa JavaScript ndani ya msimbo wa JSX kwa kutumia mabano ya mviringo, `{ }`, kwa kazi kama upatikanaji wa vigezo vinavyopitishwa, kupitisha vigezo vinavyopitishwa, upatikanaji wa state, kuingiza maoni ndani ya msimbo wako, na hivi karibuni, kupamba sehemu zako. Hizi ni matumizi ya kawaida ya kuweka JavaScript ndani ya JSX, lakini siyo njia pekee unayoweza kutumia msimbo wa JavaScript katika sehemu zako za React.

Pia unaweza kuandika JavaScript moja kwa moja katika njia za `render`, kabla ya usemi wa `return`, ***bila*** kuingiza ndani ya mabano ya mviringo. Hii ni kwa sababu bado haiko ndani ya msimbo wa JSX. Unapotaka kutumia kigezo baadaye ndani ya msimbo wa JSX *ndani* ya usemi wa `return`, unaweka jina la kigezo ndani ya mabano ya mviringo.

# --instructions--

Katika msimbo uliotolewa, njia ya `render` ina safu ya data yenye misemo 20 inayowakilisha majibu yanayopatikana katika toy ya Magic Eight Ball ya miaka ya 1980. Tukio la kubofya kitufe limefungwa kwenye njia ya `ask`, hivyo kila mara kitufe kinapobofyanwa nambari ya bahati nasibu itazalishwa na kuhifadhiwa kama `randomIndex` katika state. Katika mstari wa 52, futa mfuatano wa herufi `change me!` na weka upya const ya `answer` ili msimbo wako upate kielezo tofauti cha bahati nasibu cha safu ya data ya `possibleAnswers` kila mara sehemu inapoonyesha upya. Mwisho, weka const ya `answer` ndani ya lebo za `p`.

# --hints--

Sehemu ya `MagicEightBall` inapaswa kuwepo na ionyeshwe kwenye ukurasa.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MagicEightBall)).find('MagicEightBall')
    .length,
  1
);
```

Mtoto wa kwanza wa `MagicEightBall` anapaswa kuwa kipengele cha `input`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MagicEightBall))
    .children()
    .childAt(0)
    .name(),
  'input'
);
```

Mtoto wa tatu wa `MagicEightBall` anapaswa kuwa kipengele cha `button`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MagicEightBall))
    .children()
    .childAt(2)
    .name(),
  'button'
);
```

State ya `MagicEightBall` inapaswa kuanzishwa na sifa ya `userInput` na sifa ya `randomIndex` zote zikiwa na thamani ya mfuatano wa herufi tupu.

```js
assert(
  Enzyme.mount(React.createElement(MagicEightBall)).state('randomIndex') ===
    '' &&
    Enzyme.mount(React.createElement(MagicEightBall)).state('userInput') === ''
);
```

Wakati `MagicEightBall` inapoanzishwa kwenye DOM, inapaswa kurudisha kipengele tupu cha `p`.

```js
assert(
  Enzyme.mount(React.createElement(MagicEightBall)).find('p').length === 1 &&
    Enzyme.mount(React.createElement(MagicEightBall)).find('p').text() === ''
);
```

Wakati maandishi yanaingizwa kwenye kipengele cha `input` na kitufe kinapobofyanwa, sehemu ya `MagicEightBall` inapaswa kurudisha kipengele cha `p` chenye kipengele cha bahati nasibu kutoka kwenye safu ya data ya `possibleAnswers`.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MagicEightBall));
  const simulate = () => {
    comp.find('input').simulate('change', { target: { value: 'test?' } });
    comp.find('button').simulate('click');
  };
  const result = () => comp.find('p').text();
  const _1 = () => {
    simulate();
    return result();
  };
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _6 = () => {
    simulate();
    return result();
  };
  const _7 = () => {
    simulate();
    return result();
  };
  const _8 = () => {
    simulate();
    return result();
  };
  const _9 = () => {
    simulate();
    return result();
  };
  const _10 = () => {
    simulate();
    return result();
  };
  const _1_val = _1();
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  const _6_val = _6();
  const _7_val = _7();
  const _8_val = _8();
  const _9_val = _9();
  const _10_val = _10();
  const actualAnswers = [
    _1_val,
    _2_val,
    _3_val,
    _4_val,
    _5_val,
    _6_val,
    _7_val,
    _8_val,
    _9_val,
    _10_val
  ];
  const hasIndex = actualAnswers.filter(
    (answer, i) => possibleAnswers.indexOf(answer) !== -1
  );
  const notAllEqual = new Set(actualAnswers);
  assert(notAllEqual.size > 1 && hasIndex.length === 10);
})();
```

# --before-all--

```js
var possibleAnswers = [
  'It is certain',
  'It is decidedly so',
  'Without a doubt',
  'Yes, definitely',
  'You may rely on it',
  'As I see it, yes',
  'Outlook good',
  'Yes',
  'Signs point to yes',
  'Reply hazy try again',
  'Ask again later',
  'Better not tell you now',
  'Cannot predict now',
  'Concentrate and ask again',
  "Don't count on it",
  'My reply is no',
  'My sources say no',
  'Outlook not so good',
  'Very doubtful',
  'Most likely'
];
```

# --seed--

## --seed-contents--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class MagicEightBall extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userInput: '',
      randomIndex: ''
    };
    this.ask = this.ask.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  ask() {
    if (this.state.userInput) {
      this.setState({
        randomIndex: Math.floor(Math.random() * 20),
        userInput: ''
      });
    }
  }
  handleChange(event) {
    this.setState({
      userInput: event.target.value
    });
  }
  render() {
    const possibleAnswers = [
      'It is certain',
      'It is decidedly so',
      'Without a doubt',
      'Yes, definitely',
      'You may rely on it',
      'As I see it, yes',
      'Outlook good',
      'Yes',
      'Signs point to yes',
      'Reply hazy try again',
      'Ask again later',
      'Better not tell you now',
      'Cannot predict now',
      'Concentrate and ask again',
      "Don't count on it",
      'My reply is no',
      'My sources say no',
      'Most likely',
      'Outlook not so good',
      'Very doubtful'
    ];
    const answer = 'change me!'; // Change this line
    return (
      <div>
        <input
          type='text'
          value={this.state.userInput}
          onChange={this.handleChange}
          style={inputStyle}
        />
        <br />
        <button onClick={this.ask}>Ask the Magic Eight Ball!</button>
        <br />
        <h3>Answer:</h3>
        <p>
          {/* Change code below this line */}

          {/* Change code above this line */}
        </p>
      </div>
    );
  }
}
```

# --solutions--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class MagicEightBall extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userInput: '',
      randomIndex: ''
    };
    this.ask = this.ask.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  ask() {
    if (this.state.userInput) {
      this.setState({
        randomIndex: Math.floor(Math.random() * 20),
        userInput: ''
      });
    }
  }
  handleChange(event) {
    this.setState({
      userInput: event.target.value
    });
  }
  render() {
    const possibleAnswers = [
      'It is certain',
      'It is decidedly so',
      'Without a doubt',
      'Yes, definitely',
      'You may rely on it',
      'As I see it, yes',
      'Outlook good',
      'Yes',
      'Signs point to yes',
      'Reply hazy try again',
      'Ask again later',
      'Better not tell you now',
      'Cannot predict now',
      'Concentrate and ask again',
      "Don't count on it",
      'My reply is no',
      'My sources say no',
      'Outlook not so good',
      'Very doubtful',
      'Most likely'
    ];
    const answer = possibleAnswers[this.state.randomIndex];
    return (
      <div>
        <input
          type='text'
          value={this.state.userInput}
          onChange={this.handleChange}
          style={inputStyle}
        />
        <br />
        <button onClick={this.ask}>Ask the Magic Eight Ball!</button>
        <br />
        <h3>Answer:</h3>
        <p>{answer}</p>
      </div>
    );
  }
}
```
