---
id: 6852dd8486c7ee0171322950
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Sie benötigen eine Variable, um den geheimen Code während des Dekodierens zu speichern. Entfernen Sie daher `pass` und erstellen Sie innerhalb der Funktion eine Variable namens `secret_code`, der Sie einen leeren String zuweisen.

# --hints--

Sie sollten eine `secret_code`-Variable innerhalb der `pin_extractor`-Funktion haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

Die Variable `secret_code` sollte den Wert eines leeren Strings haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

`pass` sollte nicht mehr in Ihrem Code vorhanden sein.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
