---
id: 6852db755c9c1274d75aae9a
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop erstellen Sie einen Pin-Extraktor, wobei die Pin-Ziffern in jeder Zeile eines Gedichts versteckt sind.

Beginnen Sie damit, eine Funktion namens `pin_extractor` mit einem Parameter `poem` zu erstellen. Eine leere Funktion führt zu einem Fehler, fügen Sie daher vorübergehend `pass` innerhalb der Funktion hinzu, um sie gültig zu machen.

# --hints--

Sie sollten eine Funktion mit dem Namen `pin_extractor` haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

Die Funktion `pin_extractor` sollte einen Parameter namens `poem` haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
