---
id: 68531fbb934d735415053274
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Jetzt, wo Sie einen Funktionsaufruf hinzugefügt haben, können Sie jeden Fehler, der innerhalb der Funktion entsteht, im Terminal sehen, und Sie können jederzeit `print` verwenden, um den Wert von Dingen zu überprüfen.

In jeder Zeile ist eine Ziffer der PIN versteckt, also verwenden Sie innerhalb der Funktion die `split`-Methode, um den String in eine Liste von Zeilen zu teilen. Teilen Sie die Zeilen am Zeilenumbruchzeichen (`\n`) und weisen Sie die resultierende Liste einer Variablen namens `lines` zu.

# --hints--

Sie sollten eine Variable namens `lines` innerhalb Ihrer `pin_extractor`-Funktion haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

Die Variable `lines` sollte den Wert `poem.split('\n')` haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
