---
id: 68530231c05b93324e071104
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Während Sie an der Funktion arbeiten, wäre es hilfreich, einen Funktionsaufruf zum Debuggen zu haben. Erstellen Sie daher eine `poem`-Variable außerhalb der Funktion, die dieses Gedicht enthält:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

Dafür können Sie einen mehrzeiligen String verwenden.

Rufen Sie dann die Funktion mit der `poem`-Variable als Argument auf.

# --hints--

Sie sollten eine globale Variable mit dem Namen `poem` haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

Die `poem`-Variable sollte einen Wert haben, der `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night` entspricht.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

Sie sollten die Funktion `pin_extractor` mit `poem` als Argument aufrufen.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
