---
id: 685535bf527b2c77236b3a9a
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Die n-te Ziffer der PIN ist als Länge des n-ten Wortes in der n-ten Zeile verborgen.

Um die Länge des n-ten Wortes zu ermitteln, ist der nächste Schritt, die Zeile des Gedichts in eine Liste von Wörtern zu trennen: Erstellen Sie innerhalb der Schleife eine Variable `words` und weisen Sie ihr den Wert von `line` zu, der mit der `split`-Methode in Wörter aufgeteilt wurde.

Fügen Sie dann einen `print`-Aufruf mit `words` als Argument hinzu.

# --hints--

Sie sollten eine `words`-Variable innerhalb der Schleife haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

Die `words`-Variable sollte den Wert `line.split()` haben.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

Sie sollten `print(words)` haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
