---
id: 68557bd29f7f3c1880f98836
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Aktualisieren Sie in der letzten Zeile der Schleife den `print`-Aufruf so, dass das Wort aus der `words`-Liste am Index `line_index` ausgegeben wird.

# --hints--

Am Ende der Schleife sollten Sie `print(words[line_index])` haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words[line_index])')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(words)
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
