---
id: 685578aae65e4106ad2b4300
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Wie Sie in einer vorherigen Lektion gelernt haben, ermöglicht Ihnen die `enumerate`-Funktion, den Index beim Durchlaufen eines iterierbaren Objekts zu verfolgen:

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

Die n-te Ziffer des geheimen Codes stammt vom n-ten Wort der n-ten Zeile, daher müssen Sie wissen, welche Nummer jede Zeile hat.

Ändern Sie die Schleife so, dass sie über `enumerate(lines)` iteriert, und fügen Sie vor `line` eine weitere Schleifenvariable namens `line_index` hinzu.

Aktualisieren Sie außerdem den `print`-Aufruf zu `print(line_index, line)`.

# --hints--

Die Schleife sollte über `enumerate(lines)` iterieren.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

Die Schleife sollte `line_index, line` als Schleifenvariable haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

Sie sollten `print(line_index, line)` innerhalb der Schleife haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
