---
id: 68553423d775cc60134ac8bb
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Sie müssen jede Zeile unabhängig voneinander bearbeiten: Erstellen Sie eine Schleife über `lines`, die `line` als Schleifenvariable verwendet. Geben Sie innerhalb der Schleife `line` aus.

# --hints--

Sie sollten eine `for`-Schleife in Ihrer `pin_extractor`-Funktion haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

Die Schleife sollte `lines` durchlaufen.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

Die Schleife sollte `line` als Schleifenvariable haben (`for line in lines`).

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

Sie sollten `print(line)` innerhalb der Schleife haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
