---
id: 68559f3561d7df287586142e
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Jetzt funktioniert der Pin-Extraktor, aber er könnte effizienter sein! Er könnte den Pin aus vielen Gedichten gleichzeitig extrahieren!

Fügen Sie zwei weitere Gedichte im globalen Gültigkeitsbereich hinzu, `poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` und `poem3 = 'There\nonce\nwas\na\ndragon'`, und kommentieren Sie vorerst den `print` mit dem Aufruf von `pin_extractor` aus.


# --hints--

Die Variablen `poem2` und `poem3` sollten existieren und den korrekten Wert haben.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

Sie sollten den `print(pin_extractor(poem))`-Aufruf auskommentieren.

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

--fcc-editable-region--


print(pin_extractor(poem))
--fcc-editable-region--

```
