---
id: 6855918e87c33c6979d21d3f
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Die Funktion extrahiert jetzt gut den Code aus dem Gedicht, aber es gibt einen Sonderfall zu beachten: Die Zeile des Gedichts könnte kürzer als erwartet sein, und das würde dazu führen, dass die Funktion einen Fehler ausgibt (Sie können das ausprobieren, indem Sie das Wort `bright` aus dem Gedicht entfernen).

Setzen Sie die Zeile `secret_code += str(len(words[line_index]))` in eine `if`-Anweisung, die überprüft, ob genügend Wörter in der `words`-Liste vorhanden sind.

# --hints--

Innerhalb der Schleife sollten Sie eine `if`-Anweisung haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

Die Bedingung der if-Anweisung sollte `len(words) > line_index` sein.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

Innerhalb der `if`-Anweisung sollten Sie `secret_code += str(len(words[line_index]))` haben.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
