---
id: 685926db99041c5e619ead41
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Die `len`-Funktion gibt die Länge eines Strings zurück. Aktualisieren Sie den zweiten `print`-Aufruf, um die Länge dieses Wortes auszugeben.

# --hints--

Sie sollten `print(len(words[line_index]))` am Ende der Schleife haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(words[line_index])
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
