---
id: 6859a037b3d6721c7fbb8055
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Erstellen Sie vor der Schleife eine neue Variable `secret_codes`, die als Startwert eine leere Liste hat.

Ersetzen Sie dann die `return secret_code`-Zeile durch eine Zeile, die `secret_code` an die `secret_codes`-Liste anhängt.

# --hints--

Sie sollten eine `secret_codes`-Variable innerhalb der `pin_extractor`-Funktion haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_codes')
    `)
})
```

Die Variable `secret_codes` sollte den Wert einer leeren Liste haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_codes').is_equivalent('secret_codes = []')
    `)
})
```

Am Ende der Schleife sollten Sie `secret_codes.append(secret_code)` haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('secret_codes.append(secret_code)')
    `)
})
```

Sie sollten `secret_code` nicht zurückgeben.


```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_return('secret_code')
    `)
})
```


# --seed--

## --seed-contents--

```py
def pin_extractor(poems):
--fcc-editable-region--

    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        return secret_code
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))

```
