---
id: 646c48df8674cf2b91020ecc
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Você recebeu um boilerplate básico. Comece adicionando um elemento `link` dentro do seu elemento `head`. Para esse elemento `link`, defina o atributo `rel` como `stylesheet` e `href` como `./styles.css`.

# --hints--

Você deve ter apenas um elemento `link`.

```js
assert.exists(document.querySelector("head>link"));
```

O elemento `link` deve ter os atributos `rel` e `href` definidos como `stylesheet` e `./styles.css`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
