---
id: 6140cbeec34e970dfe75e710
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Crea un selettore per il tuo elemento `.wheel`. Inizia impostando la `border` su `2px solid black`, la `border-radius` su `50%` e la `margin-left` su `50px`.

# --hints--

Dovresti avere un selettore `.wheel`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Il tuo selettore `.wheel` dovrebbe avere una proprietà `border` impostata su `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

Il tuo selettore `.wheel` dovrebbe avere una proprietà `border-radius` impostata su `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

Il tuo selettore `.wheel` dovrebbe avere una proprietà `margin-left` impostata su `50px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
