---
id: 6140c7e645d8e905819f1dd4
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Inizia con il codice boilerplate standard. Aggiungi la tua dichiarazione `DOCTYPE`, il tuo elemento `html` con la lingua impostata su inglese, i tuoi elementi `head` e `body`.

Aggiungi il tuo elemento `meta` per il corretto `charset`, il tuo elemento `title` e un elemento `link` per il file `./styles.css`.

Imposta il `title` su `Ferris Wheel`.

# --hints--

Il tuo codice dovrebbe contenere la dichiarazione `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Dovresti includere uno spazio dopo la dichiarazione `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Dovresti definire il tipo di documento come `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un tag di chiusura `>` dopo il tipo.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Il tuo elemento `html` dovrebbe avere un tag di apertura con un attributo `lang` di `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Il tuo elemento `html` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/html\s*>/);
```

La tua dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Dovresti avere un tag di apertura `head`.

```js
assert.match(code, /<head\s*>/i);
```

Dovresti avere un tag di chiusura `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Dovresti avere un tag di apertura `body`.

```js
assert.match(code, /<body\s*>/i);
```

Dovresti avere un tag di chiusura `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

Il tuo elemento `body` dovrebbe essere il primo elemento dopo l'elemento `head`.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

I tuoi elementi `head` e `body` dovrebbero essere fratelli.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Il tuo elemento `head` dovrebbe essere il primo elemento all'interno del tuo elemento `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Il tuo elemento `head` dovrebbe essere all'interno dell'elemento `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Il tuo elemento `body` dovrebbe essere all'interno dell'elemento `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Il tuo codice dovrebbe avere un elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Il tuo elemento `meta` dovrebbe avere un attributo `charset` con il valore `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Il tuo elemento `meta` dovrebbe essere dentro l'elemento `head`.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Il tuo codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Il tuo progetto dovrebbe avere un titolo `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Ricorda, la maiuscola e l'ortografia sono importanti per il titolo.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Il tuo codice dovrebbe avere un elemento `link`.

```js
assert.match(code, /<link/);
```

Dovresti avere un elemento `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Il tuo elemento `link` dovrebbe essere all'interno del tuo elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
