---
id: 56533eb9ac21ba0edf2244b3
title: 섭씨를 화씨로 변환기 만들기
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

이 실습에서는 섭씨 온도를 화씨로 변환하는 함수를 작성할 것입니다. 섭씨를 화씨로 변환하는 공식은 다음과 같습니다:

```js
fahrenheit = celsius * (9/5) + 32
```

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `convertCtoF`라는 이름의 함수를 만들어야 합니다.
2. `convertCtoF`는 섭씨 온도를 나타내는 하나의 숫자 인수를 받아야 합니다.
3. `convertCtoF`는 숫자를 반환해야 합니다.

# --hints--

`convertCtoF`라는 이름의 함수를 만들어야 합니다.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF`는 하나의 매개변수를 받아야 합니다.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)`는 `-22` 값을 반환해야 합니다.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)`는 `14` 값을 반환해야 합니다.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)`는 `32` 값을 반환해야 합니다.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)`는 `68` 값을 반환해야 합니다.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)`는 `86` 값을 반환해야 합니다.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
