---
id: 67172b43f84bcd2dec238a3d
title: crea un organizzatore di libri
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un array di oggetti chiamato `books` dove ogni oggetto nell'array deve avere una stringa `title`, un'altra stringa `authorName` e un numero `releaseYear`.

2. Il tuo array `books` dovrebbe contenere almeno tre oggetti.

3. Dovresti avere una funzione di callback chiamata `sortByYear` che accetta due libri come parametri per ordinare l'array.

4. La funzione `sortByYear` dovrebbe restituire `-1` se il `releaseYear` del primo libro è minore di quello del secondo libro.

5. La funzione `sortByYear` dovrebbe restituire `1` se il `releaseYear` del primo libro è maggiore di quello del secondo libro.

6. La funzione `sortByYear` dovrebbe restituire `0` se entrambi i valori `releaseYear` sono uguali.

7. Dovresti filtrare i libri scritti dopo un certo anno, ad esempio il 1950, dall'array `books` e salvare l'array filtrato in un nuovo array chiamato `filteredBooks`.

8. Dovresti ordinare i libri nell'array `filteredBooks` in base al loro `releaseYear` in ordine crescente. Hai imparato in una lezione precedente che il metodo `sort()` ordina l'array in loco. Questo significa che l'array `filteredBooks` verrà modificato.

# --hints--

Dovresti avere una funzione `sortByYear` nel tuo codice.

```js
assert.isFunction(sortByYear);
```

La tua funzione `sortByYear` dovrebbe prendere due parametri.

```js
assert.lengthOf(sortByYear, 2);
```

La tua funzione `sortByYear` dovrebbe restituire `-1` se il `releaseYear` dell'oggetto `book1` è minore di quello dell'oggetto `book2`, `1` se il `releaseYear` dell'oggetto `book1` è maggiore di quello dell'oggetto `book2`, e `0` in tutti gli altri casi.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

Dovresti avere un array `books` nel tuo codice.

```js
assert.isArray(books)
```

Il tuo array `books` dovrebbe contenere almeno tre oggetti libro.

```js
assert.isAtLeast(books.length, 3)
```

Il tuo array `books` dovrebbe contenere oggetti ciascuno con una stringa `title`, un'altra stringa `authorName` e un numero `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

Dovresti avere un array `filteredBooks` nel tuo codice.

```js
assert.isArray(filteredBooks)
```

L'array `filteredBooks` dovrebbe contenere alcuni dei libri filtrati dall'array `books` e non essere vuoto.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

L'array `filteredBooks` dovrebbe contenere solo libri dall'array originale `books`.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

L'array `filteredBooks` dovrebbe includere solo libri pubblicati in o prima di un anno specificato.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

Dovresti chiamare la funzione di ordine superiore `sort` passando la funzione di callback `sortByYear` sull'array `filteredBooks`.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
