---
id: 69b279bae0f20eda63163cf8
title: Schritt 14
challengeType: 1
dashedName: step-14
---

# --description--

Die `if`- und `else`-Anweisungen, die Sie in den vorherigen Schritten erstellt haben, weisen Elemente dem sortierten Array zu.

Jedes Element, das dem sortierten Array zugewiesen wird, belegt einen Index in der Liste. Daher müssen Sie nach jeder Zuweisung zum nächsten Index im sortierten Array wechseln.

Erhöhen Sie unterhalb des `if`/`else`-Blocks, aber noch innerhalb der `while`-Schleife, `sortedIndex` um `1`.

# --hints--

Sie sollten den aktuellen Wert von `sortedIndex` nach Ihrer `1`-Klausel um `else` erhöhen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
