---
id: 69b279bae0f20eda63163cf6
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

In einem vorherigen Schritt haben Sie das Element in der `leftPart` der `array` Liste zugewiesen, weil es kleiner war. Aber das ist nicht immer der Fall. In einigen Vergleichsfällen könnte das Element rechts kleiner sein.

Erstellen Sie eine `else` Klausel, die ausgeführt wird, wenn das Element in `leftPart` nicht kleiner ist als das Element in `rightPart`.

Weisen Sie innerhalb des `else` Blocks `rightPart[rightArrayIndex]` `array[sortedIndex]` zu.

# --hints--

Sie sollten einen `else` Block erstellen, der `array[sortedIndex] = rightPart[rightArrayIndex]` enthält.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
