---
id: 69b279bae0f20eda63163cf1
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Jetzt ist es Zeit, die Listen (`leftPart` und `rightPart`) zu sortieren und in die ursprĂĽngliche `array` zusammenzufĂĽhren.

Sie kĂ¶nnen dies tun, indem Sie die Elemente beider Listen vergleichen und das kleinere Element in die Hauptliste ĂĽbernehmen. Diesen Vergleich fĂĽhren Sie fĂĽr alle Indizes in `leftPart` und `rightPart` durch.

Erstellen Sie drei Variablen: `leftArrayIndex`, `rightArrayIndex` und `sortedIndex` und setzen Sie deren Werte auf `0`. Diese Variablen helfen Ihnen, wĂ¤hrend des Sortierprozesses jeden Index im Blick zu behalten.

# --hints--

Sie sollten eine `leftArrayIndex`-Variable deklarieren, die mit `0` initialisiert wird.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

Sie sollten eine `rightArrayIndex`-Variable deklarieren, die mit `0` initialisiert wird.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

Sie sollten eine `sortedIndex`-Variable deklarieren, die mit `0` initialisiert wird.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--
  
--fcc-editable-region--
}
```
