---
id: 69b279bae0f20eda63163cf9
title: Schritt 15
challengeType: 1
dashedName: step-15
---

# --description--

Die `while`-Schleife, die Sie erstellt haben, vergleicht ein Element aus `leftPart` mit einem anderen in `rightPart` und fügt dann das kleinere Element zur Hauptliste `array` hinzu.

Diese Operation wird so lange fortgesetzt, bis keine Elemente mehr zum Vergleichen übrig sind. Aber `leftPart` kann noch Elemente enthalten, während `rightPart` keine mehr hat, und umgekehrt.

Erstellen Sie eine weitere `while`-Schleife, die ausgeführt wird, wenn `leftArrayIndex` kleiner ist als die Länge von `leftPart`. In den nächsten Schritten verwenden Sie sie, um die verbleibenden Elemente in `leftPart` in die Liste `array` zu kopieren.

# --hints--

Sie sollten eine `while`-Schleife mit der Bedingung `leftArrayIndex < leftPart.length` erstellen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
