---
id: 69b279bae0f20eda63163cf2
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Als Nächstes müssen Sie eine Schleife erstellen, die so lange läuft, wie in beiden `leftPart` und `rightPart` noch Elemente vorhanden sind.

Erstellen Sie dafür eine `while`-Schleife mit zwei Bedingungen: eine, die überprüft, ob `leftArrayIndex` kleiner als die Länge von `leftPart` ist, und eine weitere Bedingung, die überprüft, ob `rightArrayIndex` kleiner als die Länge von `rightPart` ist.

# --hints--

Sie sollten eine `while`-Schleife erstellen, die überprüft, ob `leftArrayIndex` kleiner als die Länge von `leftPart` ist und ob `rightArrayIndex` kleiner als die Länge von `rightPart` ist.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
