---
id: 69b278519ef46b2fd8bf1546
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Sie haben bereits einen `middlePoint`-Wert. Verwenden Sie ihn, um das Array in zwei Hälften zu teilen, und halten Sie den linken Teil zum Sortieren bereit.

Verwenden Sie die Slice-Syntax, um die linke Hälfte von `array` zu extrahieren und weisen Sie sie einer Variablen namens `leftPart` zu.

# --hints--

Sie sollten eine `leftPart`-Variable innerhalb der `mergeSort`-Funktion deklarieren.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

Sie sollten `array.slice(0, middlePoint)` `leftPart` zuweisen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
