---
id: 69b2713ff6c21fed2d3ad88b
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

Merge-Sort ist ein Sortieralgorithmus, der das Prinzip „Teile und herrsche“ verwendet, um Datensammlungen zu sortieren. Das bedeutet, er „teilt“ eine Sammlung in kleinere Unterteile auf, „beherrscht“ die Unterteile, indem er sie unabhängig voneinander sortiert, und führt dann die sortierten Unterteile zusammen.

Beginnen Sie damit, eine Funktion namens `mergeSort` zu erstellen, die den Parameter `array` entgegennimmt.

# --hints--

Sie sollten eine Funktion namens `mergeSort` mit einem einzelnen Parameter `array` deklarieren.

```js
assert.isFunction(mergeSort);

const explorer = await __helpers.Explorer(code);
const mergeSortFn = explorer.allFunctions.mergeSort;
assert.exists(mergeSortFn, 'You should declare a mergeSort function');

const params = mergeSortFn.parameters;
assert.lengthOf(
  params,
  1,
  'mergeSort should take a single parameter'
);
assert.isTrue(
  params[0].matches('array'),
  'The parameter of mergeSort should be named array'
);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
