---
id: 69b279bae0f20eda63163cfe
title: Schritt 20
challengeType: 1
dashedName: step-20
---

# --description--

Bevor Sie die `mergeSort()`-Funktion testen, müssen Sie einen Basisfall erstellen, der die Ausführung der Funktion stoppt, wenn die Länge von `array` kleiner oder gleich `1` ist.

Dieser Basisfall stoppt den Rekursionsaufruf. Ohne ihn würde der Merge-Sort-Vorgang weiterlaufen, selbst wenn die Liste bereits sortiert ist oder keine Elemente enthält.

Direkt nach der Deklaration der Funktion erstellen Sie eine `if`-Anweisung, die überprüft, ob `array.length` kleiner oder gleich 1 ist. Innerhalb des Körpers der `if`-Anweisung fügen Sie eine `return`-Anweisung hinzu, um die Ausführung der Funktion zu stoppen.

# --hints--

Sie sollten eine `if`-Anweisung mit der Bedingung `array.length <= 1` erstellen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

Sie sollten eine `return`-Anweisung im Körper der `if`-Anweisung haben.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }
  
  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
