---
id: 69b279bae0f20eda63163cf0
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Rufen Sie am Ende Ihres Funktionskörpers die Funktion `mergeSort()` erneut auf. Übergeben Sie diesmal `rightPart` als Argument beim Funktionsaufruf.

# --hints--

Sie sollten `mergeSort()` mit dem Argument `rightPart` am Ende Ihres Funktionskörpers aufrufen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
