---
id: 69b2762b17d23fbee40a6c00
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Der Merge-Sort-Algorithmus führt drei Aktionen aus:

- Teilen einer unsortierten Folge von Elementen in Unterteile
- Sortieren der Elemente in den Unterteilen
- Zusammenführen der sortierten Unterteile

Das oben Genannte geschieht rekursiv, bis die Unterteile zu der vollständig sortierten Folge zusammengeführt sind. Beginnen Sie mit dem Teilen der Folge.

Deklarieren Sie zuerst innerhalb der `mergeSort`-Funktion eine Variable `middlePoint` mit `const` und weisen Sie ihr `Math.floor(array.length / 2)` zu, damit Sie die Folge vor der Rekursion halbieren können.

# --hints--

Sie sollten eine `middlePoint`-Variable deklarieren und ihr `Math.floor(array.length / 2)` zuweisen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
