---
id: 69b279bae0f20eda63163cfd
title: Schritt 19
challengeType: 1
dashedName: step-19
---

# --description--

Zum Schluss sollten Sie alle verbleibenden Werte aus der rechten Hälfte in `array` kopieren.

Fügen Sie dazu eine `while`-Schleife hinzu, die so lange läuft, wie `rightArrayIndex` kleiner als `rightPart.length` ist, `rightPart[rightArrayIndex]` in `array[sortedIndex]` zuweist und sowohl `rightArrayIndex` als auch `sortedIndex` inkrementiert.

# --hints--

Sie sollten eine `while`-Schleife mit der Bedingung `rightArrayIndex < rightPart.length` erstellen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

Der Schleifenkörper sollte `rightPart[rightArrayIndex]` in `array[sortedIndex]` verschieben.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?/
);
```

Sie sollten sowohl `rightArrayIndex` als auch `sortedIndex` nach der Zuweisung inkrementieren.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*rightArrayIndex\s*(\+\+|\+=\s*1)\s*;?[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
