---
id: 69b279bae0f20eda63163cf3
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Innerhalb der Schleife beginnen Sie damit, die aktuellen Elemente aus jeder Hälfte zu vergleichen, damit Sie zuerst den kleineren Wert zusammenführen können.

Fügen Sie einen `if`-Block hinzu, der überprüft, ob das aktuelle Element in `leftPart` kleiner ist als das aktuelle Element in `rightPart`. 
# --hints--

Die `if`-Bedingung sollte prüfen, ob `leftPart[leftArrayIndex]` kleiner ist als `rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
