---
id: 69b278d82f0e514f32acb724
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Nachdem Sie die linke Hälfte ausgeschnitten haben, benötigen Sie noch die rechte Seite, um den Zusammenführungsschritt abzuschließen.

Verwenden Sie die Slice-Syntax, um die rechte Hälfte von `array` zu extrahieren und weisen Sie sie einer Variablen namens `rightPart` zu.

# --hints--

Sie sollten eine `rightPart`-Variable innerhalb der `mergeSort`-Funktion haben.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

Sie sollten die Slice-Syntax verwenden, um die rechte Hälfte des Arrays zu extrahieren und sie der `rightPart`-Variablen zuzuweisen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
