---
id: 69b279bae0f20eda63163cf4
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Wenn die `if` Bedingung zu `true` ausgewertet wird, bedeutet das, dass das Element in der `leftPart` Liste kleiner ist als das Element, mit dem es in der `rightPart` Liste verglichen wird.

In diesem Fall können Sie das an `leftArrayIndex` gefundene Element an `leftPart` innerhalb des sortierten Arrays einfügen.

Innerhalb des `if` Blocks weisen Sie `leftPart[leftArrayIndex]` `array[sortedIndex]` zu.

# --hints--

Sie sollten `leftPart[leftArrayIndex]` innerhalb des `array[sortedIndex]` an `if` zuweisen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
