---
id: 69b279bae0f20eda63163cef
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Nachdem Sie die `array` Liste in zwei separate Listen aufgeteilt haben, teilen Sie jede Liste weiter auf, bis jedes Element in seiner eigenen Liste steht. Eine Liste mit nur einer Zahl ist immer sortiert.

Dazu rufen Sie rekursiv `mergeSort` innerhalb Ihrer Funktion auf und übergeben `leftPart` als Argument an den Aufruf.

# --hints--

Sie sollten `mergeSort(leftPart)` innerhalb derselben `mergeSort` Funktion aufrufen.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
