---
id: 69b279bae0f20eda63163cf5
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Nachdem Sie das Element in `leftPart` dem sortierten Array zugewiesen haben, erhöhen Sie `leftArrayIndex` um `1`.

# --hints--

Sie sollten den aktuellen Wert von `leftArrayIndex` erhöhen, nachdem Sie das Element in das sortierte Array eingefügt haben.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftArrayIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
