---
id: 69b2713ff6c21fed2d3ad88b
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

Merge-Sort ist ein Sortieralgorithmus, der das Prinzip „Teile und herrsche“ verwendet, um Datensammlungen zu sortieren. Das bedeutet, er „teilt“ eine Sammlung in kleinere Unterteile auf, „beherrscht“ die Unterteile, indem er sie unabhängig voneinander sortiert, und führt dann die sortierten Unterteile zusammen.

Beginnen Sie damit, eine Funktion namens `mergeSort` zu erstellen, die den Parameter `array` entgegennimmt.

# --hints--

Sie sollten eine Funktion namens `mergeSort` mit einem einzelnen Parameter `array` deklarieren.

```js
assert.isFunction(mergeSort);
const regex = __helpers.functionRegex('mergeSort', ['array']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
