---
id: 69b279bae0f20eda63163cf1
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Jetzt ist es Zeit, die Listen (`leftPart` und `rightPart`) zu sortieren und in die ursprüngliche `array` zusammenzuführen.

Sie können dies tun, indem Sie die Elemente beider Listen vergleichen und das kleinere Element in die Hauptliste übernehmen. Diesen Vergleich führen Sie für alle Indizes in `leftPart` und `rightPart` durch.

Erstellen Sie drei Variablen: `leftArrayIndex`, `rightArrayIndex` und `sortedIndex` und setzen Sie deren Werte auf `0`. Diese Variablen helfen Ihnen, während des Sortierprozesses jeden Index im Blick zu behalten.

# --hints--

Sie sollten eine `leftArrayIndex`-Variable deklarieren, die mit `0` initialisiert wird.

```js
const explorer = await __helpers.Explorer(code);
const { leftArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftArrayIndex,
  'You should declare a leftArrayIndex variable inside mergeSort'
);
assert.isTrue(
  leftArrayIndex.value.matches('0'),
  'leftArrayIndex should be initialized to 0'
);
```

Sie sollten eine `rightArrayIndex`-Variable deklarieren, die mit `0` initialisiert wird.

```js
const explorer = await __helpers.Explorer(code);
const { rightArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightArrayIndex,
  'You should declare a rightArrayIndex variable inside mergeSort'
);
assert.isTrue(
  rightArrayIndex.value.matches('0'),
  'rightArrayIndex should be initialized to 0'
);
```

Sie sollten eine `sortedIndex`-Variable deklarieren, die mit `0` initialisiert wird.

```js
const explorer = await __helpers.Explorer(code);
const { sortedIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  sortedIndex,
  'You should declare a sortedIndex variable inside mergeSort'
);
assert.isTrue(
  sortedIndex.value.matches('0'),
  'sortedIndex should be initialized to 0'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
