---
id: 68507cadbf36aa089a84ce2d
title: Implementieren Sie den Luhn-Algorithmus
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

Der Luhn-Algorithmus, auch bekannt als „Modulus 10“ oder „mod 10“-Algorithmus, ist eine einfache Prüfsummenformel, die verwendet wird, um verschiedene Identifikationsnummern wie Kreditkartennummern zu validieren. Dies sind die Schritte, um eine Nummer mit dem Luhn-Algorithmus zu validieren:

- Beginnen Sie von rechts und schließen Sie die ganz rechte Ziffer (die Prüfziffer) aus, verdoppeln Sie den Wert jeder zweiten Ziffer.
- Wenn das Ergebnis der Verdopplung einer Ziffer größer als `9` ist, summieren Sie die Ziffern, um eine einzelne Ziffer zu erhalten. Alternativ können Sie `9` vom Ergebnis subtrahieren.
- Nehmen Sie die Summe aller Ziffern einschließlich der Prüfziffer.
- Wenn die Summe aller Ziffern ein Vielfaches von `10` ist, dann ist die Nummer gültig; andernfalls ist sie ungültig.

Zum Beispiel betrachten wir die Nummer `453914881`. Die Schritte, um sie mit dem Luhn-Algorithmus zu validieren, sind:

 ```md 
 Kontonummer        4   5   3   9   1   4   8   8   1 
 Verdopple jede zweite  4  10   3  18   1   8   8  16   1 
 Summe der 2-stelligen Ziffern   4   1   3   9   1   8   8   7   1 
 ``` 

- Dann summieren Sie alle Zahlen, `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`.
- Da `42` kein Vielfaches von `10` ist, ist die Nummer ungültig.

In diesem Labor bauen Sie einen Kreditkarten-Validator mit dem Luhn-Algorithmus.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `verify_card_number` definieren, die einen String von Ziffern (der eine Kartennummer darstellt) entgegennimmt und überprüft, ob sie gemäß dem Luhn-Algorithmus gültig ist.

2. Innerhalb der Funktion `verify_card_number`:

   - Sie sollten alle Bindestriche oder Leerzeichen behandeln, die in der übergebenen Kartennummer vorhanden sein können.
   - Geben Sie `VALID!` zurück, wenn die Kartennummer gültig ist; andernfalls geben Sie `INVALID!` zurück.

Wenn Sie das Projekt abschließen, sollten Sie je nach Eingabe folgende Meldungen sehen:

| Kartennummer         | Meldung  |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |


# --hints--

Sie sollten eine Funktion namens `verify_card_number` haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` sollte `VALID!` zurückgeben.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` sollte `VALID!` zurückgeben.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` sollte `INVALID!` zurückgeben.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` sollte `INVALID!` zurückgeben.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number` sollte `VALID!` zurückgeben, wenn es mit einer gültigen Kreditkartennummer aufgerufen wird.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number` sollte `INVALID!` zurückgeben, wenn es mit einer ungültigen Kreditkartennummer aufgerufen wird.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
