---
id: a0b5010f579e69b815e7c5d6
title: 建置一個智慧字替換函式（程式）
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

在本實驗中，你將創建一個函式（程式），該函式使用提供的引數對句子執行搜尋和取代，然後傳回新的句子。

**注意：** 當你替換字時，請保留原始字的第一個字元大小寫。例如，若你要將字 `Book` 替換成字 `dog`，則應替換為 `Dog`。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `myReplace` 的函式（程式）。
1. `myReplace` 函式（程式）應該接受三個引數：一個字串、一個要被取代的字，以及用來取代它的字。
1. `myReplace` 函式（程式）應該傳回取代後的新字串。
1. 你應該在替換時保留原始字的第一個字元大小寫。

# --hints--

你應該創建一個名為 `myReplace` 的函式（程式）。

```js
assert.isFunction(myReplace);
```

`myReplace` 應該接受三個引數。

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` 應該傳回字串 `Let us go to the mall`。

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` 應該傳回字串 `He is Sitting on the couch`。

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` 應該傳回字串 `I think we should look down there`。

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` 應該傳回字串 `This has a spelling error`。

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` 應該傳回字串 `His name is John`。

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` 應該傳回字串 `Let us get back to more Algorithms`。

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
