---
id: 614385513d91ae5c251c2052
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستنشئ صفحة مجلة باستخدام بيان صحفي من إصدار سابق من منهج freeCodeCamp.

ابدأ بإضافة عنصر `link` لورقة أنماط الخطوط من `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`، وعنصر `link` لورقة أنماط FontAwesome من `https://use.fontawesome.com/releases/v5.8.2/css/all.css`، وعنصر `link` لورقة أنماط `./styles.css` الخاصة بك.

ستحمّل ورقة أنماط الخطوط الخاصة بك ثلاثة خطوط منفصلة: `Anton`، و`Baskervville`، و`Raleway`.

# --hints--

يجب أن يحتوي شفرتك على ثلاثة عناصر `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

يجب أن تكون عناصر `link` داخل عنصر `head` الخاص بك.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

يجب أن تحتوي عناصر `link` الثلاثة على خاصية `rel` بقيمة `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

يجب أن يحتوي أحد عناصر الرابط على خاصية `href` مضبوطة على `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

يجب أن يحتوي أحد عناصر الرابط على خاصية `href` مضبوطة على `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

يجب أن يحتوي أحد عناصر `link` على خاصية `href` مضبوطة على `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
