---
id: 614e0e588f0e8a772a8a81a6
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

تأخذ صورتك حاليًا مساحة كبيرة. لرؤية ما تعمل عليه بشكل أفضل، أضف خاصية `width` إلى عنصر `img`، بقيمة `400`.

ستزيل هذا لاحقًا بعد أن تعمل على CSS.

# --hints--

يجب أن يحتوي عنصر `img` على خاصية `width` مضبوطة على `400`.

```js
assert.equal(document.querySelector('img')?.getAttribute('width'), '400');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
--fcc-editable-region--
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
--fcc-editable-region--
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
      </section>
    </main>
  </body>
</html>
```

```css

```
