---
id: 6165d3b702a5d92ad970b30c
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

بعد عنصر `img`، أضف عنصر `h1` مع تعيين `class` إلى `hero-title` والنص إلى `OUR NEW CURRICULUM`، يتبعه عنصر `p` مع تعيين `class` إلى `hero-subtitle` والنص إلى `Our efforts to restructure our curriculum with a more project-based focus`.

# --hints--

يجب إنشاء عنصر `h1`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يأتي عنصر `h1` بعد عنصر `img`.

```js
assert.equal(document.querySelector('h1')?.previousElementSibling?.localName, 'img');
```

يجب أن يحتوي عنصر `h1` على `class` مضبوط إلى `hero-title`.

```js
assert.equal(document.querySelector('h1')?.className, 'hero-title');
```

يجب أن يحتوي عنصر `h1` على النص `OUR NEW CURRICULUM`.

```js
assert.equal(document.querySelector('h1')?.textContent, 'OUR NEW CURRICULUM');
```

يجب إنشاء عنصر `p` جديد.

```js
assert.exists(document.querySelector('p'));
```

يجب أن يأتي عنصر `p` بعد عنصر `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

يجب أن يحتوي عنصر `p` على `class` مضبوط إلى `hero-subtitle`.

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

يجب أن يحتوي عنصر `p` على النص `Our efforts to restructure our curriculum with a more project-based focus`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </header>
      </section>
    </main>
  </body>
</html>
```

```css

```
