---
id: 66ed41f912d0bb1dc62da5dd
title: Crie um bot informativo em JavaScript
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve registrar `"Hello! I'm your coding fun fact guide!"` no console como uma mensagem de saudação para o usuário.
1. Você deve criar três variáveis: `botName`, `botLocation` e `favoriteLanguage`, que armazenam o nome do bot, de onde ele é e sua linguagem de programação favorita, respectivamente.
1. Você deve registrar `"My name is (botName) and I live on (botLocation)."` no console.
1. Você deve registrar `"My favorite programming language is (favoriteLanguage)."` no console.
1. Você deve usar `let` para criar uma variável `codingFact` e atribuir a ela uma string que seja um fato divertido sobre a linguagem de programação favorita do seu bot e incluir o uso da variável `favoriteLanguage`.
1. Você deve registrar o `codingFact` no console.
1. Você deve reatribuir a variável `codingFact` para um novo fato sobre a linguagem favorita do bot usando novamente a variável `favoriteLanguage`.
1. Você deve registrar o `codingFact` no console novamente.
1. Você deve reatribuir a variável `codingFact` novamente para outro fato novo sobre a linguagem favorita do bot usando a variável `favoriteLanguage`.
1. Você deve registrar o `codingFact` no console pela terceira vez.
1. Você deve registrar `"Foi divertido compartilhar esses fatos com você. Goodbye! - (botName) from (botLocation)."` no console como uma declaração de despedida do bot.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar `"Olá!"`Eu sou seu guia de curiosidades sobre programação!"` para o console.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

Você deve declarar uma variável `botName`. Verifique duas vezes se há erros de ortografia ou de maiúsculas e minúsculas.

```js
assert.exists(botName);
```

A variável `botName` deve ser uma string.

```js
assert.isString(botName);
```

Você deve declarar uma variável `botLocation`. Verifique duas vezes se há erros de ortografia ou maiúsculas.

```js
assert.exists(botLocation);
```

Sua variável `botLocation` deve ser uma string.

```js
assert.isString(botLocation);
```

Você deve declarar uma variável `favoriteLanguage`. Verifique duas vezes se há erros de ortografia ou de maiúsculas e minúsculas.

```js
assert.exists(favoriteLanguage);
```

Sua variável `favoriteLanguage` deve ser uma string.

```js
assert.isString(favoriteLanguage);
```

Você deve registrar no console `"Meu nome é (botName) e eu moro em (botLocation)."` e adicionar as variáveis à string.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

Você deve registrar no console `"Minha linguagem de programação favorita é (favoriteLanguage)."` e adicionar a variável à string.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

Você deve usar `let` para declarar uma nova variável `codingFact`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

Você deve atribuir a `codingFact` um valor que inclua `favoriteLanguage`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

Você deve registrar `codingFact` no console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

Você deve atribuir um novo valor a `codingFact` que também contenha `favoriteLanguage`, e registrá-lo no console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

Você deve atribuir um valor a `codingFact` pela terceira vez que também contenha `favoriteLanguage`, e registrá-lo no console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

Você deve registrar no console `"Foi divertido compartilhar esses fatos com você."`Adeus! - (botName) de (botLocation).

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
