---
id: 66ed41f912d0bb1dc62da5dd
title: Crie um bot informativo em JavaScript
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratĂ³rio.

**HistĂ³rias de UsuĂ¡rio:**

1. VocĂª deve registrar `"Hello! I'm your coding fun fact guide!"` no console como uma mensagem de saudaĂ§Ă£o para o usuĂ¡rio.
1. VocĂª deve criar trĂªs variĂ¡veis: `botName`, `botLocation` e `favoriteLanguage`, que armazenam o nome do bot, de onde ele Ă© e sua linguagem de programaĂ§Ă£o favorita, respectivamente.
1. VocĂª deve registrar `"My name is (botName) and I live on (botLocation)."` no console.
1. VocĂª deve registrar `"My favorite programming language is (favoriteLanguage)."` no console.
1. VocĂª deve usar `let` para criar uma variĂ¡vel `codingFact` e atribuir a ela uma string que seja um fato divertido sobre a linguagem de programaĂ§Ă£o favorita do seu bot e incluir o uso da variĂ¡vel `favoriteLanguage`.
1. VocĂª deve registrar o `codingFact` no console.
1. VocĂª deve reatribuir a variĂ¡vel `codingFact` para um novo fato sobre a linguagem favorita do bot usando novamente a variĂ¡vel `favoriteLanguage`.
1. VocĂª deve registrar o `codingFact` no console novamente.
1. VocĂª deve reatribuir a variĂ¡vel `codingFact` novamente para outro fato novo sobre a linguagem favorita do bot usando a variĂ¡vel `favoriteLanguage`.
1. VocĂª deve registrar o `codingFact` no console pela terceira vez.
1. VocĂª deve registrar `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` no console como uma declaraĂ§Ă£o de despedida do bot.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

VocĂª deve registrar `"Hello! I'm your coding fun fact guide!"` para o console.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

VocĂª deve declarar uma variĂ¡vel `botName`. Verifique duas vezes se hĂ¡ erros de ortografia ou de maiĂºsculas e minĂºsculas.

```js
assert.exists(botName);
```

A variĂ¡vel `botName` deve ser uma string.

```js
assert.isString(botName);
```

VocĂª deve declarar uma variĂ¡vel `botLocation`. Verifique duas vezes se hĂ¡ erros de ortografia ou maiĂºsculas.

```js
assert.exists(botLocation);
```

Sua variĂ¡vel `botLocation` deve ser uma string.

```js
assert.isString(botLocation);
```

VocĂª deve declarar uma variĂ¡vel `favoriteLanguage`. Verifique duas vezes se hĂ¡ erros de ortografia ou de maiĂºsculas e minĂºsculas.

```js
assert.exists(favoriteLanguage);
```

Sua variĂ¡vel `favoriteLanguage` deve ser uma string.

```js
assert.isString(favoriteLanguage);
```

VocĂª deve registrar no console `"My name is (botName) and I live on (botLocation)."` e adicionar as variĂ¡veis Ă  string.

```js
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

VocĂª deve registrar no console `"My favorite programming language is (favoriteLanguage)."` e adicionar a variĂ¡vel Ă  string.

```js
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

VocĂª deve usar `let` para declarar uma nova variĂ¡vel `codingFact`.

```js
const explorer = await __helpers.Explorer(code);
const { codingFact } = explorer.variables;
assert.isTrue(codingFact?.matches(`let codingFact = ${codingFact?.value?.toString()}`));
```

VocĂª deve atribuir a `codingFact` um valor que inclua `favoriteLanguage`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.codingFact);
assert.include(String(codingFact), favoriteLanguage);
```

VocĂª deve registrar `codingFact` no console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

VocĂª deve atribuir um novo valor a `codingFact` que tambĂ©m contenha `favoriteLanguage`, e registrĂ¡-lo no console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.exists(loggingCodingFacts);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
```

VocĂª deve atribuir um valor a `codingFact` pela terceira vez que tambĂ©m contenha `favoriteLanguage`, e registrĂ¡-lo no console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.lengthOf(loggingCodingFacts, 3);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.isNotEmpty(codingFact);
```

VocĂª deve registrar no console `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` e adicione os valores das variĂ¡veis.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
