---
id: 6720b3191bbafd269bae0ae8
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Erstellen Sie ein `div`-Element, das eine `id` von `greeting-card` und eine `class` von `card` hat.

Fügen Sie innerhalb des `div`-Elements ein `h1` mit dem Text `Happy Birthday!` hinzu. Fügen Sie dann ein Absatz-Element mit einer `class` namens `message` und dem Text `Wishing you all the happiness and joy on your special day!` hinzu.

# --hints--

Sie sollten ein `div`-Element haben.

```js
assert.exists(document.querySelector("div"));
```

Das `div`-Element sollte eine `id` von `greeting-card` haben.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

Das `div`-Element sollte eine `class` von `card` haben.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

Sie sollten ein `h1`-Element innerhalb des `div`-Elements haben.

```js
assert.exists(document.querySelector("div > h1"));
```

Das `h1`-Element sollte den Text `Happy Birthday!` enthalten.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

Sie sollten ein `p`-Element innerhalb des `div` haben.

```js
assert.exists(document.querySelector("div > p"));
```

Das `p`-Element sollte eine `class` von `message` haben.

```js
assert.exists(document.querySelector("div > p.message"));
```

Das `p`-Element sollte den Text `Wishing you all the happiness and joy on your special day!` enthalten.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
