---
id: 6720b7f3155a9278d6b6a8b8
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Jetzt ist es Zeit, Ihre Grußkarte zu gestalten.

Fügen Sie einen Selektor für das `body`-Element hinzu, und dann:

- ändern Sie das `font-family` so, dass es `Arial` gefolgt vom generischen `sans-serif` ist,
- geben Sie auf allen Seiten einen Innenabstand von `40px`,
- setzen Sie die `text-align`-Eigenschaft auf `center`.

# --hints--

Sie sollten einen `body`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Sie sollten `font-family` im `"Arial", sans-serif`-Selektor auf `body` setzen.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

Sie sollten `padding` im `40px`-Selektor auf `body` setzen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

Sie sollten `text-align` im `center`-Selektor auf `body` setzen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
