---
id: 6720b1f2680ec706ed534acf
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop üben Sie den Umgang mit verschiedenen Pseudoklassen und Pseudoelementen, indem Sie eine Grußkarte gestalten. Der Boilerplate-Code in HTML wurde für Sie bereitgestellt.

Beginnen Sie den Workshop, indem Sie die `styles.css`-Datei verlinken.

# --hints--

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
