---
id: 672b4434dd6d6a1edc405419
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Fügen Sie zwei `a`-Elemente innerhalb des `.card-links`-Elements hinzu.

Das erste sollte einen Text von `Send Card`, einen `class` von `send-link` und einen `href` von `#send` haben.

Das zweite sollte einen Text von `Share on Social Media`, einen `class` von `share-link` und einen `href` von `#share` haben.

# --hints--

Das erste `a`-Element innerhalb von `.card-links` sollte einen Text von `Send Card` haben.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

Das erste `a`-Element sollte ein `class` mit dem Wert `send-link` haben.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

Das erste `a`-Element sollte ein `href` mit dem Wert `#send` haben.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

Das zweite `a`-Element innerhalb von `.card-links` sollte einen Text von `Share on Social Media` haben.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

Das zweite `a`-Element sollte ein `class` mit dem Wert `share-link` haben.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

Das zweite `a`-Element sollte ein `href` mit dem Wert `#share` haben.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
