---
id: 6733352ce2a5642827b4f7f0
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Fügen Sie einen `.card-links`-Selektor hinzu und setzen Sie die `margin-top`-Eigenschaft auf `20px`.

Sie können `display: flex` hinzufügen, um ein Element mit Flexbox zu gestalten. Mehr über Flexbox lernen Sie später im Kurs; betrachten Sie dies als kleine Vorschau.

Um die beiden Links so zu verteilen, dass sie den gleichen Abstand um sich herum haben, fügen Sie eine `display`-Eigenschaft hinzu, die auf `flex` gesetzt ist, und eine `justify-content`, die auf `space-around` gesetzt ist.

# --hints--

Sie sollten einen `.card-links`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".card-links"));
```

Der `.card-links`-Selektor sollte eine `margin-top`-Eigenschaft haben, die auf `20px` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("margin-top"), "20px");
```

Der `.card-links`-Selektor sollte eine `display`-Eigenschaft haben, die auf `flex` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("display"), "flex");
```

Der `.card-links`-Selektor sollte eine `justify-content`-Eigenschaft haben, die auf `space-around` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("justify-content"), "space-around");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

--fcc-editable-region--

--fcc-editable-region--
```
