---
id: 67332db0f793bf08f2a18528
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Zeit, das zweite `section` auszufüllen!

Fügen Sie dem `h2`-Element ein `#share`-Element hinzu, das den Text `Sharing your card...` enthält, und fügen Sie dann ein `p`-Element mit dem Text `Your card was shared on social media!` hinzu.

# --hints--

Das `#share`-Element sollte ein `h2`-Element enthalten.

```js
assert.exists(document.querySelector("#share > h2"));
```

Das `h2`-Element sollte den Text `Sharing your card...` enthalten.

```js
assert.strictEqual(document.querySelector("#share > h2").innerText.trim(), "Sharing your card...");
```

Das zweite Element innerhalb von `#share` sollte ein `p`-Element sein.

```js
assert.strictEqual(document.querySelector("#share").children[1].tagName, "P");
```

Das `p`-Element sollte den Text `Your card was shared on social media!` haben.

```js
assert.strictEqual(document.querySelector("#share").children[1].innerText.trim(), "Your card was shared on social media!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
    </div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
