---
id: 67333082ab35e414eb4ea70d
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Wenn die `a`-Elemente mit der Maus berührt werden, ändert sich die Hintergrundfarbe durch eine Übergangsanimation. Sie können steuern, wie dieser Übergang mit der `transition`-Eigenschaft abläuft:

```css
a {
  transition: color 1s linear;
}
```

Die Werte, die die `transition`-Eigenschaft akzeptiert, sind der Reihe nach die Eigenschaft, auf die der Übergang angewendet wird, die Dauer des Übergangs und dann das Timing.

Wenn mehrere Eigenschaften einen Übergang haben, können Sie die Werte für jede Eigenschaft durch Kommas getrennt angeben:

```css
p {
  transition: property1 0.1s, property2 0.6s linear;
}
```

Wenn ein Wert ausgelassen wird, wie das Timing für die erste Eigenschaft, wird ein Standardwert verwendet.

Fügen Sie dem `.card`-Selektor `transition: transform 0.3s, background-color 0.3s ease` hinzu.

Probieren Sie es aus, der Hover-Übergang ist jetzt vollständig.

# --hints--

Der `.card`-Selektor sollte eine Eigenschaft `transition` mit dem Wert `transform 0.3s, background-color 0.3s ease` haben.

```js
assert.oneOf(
  new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("transition"),
  [
    "transform 0.3s ease 0s, background-color 0.3s ease 0s", // this is what comes out of `getPropertyValue`
    "transform 0.3s, background-color 0.3s" // different environment got this instead
  ]
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
--fcc-editable-region--
  
--fcc-editable-region--
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

```
