---
id: 65f569725359e10d345bc52a
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

你将编写一个程序，根据投掷的角度、速度和高度计算并绘制抛射物的轨迹。

首先导入 `math`，你将在该项目中经常使用它，因为它具有 `math.radians`、`math.cos`、`math.sin` 等有用的方法。

还需要创建这些变量以表示重力加速度的数值以及一些以后会用到的特殊符号（请使用复制和粘贴来完成这些操作）。

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

使用 `import math` 导入 `math` 模块。

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

你应该有一个 `GRAVITATIONAL_ACCELERATION` 变量，其值为 `9.81`。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

你应该有一个 `PROJECTILE` 变量，其值为 `"∙"`。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

你应该有一个 `x_axis_tick` 变量，其值为 `"T"`。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

你应该有一个 `y_axis_tick` 变量，其值为 `"⊣"`。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
