---
id: 65fd4bd84561a14d3e720061
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

类变量 `__slots__` 在 Python 类中具有特殊用途。 声明 `__slots__` 并为其分配一个字符串序列会将属性的创建限制为该序列中包含的属性。 此外，它可以防止创建 `__dict__` 特殊属性，并允许更高效的属性访问。

你应该使用类内的 `__slots__` 变量来定义该类具有哪些属性：为 `__slots__` 分配一个包含 3 个字符串的元组，每个字符串等于 `__init__` 中定义的属性名称之一。

# --hints--

你应该为 `__slots__` 分配一个包含 `'__height'`、`'__speed'` 和 `'__angle'` 的元组。

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
