---
id: 65fd4b3bedee044b4b957d33
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

创建一个 `Projectile` 类和一个 `__init__` 方法，该方法使用三个参数初始化该类：起始速度、起始高度和投掷物起始角度（按此顺序）。

在 `__init__` 方法内部，将这些参数分配给三个私有属性：`__speed`、`__height` 和 `__angle`。 提供的角度以度为单位；然而，它应该以弧度为单位进行内部存储。 为此，在将角度分配给 `__angle` 时，使用 `math.radians()` 函数将角度从度转换为弧度。

属性名称前使用两个下划线会触发 Python 中的名称修改。 这意味着不能使用给定的名称从类外部直接访问属性，并且如果外部需要则必须使用像 `ball._Projectile__height` 这样的混乱名称进行访问，表明这些属性仅供内部使用。

# --hints--

你应该使用 `class Projectile:` 来创建该类。

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

你应该在 `Projectile` 类内有一个 `__init__` 方法。

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

新的 `Projectile` 类实例应该正确初始化这三个属性的值。

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
