---
id: 660fcf3e1b9bb056b2edb567
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

此时你已准备好创建字符串表示形式。 首先创建 `Projectile` 类的实例。 在类定义下方，创建一个 `ball` 变量并使用 `10, 3, 45` 作为参数分配给 `Projectile` 的调用。

在类中定义一个 `__str__` 方法，为你的实例提供适当的字符串表示，并使其返回具有以下格式的字符串：

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

它应该以新行字符开始和结束，角度必须写为度数的整数，位移应该用一位小数打印。

你将发现 `math.degrees` 很有用，可以将角度从弧度转换为度数。

当你准备好后，你可以使用 `print(ball)` 来测试你的功能。

# --hints--

你应该有一个 `ball` 变量。

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

`ball` 变量应具有 `Projectile(10, 3, 45)` 的值。

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

`Projectile(45, 45, 45)` 的字符串表示应该是正确的。

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

对于其他实例，字符串表示也应该正确。

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

应该有一个 `print(ball)` 调用。

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
