---
id: 6616ec0bf9ee8548ce6b0f08
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

现在是时候计算轨迹的坐标了！ 创建一个名为 `__calculate_y_coordinate` 的 `Projectile` 类的方法，除了 `self` 之外，它还应该有一个 `x` 参数。

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

以上是针对任意给定的水平位置 $x$ 计算垂直位置 $y$ 的公式，其中起始角度 $\theta$、速度 $v_0$ 和高度 $y_0$。

你将需要使用 `math.tan()` 和 `math.cos()`，并记住 `x ** y` 是写 $x^y$ 的方式，并且 $g$ 的值在变量 `GRAVITATIONAL_ACCELERATION` 中。

实现该方法以便它返回$y$坐标。

如果你想在类外部测试该方法，你可以在声明 `ball` 后调用 `ball._Projectile__calculate_y_coordinate()`，并以数字作为参数。

# --hints--

您的 `Projectile` 类应该有一个名为 `__calculate_y_coordinate` 的方法。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

`__calculate_y_coordinate` 方法应该接收 `self, x` 作为参数。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

该方法应该返回正确的值。 对于 `Projectile(12, 12, 12)`，给定 4 个 $x$ 坐标 0、1、2 和 3，y 坐标应该大约为 12、12.18、12.28、12.32。

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

该方法应该返回正确的值。 对于 `Projectile(45, 12, 22)`，给定 $x$ 坐标 6、9、10 和 12，y 坐标应该大约为 14.32、15.41、15.76、16.44。

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

`__calculate_y_coordinate` 方法不应对结果进行四舍五入。

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
