---
id: 662fa453ec8033322b3ab2e3
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

现在该讨论一下封装和 getter 了。 你已使用前导双下划线将三个实例属性写为私有属性。 请注意，这些属性按照惯例被称为私有的：虽然它们仍然可以从外部访问，但都不这样做。

Getter 可用于从外部获取值。 要定义 getter，你需要定义一个返回所需属性值的方法并为其提供 `@property` 装饰器：

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

 装饰器将方法更改为属性，这意味着该方法不像常规方法那样被调用，而是像属性那样被使用：

```py
n = Nest()
print(n.number_of_eggs)
```

在上面的例子中，私有属性 `__number of eggs` 是通过 `n` 的 `number_of_eggs` 属性访问的。

创建一个名为`speed`的getter，并使其返回私有属性`__speed`的值。

# --hints--

你应该定义一个 `speed` 方法。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

你的 `speed` 方法应该接收一个参数 `self`。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

你的 `speed` 方法应该有一个 `@property` 装饰器。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

你的 `speed` 方法应该返回私有属性 `__speed` 的值。

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
