---
id: 66335223f3e2a31d62d84367
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

回到射弹轨迹计算器，现在你将创建一个新类，该类接受坐标列表并创建轨迹图。

创建一个新的类 `Graph`，它使用私有属性 `__coordinates` 来实例化，其中存储了坐标列表。 声明 `__slots__` 类变量并定义 `__init__` 和 `__repr__` 特殊方法。

另外，通过传递 `coordinates` 变量创建一个 `Graph` 实例，并将其分配给新的 `graph` 变量。

# --hints--

你应该创建一个新的 `Graph` 类。

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_class('Graph')`
    ))
})
```

你应该在类里有一个 `__slots__` 变量。

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_variable('__slots__')`
    ))
})
```

`__slots__` 变量的值应该为 `('__coordinates')` 或 `['__coordinates']` 。

```js
({
    test: () => (runPython(
        `
        slots = _Node(_code).find_class('Graph').find_variable('__slots__')

        assert any(slots.is_equivalent(str) for str in ['__slots__ = ("__coordinates")', '__slots__ = ["__coordinates"]'])`
    ))
})
```

你应该有一个 `__init__` 方法。

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('__init__')`
    ))
})
```

`__init__` 方法应该使用传入的值实例化 `__coordinates` 私有属性。

```js
({
    test: () => assert(runPython(
        `
        l = [(3, 4), (4, 1), (3, 9)]
        g = Graph(l)
        g._Graph__coordinates == l
        `
    ))
})
```

你应该有一个 `__repr__` 方法。

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('__repr__')`
    ))
})
```

`__repr__` 方法应该返回正确的字符串。

```js
({
    test: () => assert(runPython(
        `
        ball = Projectile(10, 3, 45)
        coords = ball.calculate_all_coordinates()
        repr(Graph(coords)) == f'Graph({coords})'
        `
    ))
})
```

你应该创建一个 `Graph` 实例，传入 `coordinates` 变量并将其分配给 `graph` 变量。

```js
({
    test: () => {
        runPython(`

        assert _Node(_code).has_variable('graph'), "variable graph is missing"

        assert _Node(_code).find_variable('graph').is_equivalent('graph = Graph(coordinates)'), "assignment to graph is wrong"
        `)
    }
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.speed} m/s
height: {self.height} m
angle: {self.angle}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed

    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s

    def __repr__(self):
        return f"{self.__class__.__name__}({self.speed}, {self.height}, {self.angle})"

--fcc-editable-region--


ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()


--fcc-editable-region-- 
```
