---
id: 6633b26c4f2d603072ef0818
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

使用 `__repr__` 特殊方法为类提供表示是一种很好的做法。 `__str__` 方法返回可读的字符串表示形式，旨在方便用户使用，而 `__repr__` 则是面向程序员的。 通常 `__repr__` 提供一个可用于重新创建对象的字符串。

编写 `__repr__` 方法，它将返回实例化对象所需的字符串。

你可以通过调用 `repr` 函数并向其传递类的实例来测试这一点。

# --hints--

`Projectile` 类应该有一个 `__repr__` 方法。

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Projectile').has_function('__repr__')`
    ))
})
```

实例 `Projectile(12, 12, 12)` 应具有等于 `'Projectile(12, 12, 12)'` 的表示。

```js
({
    test: () => assert(runPython(
        `repr(Projectile(12, 12, 12)) == 'Projectile(12, 12, 12)' or \\
        repr(Projectile(12, 12, 12)) == 'Projectile(12, 12, 12.0)'`
    ))
})
```

实例 `Projectile(21, 22, 45)` 应具有等于 `'Projectile(21, 22, 45)'` 的表示。

```js
({
    test: () => assert(runPython(
        `repr(Projectile(21, 22, 45)) == 'Projectile(21, 22, 45)' or \\
        repr(Projectile(21, 22, 45)) == 'Projectile(21, 22, 45.0)'`
    ))
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.speed} m/s
height: {self.height} m
angle: {self.angle}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed
    
    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()

   
```
