---
id: 66629f407d679d3105e8317f
title: 가위바위보 게임 만들기
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
게임은 컴퓨터와 대결하는 방식입니다. `"rock"`, `"paper"` 또는 `"scissors"` 중 하나를 무작위로 반환하는 함수를 작성하세요.

게임의 프런트엔드 부분은 신경 쓰지 않아도 됩니다. 게임의 로직만 작성하세요. `script.js`를 열고 지침을 따라가세요.

**사용자 스토리:**

1. `getComputerChoice`라는 이름의 함수가 있어야 합니다.

1. `getComputerChoice` 함수는 `"rock"`, `"paper"`, 또는 `"scissors"` 중 하나를 무작위로 반환해야 합니다.

**힌트:** [Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random) 메서드는 0 이상 1 미만의 무작위 숫자를 반환합니다. 이를 활용해 여러 선택지 중 하나를 조건에 따라 반환하는 방법을 생각해보세요.

게임은 사람이 플레이합니다. 사용자의 선택을 받아 반환하는 함수를 작성하세요.

1. `getHumanChoice`라는 이름의 함수를 만드세요.

1. `getHumanChoice`이 사용자의 입력에 따라 유효한 선택지 중 하나를 반환하도록 코드를 작성하세요.

**힌트:** [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt) 메서드를 사용해 사용자의 입력을 받으세요.

게임은 플레이어 점수를 기록합니다. 플레이어 점수를 기록할 변수를 작성하세요.

1. 전역 범위에 `humanScore`과 `computerScore`라는 두 변수를 만드세요.

1. 이 변수들을 `0` 값으로 초기화하세요.

게임은 라운드별로 진행됩니다. 사람과 컴퓨터 플레이어의 선택을 인수로 받아 한 라운드를 진행하고, 라운드 승자의 점수를 올리며 승자 발표를 기록하는 함수를 작성하세요.

1. `playRound`라는 새 함수를 만드세요.

1. `playRound`에 두 개의 매개변수를 정의하세요. 첫 번째 매개변수는 `humanChoice`, 두 번째는 `computerChoice`입니다. 이 두 매개변수를 사용해 사람과 컴퓨터 선택을 인수로 받으세요.

1. 함수의 `humanChoice` 매개변수가 대소문자를 구분하지 않도록 만들어 플레이어가 `"rock"`, `"ROCK"`, `"RocK"` 또는 다른 변형을 입력해도 되게 하세요.

1. `playRound` 함수 코드를 작성해 라운드 승자를 나타내는 문자열 값을 반환하세요.

- 무승부일 경우 `"It's a tie!"`을 반환해야 합니다.

- 플레이어가 이기면 `"You win! [player choice] beats [computer choice]"`를 반환해야 합니다.

- 컴퓨터가 이기면 `"You lose! [computer choice] beats [player choice]"`를 반환해야 합니다.

1. 라운드 승자에 따라 `humanScore` 또는 `computerScore` 변수를 증가시키세요.

게임은 3라운드로 진행됩니다. `playGame`라는 이름의 함수를 작성해 `playRound`를 호출해 3라운드를 진행하고 점수를 기록하며 최종 승자를 선언하세요.

1. `playGame`라는 새 함수를 만드세요.

1. 3라운드를 진행하는 반복문을 만들고, 매번 사람과 컴퓨터 선택 함수 결과를 인수로 하여 `playRound` 함수를 호출하세요.

1. 게임이 끝나면 누가 더 많은 라운드를 이겼는지에 따라 승자를 반환하세요.

- 사람이 컴퓨터보다 더 많은 라운드를 이기면 `"You win the game!"`라는 메시지를 반환하세요.

- 컴퓨터가 사람보다 더 많은 라운드를 이기면 `"You lose the game!"`라는 메시지를 반환하세요.

# --hints--

`getComputerChoice`라는 함수가 있어야 합니다.

```js
assert.isFunction(getComputerChoice);
```

`getComputerChoice` 함수는 `"rock"`, `"paper"`, 또는 `"scissors"` 중 하나를 무작위로 반환해야 합니다.

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

`getHumanChoice`라는 함수가 있어야 합니다.

```js
assert.isFunction(getHumanChoice);
```

전역 범위에 `humanScore`와 `computerScore`라는 두 변수가 있어야 합니다.

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

`playRound`라는 함수가 있어야 합니다.

```js
assert.isFunction(playRound);
```

`playRound` 함수는 `humanChoice`와 `computerChoice` 매개변수로 사람과 컴퓨터 선택을 인수로 받아야 합니다.

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

`playRound` 함수는 대소문자를 구분하지 않아 플레이어가 `"rock"`, `"ROCK"`, `"RocK"` 또는 다른 변형을 입력해도 됩니다.

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

`playRound` 함수는 사람과 컴퓨터 선택이 같으면 `"It's a tie!"` 문자열을 반환해야 합니다.

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

`playRound` 함수는 플레이어가 이기면 `"You win! [player choice] beats [computer choice]"` 문자열을 반환해야 합니다.

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

`playRound` 함수는 컴퓨터가 이기면 `"You lose! [computer choice] beats [player choice]"` 문자열을 반환해야 합니다.

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

`playRound` 함수는 라운드 승자에 따라 `humanScore` 또는 `computerScore` 변수를 증가시켜야 합니다.

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

`playGame`라는 함수가 있어야 합니다.

```js
assert.isFunction(playGame);
```

3라운드를 진행하는 반복문을 사용해야 합니다.

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

누가 더 많은 라운드를 이겼는지에 따라 게임 승자를 반환해야 합니다.

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```
