---
id: 63ee3f71381756f9716727ef
title: esercizio sulle basi di CSS A
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**Obiettivo:**  
In questo esercizio, praticherai ad aggiungere CSS a un file HTML usando tutti e tre i metodi: CSS esterno, CSS interno e CSS in linea. Dovresti usare solo selettori di tipo per questo esercizio quando aggiungi stili tramite i metodi esterno e interno. Dovresti anche usare parole chiave per i colori (ad esempio "blue") invece di usare valori `RGB` o `HEX`.

## Storie utente

- Dovresti vedere un elemento `div` con del testo.  
  - Dovrebbe avere uno sfondo `red`, testo `white`, una dimensione del carattere di `32px`, testo centrato e `bold`.  
  - Il CSS per l'elemento `div` dovrebbe essere aggiunto esternamente, usando un selettore di tipo.  
- Dovresti vedere un elemento `p` con del testo.  
  - Dovrebbe avere uno sfondo `green`, testo `white` e una dimensione del carattere di `18px`.  
  - Il CSS per l'elemento `p` dovrebbe essere aggiunto internamente, usando un selettore di tipo.  
- Dovresti vedere un elemento `button` con del testo.  
  - L'elemento `button` dovrebbe avere uno sfondo `orange` e una dimensione del carattere di `18px`.  
  - Il CSS per l'elemento `button` dovrebbe essere aggiunto usando stili in linea.

# --hints--

Dovresti avere un elemento `div` contenente del testo.

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

Dovresti avere un foglio di stile esterno contenente gli stili per l'elemento `div`.

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

Il tuo elemento `div` non dovrebbe avere il CSS aggiunto usando stili interni o in linea.

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

Il tuo elemento `div` dovrebbe avere uno sfondo `background-color` di `red` e un testo `color` di `white`.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

Il tuo elemento `div` dovrebbe avere `font-weight` impostato su `bold`, `font-size` impostato su `32px` e `text-align` impostato su `center`.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

Dovresti avere un elemento `p` e dovrebbe contenere del testo.

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

Il tuo elemento `p` dovrebbe avere i suoi stili aggiunti internamente usando un elemento `style`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

Il tuo elemento `p` dovrebbe avere uno sfondo `font-size` di `18px` e avere `color` impostato su `white`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

Dovresti avere un elemento `button` contenente del testo.

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

Il tuo elemento `button` dovrebbe avere uno stile in linea.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```
 
Il tuo elemento `button` dovrebbe avere `background-color` impostato su `orange`.

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

Il tuo elemento `button` dovrebbe avere `font-size` impostato su `18px`.

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
