---
id: 646c5d7057c45f432fcdd46c
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Usando un selettore di classe, dai all'elemento `.cat-head` una larghezza di `205px` e un'altezza di `180px`. Inoltre, dagli un bordo di `1px solid #000` e un `border-radius` di `46%`. 

# --hints--

Dovresti avere un selettore `.cat-head`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Il tuo selettore `.cat-head` dovrebbe avere un `width` impostato su `205px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

Il tuo selettore `.cat-head` dovrebbe avere un `height` impostato su `180px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

Il tuo selettore `.cat-head` dovrebbe avere un `border` impostato su `1px solid #000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

Il tuo selettore `.cat-head` dovrebbe avere un `border-radius` impostato su `46%`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
