---
id: 646ceb843412c74edee27a79
title: Passaggio 22
challengeType: 0
dashedName: step-22
---

# --description--

Ora imparerai un trucco CSS per disegnare triangoli.

Questo servirà per disegnare le orecchie del gatto.

Usando un selettore di classe, dai all'elemento `.cat-right-ear` le proprietà `height` e `width` impostate su `100px`. Imposta la proprietà `background-color` su `white`. Imposta i bordi sinistro e destro su `35px solid blue`. Imposta i bordi superiore e inferiore su `35px solid red`.

# --hints--

Dovresti avere un selettore `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `height` impostata su `100px`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `width` impostata su `100px`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `background-color` impostata su `white`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `border-left` impostata su `35px solid blue`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `border-right` impostata su `35px solid blue`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `border-top` impostata su `35px solid red`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

Il tuo selettore `.cat-right-ear` dovrebbe avere una proprietà `border-bottom` impostata su `35px solid red`. Non dimenticare di aggiungere un punto e virgola.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
