---
id: 646c48df8674cf2b91020ecc
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Aggiungi un elemento `link` all’interno del tuo elemento `head`. Per quell’elemento `link`, imposta l’attributo `rel` su `stylesheet` e il `href` su `./styles.css`.

# --hints--

Dovresti avere un elemento `link`.

```js
assert.match(code, /<link/)
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` e `href` impostati su `stylesheet` e `./styles.css`.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
