---
id: 646ce8bb4b35544d501c7184
title: Passaggio 20
challengeType: 0
dashedName: step-20
---

# --description--

All'interno del tuo elemento `.cat-ears`, crea due elementi `div` con le classi `cat-left-ear` e `cat-right-ear` rispettivamente.

# --hints--

Non dovresti modificare l'elemento `div` esistente con la classe `cat-ears`.

```js
assert(document.querySelectorAll('div.cat-ears').length === 1);
```

Dovresti creare due elementi `div` all'interno del tuo elemento `.cat-ears`.

```js
assert(document.querySelectorAll('.cat-ears div').length === 2);
```

Il tuo primo elemento `div` dovrebbe avere la classe `cat-left-ear`.

```js
assert(document.querySelectorAll('.cat-ears div')[0]?.classList.contains('cat-left-ear'));
```

Il tuo secondo elemento `div` dovrebbe avere la classe `cat-right-ear`.

```js
assert(document.querySelectorAll('.cat-ears div')[1]?.classList.contains('cat-right-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        --fcc-editable-region--
        <div class="cat-ears">
          
        </div>
        --fcc-editable-region--
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
