---
id: 646c5ace05e4be4211407935
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Per vedere l'elemento `.cat-head`, dagli uno sfondo con un gradiente lineare con `#5e5e5e` a `85%` e `#45454f` a `100%`.

Potresti non notare la differenza tra questi due colori, ma ci sono.

# --hints--

Il tuo selettore `.cat-head` dovrebbe avere una proprietà `background`.

```js
assert.match(code, /background:/)
```

La tua proprietà `background` dovrebbe usare la funzione `linear-gradient`.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

La tua funzione `linear-gradient` dovrebbe impostare il primo colore come `#5e5e5e` a `85%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

La tua funzione `linear-gradient` dovrebbe impostare il secondo colore come `#45454f` a `100%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
