---
id: 646c48df8674cf2b91020ecb
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

All’interno del tuo elemento `head`, aggiungi un tag `meta` con l’attributo `charset` impostato su `utf-8`. Aggiungi anche un elemento `title` con il testo `fCC Cat Painting`.

# --hints--

Dovresti aggiungere un solo elemento `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Il tuo elemento `meta` dovrebbe avere un attributo `charset` impostato su `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Dovresti aggiungere un solo elemento `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

Il tuo elemento `title` dovrebbe contenere il testo `fCC Cat Painting`. Ricorda che ortografia e maiuscole/minuscole sono importanti.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
