---
id: 646c47867800472a4ed5d2ea
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Inizia con la struttura HTML di base. Aggiungi un riferimento `DOCTYPE` di `html` e un elemento `html` con il suo attributo `lang` impostato su `en`. Inoltre, aggiungi un elemento `head` e un elemento `body` all'interno dell'elemento `html`.

# --hints--

Dovresti avere la dichiarazione `DOCTYPE` di `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

La tua dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dovresti avere un tag di apertura `html` con un attributo `lang` di `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dovresti avere un tag di chiusura `html`.

```js
assert(code.match(/<\/html>/i));
```

Dovresti avere un tag di apertura `head`.

```js
assert(code.match(/<head>/i));
```

Dovresti avere un tag di chiusura `head`.

```js
assert(code.match(/<\/head>/i));
```

Dovresti avere un tag di apertura `body`.

```js
assert(code.match(/<body>/i));
```

Dovresti avere un tag di chiusura `body`.

```js
assert(code.match(/<\/body>/i));
```

Il tuo elemento `body` dovrebbe venire dopo l'elemento `head`.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

I tuoi elementi `head` e `body` dovrebbero essere all'interno dell'elemento `html`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
