---
id: 5f3313e74582ad9d063e3a38
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте елемент `head` всередину елемента `html`, щоб ви могли додати елемент `title`. Текст елемента `title` має бути `Cafe Menu`.

# --hints--

У вас має бути відкриваючий тег `<head>`.

```js
assert(code.match(/<head>/i));
```

У вас має бути закриваючий тег `</head>`.

```js
assert(code.match(/<head>/i));
```

У вас має бути відкриваючий тег `<title>`.

```js
assert(code.match(/<title>/i));
```

У вас має бути закриваючий тег `</title>`.

```js
assert(code.match(/<\/title>/i));
```

Ваш елемент `title` має бути вкладений у ваш елемент `head`.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Ваш елемент `title` має містити текст `Cafe Menu`. Можливо, потрібно перевірити орфографію.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
