---
id: 5f332b23c2045fb843337579
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Щоб повідомити відвідувачам, що кафе було засноване у 2020 році, додайте елемент `p` під елементом `h1` з текстом `Est. 2020`.

# --hints--

У вас має бути відкриваючий тег `<p>`.

```js
assert(code.match(/<p>/i));
```

У вас має бути кінцевий тег `</p>`.

```js
assert(code.match(/<\/p>/i));
```

Ви не повинні змінювати свій існуючий елемент `h1`. Переконайтеся, що ви не видалили кінцевий тег.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

Ваш елемент `p` має бути під елементом `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Ваш елемент `p` має містити текст `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```
