---
id: 5f33071498eb2472b87ddee4
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому проєкті ви вивчите основи CSS (каскадних таблиць стилів), створюючи меню кафе. CSS — це мова, яка використовується для стилізації HTML-документа. Вона описує, як HTML-елементи мають відображатися на екрані.

Як ви дізналися в останніх кроках застосунку Cat Photo App, існує базова структура, необхідна для початку створення вашої вебсторінки. Кожен HTML-документ повинен мати оголошення `DOCTYPE` і елемент `html`. `DOCTYPE` повідомляє браузеру, у якій версії HTML написаний документ. А елемент `html` представляє кореневий елемент, який містить усі інші елементи.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Додайте тег `<!DOCTYPE html>` і елемент `html` з атрибутом `lang`, значення якого `en`.

# --hints--

Ви повинні мати оголошення `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Ви повинні мати відкриваючий тег `<html lang="en">` з мовою, встановленою на англійську.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ви повинні мати закриваючий тег `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
