---
id: 5f332a88dc25a0fd25c7687a
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Назва кафе — `CAMPER CAFE`. Додайте елемент `h1` всередину вашого елемента `main`. Напишіть у ньому назву кафе великими літерами, щоб вона виділялася.

# --hints--

У вас має бути відкриваючий тег `<h1>`.

```js
assert(code.match(/<h1>/i));
```

У вас має бути закриваючий тег `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Не змінюйте існуючий елемент `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Ваш елемент `h1` має бути вкладений у ваш елемент `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Ваш елемент `h1` має містити текст `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```
