---
id: 5f3326b143638ee1a09ff1e3
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Щоб підготуватися до створення реального вмісту, додайте елемент `body` під елементом `head`.

# --hints--

У вас має бути відкриваючий тег `<body>`.

```js
assert(code.match(/<body>/i));
```

У вас має бути кінцевий тег `</body>`.

```js
assert(code.match(/<\/body>/i));
```

Ви не повинні змінювати свій елемент `head`. Переконайтеся, що ви не видалили кінцевий тег.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

Ваш елемент `body` має йти після елемента `head`.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
