---
id: 5f33294a6af5e9188dbdb8f3
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Настав час додати вміст меню. Додайте елемент `main` всередину існуючого елемента `body`. Він згодом міститиме інформацію про ціни на каву та десерти, які пропонує кафе.

# --hints--

Ваш код має містити відкриваючий тег `<main>`.

```js
assert(code.match(/<main>/i));
```

Ваш код має містити кінцевий тег `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Ви не повинні змінювати свій елемент `body`. Переконайтеся, що випадково не видалили кінцевий тег.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Ваш тег `main` має бути всередині тегу `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
