---
id: 5f331e55dfab7a896e53c3a1
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` — це один із кількох елементів, які надають додаткову інформацію, що не відображається на вебсторінці, але корисна для пошукових систем або для того, як сторінка відображається.

Всередині елемента `head` вкладіть елемент `meta` з атрибутом `charset`, встановленим у значення `utf-8`, щоб повідомити браузеру, як кодувати символи для сторінки.

Пам’ятайте, що елементи `meta` — це порожні елементи, які не потребують закриваючого тегу.

# --hints--

Ви повинні мати тег `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

Елемент `meta` — це порожній елемент, він не повинен мати закриваючого тегу `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Ваш тег `meta` повинен мати атрибут `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Ваш атрибут `charset` повинен мати значення `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
