---
id: 5f344f9c805cd193c33d829c
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Ви можете додати стиль до елемента, вказавши його в `style` елементі та встановивши для нього властивість ось так:

```css
element {
 property: value;
}
```

Вирівняйте вміст `h1` елемента по центру, встановивши його `text-align` властивість у значення `center`.

# --hints--

У вашому `h1` елементі має бути селектор `style`.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

Ваша `text-align` властивість має мати значення `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

Ваш селектор `h1` має встановити властивість `text-align` у значення `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
