---
id: 5f3477cb2e27333b1ab2b955
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Тепер потрібно додати посилання на файл `styles.css`, щоб стилі знову застосувалися. Всередині елемента `head` додайте елемент `link`. Додайте йому атрибут `rel` зі значенням `"stylesheet"` та атрибут `href` зі значенням `"styles.css"`.

Зверніть увагу, що елемент `link` є порожнім елементом.

# --hints--

Ваш код має містити елемент `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

Елемент `link` є порожнім елементом, він не повинен мати кінцевого тегу `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Ви не повинні змінювати існуючий елемент `head`. Переконайтеся, що ви не видалили кінцевий тег.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

У вас має бути один елемент `link`.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Ваш елемент `link` має бути всередині елемента `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Ваш елемент `link` має мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Ваш елемент `link` має мати атрибут `href` зі значенням `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
