---
id: 5f344fc1520b6719f2e35605
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

У меню буде два розділи: один для кави, інший — для десертів. Додайте елемент `section` всередину елемента `main`, щоб мати місце для всіх доступних кав.

# --hints--

Ви повинні мати відкриваючий тег `<section>`.

```js
assert(code.match(/<section\s*>/i));
```

Ви повинні мати кінцевий тег `</section>`.

```js
assert(code.match(/<\/section\s*>/i));
```

Ви не повинні змінювати свій існуючий елемент `main`. Переконайтеся, що ви не видалили кінцевий тег.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Ваш елемент `section` має бути всередині вашого елемента `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
