---
id: 5f356ed69db0a491745e2bb6
title: Крок 27
challengeType: 0
dashedName: step-27
---

# --description--

Щоб застосувати стилі класу до елемента `div`, видаліть атрибут `id` і додайте атрибут `class` до початкового тегу елемента `div`. Переконайтеся, що значення класу встановлено на `menu`.

# --hints--

Ваш `div` має і надалі відтворюватися. Переконайтеся, що ви не пошкодили тег `<div>`.

```js
assert.lengthOf(document.querySelectorAll('div'),1);
```

Ваш елемент `div` має мати клас `menu`.

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

Ваш елемент `div` більше не повинен мати ідентифікатор `menu`.

```js
assert.lengthOf(document.querySelectorAll('div#menu'),0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
