---
id: 5f356ed63c7807a4f1e6d054
title: Крок 21
challengeType: 0
dashedName: step-21
---

# --description--

Зараз мета — зробити так, щоб `div` не займав всю ширину сторінки. CSS-властивість `width` ідеально підходить для цього.

Ви можете використати селектор `id`, щоб націлитися на конкретний елемент з атрибутом `id`. Селектор `id` визначається шляхом розміщення символу решітки `#` безпосередньо перед значенням `id` елемента. Наприклад, якщо елемент має ідентифікатор `id` зі значенням `cat`, тоді ви націлюєтеся на цей елемент так:

```css 
#cat {
  width: 250px;
}
```

Використайте селектор `#menu`, щоб задати вашому елементу ширину `300px`.

# --hints--

Ви повинні мати селектор `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

Ви повинні встановити властивість `width` у значення `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

Ваш `div` має мати ширину 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```
