---
id: 69de4308e652aaeb5ccdb03e
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 Python의 문자열 메서드를 다루는 연습을 하세요.

`email`이라는 변수를 만들고 문자열 `alice.johnson@company.com`을 할당하는 것부터 시작하세요.

# --hints--

`email`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

`email`에 문자열 `alice.johnson@company.com`을 할당해야 합니다. 텍스트를 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
