---
id: 69e20876baed33de636b989a
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

때때로 문자열 안에서 특정 문자의 위치를 찾고 싶을 때가 있습니다. Python은 이를 위해 내장된 `.find()` 메서드를 제공합니다. 이 메서드는 주어진 부분 문자열이 처음 나타나는 위치인 인덱스(위치)를 나타내는 정수를 반환합니다.

예를 들어 다음과 같습니다:

```py
text = 'hello'
print(text.find('e'))  # 1
```

이제 `email`에서 `.find()` 메서드를 호출하고 `'@'`를 전달하여 `email` 문자열에서 골뱅이(@) 기호의 위치를 찾으세요. 결과를 `at_position`이라는 변수에 할당하세요.

그다음 `at_position`을 출력하여 결과가 `Position of @: 13`이 되도록 하세요.

# --hints--

`at_position`이라는 변수를 만드세요.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

`at_position`에 `email.find('@')` 값을 할당하세요.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

`at_position`을 출력하여 결과가 `Position of @: 13`이 되도록 하세요.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
