---
id: 69e20a0451c7a3ef189514d3
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이제 이메일에서 `@`의 위치를 알았으니, `@` 기호 앞에 있는 사용자 이름을 추출할 수 있습니다.

이전 워크숍에서 슬라이싱을 사용해 문자열의 일부를 추출했습니다. 참고로 구문은 다음과 같습니다:

```python
text = 'hello world'
print(text[:5])  # hello
```

`email`에 슬라이싱을 사용해 `at_position` 이전의 모든 내용을 추출하고, 결과를 `username`이라는 변수에 저장하세요.

그런 다음 `username`을 출력하여 결과가 `Username: alice.johnson`이 되도록 하세요.

# --hints--

`username`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

`email`에 `at_position`을 사용한 슬라이싱으로 `username`에 값을 할당해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

`username`을 출력하여 결과가 `Username: alice.johnson`이 되도록 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
