---
id: 69e6299710319e7cb8ebdc37
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

실제 데이터는 종종 여분의 공백과 일관되지 않은 대소문자로 인해 지저분합니다.

`messy_name` 변수를 만들고 두 개의 앞뒤 공백이 있는 `sARaH dAVis` 문자열을 할당하세요. 그런 다음 출력이 다음과 같이 나오도록 출력하세요:

```md
Original name: '  sARaH dAVis  '
```

다음 단계에서 이것을 정리할 것입니다.

# --hints--

`messy_name`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

`messy_name`에 <code>'  sARaH dAVis  '</code>를 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

`messy_name`을 출력하여 출력이 `Original name: '  sARaH dAVis  '`가 되도록 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
