---
id: 69e6284f3b7659c5a05c2a79
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`in` 키워드를 사용해 문자열 안에 부분 문자열이 있는지 확인할 수 있습니다. 부분 문자열이 있으면 `True`를, 없으면 `False`를 반환합니다. 예시는 다음과 같습니다:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

`email`에 문자 `@`가 있는지 확인하는 `has_at` 변수를 만드세요.

`has_at`을 출력해서 결과가 `Has @ symbol: True`가 되도록 하세요.

# --hints--

`has_at`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

`has_at` 변수는 `'@' in email` 값을 가져야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

`has_at`을 출력해서 결과가 `Has @ symbol: True`가 되도록 하세요.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
