---
id: 69e62ce96df237884da25282
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`.lower()` 메서드는 문자열의 모든 문자를 소문자로 변환합니다. 이는 이메일 주소를 표준화할 때 유용합니다.

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

문자열에서 `.strip()`과 `.lower()`를 이렇게 모두 호출할 수 있습니다:

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

`messy_email` 변수를 만들고 두 개의 앞뒤 공백이 있는 `SARAH.DAVIS@COMPANY.COM` 문자열을 할당하세요.

`messy_email`에 `.strip().lower()`를 호출하고 결과를 `clean_email` 변수에 할당하세요.

그런 다음 두 변수를 출력하여 결과가 다음과 같도록 하세요:

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

`messy_email`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

`messy_email`에 <code>'  SARAH.DAVIS@COMPANY.COM  '</code>를 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

`clean_email`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

`messy_email`에 `.strip().lower()`를 호출하고 결과를 `clean_email` 변수에 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

`messy_email`을 출력하여 결과가 `Original email: '  SARAH.DAVIS@COMPANY.COM  '`가 되도록 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

`clean_email`을 출력하여 결과가 `Cleaned email: 'sarah.davis@company.com'`가 되도록 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
