---
id: 69e76d1997661be7f0040e2d
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

`.split(separator)` 메서드는 문자열을 `separator`가 나타날 때마다 하위 문자열 목록으로 나눕니다. 예를 들어:

```py
fruits = 'apple, banana, cherry'
parts = fruits.split(', ')
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

`full_address`라는 변수를 만들고 문자열 `123 Main Street, Springfield, IL`을 할당하세요. 그런 다음 쉼표와 한 칸 공백으로 나누어 `address_parts`를 만드세요.

그다음 `address_parts`를 출력하여 결과가 `Address parts: ['123 Main Street', 'Springfield', 'IL']`가 되도록 하세요.

# --hints--

`full_address`라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('full_address')
`) })
```

`full_address`에 문자열 `123 Main Street, Springfield, IL`을 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('full_address').is_equivalent("full_address = '123 Main Street, Springfield, IL'")
`) })
```

`address_parts`라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('address_parts')
`) })
```

`full_address`에 `.split(', ')`를 사용하고 결과를 `address_parts` 변수에 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('address_parts').is_equivalent("address_parts = full_address.split(', ')")
`) })
```

`address_parts`를 출력하여 결과가 `Address parts: ['123 Main Street', 'Springfield', 'IL']`가 되도록 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Address parts: {address_parts}')",
    "print('Address parts: ' + str(address_parts))",
    "print('Address parts:', address_parts)",
    "print('Address parts:', str(address_parts))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

--fcc-editable-region--

--fcc-editable-region--
```
