---
id: 69e73a141addc6f442924b0d
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

이미 `email` 변수에 `.find()`와 슬라이싱을 사용했어요. 이제 같은 기법을 `clean_email`에 적용할 거예요.

`clean_email`에서 `@` 문자의 위치를 찾아 `at_pos`에 저장하세요. 그런 다음 슬라이싱을 사용해 `@` 앞의 모든 내용을 추출해 `email_user`에 저장하세요.

그다음 `clean_email`과 `email_user`를 출력해서 결과가 다음과 같도록 하세요:

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

`at_pos`라는 변수를 만들어야 해요.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

`clean_email`에서 `.find('@')`를 호출하고 결과를 `at_pos` 변수에 할당해야 해요.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

`email_user`라는 변수를 만들어야 해요.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

`clean_email`에서 `at_pos`를 사용해 슬라이싱하고 결과를 `email_user` 변수에 할당해야 해요.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

`clean_email`을 출력해서 결과가 `Email: sarah.davis@company.com`이 되도록 해야 해요.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

`email_user`를 출력해서 결과가 `Username: sarah.davis`가 되도록 해야 해요.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
