---
id: 69e73f4f5f986657647bbd63
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`.replace(old, new)` 메서드는 문자열 `old`의 모든 발생을 다른 문자열 `new`로 바꿉니다. 이 메서드는 다른 메서드와 연결해서 사용할 수도 있습니다. 예를 들어:

```py
text = 'john.doe'
name = text.replace('.', ' ').title()
print(name)  # Output: John Doe
```

`email_user`에서 `.`를 한 칸 공백으로 바꾸고, 그 결과에 `.title()`을 호출하세요. 결과를 `display_name`이라는 변수에 할당하세요.

`display_name`을 출력해서 결과가 `Display name: Sarah Davis`가 되도록 하세요.

# --hints--

`display_name`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('display_name')
`) })
```

`email_user`에 `.replace('.', ' ')`와 `.title()`을 사용하고, 그 결과를 `display_name` 변수에 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('display_name').is_equivalent("display_name = email_user.replace('.', ' ').title()")
`) })
```

`display_name`을 출력해서 결과가 `Display name: Sarah Davis`가 되도록 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Display name: {display_name}')",
    "print('Display name: ' + display_name)",
    "print('Display name:', display_name)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

--fcc-editable-region--

--fcc-editable-region--
```
