---
id: 69e75dc335b6ec9903a87d3e
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

원하지 않는 문자를 제거하려면 `.replace()`를 사용할 수도 있습니다. 예를 들어, 전화번호에서 대시를 제거하는 경우입니다:

```py
phone = '456-147'
clean = phone.replace('-', '')
print(clean)  # Output: 456147
```

`phone` 변수를 만들고 문자열 `123-456-7890`을 할당하세요. 그런 다음 `-`를 빈 문자열로 바꾸고 결과를 `phone_clean`이라는 변수에 할당하세요.

그런 다음 두 변수를 출력하여 결과가 다음과 같이 나오도록 하세요:

```md
Phone with dashes: 123-456-7890
Phone without dashes: 1234567890
```

# --hints--

`phone`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone')
`) })
```

`phone`에 문자열 `123-456-7890`을 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone').is_equivalent("phone = '123-456-7890'")
`) })
```

`phone_clean`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone_clean')
`) })
```

`phone`에 `.replace('-', '')`를 사용하고 결과를 `phone_clean` 변수에 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone_clean').is_equivalent("phone_clean = phone.replace('-', '')")
`) })
```

`phone`을 출력하여 결과가 `Phone with dashes: 123-456-7890`이 되도록 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone with dashes: {phone}')",
    "print('Phone with dashes: ' + phone)",
    "print('Phone with dashes:', phone)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

`phone_clean`을 출력하여 결과가 `Phone without dashes: 1234567890`이 되도록 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone without dashes: {phone_clean}')",
    "print('Phone without dashes: ' + phone_clean)",
    "print('Phone without dashes:', phone_clean)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

--fcc-editable-region--

--fcc-editable-region--
```
