---
id: 69e76e91402d7850510a5f18
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

`.join(list)` 메서드는 `.split()`의 반대입니다. 구분자를 사용해 문자열 목록을 하나의 문자열로 합칩니다. 예를 들어:

```py
parts = ['apple', 'banana', 'cherry']
joined = ' - '.join(parts)
print(joined)  # Output: apple - banana - cherry
```

`address_parts`의 요소들을 공백으로 둘러싸인 `|` 문자로 연결하고 결과를 `rejoined`라는 변수에 할당하세요.

그다음 `rejoined`를 출력해서 결과가 `Rejoined: 123 Main Street | Springfield | IL`가 되도록 하세요.

# --hints--

`rejoined`라는 변수를 만드세요.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('rejoined')
`) })
```

`rejoined`에 `' | '.join(address_parts)` 값을 할당하세요.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('rejoined').is_equivalent("rejoined = ' | '.join(address_parts)")
`) })
```

`rejoined`를 출력해서 결과가 `Rejoined: 123 Main Street | Springfield | IL`가 되도록 하세요.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Rejoined: {rejoined}')",
    "print('Rejoined: ' + rejoined)",
    "print('Rejoined:', rejoined)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

--fcc-editable-region--

--fcc-editable-region--
```
