---
id: 69e736c39dfa403496e8071e
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

`.upper()` 메서드는 모든 문자를 대문자로 변환합니다. 이는 이름표나 제목에 유용합니다:

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

`name_badge`라는 변수를 만들고 `clean_name`에 `.upper()` 메서드를 호출한 결과를 할당하세요.

그런 다음 `name_badge`를 출력하여 결과가 `Name badge: SARAH DAVIS`가 되도록 하세요.

# --hints--

`name_badge`라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('name_badge')
`) })
```

`clean_name`에 `.upper()`를 호출하고 그 결과를 `name_badge` 변수에 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`) })
```

`name_badge`를 출력하여 결과가 `Name badge: SARAH DAVIS`가 되도록 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Name badge: {name_badge}')",
    "print('Name badge: ' + name_badge)",
    "print('Name badge:', name_badge)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
