---
id: 69eb3158c33d9a229a79951c
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

`.count(substring)` 메서드는 문자열에서 부분 문자열이 나타나는 횟수를 반환합니다. 예를 들어:

```py
text = 'banana'
print(text.count('a'))  # Output: 3
```

`dot_count`와 `dash_count`라는 두 변수를 만드세요. `dot_count`는 `clean_email`에 나타나는 `.` 문자의 개수를 세고, `dash_count`는 `employee_code`에 나타나는 `-` 문자의 개수를 세야 합니다.

# --hints--

`dot_count`라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('dot_count')
`) })
```

`clean_email`에서 `.count('.')`를 사용해 `dot_count`에 값을 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('dot_count').is_equivalent("dot_count = clean_email.count('.')")
`) })
```

`dash_count`라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('dash_count')
`) })
```

`employee_code`에서 `.count('-')`를 사용해 `dash_count`에 값을 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('dash_count').is_equivalent("dash_count = employee_code.count('-')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')
ends_with_com = clean_email.endswith('.com')

print(f'Code: {employee_code}')
print(f'Starts with EMP: {starts_with_emp}')
print(f'Clean email: {clean_email}')
print(f'Ends with .com: {ends_with_com}')

--fcc-editable-region--

--fcc-editable-region--
```
