---
id: 6a71a8df10e823e77358d896
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

`.com`이 `email`에 있는지 확인하는 `has_com`이라는 변수를 하나 더 만드세요.

그런 다음 `has_com`을 출력해서 결과가 `Has .com: True`가 되게 하세요.

# --hints--

`has_com`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_com')
`) })
```

`has_com` 변수는 `'.com' in email` 값을 가져야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_com').is_equivalent("has_com = '.com' in email")
`) })
```

`has_com`을 출력해서 결과가 `Has .com: True`가 되게 해야 합니다.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has .com: {has_com}')",
    "print('Has .com: ' + str(has_com))",
    "print('Has .com:', has_com)",
    "print('Has .com:', str(has_com))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')
--fcc-editable-region--

--fcc-editable-region--
```
