---
id: 6a71b3e49668c3be461257b5
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

`.title()` 메서드는 각 단어의 첫 글자를 대문자로 만듭니다. 예시는 다음과 같습니다:

```py
lowercase = 'hello world'
title_case = lowercase.title()
print(title_case)  # Output: Hello World
```

`clean_name`이라는 변수를 만들고 `messy_name`에 `.title()`을 호출한 결과를 할당하세요.

# --hints--

`clean_name`이라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_name')
`) })
```

`messy_name`에 `.title()`을 사용하여 `clean_name`에 값을 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.title()")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
