---
id: 6a86be3e8cc29e9ae239e70a
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

문자열이 어떻게 끝나는지 확인해야 할 때도 있습니다. Python은 이를 위해 `.endswith()` 메서드를 제공합니다. 문자열이 주어진 접미사로 끝나면 `True`를 반환하고, 그렇지 않으면 `False`를 반환합니다.

예를 들어 다음과 같습니다:

```python
code = 'EMP-001'
print(code.endswith('EMP')) # False
print(code.endswith('001')) # True
```

`clean_email`이 `.com`으로 끝나는지 확인하는 `ends_with_com` 변수를 만드세요.

# --hints--

`ends_with_com`라는 변수를 만들어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('ends_with_com')
`) })
```

`clean_email`에 `.endswith('.com')`을 사용하고 그 결과를 `ends_with_com` 변수에 할당해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('ends_with_com').is_equivalent("ends_with_com = clean_email.endswith('.com')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')

--fcc-editable-region--

--fcc-editable-region--
```
