---
id: 6a86b9e64da746838e67fb13
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

`.strip()` 메서드는 문자열 양쪽 끝의 공백을 제거합니다.

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

`.strip()`과 `.title()`을 연결해서 문자열 양쪽 끝의 공백을 제거하고 각 단어의 첫 글자를 대문자로 만들 수 있습니다:

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

`clean_name` 값을 `messy_name`에 `.title()`을 호출하기 전에 `.strip()`을 호출하도록 업데이트하세요.

그다음 `clean_name`을 출력해서 결과가 `Cleaned name: 'Sarah Davis'`가 되도록 하세요.

# --hints--

`clean_name`을 `messy_name`에 `.title()`을 호출하기 전에 `.strip()`을 호출하도록 업데이트하세요.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.strip().title()")
`) })
```

`clean_name`을 출력해서 결과가 `Cleaned name: 'Sarah Davis'`가 되도록 하세요.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned name: '{clean_name}'\\")",
    "print(\\"Cleaned name: '\\" + clean_name + \\"'\\")",
    "print(\\"Cleaned name: '\\", clean_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
