---
id: 6a870cc60c23f0332a2fe788
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

이전 워크숍에서 문자열 안에 변수를 포함하려면, 여는 따옴표 앞에 `f`를 쓰고 변수를 `{}`로 감싸는 f-string을 사용할 수 있다는 것을 기억하세요.

작동 방식에 대한 간단한 복습입니다:

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

이제 `email`을 출력해서 결과가 `Email: alice.johnson@company.com`이 되도록 하세요.

# --hints--

`email`을 출력해서 결과가 `Email: alice.johnson@company.com`이 되도록 하세요.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
