---
id: 5dc17d3bf86c76b9248c6eb4
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Елемент `p` використовується для створення абзацу тексту на вебсайтах. Створіть елемент `p` під вашим елементом `h2` і додайте до нього такий текст:

`Everyone loves cute cats online!`

# --hints--

Ваш елемент `p` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('p'));
```

Ваші теги `p` мають бути в нижньому регістрі. За конвенцією всі HTML-теги пишуться в нижньому регістрі.

```js
assert.notMatch(code, /<\/?P>/);
```

Ваш елемент `p` має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/p\>/));
```

Текст елемента `p` має бути `Everyone loves cute cats online!`. Ви або пропустили текст, або зробили помилку.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Ваш елемент `p` має бути розташований нижче елемента `h2`. Ви розташували їх у неправильному порядку.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```
