---
id: 5dc1798ff86c76b9248c6eb3
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Елементи заголовків `h1` до `h6` використовуються, щоб позначити важливість вмісту, що знаходиться під ними. Чим менше число, тим вища важливість, тому елементи `h2` мають меншу важливість, ніж елементи `h1`.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Використовуйте лише один елемент `h1` на сторінці і розміщуйте заголовки з меншою важливістю під заголовками з більшою важливістю.

Під елементом `h1` додайте елемент `h2` з таким текстом:

`Cat Photos`

# --hints--

Ваш елемент `h1` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Ваш елемент `h1` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/h1\>/));
```

У вас повинен бути лише один елемент `h1`. Видаліть зайвий.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Текст вашого елемента `h1` повинен бути `CatPhotoApp`. Ви або пропустили текст, або зробили помилку.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Ваші теги `h2` повинні бути у нижньому регістрі. За конвенцією всі HTML-теги пишуться у нижньому регістрі.

```js
assert.notMatch(code, /<\/?H2>/);
```

Ваш елемент `h2` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

Ваш елемент `h2` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/h2\>/));
```

Текст вашого елемента `h2` повинен бути `Cat Photos`. Поміщайте текст `Cat Photos` лише між початковим і кінцевим тегами `h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

Ваш елемент `h2` повинен бути розташований нижче елемента `h1`. Елемент `h1` має більшу важливість і повинен бути вище елемента `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```
