---
id: 5dc174fcf86c76b9248c6eb2
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML-елементи мають початкові теги, як `<h1>`, і кінцеві теги, як `</h1>`. Текст, який елемент відображатиме, розміщується між його початковим і кінцевим тегами.

Змініть текст елемента `h1` нижче з `Hello World` на `CatPhotoApp` і подивіться на зміну у попередньому перегляді браузера.

Коли закінчите, натисніть кнопку "Перевірити код", щоб дізнатися, чи все правильно.

# --hints--

Текст `CatPhotoApp` має бути присутнім у коді. Можливо, варто перевірити орфографію.

```js
assert(code.match(/catphotoapp/i));
```

Ваш елемент `h1` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Ваші теги `h1` мають бути в нижньому регістрі. За конвенцією всі HTML-теги пишуться в нижньому регістрі.

```js
assert.notMatch(code, /<\/?H1>/);
```

Ваш елемент `h1` має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/h1\>/));
```

У вас більше одного елемента `h1`. Видаліть зайвий елемент `h1`.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Текст вашого елемента `h1` має бути `CatPhotoApp`. Ви або пропустили текст, або зробили помилку, або він не розміщений між початковим і кінцевим тегами елемента `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```
