---
id: 5dc174fcf86c76b9248c6eb2
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Елементи HTML мають початкові теги, як-от `<h1>`, та кінцеві теги, як-от `</h1>`. Текст, який відтворюватиме елемент, пишуть між його початковим та кінцевим тегами.

Змініть текст елемента `h1` нижче з `Hello World` на `CatPhotoApp` та перегляньте зміну в попередньому перегляді браузера.

Коли закінчите, натисніть на кнопку «Перевірити код», щоб переконатись, що все правильно.

# --hints--

Код повинен містити текст `CatPhotoApp`. Можливо, потрібно перевірити написання.

```js
assert(code.match(/catphotoapp/i));
```

Елемент `h1` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Your `h1` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
assert.notMatch(code, /<\/?H1>/);
```

Елемент `h1` повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/h1\>/));
```

Ви маєте більше ніж один елемент `h1`. Видаліть зайвий елемент `h1`.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Текстом елемента `h1` має бути `CatPhotoApp`. Ви маєте друкарську помилку, пропустили текст або він відсутній між тегами елемента `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

