---
id: 5dc23f9bf86c76b9248c6eba
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Ви можете додати зображення на свій вебсайт, використовуючи елемент `img`. Елементи `img` мають початковий тег без кінцевого тегу. Елемент без кінцевого тегу називається <dfn>порожнім елементом</dfn>.

Додайте елемент `img` під елементом `p`. На цьому етапі зображення не відобразиться у браузері.

# --hints--

Ваш елемент `img` повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Ваш елемент `img` не повинен мати кінцевого тегу `</img>`.

```js
assert(!code.match(/<\/img\>/));
```

У вас повинен бути лише один елемент `img`. Видаліть усі зайві.

```js
assert(document.querySelectorAll('img').length === 1);
```

Ваш елемент `img` повинен бути розташований під елементом `p`. Ви розташували їх у неправильному порядку.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
