---
id: 5dc24073f86c76b9248c6ebb
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>атрибути</dfn> — це спеціальні слова, які використовуються всередині початкового тегу елемента, щоб керувати поведінкою цього елемента. Атрибут `src` у елементі `img` вказує URL зображення (де знаходиться зображення).

Ось приклад елемента `img` з атрибутом `src`, який вказує на логотип freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Всередині існуючого елемента `img` додайте атрибут `src` з таким URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Ваш код має містити елемент `img`. Можливо, ви видалили елемент `img` або не взяли значення атрибута `src` у лапки.

```js
assert(document.querySelector('img'));
```

Ваш елемент `img` має містити атрибут `src`. Ви або пропустили атрибут, або зробили помилку в написанні. Переконайтеся, що між назвою елемента і назвою атрибута є пробіл.

```js
assert(document.querySelector('img').src);
```

Атрибут `img` елемента `src` має бути встановлений у '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Ви або пропустили URL, або зробили помилку в написанні. Регістр URL має значення.

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Хоча ви встановили атрибут `img` елемента `src` на правильний URL, рекомендується завжди брати значення атрибута в лапки.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Ваш елемент `img` має бути закритий або `>`, або `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
