---
id: 5dc24614f86c76b9248c6ebd
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Ви можете направити посилання на іншу сторінку за допомогою елемента посилання (`a`).

Ось приклад посилання на `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

Додайте елемент посилання після абзацу, який направляє на `https://freecatphotoapp.com`. На цьому етапі посилання не відобразиться у попередньому перегляді.

# --hints--

Ваш елемент посилання (`a`) повинен мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Ваш елемент посилання (`a`) повинен мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/a\>/));
```

Ваш елемент посилання (`a`) повинен бути розташований нижче елемента `p`. Ви розташували їх у неправильному порядку.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

Ваш елемент посилання (`a`) не має атрибута `href`. Перевірте, чи є пробіл після назви початкового тегу і/або чи є пробіли перед усіма назвами атрибутів.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

Ваш елемент посилання (`a`) повинен направляти на `https://freecatphotoapp.com`. Ви або пропустили URL, або зробили помилку.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

Хоча ви встановили атрибут `a` елемента посилання (`href`) на правильне посилання, рекомендується завжди брати значення атрибута в лапки.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
