---
id: 5dc23991f86c76b9248c6eb8
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

У попередньому кроці ви розмістили елементи `h1`, `h2`, коментар і `p` всередині елемента `main`. Це називається *вкладенням*. Вкладені елементи слід розміщувати на два пробіли праворуч від елемента, у який вони вкладені. Такий відступ називається індентацією і використовується для полегшення читання HTML.

Ось приклад вкладення та індентації:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Елементи `h1`, `h2` і коментар відступають на два пробіли більше, ніж елемент `main` у наведеному нижче коді. Використайте пробіл на клавіатурі, щоб додати два пробіли перед елементом `p`, щоб він також мав правильний відступ.

# --hints--

Ваш код має містити елемент `h2` з текстом `Cat Photos`. Можливо, ви випадково його видалили, у нього відсутні як початковий, так і кінцевий теги, або текст змінився.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

Вам не слід додавати елементи `ul` або `li` з прикладу.

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

Не змінюйте індентацію в рядку з вашим елементом `h2`. Його початковий тег має починатися на 6 пробілів від початку рядка. Ви можете перезапустити крок, щоб відновити початкову індентацію.

```js
assert(code.toLowerCase().match(/<\/h1\>\s*\n\s{6}<h2>/));
```

Ваш код має містити коментар. Ви видалили коментар з попереднього кроку.

```js
assert(code.match(/<!--.*-->/));
```

Текст коментаря має бути `TODO: Add link to cat photos`. Не змінюйте текст або відступи коментаря.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

Не змінюйте індентацію в рядку з вашим коментарем. Його початковий тег має починатися на 6 пробілів від початку рядка. Ви можете перезапустити крок, щоб відновити початкову індентацію.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

Ваш код має містити елемент `p`. Ви видалили елемент `p` з попереднього кроку.

```js
assert(document.querySelector('p'));
```

Текст елемента `p` має бути `Everyone loves cute cats online!`. Не змінюйте текст, відступи або пунктуацію елемента `p`.

```js
assert.match(
  document
    .querySelector('p')
    .innerText.toLowerCase(),
    /everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

Початковий тег `p` має мати індентацію, що відповідає вашим елементам `h2` і коментарю. Його початковий тег має починатися на 6 пробілів від початку рядка.

```js
assert(code.toLowerCase().match(/-->\s*\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
