---
id: 5ddbd81294d8ddc1510a8e56
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Текст посилання має бути розміщений між початковим і кінцевим тегами елемента посилання (`a`).

Ось приклад посилання з текстом `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Додайте текст посилання `cat photos` до елемента посилання. Це стане текстом посилання.

# --hints--

Ваш елемент посилання (`a`) має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Ваш елемент посилання (`a`) має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/a\>/));
```

Текст вашого елемента посилання (`a`) має бути `cat photos`. Переконайтеся, що текст посилання розміщено між початковим і кінцевим тегами елемента посилання (`a`).

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
