---
id: 5dfa3589eacea3f48c6300ae
title: Крок 20
challengeType: 0
dashedName: step-20
---

# --description--

У другому елементі `section` додайте новий елемент `h2` з текстом `Cat Lists`.

# --hints--

Ваш елемент `section` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(
  document.querySelectorAll('section').length === 2 &&
    code.match(/<\/section>/g).length === 2
);
```

Ваш елемент `h2` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelectorAll('h2').length === 2);
```

Ваш елемент `h2` має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/h2\>/g).length === 2);
```

Ваш другий елемент `h2` має бути безпосередньо над кінцевим тегом другого елемента `section`. Він розташований не на правильному місці.

```js
const secondSection = document.querySelectorAll('section')[1];
assert(secondSection.lastElementChild.nodeName === 'H2');
```

Другий елемент `h2` має містити текст `Cat Lists`. Ви або пропустили текст, або зробили помилку.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase() === 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
