---
id: 5dfa2407b521be39a3de7be1
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Щоб відкривати посилання в новій вкладці, ви можете використати атрибут `target` у елемента якоря (`a`).

Атрибут `target` визначає, де відкривати пов’язаний документ. `target="_blank"` відкриває пов’язаний документ у новій вкладці або вікні.

Ось базовий синтаксис для елемента `a` з атрибутом `target`:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Додайте атрибут `target` зі значенням `_blank` у початковий тег елемента якоря (`a`), щоб посилання відкривалося в новій вкладці.

# --hints--

Ваш елемент `p` має містити вкладений елемент якоря (`a`) з текстом `cat photos`. Можливо, ви його видалили або допустили помилку.

```js
const anchor = document.querySelectorAll('p > a');
assert.isNotEmpty(anchor);
assert.strictEqual(anchor[1]?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos')
```

Ваш елемент якоря (`a`) не має атрибута `target`. Перевірте, чи є пробіл після назви початкового тегу та/або чи є пробіли перед усіма назвами атрибутів.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

Значення атрибута `target` має бути `_blank`. Ви або пропустили значення, або допустили помилку. Пам’ятайте, що значення атрибутів повинні бути взяті в лапки.

```js
assert.strictEqual(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
