---
id: 5ef9b03c81a63668521804db
title: Крок 43
challengeType: 0
dashedName: step-43
---

# --description--

Щоб запобігти відправленню форми користувачем, коли обов’язкова інформація відсутня, потрібно додати атрибут `required` до елемента `input`.

Ось приклад поля введення з атрибутом `required`:

```html
<input type="text" name="firstName" required>
```

Не потрібно задавати значення атрибуту `required`. Замість цього просто додайте слово `required` до елемента `input`, переконавшись, що між ним і іншими атрибутами є пробіл.

# --hints--

Ви або видалили свій елемент `input`, або у нього неправильний синтаксис. Значення всіх атрибутів повинні бути взяті в лапки.

```js
assert(document.querySelectorAll('input').length);
```

Ваш `form` повинен містити лише елемент `input`. Видаліть будь-які додаткові HTML-елементи або текст всередині елемента `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Ваш елемент `input` повинен мати атрибут `required`. Пам’ятайте, що ви просто додаєте слово `required` всередину тегу елемента `input`.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('required'));
```

Не слід задавати значення атрибуту `required`.

```js
assert.strictEqual(document.querySelector('input')?.getAttribute('required'), '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
