---
id: 5ef9b03c81a63668521804d5
title: Крок 36
challengeType: 0
dashedName: step-36
---

# --description--

Всередині третього елемента `section` додайте елемент `h2` з текстом:

`Cat Form`

# --hints--

Не вдається знайти третій елемент `section`. Можливо, ви випадково його видалили або початковий тег чи кінцевий тег.

```js
assert(
  document.querySelectorAll('section').length === 3 &&
    code.match(/<\/section>/g).length === 3
);
```

Ваш елемент `h2` повинен мати початковий тег і кінцевий тег. Можливо, бракує одного або обох необхідних тегів.

```js
assert(
  document.querySelectorAll('h2').length >= 3 &&
    code.match(/<\/h2>/g).length >= 3
);
```

Ви повинні додати лише один елемент `h2`. Будь ласка, видаліть усі зайві.

```js
assert(document.querySelectorAll('h2').length === 3);
```

Новий елемент `h2` має розташовуватися безпосередньо над кінцевим тегом третього елемента `section`.

```js
const thirdSection = document.querySelectorAll('section')[2];
assert(thirdSection.lastElementChild.nodeName === 'H2');
```

Текст вашого елемента `h2` має бути `Cat Form`.

```js
const thirdSection = document.querySelectorAll('section')[2];
assert(
  thirdSection
    .querySelector('h2')
    .innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'cat form'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>
        </figure>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
