---
id: 5ef9b03c81a63668521804d9
title: Крок 42
challengeType: 0
dashedName: step-42
---

# --description--

Текст заповнювача використовується, щоб підказати людям, який тип інформації потрібно ввести у вхідні дані.

Ось приклад елемента `input` із заповнювачем, встановленим на `e.g. Jane Doe`:

```html
<input type="text" placeholder="e.g. Jane Doe">
```

Додайте текст заповнювача `cat photo URL` до вашого елемента `input`.

# --hints--

Ви або видалили свій елемент `input`, або він має неправильний синтаксис. Значення всіх атрибутів повинні бути взяті в лапки.

```js
assert(document.querySelectorAll('input').length);
```

Ваш `form` повинен містити лише елемент `input`. Видаліть будь-які додаткові HTML-елементи або текст всередині елемента `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Ваш елемент `input` не має атрибута `placeholder`. Перевірте, чи є пробіл після назви початкового тегу і/або чи є пробіли перед усіма назвами атрибутів.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('placeholder'));
```

Ваш елемент `input` повинен мати атрибут `placeholder` зі значенням `cat photo URL`. Ви або пропустили значення, або зробили помилку. Пам’ятайте, що значення атрибутів повинні бути взяті в лапки.

```js
assert(
  document.querySelector('input')
    ?.getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

Хоча ви встановили атрибут `input` елемента `placeholder` у `cat photo URL`, рекомендується завжди брати значення атрибута в лапки.

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
