---
id: 5ef9b03c81a63668521804d7
title: Крок 38
challengeType: 0
dashedName: step-38
---

# --description--

Атрибут `action` вказує, куди слід надсилати дані форми.

Ось приклад елемента `form` з атрибутом `action`:

```html
<form action="/submit-url"></form>
```

У прикладі `action="/submit-url"` повідомляє браузеру, що дані форми слід надсилати за шляхом `/submit-url`.

Додайте атрибут `action` зі значенням `https://freecatphotoapp.com/submit-cat-photo` до елемента `form`.

# --hints--

Ваш елемент `form` повинен мати початковий тег і кінцевий тег у правильному порядку. Можливо, у вас відсутній один або обидва необхідні теги, або вони розташовані в неправильному порядку.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Ваш вкладений елемент `form` у останньому елементі `section` повинен бути розташований нижче елемента `h2`. У вас елемент `h2` і елемент `form` розташовані в неправильному порядку.

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

Ваш елемент `form` не повинен містити вмісту. Видаліть будь-які HTML-елементи або текст між тегами елемента `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

Ваш елемент `form` не має атрибута `action`. Перевірте, чи є пробіл після назви початкового тегу і/або чи є пробіли перед усіма назвами атрибутів.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

Ваш елемент `form` повинен мати атрибут `action` зі значенням `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Хоча ви встановили атрибут `action` на правильний URL, рекомендується завжди брати значення атрибута в лапки.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
