---
id: 5ef9b03c81a63668521804da
title: Крок 44
challengeType: 0
dashedName: step-44
---

# --description--

Елемент `button` використовується для створення активної кнопки.

Додайте елемент `button` з текстом `Submit` під елементом `input`. За замовчуванням натискання кнопки форми без будь-яких атрибутів відправляє форму за адресою, вказаною в атрибуті `action` форми.

# --hints--

Ваш елемент `button` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Ваш елемент `button` має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/button\>/));
```

Текст вашого елемента `button` має бути `Submit`. Ви або пропустили текст, або зробили помилку.

```js
assert(document.querySelector('button').innerText.toLowerCase() === 'submit');
```

Ваш елемент `button` має бути розташований під елементом `input`. Ви розташували їх у неправильному порядку.

```js
const collection = [...document.querySelectorAll('input, button')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('INPUT') < collection.indexOf('BUTTON'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
