---
id: 5ef9b03c81a63668521804d8
title: Крок 39
challengeType: 0
dashedName: step-39
---

# --description--

Елемент `input` дозволяє кількома способами збирати дані з веб-форми. Як і елементи `img`, елементи `input` є <dfn>порожніми елементами</dfn> і не потребують кінцевих тегів.

Вкладіть елемент `input` у елемент `form`.

# --hints--

Ваш елемент `form` має мати початковий тег і кінцевий тег у правильному порядку. Можливо, у вас відсутній один або обидва необхідні теги, або вони розташовані в неправильному порядку.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Початковий тег вашого елемента `form` повинен містити лише атрибут `action`. Видаліть усе інше, що ви могли в нього ввести.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

Ви маєте створити елемент `input`. Перевірте синтаксис.

```js
assert(document.querySelector('input'));
```

Ваш елемент `input` має мати початковий тег, але не кінцевий тег.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

Ваш елемент `input` має бути вкладений у елемент `form`.

```js
assert(document.querySelector('form > input'));
```

Ваш `form` має містити лише елемент `input`. Видаліть будь-які HTML-елементи або текст між тегами елемента `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
