---
id: 5ef9b03c81a63668521804e2
title: Крок 56
challengeType: 0
dashedName: step-56
---

# --description--

Форми зазвичай використовують прапорці для питань, на які може бути більше однієї відповіді. Елемент `input` з атрибутом `type`, встановленим у `checkbox`, створює прапорець.

Під елементом `legend`, який ви щойно додали, додайте `input` з атрибутом `type`, встановленим у `checkbox`, і текстом:

` Loving`

# --hints--

Елемент `input` для вашого прапорця повинен мати початковий тег, але не кінцевий тег.

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

Ви повинні додати лише один елемент input для вашого прапорця. Видаліть усі зайві.

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

Ваш новий елемент `input` має бути розташований під елементом `legend` з текстом `What's your cat's personality?`. У вас вони в неправильному порядку.

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

Ваш новий елемент `input` не має атрибута `type`. Перевірте, чи є пробіл після назви початкового тегу.

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

Ваш новий елемент `input` повинен мати атрибут `type` зі значенням `checkbox`. Ви або пропустили значення, або зробили помилку. Пам’ятайте, що значення атрибутів повинні бути в лапках.

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

Хоча ви встановили атрибут `input` елемента `type` у `checkbox`, рекомендується завжди брати значення атрибута в лапки.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

Текст ` Loving` має бути розташований безпосередньо праворуч від вашого прапорця. Переконайтеся, що між елементом і текстом є пробіл. Ви або пропустили текст, або зробили помилку.

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>
            
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
