---
id: 5ef9b03c81a63668521804dc
title: Крок 46
challengeType: 0
dashedName: step-46
---

# --description--

Ви можете використовувати радіокнопки для питань, де потрібно вибрати лише одну відповідь із кількох варіантів.

Ось приклад радіокнопки з текстом, встановленим як `cat`:

```html
<input type="radio"> cat
```

Пам’ятайте, що елемент `input` є порожнім елементом.

Перед текстовим полем додайте радіокнопку з текстом, встановленим як:

`Indoor`

# --hints--

Ви повинні створити елемент `input` для своєї радіокнопки. Перевірте синтаксис.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

Ваш елемент `input` повинен мати початковий тег, але не кінцевий тег.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

Ви повинні додати лише один елемент `input` для своєї радіокнопки. Видаліть зайві.

```js
assert(document.querySelectorAll('form > input'), 2);
```

Ваш новий елемент `input` має бути вище за існуючий `input` з атрибутом `type`, встановленим у `text`. Ви розташували їх у неправильному порядку.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Ваш новий елемент `input` не має атрибута `type`. Перевірте, чи є пробіл після назви початкового тегу.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Ваш новий елемент `input` повинен мати лише один атрибут `type`. Видаліть зайві.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Ваш новий елемент `input` повинен мати атрибут `type` зі значенням `radio`. Ви або пропустили значення, або зробили помилку. Пам’ятайте, що значення атрибутів повинні бути в лапках.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

Хоча ви встановили атрибут `input` елемента `type` у `radio`, рекомендується завжди брати значення атрибутів у лапки.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Текст `radio` кнопки ` Indoor` повинен розташовуватися після неї, а не перед нею.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

Текст ` Indoor` повинен розташовуватися безпосередньо праворуч від вашої кнопки `radio`. Ви або пропустили текст, або зробили помилку.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
