---
id: 5ef9b03c81a63668521804df
title: Крок 48
challengeType: 0
dashedName: step-48
---

# --description--

Атрибут `id` використовується для ідентифікації конкретних HTML-елементів. Значення кожного атрибута `id` має бути унікальним серед усіх інших значень `id` на всій сторінці.

Ось приклад елемента `input` з атрибутом `id`:

```html
<input id="email">
```

Додайте атрибут `id` зі значенням `indoor` до радіокнопки. Коли елементи мають кілька атрибутів, порядок атрибутів не має значення.

# --hints--

Ваша радіокнопка має залишатися між початковим і кінцевим тегами елемента `label`.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')?.[0]?.childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

Ваша радіокнопка має мати атрибут `id`. Перевірте, що після назви початкового тегу є пробіл і/або перед усіма назвами атрибутів є пробіли.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('id'));
```

Ваша радіокнопка має мати атрибут `type`. Перевірте, що після назви початкового тегу є пробіл і/або перед усіма назвами атрибутів є пробіли.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Ваша радіокнопка має мати атрибут `id` зі значенням `indoor`. Ви або пропустили значення, або зробили помилку. Пам’ятайте, що значення атрибутів мають бути взяті в лапки.

```js
assert.match(document.querySelector('input')?.id, /^indoor$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
