---
id: 5efb2c990dc218d6c85f89b2
title: Крок 45
challengeType: 0
dashedName: step-45
---

# --description--

Навіть якщо ви додали кнопку під текстовим полем, вони відображаються поруч на сторінці. Це тому, що обидва елементи `input` і `button` є <dfn>рядковими елементами</dfn>, які не з’являються на нових лініях.

Кнопка, яку ви додали, за замовчуванням відправляє форму. Однак покладатися на поведінку за замовчуванням може викликати плутанину. Додайте атрибут `type` зі значенням `submit` до `button`, щоб було зрозуміло, що це кнопка відправки.

# --hints--

Ваш елемент `button` має мати початковий тег. Початкові теги мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Ваш елемент `button` має мати кінцевий тег. Кінцеві теги мають `/` одразу після символу `<`.

```js
assert(code.match(/<\/button\>/));
```

Ваш елемент `button` не має атрибута `type`. Перевірте, чи є пробіл після назви початкового тегу.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('type'));
```

Ваш елемент `button` має мати атрибут `type` зі значенням `submit`. Ви або пропустили значення, або зробили помилку. Пам’ятайте, що значення атрибутів мають бути в лапках.

```js
assert.match(
  document.querySelector('button')?.getAttribute('type'),
  /^submit$/i
);
```

Хоча ви встановили атрибут `button` елемента `type` у значення `submit`, рекомендується завжди брати значення атрибута в лапки.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
