---
id: 5f05a1d8e233dff4a68508d8
title: Крок 49
challengeType: 0
dashedName: step-49
---

# --description--

Створіть ще одну радіокнопку під першою. Вкладіть її всередину елемента `label` з `Outdoor` як текстом `label`. Додайте радіокнопці атрибут `id` зі значенням `outdoor`.

# --hints--

Тільки оригінальна радіокнопка `Indoor` повинна мати множину `id` зі значенням `indoor`. Ви можете перезапустити крок, щоб повернути оригінальний HTML, якщо потрібно.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

Ви не повинні вносити жодних змін до радіокнопки `Indoor`. Ви можете перезапустити крок, щоб повернути оригінальний HTML, якщо потрібно.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

Ви повинні додати рівно один новий елемент `input`, вкладений у новий елемент `label`. Переконайтеся, що ваш новий `label` має як відкриваючий, так і кінцевий тег.

```js
assert(document.querySelectorAll('label input').length === 2);
```

Ви не повинні додавати жодних атрибутів до відкриваючого тегу `label`.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

Ваш новий `input` повинен мати атрибут `id` зі значенням `outdoor`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

Ваш новий `input` повинен мати атрибут `type` зі значенням `radio`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

Ви не повинні додавати жодних нових елементів, окрім `input`, вкладеного в `label`.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

Зліва від вашого нового `input` не повинно бути тексту.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

Текст мітки для вашої нової радіокнопки має бути точно `Outdoor`. Це включає велику літеру.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

Ваша нова радіокнопка та мітка повинні бути безпосередньо під/після радіокнопки та мітки `Indoor`. Між ними не повинно бути інших тегів.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
