---
id: 671141f948cbab359e74cc93
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Додайте теги `p`, щоб перетворити `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` на абзац.

# --hints--

Додайте елемент `p` навколо `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.`.

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P);
assert.strictEqual(P.innerText.trim(), "See more cat photos in our gallery.");
```

У новому абзаці досі має бути елемент посилання.

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P.querySelector('a'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
