---
id: 7cf9b03d81a65668421804c3
title: Крок 41
challengeType: 0
dashedName: step-41
---

# --description--

Щоб дані форми були доступні за адресою, вказаною в атрибуті `action`, потрібно додати текстовому полю атрибут `name` і призначити йому значення, яке представлятиме дані, що відправляються.

Ось приклад елемента `input` з атрибутом `name`:

```html
<input type="text" name="name">
```

Додайте атрибут `name` зі значенням `catphotourl` до вашого текстового поля.

# --hints--

Ви або видалили свій елемент `input`, або у нього неправильний синтаксис. Значення всіх атрибутів повинні бути взяті в лапки.

```js
assert(document.querySelectorAll('input').length);
```

Ваш `form` повинен містити лише елемент `input`. Видаліть будь-які додаткові HTML-елементи або текст всередині елемента `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Ваш елемент `input` не має атрибута `name`. Перевірте, чи є пробіл після назви початкового тегу і/або пробіли перед усіма іменами атрибутів.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

Ваш елемент `input` повинен мати атрибут `name` зі значенням `catphotourl`. Ви або пропустили значення, або допустили помилку. Пам’ятайте, що значення атрибутів повинні бути взяті в лапки.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

Хоча ви призначили атрибуту `input` елемента `name` значення `catphotourl`, рекомендується завжди брати значення атрибутів у лапки.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
