---
id: 69de4308e652aaeb5ccdb03e
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, eserciterai l'uso dei metodi per le stringhe in Python.

Inizia creando una variabile chiamata `email` e assegnale la stringa `alice.johnson@company.com`.

# --hints--

Devi creare una variabile chiamata `email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

Devi assegnare la stringa `alice.johnson@company.com` a `email`. Ricorda di racchiudere il testo tra virgolette singole o virgolette doppie.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
