---
id: 69e20876baed33de636b989a
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

A volte, potresti voler trovare la posizione di un carattere specifico all'interno di una stringa. Python fornisce il metodo integrato `.find()` per questo scopo. Restituisce un intero che rappresenta l'indice (posizione) della prima occorrenza della sottostringa data.

Ecco un esempio:

```py
text = 'hello'
print(text.find('e'))  # 1
```

Ora, chiama il metodo `.find()` su `email` e passagli `'@'` per trovare la posizione del simbolo at nella stringa `email`. Assegna il risultato a una variabile chiamata `at_position`.

Poi stampa `at_position` in modo che l'output sia: `Position of @: 13`.

# --hints--

Devi creare una variabile chiamata `at_position`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

Devi assegnare a `at_position` il valore `email.find('@')`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

Devi stampare `at_position` in modo che l'output sia `Position of @: 13`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
