---
id: 69e20a0451c7a3ef189514d3
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora che conosci la posizione di `@` nell'email, puoi estrarre il nome utente, che si trova prima del simbolo `@`.

Nel workshop precedente, hai estratto una porzione di una stringa usando lo slicing. Come promemoria, la sintassi è:

```python
text = 'hello world'
print(text[:5])  # hello
```

Usa lo slicing su `email` per estrarre tutto ciò che sta prima di `at_position` e archivia il risultato in una variabile chiamata `username`.

Poi stampa `username` in modo che l'output sia `Username: alice.johnson`.

# --hints--

Dovresti creare una variabile chiamata `username`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

Dovresti usare lo slicing con `at_position` su `email` per assegnare un valore a `username`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

Dovresti stampare `username` in modo che l'output sia `Username: alice.johnson`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
