---
id: 69e6284f3b7659c5a05c2a79
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Puoi verificare se una sottostringa esiste all'interno di una stringa usando la parola chiave `in`. Restituisce `True` se la sottostringa esiste e `False` altrimenti. Ecco un esempio:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

Crea una variabile chiamata `has_at` che verifica se il carattere `@` è in `email`.

Poi stampa `has_at` in modo che l'output sia: `Has @ symbol: True`.

# --hints--

Dovresti creare una variabile chiamata `has_at`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

La tua variabile `has_at` dovrebbe avere il valore `'@' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

Dovresti stampare `has_at` in modo che l'output sia `Has @ symbol: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
