---
id: 69e62ce96df237884da25282
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Il metodo `.lower()` converte tutti i caratteri di una stringa in minuscolo. Questo è utile per normalizzare gli indirizzi email:

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

Puoi chiamare sia `.strip()` che `.lower()` su una stringa in questo modo:

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

Crea una variabile `messy_email` e assegnale la stringa `SARAH.DAVIS@COMPANY.COM` con due spazi all'inizio e due alla fine.

Chiama anche `.strip().lower()` su `messy_email` e assegna il risultato a una variabile `clean_email`.

Poi stampa entrambe in modo che l'output sia:

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

Dovresti creare una variabile chiamata `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

Dovresti assegnare <code>'  SARAH.DAVIS@COMPANY.COM  '</code> a `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

Dovresti creare una variabile chiamata `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

Dovresti chiamare `.strip().lower()` su `messy_email` e assegnare il risultato a una variabile `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

Dovresti stampare `messy_email` in modo che l'output sia `Original email: '  SARAH.DAVIS@COMPANY.COM  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Dovresti stampare `clean_email` in modo che l'output sia `Cleaned email: 'sarah.davis@company.com'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
