---
id: 69e6299710319e7cb8ebdc37
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

I dati reali sono spesso disordinati, con spazi extra e maiuscole/minuscole incoerenti.

Crea una variabile `messy_name` e assegna la stringa `sARaH dAVis` con due spazi iniziali e due spazi finali. Poi stampala in modo che l'output sia:

```md
Original name: '  sARaH dAVis  '
```

La sistemerai nei passaggi successivi.

# --hints--

Devi creare una variabile chiamata `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

Devi assegnare <code>'&nbsp;&nbsp;sARaH dAVis&nbsp;&nbsp;'</code> a `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

Devi stampare `messy_name` in modo che l'output sia `Original name: '  sARaH dAVis  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
