---
id: 69e73a141addc6f442924b0d
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Hai già usato `.find()` e lo slicing sulla variabile `email`. Ora userai la stessa tecnica su `clean_email`.

Trova la posizione del carattere `@` in `clean_email` e archiviala in `at_pos`. Poi usa lo slicing per estrarre tutto ciò che sta prima di `@` e archiviarlo in `email_user`.

Poi stampa `clean_email` e `email_user` in modo che l'output sia:

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

Dovresti creare una variabile chiamata `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

Dovresti chiamare `.find('@')` su `clean_email` e assegnare il risultato alla variabile `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

Dovresti creare una variabile chiamata `email_user`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

Dovresti usare lo slicing con `at_pos` su `clean_email` e assegnare il risultato alla variabile `email_user`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

Dovresti stampare `clean_email` in modo che l'output sia `Email: sarah.davis@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Dovresti stampare `email_user` in modo che l'output sia `Username: sarah.davis`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
