---
id: 69e736c39dfa403496e8071e
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Il metodo `.upper()` converte tutti i caratteri in maiuscolo. Questo è utile per i cartellini con il nome o le intestazioni:

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

Crea una variabile chiamata `name_badge` e assegna il risultato della chiamata al metodo `.upper()` su `clean_name`.

Poi stampa `name_badge` in modo che l'output sia: `Name badge: SARAH DAVIS`.

# --hints--

Dovresti creare una variabile chiamata `name_badge`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('name_badge')
`) })
```

Dovresti chiamare `.upper()` su `clean_name` e assegnare il risultato alla variabile `name_badge`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`) })
```

Dovresti stampare `name_badge` in modo che l'output sia `Name badge: SARAH DAVIS`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Name badge: {name_badge}')",
    "print('Name badge: ' + name_badge)",
    "print('Name badge:', name_badge)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
