---
id: 69e75dc335b6ec9903a87d3e
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Puoi anche usare `.replace()` per rimuovere caratteri indesiderati. Per esempio, rimuovere i trattini da un numero di telefono:

```py
phone = '456-147'
clean = phone.replace('-', '')
print(clean)  # Output: 456147
```

Crea una variabile `phone` e assegna la stringa `123-456-7890`. Poi sostituisci `-` con una stringa vuota e assegna il risultato a una variabile chiamata `phone_clean`.

Poi stampa entrambe in modo che l'output sia:

```md
Phone with dashes: 123-456-7890
Phone without dashes: 1234567890
```

# --hints--

Dovresti creare una variabile chiamata `phone`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone')
`) })
```

Dovresti assegnare la stringa `123-456-7890` a `phone`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone').is_equivalent("phone = '123-456-7890'")
`) })
```

Dovresti creare una variabile chiamata `phone_clean`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone_clean')
`) })
```

Dovresti usare `.replace('-', '')` su `phone` e assegnare il risultato alla variabile `phone_clean`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone_clean').is_equivalent("phone_clean = phone.replace('-', '')")
`) })
```

Dovresti stampare `phone` in modo che l'output sia `Phone with dashes: 123-456-7890`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone with dashes: {phone}')",
    "print('Phone with dashes: ' + phone)",
    "print('Phone with dashes:', phone)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Dovresti stampare `phone_clean` in modo che l'output sia `Phone without dashes: 1234567890`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone without dashes: {phone_clean}')",
    "print('Phone without dashes: ' + phone_clean)",
    "print('Phone without dashes:', phone_clean)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

--fcc-editable-region--

--fcc-editable-region--
```
