---
id: 6a71a8df10e823e77358d896
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Crea un'altra variabile chiamata `has_com` che verifica se `.com` è in `email`.

Poi stampa `has_com` in modo che l'output sia: `Has .com: True`.

# --hints--

Dovresti creare una variabile chiamata `has_com`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_com')
`) })
```

La tua variabile `has_com` dovrebbe avere il valore `'.com' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_com').is_equivalent("has_com = '.com' in email")
`) })
```

Dovresti stampare `has_com` in modo che l'output sia `Has .com: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has .com: {has_com}')",
    "print('Has .com: ' + str(has_com))",
    "print('Has .com:', has_com)",
    "print('Has .com:', str(has_com))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')
--fcc-editable-region--

--fcc-editable-region--
```
