---
id: 6a86b9e64da746838e67fb13
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Il metodo `.strip()` rimuove gli spazi bianchi da entrambe le estremità di una stringa.

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

Puoi concatenare `.strip()` e `.title()` per rimuovere gli spazi bianchi da entrambe le estremità di una stringa e mettere in maiuscolo la prima lettera di ogni parola:

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

Aggiorna il valore di `clean_name` chiamando `.strip()` su `messy_name` prima di `.title()`.

Dopodiché, stampa `clean_name` in modo che l'output sia `Cleaned name: 'Sarah Davis'`.

# --hints--

Dovresti aggiornare `clean_name` chiamando `.strip()` su `messy_name` prima di `.title()`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.strip().title()")
`) })
```

Dovresti stampare `clean_name` in modo che l'output sia `Cleaned name: 'Sarah Davis'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned name: '{clean_name}'\\")",
    "print(\\"Cleaned name: '\\" + clean_name + \\"'\\")",
    "print(\\"Cleaned name: '\\", clean_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
