---
id: 6a870cc60c23f0332a2fe788
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Ricorda da un workshop precedente che puoi usare una f-stringa per includere una variabile all'interno di una stringa, scrivendo una `f` prima della virgoletta di apertura e racchiudendo la variabile tra `{}`.

Ecco un promemoria di come funziona:

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

Ora, stampa `email` in modo che l'output sia: `Email: alice.johnson@company.com`.

# --hints--

Dovresti stampare `email` in modo che l'output sia `Email: alice.johnson@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
