---
id: 68e45a2b4ab9b2f48b1cc6df
title: 設計價格方案佈局頁面
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事**

1. 你的頁面應該有一個 `h1` 元素，且其 `text` 為 `Pricing Plans`。
2. 你的頁面應該在 `h1` 元素下方有一個帶有類別 `pricing-container` 的 `div` 元素。
    - `.pricing-container` 選擇器應該有一個 `display` 屬性，值為 `flex`，以及一個 `flex-wrap` 屬性，值為 `wrap`。
3. 在 `.pricing-container` 元素中，你應該有三個帶有類別 `pricing-card` 的 `div` 元素來表述價格方案。
    - 其中一個 `.pricing-card` 元素除了 `pricing-card` 類別外，還應該有 `basic-plan` 類別。
    - 其中一個 `.pricing-card` 元素除了 `pricing-card` 類別外，還應該有 `pro-plan` 類別。
    - 其中一個 `.pricing-card` 元素除了 `pricing-card` 類別外，還應該有 `premium-plan` 類別。
4. 你的 `.basic-plan` 元素應該有一個 `h2` 元素，且其文字為 `Basic`。
5. 你的 `.basic-plan` 元素應該有一個 `p` 元素，裡面有文字 `$9/month`。
6. 你的 `.pro-plan` 元素應該包含一個 `h2` 元素，且該元素的文字為 `Pro`。
7. 你的 `.pro-plan` 元素應該有一個 `p` 元素，裡面包含文字 `$19/month`。
8. 你的 `.premium-plan` 元素應該包含一個帶有文字 `Premium` 的 `h2` 元素。
9. 你的 `.premium-plan` 元素應該有一個 `p` 元素，裡面包含文字 `$29/month`。
10. 你的每個 `.pricing-card` 元素應該：
    - 使用 Flexbox 將其內容堆疊成一個行，並使用 `space-between` 對子項目之間的空間進行對齊。
    - 將 `flex` 屬性設定為 `0 0 200px`，以賦予其一致的寬度並防止其在佈局中成長或縮小。
    - 將 `border` 屬性設為 `2px solid black`，以查看不同卡片佔用的空間。
11. `.basic-plan` 元素應該在版面配置中首先出現。你應該使用 `order` 屬性來達成這個目的。
12. `.pro-plan` 元素應該在版面配置中出現第二。你應該使用 `order` 屬性並將其 `flex-grow` 屬性設為 `2`，以便它佔用比其他方案更多的空間。
13. `.premium-plan` 元素應該在版面配置中最後出現。你應該使用 `order` 屬性來達成這個目的。

**注意：** 請務必在你的 HTML 中聯結你的樣式表並套用你的 CSS。

# --hints--

你應該有一個 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的頁面應該有一個 `h1` 元素，且其 `text` 為 `Pricing Plans`。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

你應該在 `h1` 元素下方有一個 `div` 元素。

```js
assert.exists(document.querySelector('h1 + div'));
```

你的 `div` 元素應該有類別 `pricing-container`。

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

你應該有一個 `.pricing-container` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

你的 `pricing-container` 選擇器應該有屬性 `display`，其值為 `flex`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

你的 `.pricing-container` 選擇器應該有屬性 `flex-wrap`，其值為 `wrap`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

你應該在 `.pricing-container` 元素內有三個帶有 `pricing-card` 類別的 `div` 元素。

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

你的其中一個 `.pricing-card` 元素應該有 `basic-plan` 類別。

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

你的其中一個 `.pricing-card` 元素應該有 `premium-plan` 類別。

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

你的其中一個 `.pricing-card` 元素應該有 `pro-plan` 類別。

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

你應該有一個 `.pricing-card` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

你的 `.pricing-card` 選擇器應該有屬性 `display`，其值為 `flex`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

你的 `.pricing-card` 選擇器應該有屬性 `flex-direction`，其值為 `column`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

你的 `.pricing-card` 選擇器應該有屬性 `justify-content`，其值為 `space-between`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

你的 `.pricing-card` 選擇器應該有屬性 `flex`，其值為 `0 0 200px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

你的 `.pricing-card` 選擇器應該有 `border` 屬性，值為 `2px solid black`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

你的 `.basic-plan` 選擇器應該有 `order` 屬性，且值為 `0`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

你的 `.basic-plan` 元素應該有一個 `h2` 元素，且其文字為 `Basic`。

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

你的 `.basic-plan` 元素應該有一個 `p` 元素，裡面有文字 `$9/month`。

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

你的 `.pro-plan` 元素應該包含一個 `h2` 元素，且該元素的文字為 `Pro`。

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

你的 `.pro-plan` 元素應該有一個 `p` 元素，裡面包含文字 `$19/month`。

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

你的 `.pro-plan` 元素應該有 `order` 屬性，且值為 `1`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

你的 `.pro-plan` 元素應該有 `flex-grow` 屬性，值為 `2`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

你的 `.premium-plan` 元素應該包含一個帶有文字 `Premium` 的 `h2` 元素。

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

你的 `.premium-plan` 元素應該有一個 `p` 元素，裡面包含文字 `$29/month`。

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

你的 `premium-plan` 元素應該是版面配置中的最後一個元素，並且具有屬性 `order` 和值 `2`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```

