---
id: 66def5467aee701733aaf8cc
title: Створення застосунку для керування закладками
challengeType: 25
dashedName: build-a-bookmark-manager-app
demoType: onClick
---

# --description--

Для цієї лабораторної роботи вам надано весь HTML і CSS. Ви використаєте JavaScript, щоб завершити застосунок для керування закладками, щоб він міг зберігати колекцію закладок у локальному сховищі та читати дані з нього.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути `getBookmarks` функція, яка повертає `bookmarks` масив, збережений у локальному сховищі. Якщо ключ `bookmarks` ще не встановлено або він не містить дійсний масив із закладками, функція `getBookmarks` має повертати порожній масив.
1. Ключ `bookmarks`, збережений у локальному сховищі, має бути масивом об’єктів закладок. Кожен об’єкт закладки має містити три ключі: `name`, `category` і `url`.
1. У вас має бути функція з іменем `displayOrCloseForm`, яка перемикає клас `hidden` на `#main-section` і `#form-section`.
1. Коли ви натискаєте на `#add-bookmark-button`, ви маєте оновити внутрішній текст `.category-name`, щоб він дорівнював значенню вибраної опції з `#category-dropdown`, і викликати `displayOrCloseForm`, щоб показати розділ форми та приховати головний розділ.
1. Коли ви натискаєте на `#close-form-button`, ви маєте запустити свою функцію, щоб приховати розділ форми і показати головний розділ.
1. Коли ви натискаєте на `#add-bookmark-button-form`, ви маєте оновити ключ `bookmarks`, збережений у локальному сховищі, додавши об’єкт закладки в кінець масиву. Об’єкт має мати `name`, встановлений у значення `#name` вхідних даних, `category`, встановлений у значення вибраної опції з випадаючого списку категорій, і `url`, встановлений у значення `#url` вхідних даних.
1. Після оновлення ключа `bookmarks` ви маєте скинути значення `#name` і `#url` до порожнього рядка перед запуском функції, щоб приховати розділ форми і показати головний розділ.
1. У вас має бути функція з іменем `displayOrHideCategory`, яка перемикає клас `hidden` на `#main-section` і `#bookmark-list-section`.
1. Коли ви натискаєте на `#view-category-button`, ви маєте оновити внутрішній текст `.category-name`, щоб він дорівнював значенню вибраної опції з `#category-dropdown`, змінити внутрішній HTML `#category-list` відповідно до наведених нижче історій користувача і викликати функцію `displayOrHideCategory`.
1. Якщо жодна з закладок у локальному сховищі не має цієї категорії, ви маєте встановити внутрішній HTML `#category-list` у елемент `p` із текстом `No Bookmarks Found`.
1. Якщо одна або кілька закладок у локальному сховищі мають вибрану категорію, додайте радіокнопку для кожної закладки у внутрішній HTML `#category-list`. Кожна радіокнопка має мати `id` і `value`, встановлені в ім’я закладки, і `name`, однаковий для всіх кнопок.
1. Кожна радіокнопка має мати відповідну мітку, що містить елемент посилання з ім’ям закладки та атрибутом `href`, встановленим у URL закладки.
1. Коли ви натискаєте на `#close-list-button`, ви маєте запустити свою функцію, щоб приховати `#bookmark-list-section` і показати головний розділ.
1. Коли ви натискаєте на `#delete-bookmark-button`, ви маєте видалити закладку, що відповідає вибраній радіокнопці та відповідній категорії, з локального сховища і оновити відображений список закладок.

# --before-all--

```js
const getHidden = (notMain) => {
  const mainSection = document.querySelector("#main-section");
  const formSection = document.querySelector("#form-section");
  const bookmarkListSection = document.querySelector("#bookmark-list-section");
  if (notMain === "form") {
    if (mainSection.classList.contains("hidden") && !formSection.classList.contains("hidden")) {
      return "main"
    } else if (!mainSection.classList.contains("hidden") && formSection.classList.contains("hidden")) {
      return notMain;
    }
  } else if (notMain === "bookmark list") {
    if (mainSection.classList.contains("hidden") && !bookmarkListSection.classList.contains("hidden")) {
      return "main"
    } else if (!mainSection.classList.contains("hidden") && bookmarkListSection.classList.contains("hidden")) {
      return notMain;
    }
  }  
}

const originalLocalStorage = JSON.parse(localStorage.getItem("bookmarks")) || [];
const temp = originalLocalStorage;
function resetLocalStorage() {
  localStorage.setItem("bookmarks", JSON.stringify(temp));
}

function setLocalStorage() {
  localStorage.setItem("bookmarks", JSON.stringify([{name: "example1", category: "news", url: "example1.com"}, {name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}]));
}

function clearCategoryList() {
  const categoryList = document.getElementById("category-list");
  categoryList.innerHTML = "";
}
```

# --hints--

У вас має бути `getBookmarks` функція.

```js
assert.isFunction(getBookmarks);
```

Ваша функція `getBookmarks` має повертати масив.

```js
assert.isArray(getBookmarks());
```

Ваша функція `getBookmarks` має повертати ключ `bookmarks` з `localStorage`.

```js
try {
  setLocalStorage()
  const testArr = getBookmarks();
  const expected = [{name: "example1", category: "news", url: "example1.com"}, {name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}];
  assert.deepEqual(testArr, expected);
} finally {
  resetLocalStorage();
}
```

Коли ключ `bookmarks` не встановлено в `localStorage` або він є порожнім масивом, функція `getBookmarks` має повертати порожній масив.

```js
try {
  localStorage.setItem("bookmarks", JSON.stringify([]));
  assert.isEmpty(getBookmarks());
  localStorage.removeItem("bookmarks");
  assert.isEmpty(getBookmarks());
} finally {
  resetLocalStorage();
}
```

Коли ключ `bookmarks` у `localStorage` не містить дійсний масив об’єктів закладок, функція `getBookmarks` має повертати порожній масив.

```js
try {
  localStorage.setItem("bookmarks", 'invalid');
  const arrayFromInvalidValue = getBookmarks();
  assert.isArray(arrayFromInvalidValue);
  assert.isEmpty(arrayFromInvalidValue);

  localStorage.setItem("bookmarks", JSON.stringify([{name: "example1", category: "news", notUrl: "example1.com"}]));
  const arrayFromInvalidArray = getBookmarks();
  assert.isArray(arrayFromInvalidArray)
  assert.isEmpty(arrayFromInvalidArray)
} finally {
  resetLocalStorage();
}
```

У вас має бути функція з іменем `displayOrCloseForm`.

```js
assert.isFunction(displayOrCloseForm);
```

Ваша функція `displayOrCloseForm` має перемикати клас `hidden` на `#main-section` і `#form-section`.

```js
const hidden = getHidden("form");
displayOrCloseForm();
const hiddenAfterCall = getHidden("form");
assert.isDefined(hidden, hiddenAfterCall);
assert.notEqual(hidden, hiddenAfterCall);
```

Коли ви натискаєте на `#add-bookmark-button`, ви маєте оновити внутрішній текст `.category-name`, щоб він дорівнював значенню вибраної опції з `#category-dropdown`.

```js
const addBookMarkButtonTest = document.getElementById("add-bookmark-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
addBookMarkButtonTest.dispatchEvent(new Event("click"));
const categoryName = document.querySelector(".category-name").innerText.trim();
assert.strictEqual(categoryName.toLowerCase(), categoryDropdownTest.value.toLowerCase());
```

Коли ви натискаєте на `#add-bookmark-button`, ви маєте викликати `displayOrCloseForm`, щоб показати розділ форми і приховати головний розділ.

```js
const addBookMarkButtonTest = document.getElementById("add-bookmark-button");
assert.strictEqual(getHidden("form"), "form");
addBookMarkButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("form"), "main");
```

Коли ви натискаєте на `#close-form-button`, ви маєте викликати `displayOrCloseForm`, щоб приховати розділ форми і показати головний розділ.

```js
const closeFormButtonTest = document.getElementById("close-form-button");
assert.strictEqual(getHidden("form"), "main");
closeFormButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("form"), "form");
```

Коли ви натискаєте на `#add-bookmark-button-form`, ви маєте оновити ключ `bookmarks`, збережений у локальному сховищі, додавши об’єкт у кінець масиву. Доданий об’єкт має мати `name`, встановлений у значення `#name` вхідних даних, `category`, встановлений у значення вибраної опції з випадаючого списку категорій, і `url`, встановлений у значення `#url` вхідних даних.

```js
setLocalStorage();
const addBookmarkButtonFormTest = document.getElementById("add-bookmark-button-form");
const categoryDropdownTest = document.getElementById("category-dropdown");
const bookmarkNameTest = document.getElementById("name");
const bookmarkURLTest = document.getElementById("url");
try {
  bookmarkNameTest.value = "test example5";
  bookmarkURLTest.value = "test example5.com";
  const bookmarkObject = {
    name: bookmarkNameTest.value,
    category: categoryDropdownTest.value,
    url: bookmarkURLTest.value
  };
  const expected = [{name: "example1", category: "news", url: "example1.com"}, {name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}, bookmarkObject];
  addBookmarkButtonFormTest.dispatchEvent(new Event("click"));
  assert.deepEqual(getBookmarks(), expected);
} finally {
  bookmarkNameTest.value = "";
  bookmarkURLTest.value = "";
  resetLocalStorage();
}
```

Коли ви натискаєте на `#add-bookmark-button-form`, ви маєте скинути значення `#name` і `#url` до порожнього рядка.

```js
const addBookmarkButtonFormTest = document.getElementById("add-bookmark-button-form");
const bookmarkNameTest = document.getElementById("name");
const bookmarkURLTest = document.getElementById("url");
try {
  bookmarkNameTest.value = "test example2";
  bookmarkURLTest.value = "test example2.com";
  addBookmarkButtonFormTest.dispatchEvent(new Event("click"));
  assert.strictEqual(bookmarkNameTest.value, "");
  assert.strictEqual(bookmarkURLTest.value, "");
} finally {
  resetLocalStorage()  
  bookmarkNameTest.value = "";
  bookmarkURLTest.value = "";
}
```

Коли ви натискаєте на `#add-bookmark-button-form`, ви маєте запустити `displayOrCloseForm`, щоб приховати розділ форми і показати головний розділ.

```js
const addBookmarkButtonFormTest = document.getElementById("add-bookmark-button-form");
const addBookMarkButtonTest = document.getElementById("add-bookmark-button");
const bookmarkNameTest = document.getElementById("name");
const bookmarkURLTest = document.getElementById("url");
try {
  bookmarkNameTest.value = "test example3";
  bookmarkURLTest.value = "test example3.com";  
  addBookMarkButtonTest.dispatchEvent(new Event("click"));  
  assert.strictEqual(getHidden("form"), "main");
  addBookmarkButtonFormTest.dispatchEvent(new Event("click"));
  assert.strictEqual(getHidden("form"), "form");
} finally {
  resetLocalStorage()  
}
```

У вас має бути функція з іменем `displayOrHideCategory`.

```js
assert.isFunction(displayOrHideCategory);
```

Ваша функція `displayOrHideCategory` має перемикати клас `hidden` на `#main-section` і `#bookmark-list-section`.

```js
const hidden = getHidden("bookmark list");
displayOrHideCategory();
const hiddenAfterCall = getHidden("bookmark list");
assert.isDefined(hidden, hiddenAfterCall)
assert.notEqual(hidden, hiddenAfterCall)
```

Коли ви натискаєте на `#view-category-button`, ви маєте оновити внутрішній текст `.category-name`, щоб він дорівнював значенню вибраної опції з `#category-dropdown`.

```js
const viewCategoryButtonTest = document.getElementById("view-category-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
viewCategoryButtonTest.dispatchEvent(new Event("click"));
const categoryName = document.querySelector(".category-name").innerText.trim();
assert.strictEqual(categoryName.toLowerCase(), categoryDropdownTest.value.toLowerCase());
```

Коли ви натискаєте на `#view-category-button`, ви маєте додати елемент `p` із текстом `No Bookmarks Found` до внутрішнього HTML `#category-list`, якщо жодна з закладок у локальному сховищі не має вибраної категорії.

```js
try {
  const viewCategoryButtonTest = document.getElementById("view-category-button");
  const categoryDropdownTest = document.getElementById("category-dropdown");
  const categoryListSection = document.getElementById("bookmark-list-section");
  const mainSection = document.getElementById("main-section");
  const categoryList = document.getElementById("category-list");
  mainSection.classList.remove("hidden");
  categoryListSection.classList.add("hidden");
  categoryList.innerHTML = "";
  setLocalStorage()
  categoryDropdownTest.value = "miscellaneous";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  assert.strictEqual(categoryList.innerHTML,"<p>No Bookmarks Found</p>")
  assert.isFalse(categoryListSection.classList.contains("hidden"));
  assert.isTrue(mainSection.classList.contains("hidden"));
} finally {
  resetLocalStorage()
  clearCategoryList()
}
```

Коли ви натискаєте на `#view-category-button`, ви маєте змінити внутрішній HTML елемента `#category-list`, додавши радіокнопку. Радіокнопка має мати атрибути `id` і `value`, встановлені в ім’я закладки для кожної закладки у вибраній категорії. Крім того, кожна радіокнопка має мати однаковий атрибут `name`.

```js
try {
  const viewCategoryButtonTest = document.getElementById("view-category-button");
  const categoryDropdownTest = document.getElementById("category-dropdown");
  setLocalStorage();
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const bookmarksDisplayed = document.querySelectorAll('#category-list input[type="radio"]');
  assert.lengthOf(bookmarksDisplayed, 2);
  assert.strictEqual(bookmarksDisplayed[0].id, "example1");
  assert.strictEqual(bookmarksDisplayed[0].value, "example1");
  assert.strictEqual(bookmarksDisplayed[1].id, "example4");
  assert.strictEqual(bookmarksDisplayed[1].value, "example4");
  assert.isOk(bookmarksDisplayed[0].name)
  assert.strictEqual(bookmarksDisplayed[0].name, bookmarksDisplayed[1].name);
} finally {
  resetLocalStorage();
}
```

Кожна радіокнопка, додана до внутрішнього HTML `#category-list`, має мати відповідну мітку, що містить елемент посилання з ім’ям закладки та атрибутом `href`, встановленим у URL закладки.

```js
setLocalStorage();
const viewCategoryButtonTest = document.getElementById("view-category-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
try {
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const bookmarkLabelsDisplayed = document.querySelectorAll('#category-list label');
  assert.lengthOf(bookmarkLabelsDisplayed, 2);
  assert.strictEqual(bookmarkLabelsDisplayed[0].htmlFor, "example1");
  assert.strictEqual(bookmarkLabelsDisplayed[1].htmlFor, "example4");
} finally {
  resetLocalStorage();
}
```

Кожен елемент `label` має містити елемент посилання з ім’ям закладки як текстом і атрибутом `href`, встановленим у URL закладки.

```js
setLocalStorage();
const viewCategoryButtonTest = document.getElementById("view-category-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
try {
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const bookmarkAnchorsDisplayed = document.querySelectorAll('#category-list label>a');
  assert.lengthOf(bookmarkAnchorsDisplayed, 2);
  assert.strictEqual(bookmarkAnchorsDisplayed[0].innerText.trim(), "example1");
  assert.isTrue(bookmarkAnchorsDisplayed[0].href.endsWith("example1.com"));
  assert.strictEqual(bookmarkAnchorsDisplayed[1].innerText.trim(), "example4");
  assert.isTrue(bookmarkAnchorsDisplayed[1].href.endsWith("example4.com"));
} finally {
  resetLocalStorage();
}
```

Коли ви натискаєте на `#close-list-button`, ви маєте приховати `#bookmark-list-section` і показати головний розділ.

```js
const viewCategoryButtonTest = document.getElementById("view-category-button");
const closeListButtonTest = document.getElementById("close-list-button");
viewCategoryButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("bookmark list"), "main");
closeListButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("bookmark list"), "bookmark list");
```

Коли ви натискаєте на `#close-list-button` і потім відкриваєте будь-яку категорію, `#category-list` має містити лише дані, що стосуються вибраної категорії, без дублювання записів.

```js
setLocalStorage();
const viewCategoryButtonTest = document.getElementById("view-category-button");
const closeListButtonTest = document.getElementById("close-list-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
try {
  categoryDropdownTest.value = "miscellaneous";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const categoryList = document.getElementById("category-list").innerHTML;
  assert.strictEqual(categoryList, "<p>No Bookmarks Found</p>");

  closeListButtonTest.dispatchEvent(new Event("click"));
  categoryDropdownTest.value = "news";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));

  const newsBookmarks = document.querySelectorAll('#category-list input[type="radio"]');
  assert.lengthOf(newsBookmarks, 2);
  assert.strictEqual(newsBookmarks[0].id, "example1");
  assert.strictEqual(newsBookmarks[0].value, "example1");
  assert.strictEqual(newsBookmarks[1].id, "example4");
  assert.strictEqual(newsBookmarks[1].value, "example4");

  closeListButtonTest.dispatchEvent(new Event("click"));
  categoryDropdownTest.value = "work";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));

  const workBookmarks = document.querySelectorAll('#category-list input[type="radio"]');
  assert.lengthOf(workBookmarks, 1);
  assert.strictEqual(workBookmarks[0].id, "example3");
  assert.strictEqual(workBookmarks[0].value, "example3");

  categoryDropdownTest.value = "miscellaneous";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));

  const paragraphs = document.querySelectorAll("#category-list p");
  assert.lengthOf(paragraphs, 1);
  assert.strictEqual(paragraphs[0].innerText.trim(), "No Bookmarks Found");
} finally {
  resetLocalStorage();
  clearCategoryList();
}
```

Коли ви натискаєте на `#delete-bookmark-button`, ви маєте видалити закладку, що відповідає вибраній радіокнопці та відповідній категорії, з локального сховища і оновити відображений список закладок.

```js
setLocalStorage();
const viewCategoryButtonTest = document.getElementById("view-category-button");
const deleteBookmarkButtonTest = document.getElementById("delete-bookmark-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
try {
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const firstRadioButton = document.querySelector('#category-list input[type="radio"]');
  firstRadioButton.checked = true;
  deleteBookmarkButtonTest.dispatchEvent(new Event("click"));
  let bookmarksDisplayed = document.querySelectorAll('#category-list input[type="radio"]');
  let bookmarkLabelsDisplayed = document.querySelectorAll('#category-list label');
  assert.lengthOf(bookmarksDisplayed, 1);
  assert.strictEqual(bookmarksDisplayed[0].id, "example4");
  assert.strictEqual(bookmarksDisplayed[0].value, "example4");
  assert.lengthOf(bookmarkLabelsDisplayed, 1);
  assert.strictEqual(bookmarkLabelsDisplayed[0].htmlFor, "example4");
  const expected = [{name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}];
  assert.deepEqual(getBookmarks(), expected);

  localStorage.setItem("bookmarks", JSON.stringify([{name: "duplicated-name", category: "news", url: "example1.com"}, {name: "duplicated-name", category: "entertainment", url: "example2.com"}]));

  categoryDropdownTest.value = "entertainment";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const radioForDuplicate = document.querySelector('#category-list input[type="radio"]');
  radioForDuplicate.checked = true;

  deleteBookmarkButtonTest.dispatchEvent(new Event("click"));

  assert.deepEqual(getBookmarks(), [{name: "duplicated-name", category: "news", url: "example1.com"}])
} finally {
  resetLocalStorage();
}
```

# --seed--

## --seed-contents--

```js

```

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=0.0">
    <title>Bookmark Manager</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <section id="main-section">
            <h1>Bookmark Manager</h1>
            <div id="dropdown">
                <label for="category-dropdown">Select a category:</label>
                <select id="category-dropdown" name="options">
                    <option value="news" selected>News</option>
                    <option value="entertainment">Entertainment</option>
                    <option value="work">Work</option>
                    <option value="miscellaneous">Miscellaneous</option>
                </select>
            </div>
            <div id="buttons">
                <button type="button" id="view-category-button">View Category</button>
                <button type="button" id="add-bookmark-button">Add Bookmark</button>
            </div>
        </section>

        <section id="form-section" class="hidden">
            <form>
                <h2 class="category-name"></h2>
                <div>
                    <label for="name">Name:</label>
                    <input type="text" id="name">
                </div>
                <div>
                    <label for="url">URL:</label>
                    <input type="text" id="url">
                </div>
                <div>
                    <button type="button" id="close-form-button">Go Back</button>
                    <button type="button" id="add-bookmark-button-form">Add Bookmark</button>
                </div>
            </form>
        </section>

        <section id="bookmark-list-section" class="hidden">
            <h2 class="category-name"></h2>
            <div id="category-list">
            </div>
            <div>
                <button type="button" id="close-list-button">Go Back</button>
                <button type="button" id="delete-bookmark-button">Delete Bookmark</button>
            </div>
        </section>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-grey: #0a0a23;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

main {
  display: flex;
  justify-content: center;
}

body {
  background-color: var(--dark-grey);
}

.hidden {
  display: none;
}

section {
  display: flex;
  flex-direction: column;
  justify-content: center;
}

select,
input,
label {
  margin-left: 10px;
}

div {
  padding: 30px;
  display: flex;
  justify-content: center;
}

.close-form-button {
  background: none;
  border: none;
  cursor: pointer;
}

h1, h2 {
  margin-top: 20px;
  text-align: center;
}

#category-list {
  text-align: center;
  display: flex;
  flex-direction: column;
  justify-content: flex-start;
  background-color: var(--light-grey);
  align-self: center;
  width: 80%;
  margin-top: 15px;
  border-radius: 10px;
}

#category-list,
h1,
h2,
label {
  color: var(--light-grey);
}

#category-list p {
  color: var(--dark-grey);
}

button {
  cursor: pointer;
  padding: 5px;
  width: 100px;
  margin: 10px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

section {
  margin-top: 60px;
  border: 2px solid var(--golden-yellow);
  width: fit-content;
  border-radius: 10px;
}
```

# --solutions--

```js
const categoryDropdown = document.getElementById("category-dropdown");
const viewCategoryButton = document.getElementById("view-category-button");
const addBookmarkButton = document.getElementById("add-bookmark-button");
const mainSection = document.getElementById("main-section");
const formSection = document.getElementById("form-section");
const bookmarkListSection = document.getElementById("bookmark-list-section");
const categoryHeading = document.querySelectorAll(".category-name");
const closeFormButton = document.getElementById("close-form-button");
const bookmarkName = document.getElementById("name");
const bookmarkURL = document.getElementById("url");
const addBookmarkButtonForm = document.getElementById("add-bookmark-button-form");
const categoryList = document.getElementById("category-list");
const closeListButton = document.getElementById("close-list-button");
const deleteBookmarkButton = document.getElementById("delete-bookmark-button");

const getBookmarks = () => {
  try {
    const parsed = JSON.parse(localStorage.getItem("bookmarks"));
    if (
      Array.isArray(parsed) &&
      parsed.every(
        (item) =>
          item.hasOwnProperty("category") &&
          item.hasOwnProperty("name") &&
          item.hasOwnProperty("url"),
      )
    ) {
      return parsed;
    }
    return [];
  } catch {
    return [];
  }
};

const displayOrCloseForm = () => {
    mainSection.classList.toggle("hidden");
    formSection.classList.toggle("hidden");
}

addBookmarkButton.addEventListener("click", () => {
    categoryHeading[0].innerText = categoryDropdown.value.charAt(0).toUpperCase()
        + categoryDropdown.value.slice(1);
    displayOrCloseForm();
});

closeFormButton.addEventListener("click", displayOrCloseForm);

const addBookmark = () => {
    const bookmarkObject = {
        name: bookmarkName.value,
        category: categoryDropdown.value,
        url: bookmarkURL.value
    };
    if (bookmarkName.value.length && bookmarkURL.value.length) {        
        localStorage.setItem("bookmarks", JSON.stringify(getBookmarks().concat(bookmarkObject)))
        bookmarkName.value = "";
        bookmarkURL.value = "";
        displayOrCloseForm()
    } else {
        alert("Please, provide valid name and URL.")
    }
}

addBookmarkButtonForm.addEventListener("click", addBookmark);

const displayOrHideCategory = () => {
    mainSection.classList.toggle("hidden");
    bookmarkListSection.classList.toggle("hidden");
}

const fillBookmarkList = () => {
    categoryHeading[1].innerText = categoryDropdown.value.charAt(0).toUpperCase()
        + categoryDropdown.value.slice(1);
    const bookmarksToDisplay = getBookmarks().filter((i) => i.category === categoryDropdown.value);
    if (bookmarksToDisplay.length) {
        categoryList.innerHTML = '';
        for (const bookmark of bookmarksToDisplay) {
            categoryList.innerHTML += `
            <div>
              <input type="radio" name="bookmarks" id="${bookmark.name}" value="${bookmark.name}">
              <label for="${bookmark.name}">
                <a href="${bookmark.url}">${bookmark.name}</a>
              </label>
            </div>`
        }
    } else {
        categoryList.innerHTML = '<p>No Bookmarks Found</p>';
    }
}

viewCategoryButton.addEventListener("click", () => {
    fillBookmarkList();
    displayOrHideCategory();
})

closeListButton.addEventListener("click", displayOrHideCategory)

const deleteBookmark = () => {
    const radioBookmarks = document.querySelectorAll('input[type="radio"]');
    for (const radioBookmark of radioBookmarks) {
        if (radioBookmark.checked) {
            const bookmarks = getBookmarks();
            const indexToRemove = bookmarks.findIndex(i => i.name == radioBookmark.value && i.category === categoryDropdown.value);
            bookmarks.splice(indexToRemove, 1);
            localStorage.setItem("bookmarks", JSON.stringify(bookmarks));
            return
        }
    }
}

deleteBookmarkButton.addEventListener("click", () => {
    deleteBookmark();
    fillBookmarkList();
})
```

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=0.0">
    <title>Bookmark Manager</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <section id="main-section">
            <h1>Bookmark Manager</h1>
            <div id="dropdown">
                <label for="category-dropdown">Select a category:</label>
                <select id="category-dropdown" name="options">
                    <option value="news" selected>News</option>
                    <option value="entertainment">Entertainment</option>
                    <option value="work">Work</option>
                    <option value="miscellaneous">Miscellaneous</option>
                </select>
            </div>
            <div id="buttons">
                <button type="button" id="view-category-button">View Category</button>
                <button type="button" id="add-bookmark-button">Add Bookmark</button>
            </div>
        </section>

        <section id="form-section" class="hidden">
            <form>
                <h2 class="category-name"></h2>
                <div>
                    <label for="name">Name:</label>
                    <input type="text" id="name">
                </div>
                <div>
                    <label for="url">URL:</label>
                    <input type="text" id="url">
                </div>
                <div>
                    <button type="button" id="close-form-button">Go Back</button>
                    <button type="button" id="add-bookmark-button-form">Add Bookmark</button>
                </div>
            </form>
        </section>

        <section id="bookmark-list-section" class="hidden">
            <h2 class="category-name"></h2>
            <div id="category-list">
            </div>
            <div>
                <button type="button" id="close-list-button">Go Back</button>
                <button type="button" id="delete-bookmark-button">Delete Bookmark</button>
            </div>
        </section>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-grey: #0a0a23;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

main {
  display: flex;
  justify-content: center;
}

body {
  background-color: var(--dark-grey);
}

.hidden {
  display: none;
}

section {
  display: flex;
  flex-direction: column;
  justify-content: center;
}

select,
input,
label {
  margin-left: 10px;
}

div {
  padding: 30px;
  display: flex;
  justify-content: center;
}

.close-form-button {
  background: none;
  border: none;
  cursor: pointer;
}

h1, h2 {
  margin-top: 20px;
  text-align: center;
}

#category-list {
  text-align: center;
  display: flex;
  flex-direction: column;
  justify-content: flex-start;
  background-color: var(--light-grey);
  align-self: center;
  width: 80%;
  margin-top: 15px;
  border-radius: 10px;
}

#category-list,
h1,
h2,
label {
  color: var(--light-grey);
}

#category-list p {
  color: var(--dark-grey);
}

button {
  cursor: pointer;
  padding: 5px;
  width: 100px;
  margin: 10px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

section {
  margin-top: 60px;
  border: 2px solid var(--golden-yellow);
  width: fit-content;
  border-radius: 10px;
}
```
