---
id: 657a0ea50da0c8d9d6d7950a
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este proyecto, aprenderás el pensamiento algorítmico construyendo un juego de dados. Hay un total de 6 rondas y para cada ronda, el jugador puede lanzar los dados hasta 3 veces y recoger una puntuación. 

El HTML y CSS han sido proporcionados para ti. Siéntete libre de explorarlos. Cuando estés listo, necesitarás configurar tus variables HTML. Obtén todos tus elementos `.die` y asígnalos a una variable `listOfAllDice`. Obtén tus entradas de puntuación (los elementos `input` en tu `#score-options`) y los span de puntuación, y asígnalos a `scoreInputs` y `scoreSpans`. Asigna el elemento `#current-round` a `roundElement` y el elemento `#current-round-rolls` a `rollsElement`, luego haz lo mismo para tus elementos `#total-score` y `#score-history`. Asigna tu `#roll-dice-btn`, `#keep-score-btn`, `#rules-btn` y `.rules-container` a variables con nombres correctamente formateados.

Cuando el usuario haga clic en el botón `Show rules`, debería poder alternar entre mostrar y ocultar las reglas del juego. Crea una variable `isModalShowing` para rastrear el estado de las reglas del juego.

Cada vez que el usuario lanza los dados, necesitas hacer un seguimiento de todos los valores de los dados. Crea una variable `diceValuesArr` para rastrear esto.

A lo largo del juego, necesitarás hacer un seguimiento de la puntuación actual, la puntuación total, número de lanzamientos y qué ronda está el jugador. Declara las variables `rolls`, `score` y `round` para este propósito.

Piensa en cuál debería ser el valor inicial de cada una de estas variables. Todas estas variables deberían poder reasignarse.

# --hints--

Deberías asignar tus elementos `.die` a `listOfAllDice`.

```js
assert.deepEqual(listOfAllDice, document.querySelectorAll(".die"));
```

Deberías asignar tus elementos `#score-options input` a `scoreInputs`.

```js
assert.deepEqual(scoreInputs, document.querySelectorAll("#score-options input"));
```

Deberías asignar tus elementos `#score-options span` a `scoreSpans`.

```js
assert.deepEqual(scoreSpans, document.querySelectorAll("#score-options span"));
```

Deberías asignar tu elemento `#current-round` a `roundElement`.

```js
assert.deepInclude([document.getElementById("current-round"), document.querySelector("#current-round")], roundElement);
```

Deberías asignar tu elemento `#current-round-rolls` a `rollsElement`.

```js
assert.deepInclude([document.getElementById("current-round-rolls"), document.querySelector("#current-round-rolls")], rollsElement);
```

Deberías asignar tu elemento `#total-score` a `totalScoreElement`.

```js
assert.deepInclude([document.getElementById("total-score"), document.querySelector("#total-score")], totalScoreElement);
```

Deberías asignar tu elemento `#score-history` a `scoreHistory`.

```js
assert.deepInclude([document.getElementById("score-history"), document.querySelector("#score-history")], scoreHistory);
```

Deberías asignar tu elemento `#roll-dice-btn` a `rollDiceBtn`.

```js
assert.deepInclude([document.getElementById("roll-dice-btn"), document.querySelector("#roll-dice-btn")], rollDiceBtn);
```

Deberías asignar tu elemento `#keep-score-btn` a `keepScoreBtn`.

```js
assert.deepInclude([document.getElementById("keep-score-btn"), document.querySelector("#keep-score-btn")], keepScoreBtn);
```

Deberías asignar tu elemento `#rules-btn` a `rulesBtn`.

```js
assert.deepInclude([document.getElementById("rules-btn"), document.querySelector("#rules-btn")], rulesBtn);
```

Deberías asignar tu elemento `.rules-container` a `rulesContainer`.

```js
assert.deepEqual(rulesContainer, document.querySelector(".rules-container"));
```

Tu variable `isModalShowing` debería tener el valor `false`.

```js
assert.isFalse(isModalShowing);
```

Tu `diceValuesArr` debería tener un array vacío como el valor.

```js
assert.isArray(diceValuesArr);
assert.isEmpty(diceValuesArr);
```

Tu variable `rolls` debería tener el valor `0`.

```js
assert.strictEqual(rolls, 0);
```

Tu variable `score` debería tener el valor `0`.

```js
assert.strictEqual(score, 0);
```

Tu variable `round` debería tener el valor `1`.

```js
assert.strictEqual(round, 1);
```

Todas tus nuevas variables deberían ser declaradas con `let`.

```js
assert.match(code, /let isModalShowing/);
assert.match(code, /let diceValuesArr/);
assert.match(code, /let rolls/);
assert.match(code, /let score/);
assert.match(code, /let round/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Advanced Dice Game</title>
  <link rel="stylesheet" href="styles.css" />
</head>

<body>
  <header>
    <h1>Advanced Dice Game</h1>
    <button class="btn" id="rules-btn" type="button">Show rules</button>
    <div class="rules-container">
      <h2>Rules</h2>
      <ul>
        <li>There are total of six rounds</li>
        <li>You can only roll the dice three times per round</li>
        <li>To start the game, roll the dice</li>
        <li>
          Then, choose from one of the selected scores or roll the dice again
        </li>
        <li>
          If you choose a selected score, then you will move to the next round
        </li>
        <li>
          If you decline to choose a selected score, then you can roll the
          dice again two more times
        </li>
      </ul>
      <h2 class="points">Points</h2>
      <ul>
        <li>Three of a kind: Sum of all five dice</li>
        <li>Four of a kind: Sum of all five dice</li>
        <li>Full house: Three of a kind and a pair - 25 points</li>
        <li>
          Small straight: Four of the dice have consecutive values - 30 points
        </li>
        <li>
          Large straight: All five dice have consecutive values - 40 points
        </li>
      </ul>
    </div>
  </header>

  <main>
    <form id="game">
      <div id="dice">
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
      </div>

      <p class="rounds-text">
        <strong>Rolls:</strong> <span id="current-round-rolls">0</span> |
        <strong>Round:</strong> <span id="current-round">1</span>
      </p>

      <div id="score-options">
        <div>
          <input type="radio" name="score-options" id="three-of-a-kind" value="three-of-a-kind" disabled />
          <label for="three-of-a-kind">Three of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="four-of-a-kind" value="four-of-a-kind" disabled />
          <label for="four-of-a-kind">Four of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="full-house" value="full-house" disabled />
          <label for="full-house">Full house<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="small-straight" value="small-straight" disabled />
          <label for="small-straight">Small straight<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="large-straight" value="large-straight" disabled />
          <label for="large-straight">Large straight<span></span></label>
        </div>

        <div>
          <input type="radio" name="score-options" id="none" value="none" disabled />
          <label for="none">None of the above<span></span></label>
        </div>
      </div>

      <button class="btn" id="keep-score-btn" type="button">
        Keep the above selected score
      </button>
      <button class="btn" id="roll-dice-btn" type="button">
        Roll the dice
      </button>
    </form>

    <div id="scores">
      <h3>Score history (Total score: <span id="total-score">0</span>)</h3>
      <ol id="score-history"></ol>
    </div>
  </main>
  <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

header {
  color: var(--light-grey);
  text-align: center;
}

h1 {
  font-size: 2.5rem;
  margin: 25px 0;
}

.rules-container {
  display: none;
  background-color: var(--light-grey);
  color: var(--black);
  width: 50%;
  margin: 20px auto;
  height: 300px;
  border-radius: 10px;
  overflow-y: scroll;
}

.rules-container ul {
  list-style-type: none;
}

.points {
  margin-top: 15px;
}

main {
  background-color: var(--light-grey);
  padding: 20px;
  display: grid;
  grid-template-columns: repeat(auto-fit, minmax(500px, 1fr));
  gap: 1rem;
  margin: auto;
  justify-items: center;
  width: 50%;
  border-radius: 10px;
}

#dice {
  display: flex;
  justify-content: space-around;
  margin-bottom: 15px;
}

.die {
  width: 40px;
  height: 40px;
  text-align: center;
  margin-right: 15px;
  border: 4px solid var(--black);
  padding: 10px;
}

.rounds-text {
  text-align: center;
}

input[type="radio"]:disabled + label {
  color: var(--grey);
}

#score-history {
  margin-top: 15px;
  text-align: center;
  list-style-position: inside;
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px 0 10px 0.5rem;
  color: var(--black);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover:enabled {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

@media (max-width: 992px) {
  main {
    width: 100%;
  }
}

@media (max-width: 500px) {
  .btn {
    width: 120px;
  }
}

```

```js
--fcc-editable-region--

--fcc-editable-region--
```
