---
id: 687ad7be1ff45032ccf1d92f
title: بناء مبدل ثيم
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

في هذا المختبر، ستبني تطبيقًا يبدل بين ثيمات مختلفة. عند التبديل إلى ثيم مختلف، سيتغير لون الخلفية ويُعرض رسالة مختلفة.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يحتوي التطبيق على عنصر `button` يحمل النص `Switch Theme`.
2. يجب أن يحتوي عنصر `button` على الخصائص التالية:
   - خاصية `id` مضبوطة على `"theme-switcher-button"`.
   - خاصية `aria-haspopup` مضبوطة على `"true"`.
   - خاصية `aria-expanded` مضبوطة على `"false"`.
   - خاصية `aria-controls` مضبوطة على `"theme-dropdown"`.
3. يجب أن يحتوي التطبيق على عنصر `ul` بهذه الخصائص:
   - خاصية `id` مضبوطة على `"theme-dropdown"`.
   - خاصية `role` مضبوطة على `"menu"`.
   - خاصية `aria-labelledby` مضبوطة على `"theme-switcher-button"`.
   - خاصية `hidden`.
4. يجب أن يحتوي عنصر `ul` على عنصرين على الأقل من نوع `li`، كل واحد منها يحمل خاصية `role` مضبوطة على `"menuitem"` ونص يمثل سمة مختلفة من اختيارك.
5. يجب أن يحمل كل عنصر من عناصر `li` خاصية `id` تبدأ بـ `theme-` وتنتهي باسم السمة التي يمثلها نص عنصر `li`. على سبيل المثال، إذا كانت إحدى السمات هي `Light` فيجب أن يكون `id` هو `theme-light`.
6. يجب أن يحتوي التطبيق على عنصر يحمل خاصية `aria-live` مضبوطة على `"polite"`.
7. يجب أن يكون لديك مصفوفة `themes` تحتوي على كائنين على الأقل، كل كائن فيه خاصيتا `name` و `message`. خاصية `name` تمثل اسم السمة، وخاصية `message` تظهر عند التبديل بين السمات. يمكنك اختيار القيم التي تريدها لـ `name` و `message`، لكن يجب أن تتطابق قيم `name` مع النص الموجود في عناصر `li`.
8. عند نقر المستخدم على `#theme-switcher-button`، يجب أن يظهر عنصر `#theme-dropdown` ويعرض السمات المتاحة.
9. عند عرض عنصر `#theme-dropdown`، يجب ضبط خاصية `aria-expanded` على `"true"` لعنصر `button`.
10. عند نقر المستخدم على `#theme-switcher-button` أثناء عرض عنصر `#theme-dropdown`، يجب إغلاق القائمة وإضافة خاصية `hidden` وضبط خاصية `aria-expanded` على `"false"`.
11. عند نقر المستخدم على `#theme-switcher-button` واختيار سمة، يجب إضافة الصنف `theme-<name>` إلى عنصر `body`، ويجب عرض رسالة السمة المرتبطة التي وضعتها في مصفوفة `themes` داخل العنصر الذي يحمل الخاصية `aria-live="polite"`.
12. عند نقر المستخدم على `#theme-switcher-button` واختيار سمة مختلفة عن الحالية، يجب تحديث صنف السمة في عنصر `body` ورسالة الحالة لتطابق السمة الجديدة.

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

يجب أن يكون لديك عنصر `button`.

```js
assert.exists(document.querySelector("button"));
```

يجب أن يحتوي عنصر `button` على النص `Switch Theme`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

يجب أن يحتوي عنصر `button` على خاصية `id` بقيمة `"theme-switcher-button"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

يجب أن يحتوي عنصر `button` على خاصية `aria-haspopup` بقيمة `"true"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

يجب أن يحتوي عنصر `button` على خاصية `aria-expanded` بقيمة `"false"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

يجب أن يحتوي عنصر `button` على خاصية `aria-controls` بقيمة `"theme-dropdown"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

يجب أن يكون لديك عنصر `ul`.

```js
assert.exists(document.querySelector("ul"));
```

يجب أن يحتوي عنصر `ul` على خاصية `id` بقيمة `"theme-dropdown"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

يجب أن يحتوي عنصر `ul` على خاصية `role` بقيمة `"menu"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

يجب أن يحتوي عنصر `ul` على خاصية `aria-labelledby` بقيمة `"theme-switcher-button"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

يجب أن يحتوي عنصر `ul` على خاصية `hidden`.

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

يجب أن يحتوي عنصر `ul` على عنصرين `li` على الأقل.

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

يجب أن يحتوي كل عنصر `li` على خاصية `role` بقيمة `"menuitem"`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

يجب أن يحتوي كل عنصر `li` على نص يمثل ثيمًا.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

يجب أن يحتوي كل عنصر `li` على خاصية `id` تبدأ بـ `theme-` وتنتهي باسم الثيم المقابل لذلك العنصر. تأكد من أن قيمة `id` كلها أحرف صغيرة.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

يجب أن يكون لديك عنصر يحتوي على خاصية `aria-live` بقيمة `"polite"`.

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

يجب أن يكون لديك مصفوفة `themes`.

```js
assert.isArray(themes);
```

يجب أن تحتوي مصفوفة `themes` على كائنين على الأقل.

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

يجب أن تحتوي مصفوفة `themes` على كائنين على الأقل، كل منهما يحتوي على خاصيتي `name` و`message`.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

يجب أن تكون قيم خاصيتي `name` و`message` عبارة عن سلاسل نصية غير فارغة.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

يجب أن تتطابق كل قيمة من خاصية `name` مع نص أحد عناصر `li`. تأكد من أن هذه القيم كلها أحرف صغيرة.

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

عند نقر المستخدم على `#theme-switcher-button` لعرض خيارات الثيم، يجب إزالة خاصية `hidden` من عنصر `#theme-dropdown`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

عندما يُعرض العنصر `#theme-dropdown`، يجب تعيين الخاصية `aria-expanded` إلى القيمة `"true"` لعنصر الـ `button`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

عند نقر المستخدم على `#theme-switcher-button` أثناء عرض عنصر `#theme-dropdown`، يجب إعادة تعيين القائمة المنسدلة إلى `hidden`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

عند نقر المستخدم على `#theme-switcher-button` أثناء عرض عنصر `#theme-dropdown`، يجب تعيين خاصية `aria-expanded` إلى `"false"`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

عند نقر المستخدم على `#theme-switcher-button` واختيار ثيم، يجب إضافة الصنف `theme-<name>` المقابل إلى عنصر `body`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

عندما ينقر المستخدم على `#theme-switcher-button` ويختار موضوعًا، يجب عرض `message` المرتبط بـ `theme` المختار من مصفوفة `themes` في العنصر الذي يحتوي على الخاصية `aria-live="polite"`.


```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

عند نقر المستخدم على `#theme-switcher-button` واختيار ثيمين مختلفين، يجب تحديث صنف `body` ورسالة الحالة لتتطابق مع كل ثيم مختار.

```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[0]?.dispatchEvent(new Event("click", { bubbles: true }));
const firstClass = document.body.className;
const firstMessage = liveRegion?.textContent;

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", { bubbles: true }));
const secondClass = document.body.className;
const secondMessage = liveRegion?.textContent;

assert.equal(firstClass, listItems[0]?.id);
assert.equal(secondClass, listItems[1]?.id);
assert.notEqual(firstClass, secondClass);
assert.notEqual(firstMessage, secondMessage);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
