---
id: 67172b43f84bcd2dec238a3d
title: Створіть органайзер для книг
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути масив об’єктів з назвою `books`, де кожен об’єкт у масиві повинен мати рядок `title`, інший рядок `authorName` і число `releaseYear`.

2. Ваш масив `books` повинен містити щонайменше три об’єкти.

3. У вас має бути функція зворотного виклику з назвою `sortByYear`, яка приймає два книги як параметри для сортування масиву.

4. Функція `sortByYear` має повертати `-1`, якщо `releaseYear` першої книги менше, ніж у другої.

5. Функція `sortByYear` має повертати `1`, якщо `releaseYear` першої книги більше, ніж у другої.

6. Функція `sortByYear` має повертати `0`, якщо обидва значення `releaseYear` рівні.

7. Ви маєте відфільтрувати книги, написані після певного року, наприклад 1950, з масиву `books` і зберегти відфільтрований масив у новому масиві з назвою `filteredBooks`.

8. Ви маєте відсортувати книги в масиві `filteredBooks` за їх `releaseYear` у зростаючому порядку. Ви вже вивчали на попередньому уроці, що метод `sort()` сортує масив на місці. Це означає, що масив `filteredBooks` буде змінено.

# --hints--

У вашому коді має бути функція `sortByYear`.

```js
assert.isFunction(sortByYear);
```

Ваша функція `sortByYear` має приймати два параметри.

```js
assert.lengthOf(sortByYear, 2);
```

Ваша функція `sortByYear` має повертати `-1`, якщо `releaseYear` об’єкта `book1` менше, ніж у об’єкта `book2`, `1`, якщо `releaseYear` об’єкта `book1` більше, ніж у об’єкта `book2`, і `0` у всіх інших випадках.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

У вашому коді має бути масив `books`.

```js
assert.isArray(books)
```

Ваш масив `books` має містити щонайменше три об’єкти книг.

```js
assert.isAtLeast(books.length, 3)
```

Ваш масив `books` має містити об’єкти, кожен з яких має рядок `title`, інший рядок `authorName` і число `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

У вашому коді має бути масив `filteredBooks`.

```js
assert.isArray(filteredBooks)
```

Масив `filteredBooks` має містити деякі книги, відфільтровані з масиву `books`, і не бути порожнім.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

Масив `filteredBooks` має містити лише книги з оригінального масиву `books`.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

Масив `filteredBooks` має містити лише книги, випущені в або до певного року.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

Ви маєте викликати функцію вищого порядку `sort`, передаючи їй функцію зворотного виклику `sortByYear` для масиву `filteredBooks`.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
