---
id: 594810f028c0303b75339acf
title: Функція Акермана
challengeType: 1
forumTopicId: 302223
dashedName: ackermann-function
---

# --description--

Функція Акермана — класичний приклад рекурсивної функції, особливо відома тим, що вона не є примітивно рекурсивною функцією. Вона дуже швидко зростає за значенням, як і розмір її дерева викликів.

Функція Акермана зазвичай визначається так:

$A(m, n) = \\begin{cases} n+1 & \\mbox{if } m = 0 \\\\ A(m-1, 1) & \\mbox{if } m > 0 \\mbox{ and } n = 0 \\\\ A(m-1, A(m, n-1)) & \\mbox{if } m > 0 \\mbox{ and } n > 0. \\end{cases}$

Її аргументи ніколи не від’ємні, і вона завжди завершується.

# --instructions--

Напишіть функцію, яка повертає значення $A(m, n)$. Бажано використовувати довільну точність (оскільки функція дуже швидко зростає), але це не обов’язково.

# --hints--

`ack` має бути функцією.

```js
assert(typeof ack === 'function');
```

`ack(0, 0)` має повертати 1.

```js
assert(ack(0, 0) === 1);
```

`ack(1, 1)` має повертати 3.

```js
assert(ack(1, 1) === 3);
```

`ack(2, 5)` має повертати 13.

```js
assert(ack(2, 5) === 13);
```

`ack(3, 3)` має повертати 61.

```js
assert(ack(3, 3) === 61);
```

# --seed--

## --seed-contents--

```js
function ack(m, n) {

}
```

# --solutions--

```js
function ack(m, n) {
  return m === 0 ? n + 1 : ack(m - 1, n === 0 ? 1 : ack(m, n - 1));
}
```
