---
id: 594810f028c0303b75339ace
title: Фабрика акумуляторів
challengeType: 1
forumTopicId: 302222
dashedName: accumulator-factory
---

# --description--

Завдання, запропоноване Полом Гремом, полягає у створенні функції, яка приймає один (числовий) аргумент і повертає іншу функцію — акумулятор. Повернена функція-акумулятор також приймає один числовий аргумент і повертає суму всіх числових значень, переданих цьому акумулятору до цього моменту (включно з ініціальним значенням, переданим під час створення акумулятора).

# --instructions--

Створіть функцію, яка приймає число $n$ і генерує функції-акумулятори, що повертають суму всіх чисел, які їм коли-небудь передавали.

**Правила:**

Не використовуйте глобальні змінні.

**Підказка:**

Замикання зберігають зовнішній стан.

# --hints--

`accumulator` має бути функцією.

```js
assert(typeof accumulator === 'function');
```

`accumulator(0)` має повертати функцію.

```js
assert(typeof accumulator(0) === 'function');
```

`accumulator(0)(2)` має повертати число.

```js
assert(typeof accumulator(0)(2) === 'number');
```

Передача значень 3, -4, 1.5 і 5 має повертати 5.5.

```js
assert(testFn && testFn(5) === 5.5);
```

# --before-each--

```js
const testFn = typeof accumulator === 'function' && accumulator(3);
if (testFn) {
  testFn(-4);
  testFn(1.5);
}
```

# --seed--

## --seed-contents--

```js
function accumulator(sum) {

}
```

# --solutions--

```js
function accumulator(sum) {
  return function(n) {
    return sum += n;
  };
}
```
