---
id: 594810f028c0303b75339acd
title: 'Класифікації чисел: надлишкові, дефіцитні та досконалі'
challengeType: 1
forumTopicId: 302221
dashedName: abundant-deficient-and-perfect-number-classifications
---

# --description--

Це визначає три класифікації натуральних чисел на основі їх власних дільників.

Нехай $P(n)$ — сума власних дільників `n`, де власні дільники — це всі натуральні числа `n`, крім самого `n`.

Якщо `P(n) < n`, тоді `n` належить до класу `deficient`.

Якщо `P(n) === n`, тоді `n` належить до класу `perfect`.

Якщо `P(n) > n`, тоді `n` належить до класу `abundant`.

**Приклад**: `6` має власні дільники `1`, `2` та `3`. `1 + 2 + 3 = 6`, тому `6` належить до досконалих чисел.

# --instructions--

Реалізуйте функцію, яка рахує, скільки цілих чисел від `1` до `num` (включно) належать до кожного з трьох класів. Виведіть результат у вигляді масиву у форматі `[deficient, perfect, abundant]`.

# --hints--

`getDPA` має бути функцією.

```js
assert(typeof getDPA === 'function');
```

`getDPA(5000)` має повертати масив.

```js
assert(Array.isArray(getDPA(5000)));
```

`getDPA(5000)` повернутий масив має мати довжину `3`.

```js
assert(getDPA(5000).length === 3);
```

`getDPA(5000)` має повертати `[3758, 3, 1239]`.

```js
assert.deepEqual(getDPA(5000), [3758, 3, 1239]);
```

`getDPA(10000)` має повертати `[7508, 4, 2488]`.

```js
assert.deepEqual(getDPA(10000), [7508, 4, 2488]);
```

`getDPA(20000)` має повертати `[15043, 4, 4953]`.

```js
assert.deepEqual(getDPA(20000), [15043, 4, 4953]);
```

# --seed--

## --seed-contents--

```js
function getDPA(num) {

}
```

# --solutions--

```js
function getDPA(num) {
  const dpa = [1, 0, 0];
  for (let n = 2; n <= num; n += 1) {
    let ds = 1;
    const e = Math.sqrt(n);
    for (let d = 2; d < e; d += 1) {
      if (n % d === 0) {
        ds += d + (n / d);
      }
    }
    if (n % e === 0) {
      ds += e;
    }
    dpa[ds < n ? 0 : ds === n ? 1 : 2] += 1;
  }
  return dpa;
}
```
