---
id: 594810f028c0303b75339ad7
title: алгоритм витоншення Чжан-Суена
challengeType: 1
forumTopicId: 302347
dashedName: zhang-suen-thinning-algorithm
---

# --description--

Це алгоритм, який використовується для витоншення чорно-білих, тобто з одним бітом на піксель, зображень. Наприклад, для вхідного зображення:

```js
const testImage1 = [
 '                               ',
 '#########       ########       ',
 '###   ####     ####  ####      ',
 '###    ###     ###    ###      ',
 '###   ####     ###             ',
 '#########      ###             ',
 '### ####       ###    ###      ',
 '###  ####  ### ####  #### ###  ',
 '###   #### ###  ########  ###  ',
 '                               '
];
```

Він створює витоншене вихідне зображення:

```js
[ '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               ' ];
```

## Алгоритм

Припустимо, що чорні пікселі — це одиниці, а білі — нулі, і що вхідне зображення — це прямокутний N на M масив одиниць і нулів. Алгоритм працює з усіма чорними пікселями P1, які можуть мати вісім сусідів. Сусіди розташовані в такому порядку:

$$\begin{array}{|c|c|c|}
  \\hline
  P9 & P2              & P3\\\\ \\hline
  P8 & \boldsymbol{P1} & P4\\\\ \\hline
  P7 & P6              & P5\\\\ \\hline
\end{array}$$

Очевидно, що пікселі на межі зображення не можуть мати повних восьми сусідів.

- Визначимо $A(P1)$ як кількість переходів від білого до чорного ($0 \to 1$) у послідовності P2, P3, P4, P5, P6, P7, P8, P9, P2. (Зверніть увагу на додатковий P2 в кінці — послідовність циклічна).
- Визначимо $B(P1)$ як кількість чорних сусідів пікселя P1. ($= \\sum(P2 \ldots P9)$)

**Крок 1:**

Всі пікселі перевіряються, і пікселі, які одночасно задовольняють усі наступні умови, лише позначаються на цьому етапі.

1. Піксель чорний і має вісім сусідів
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. Принаймні один із $P2$, $P4$ та $P6$ білий
5. Принаймні один із $P4$, $P6$ та $P8$ білий

Після проходження по зображенню та збору всіх пікселів, що задовольняють умови кроку 1, усі ці пікселі встановлюються в білий колір.

**Крок 2:**

Знову перевіряються всі пікселі, і пікселі, які одночасно задовольняють усі наступні умови, лише позначаються на цьому етапі.

1. Піксель чорний і має вісім сусідів
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. Принаймні один із $P2$, $P4$ та $P8$ білий
5. Принаймні один із $P2$, $P6$ та $P8$ білий

Після проходження по зображенню та збору всіх пікселів, що задовольняють умови кроку 2, усі ці пікселі знову встановлюються в білий колір.

**Ітерація:**

Якщо в цьому раунді на кроці 1 або кроці 2 були змінені якісь пікселі, то всі кроки повторюються, доки жоден піксель зображення не зміниться.

# --instructions--

Напишіть рутину для виконання витоншення Чжан-Суена на наданому `image`, масиві рядків, де кожен рядок представляє один рядок зображення. У рядку `#` означає чорний піксель, а пробіл — білий піксель. Функція має повертати витоншене зображення, використовуючи той самий спосіб представлення.

# --before-each--

```js
const _testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
const expected1 = [
  '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               '
];
const _testImage2 = [
  '                                                          ',
  ' #################                   #############        ',
  ' ##################               ################        ',
  ' ###################            ##################        ',
  ' ########     #######          ###################        ',
  '   ######     #######         #######       ######        ',
  '   ######     #######        #######                      ',
  '   #################         #######                      ',
  '   ################          #######                      ',
  '   #################         #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######         #######       ######        ',
  ' ########     #######          ###################        ',
  ' ########     ####### ######    ################## ###### ',
  ' ########     ####### ######      ################ ###### ',
  ' ########     ####### ######         ############# ###### ',
  '                                                          '
];
const expected2 = [
  '                                                          ',
  '                                                          ',
  '    # ##########                       #######            ',
  '     ##        #                   ####       #           ',
  '     #          #                 ##                      ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     ############               #                         ',
  '     #          #               #                         ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #                            ##                      ',
  '     #                             ############           ',
  '                       ###                          ###   ',
  '                                                          ',
  '                                                          '
];
```

# --hints--

`thinImage` має бути функцією.

```js
assert.equal(typeof thinImage, 'function');
```

`thinImage` має повертати масив.

```js
assert(Array.isArray(thinImage(_testImage1)));
```

`thinImage` має повертати масив рядків.

```js
assert.equal(typeof thinImage(_testImage1)[0], 'string');
``` 

`thinImage(testImage1)` має повертати витоншене зображення, як у прикладі.

```js
assert.deepEqual(thinImage(_testImage1), expected1);
``` 

`thinImage(testImage2)` має повертати витоншене зображення.

```js
assert.deepEqual(thinImage(_testImage2), expected2);
```

# --seed--

## --seed-contents--

```js
function thinImage(image) {

}

const testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
```

# --solutions--

```js
function Point(x, y) {
  this.x = x;
  this.y = y;
}

const ZhangSuen = (function () {
  function ZhangSuen() {
  }

  ZhangSuen.nbrs = [[0, -1], [1, -1], [1, 0], [1, 1], [0, 1], [-1, 1], [-1, 0], [-1, -1], [0, -1]];

  ZhangSuen.nbrGroups = [[[0, 2, 4], [2, 4, 6]], [[0, 2, 6], [0, 4, 6]]];

  ZhangSuen.toWhite = [];

  ZhangSuen.main = function (image) {
    ZhangSuen.grid = new Array(image);
    for (let r = 0; r < image.length; r++) {
      ZhangSuen.grid[r] = image[r].split('');
    }
    ZhangSuen.thinImage();
    return ZhangSuen.getResult();
  };

  ZhangSuen.thinImage = function () {
    let firstStep = false;
    let hasChanged;
    do {
      hasChanged = false;
      firstStep = !firstStep;
      for (let r = 1; r < ZhangSuen.grid.length - 1; r++) {
        for (let c = 1; c < ZhangSuen.grid[0].length - 1; c++) {
          if (ZhangSuen.grid[r][c] !== '#') {
            continue;
          }
          const nn = ZhangSuen.numNeighbors(r, c);
          if (nn < 2 || nn > 6) {
            continue;
          }
          if (ZhangSuen.numTransitions(r, c) !== 1) {
            continue;
          }
          if (!ZhangSuen.atLeastOneIsWhite(r, c, firstStep ? 0 : 1)) {
            continue;
          }
          ZhangSuen.toWhite.push(new Point(c, r));
          hasChanged = true;
        }
      }
      for (let i = 0; i < ZhangSuen.toWhite.length; i++) {
        const p = ZhangSuen.toWhite[i];
        ZhangSuen.grid[p.y][p.x] = ' ';
      }
      ZhangSuen.toWhite = [];
    } while ((firstStep || hasChanged));
  };

  ZhangSuen.numNeighbors = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === '#') {
        count++;
      }
    }
    return count;
  };

  ZhangSuen.numTransitions = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === ' ') {
        if (ZhangSuen.grid[r + ZhangSuen.nbrs[i + 1][1]][c + ZhangSuen.nbrs[i + 1][0]] === '#') {
          count++;
        }
      }
    }
    return count;
  };

  ZhangSuen.atLeastOneIsWhite = function (r, c, step) {
    let count = 0;
    const group = ZhangSuen.nbrGroups[step];
    for (let i = 0; i < 2; i++) {
      for (let j = 0; j < group[i].length; j++) {
        const nbr = ZhangSuen.nbrs[group[i][j]];
        if (ZhangSuen.grid[r + nbr[1]][c + nbr[0]] === ' ') {
          count++;
          break;
        }
      }
    }
    return count > 1;
  };

  ZhangSuen.getResult = function () {
    const result = [];
    for (let i = 0; i < ZhangSuen.grid.length; i++) {
      const row = ZhangSuen.grid[i].join('');
      result.push(row);
    }
    return result;
  };
  return ZhangSuen;
}());

function thinImage(image) {
  return ZhangSuen.main(image);
}
```
