---
id: 594810f028c0303b75339ad4
title: Перенесення слів
challengeType: 1
forumTopicId: 302344
dashedName: word-wrap
---

# --description--

Навіть сьогодні, з пропорційними шрифтами та складними макетами, все ще є випадки, коли потрібно переносити текст на вказаній колонці. Основне завдання — просто перенести абзац тексту.

# --instructions--

Напишіть функцію, яка може переносити цей текст на будь-яку кількість символів. Зверніть увагу, що вхідний текст уже містить розриви рядків, які ваша функція має коректно обробляти. Наприклад, текст, перенесений на 80 символів, має виглядати так:

<pre>
Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX
algorithm. If your language provides this, you get easy extra credit, but you
must reference documentation indicating that the algorithm is something better
than a simple minimum length algorithm.
</pre>

# --before-each--

```js
const text =
`Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX algorithm.
If your language provides this, you get easy extra credit,
but you ''must reference documentation'' indicating that the algorithm
is something better than a simple minimum length algorithm.`;

const wrapped80 = wrap(text, 80);
const wrapped42 = wrap(text, 42);

const firstRow80 =
    'Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX';

const firstRow42 = 'Wrap text using a more sophisticated';
```

# --hints--

`wrap` має бути функцією.

```js
assert.equal(typeof wrap, 'function');
```

`wrap` має повертати рядок.

```js
assert.equal(typeof wrap('abc', 10), 'string');
```

`wrap(text,80)` має повертати 4 рядки.

```js
assert(wrapped80.split('\n').length === 4);
```

Ваша функція `wrap` має повертати очікуваний текст.

```js
assert.equal(wrapped80.split('\n')[0], firstRow80);
```

`wrap(text,42)` має повертати 7 рядків.

```js
assert(wrapped42.split('\n').length === 7);
```

Ваша функція `wrap` має повертати очікуваний текст.

```js
assert.equal(wrapped42.split('\n')[0], firstRow42);
```

# --seed--

## --seed-contents--

```js
function wrap(text, limit) {
  return text;
}
```

# --solutions--

```js
function wrap(text, limit) {
  const noNewlines = text.replace('\n', '');
  if (noNewlines.length > limit) {
    // find the last space within limit
    const edge = noNewlines.slice(0, limit).lastIndexOf(' ');
    if (edge > 0) {
      const line = noNewlines.slice(0, edge);
      const remainder = noNewlines.slice(edge + 1);
      return line + '\n' + wrap(remainder, limit);
    }
  }
  return text;
}
```
