---
id: 594810f028c0303b75339ad5
title: Y-комбінатор
challengeType: 1
forumTopicId: 302345
dashedName: y-combinator
---

# --description--

У строгому <a href="https://www.freecodecamp.org/news/the-principles-of-functional-programming/" target="_blank" rel="noopener noreferrer nofollow">функціональному програмуванні</a> та лямбда-численні функції (лямбда-вирази) не мають стану і можуть посилатися лише на аргументи зовнішніх функцій. Це виключає звичайне визначення рекурсивної функції, у якому функція пов’язана зі станом змінної, і цей стан змінної використовується в тілі функції.

Y-комбінатор — це сама по собі функція без стану, яка, застосована до іншої функції без стану, повертає рекурсивну версію цієї функції. Y-комбінатор є найпростішим із класу таких функцій, які називають фіксованими точковими комбінаторами.

# --instructions--

Визначте функцію Y-комбінатора без стану і використайте її для обчислення факторіалів. Функція `factorial(N)` вже надана вам.

# --hints--

Y має повертати функцію.

```js
assert.equal(typeof Y((f) => (n) => n), 'function');
```

factorial(1) має повертати 1.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(1), 1);
```

factorial(2) має повертати 2.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(2), 2);
```

factorial(3) має повертати 6.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(3), 6);
```

factorial(4) має повертати 24.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(4), 24);
```

factorial(10) має повертати 3628800.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(10), 3628800);
```

# --seed--

## --seed-contents--

```js
function Y(f) {
  return function() {

  };
}

var factorial = Y(function(f) {
  return function (n) {
    return n > 1 ? n * f(n - 1) : 1;
  };
});
```

# --solutions--

```js
var Y = f => (x => x(x))(y => f(x => y(y)(x)));
```
