---
id: 594810f028c0303b75339ad2
title: Векторний векторний добуток
challengeType: 1
forumTopicId: 302342
dashedName: vector-cross-product
---

# --description--

Вектор визначається як тривимірний об’єкт, представлений впорядкованою колекцією з трьох чисел: (X, Y, Z).

# --instructions--

Напишіть функцію, яка приймає два вектори (масиви) як вхідні дані і обчислює їхній векторний добуток. Ваша функція має повертати `null` у разі недійсних вхідних даних, наприклад, якщо вектори мають різну довжину.

# --hints--

`crossProduct` має бути функцією.

```js
assert.equal(typeof crossProduct, 'function');
```

`crossProduct()` має повертати null.

```js
assert.equal(crossProduct(), null);
```

`crossProduct([1, 2, 3], [4, 5, 6])` має повертати `[-3, 6, -3]`.

```js
assert.deepEqual(res12, exp12);
```

# --before-each--

```js
const tv1 = [1, 2, 3];
const tv2 = [4, 5, 6];
const res12 = crossProduct(tv1, tv2);
const exp12 = [-3, 6, -3];
```

# --seed--

## --seed-contents--

```js
function crossProduct(a, b) {

}
```

# --solutions--

```js
function crossProduct(a, b) {
  if (!a || !b) {
    return null;
  }

  // Check lengths
  if (a.length !== 3 || b.length !== 3) {
    return null;
  }

  return [
    (a[1] * b[2]) - (a[2] * b[1]),
    (a[2] * b[0]) - (a[0] * b[2]),
    (a[0] * b[1]) - (a[1] * b[0])
  ];
}
```
