---
id: 594810f028c0303b75339ad1
title: Щасливі числа
challengeType: 1
forumTopicId: 302280
dashedName: happy-numbers
---

# --description--

Щасливе число визначається таким процесом:

Починаючи з будь-якого додатного цілого числа, замініть число сумою квадратів його цифр і повторюйте цей процес, доки число не стане `1` (і залишатиметься таким), або поки не почнеться нескінченний цикл, який не включає `1`. Числа, для яких цей процес закінчується `1`, є щасливими числами, а ті, що не закінчуються `1`, — нещасливими.

# --instructions--

Реалізуйте функцію, яка повертає true, якщо число щасливе, або false, якщо ні.

# --hints--

`happy` має бути функцією.

```js
assert(typeof happy === 'function');
```

`happy(1)` має повертати булеве значення.

```js
assert(typeof happy(1) === 'boolean');
```

`happy(1)` має повертати `true`.

```js
assert(happy(1));
```

`happy(2)` має повертати `false`.

```js
assert(!happy(2));
```

`happy(7)` має повертати `true`.

```js
assert(happy(7));
```

`happy(10)` має повертати `true`.

```js
assert(happy(10));
```

`happy(13)` має повертати `true`.

```js
assert(happy(13));
```

`happy(19)` має повертати `true`.

```js
assert(happy(19));
```

`happy(23)` має повертати `true`.

```js
assert(happy(23));
```

`happy(28)` має повертати `true`.

```js
assert(happy(28));
```

`happy(31)` має повертати `true`.

```js
assert(happy(31));
```

`happy(32)` має повертати `true`.

```js
assert(happy(32));
```

`happy(33)` має повертати `false`.

```js
assert(!happy(33));
```

# --seed--

## --seed-contents--

```js
function happy(number) {

}
```

# --solutions--

```js
function happy (number) {
  let m;
  let digit;
  const cycle = [];

  while (number !== 1 && cycle[number] !== true) {
    cycle[number] = true;
    m = 0;
    while (number > 0) {
      digit = number % 10;
      m += Math.pow(digit, 2);
      number = (number - digit) / 10;
    }
    number = m;
  }
  return (number === 1);
}
```
