---
id: 594810f028c0303b75339acc
title: Завдання з блоками ABC
challengeType: 1
forumTopicId: 302220
dashedName: abc-problem
---

# --description--

Вам надано набір блоків ABC (наприклад, дитячі алфавітні блоки). Є 20 блоків, на кожному з яких по дві літери. Повний алфавіт гарантовано присутній серед усіх боків блоків. Приклад набору блоків:

<pre>(B O)
(X K)
(D Q)
(C P)
(N A)
(G T)
(R E)
(T G)
(Q D)
(F S)
(J W)
(H U)
(V I)
(A N)
(O B)
(E R)
(F S)
(L Y)
(P C)
(Z M)
</pre>

# --instructions--

Реалізуйте функцію, яка приймає рядок (слово) і визначає, чи можна це слово скласти з наданого набору блоків.

Декілька правил, які слід пам’ятати:

<ul>
  <li>Якщо літера на блоці вже використана, цей блок більше не можна використовувати.</li>
  <li>Функція має бути нечутливою до регістру.</li>
</ul>

# --before-each--

```js
const words = ['bark', 'BooK', 'TReAT', 'COMMON', 'squAD', 'conFUSE'];
```

# --hints--

`canMakeWord` має бути функцією.

```js
assert(typeof canMakeWord === 'function');
```

`canMakeWord` має повертати булеве значення.

```js
assert(typeof canMakeWord('hi') === 'boolean');
```

`canMakeWord("bark")` має повертати true.

```js
assert(canMakeWord(words[0]));
```

`canMakeWord("BooK")` має повертати false.

```js
assert(!canMakeWord(words[1]));
```

`canMakeWord("TReAT")` має повертати true.

```js
assert(canMakeWord(words[2]));
```

`canMakeWord("COMMON")` має повертати false.

```js
assert(!canMakeWord(words[3]));
```

`canMakeWord("squAD")` має повертати true.

```js
assert(canMakeWord(words[4]));
```

`canMakeWord("conFUSE")` має повертати true.

```js
assert(canMakeWord(words[5]));
```

# --seed--

## --seed-contents--

```js
function canMakeWord(word) {

}
```

# --solutions--

```js
function canMakeWord(word) {
  const characters = 'BO XK DQ CP NA GT RE TG QD FS JW HU VI AN OB ER FS LY PC ZM';
  const blocks = characters.split(' ').map(pair => pair.split(''));

  const letters = [...word.toUpperCase()];
  let length = letters.length;
  const copy = new Set(blocks);

  letters.forEach(letter => {
    for (let block of copy) {
      const index = block.indexOf(letter);

      if (index !== -1) {
        length--;
        copy.delete(block);
        break;
      }
    }
  });
  return !length;
}
```
