---
id: 594810f028c0303b75339acb
title: 100 дверей
challengeType: 1
forumTopicId: 302217
dashedName: 100-doors
---

# --description--

Є 100 дверей у ряд, які спочатку всі зачинені. Ви проходите повз двері 100 разів. Першого разу відвідайте кожні двері і «переключіть» їх (якщо двері зачинені, відкрийте їх; якщо відкриті — зачиніть). Другого разу відвідайте лише кожні другі двері (тобто двері №2, №4, №6, ...) і переключіть їх. Третього разу відвідайте кожні треті двері (тобто двері №3, №6, №9, ...) і так далі, доки не відвідаєте лише 100-ті двері.

# --instructions--

Реалізуйте функцію, щоб визначити стан дверей після останнього проходу. Поверніть фінальний результат у масиві, включаючи лише номер дверей, якщо вони відкриті.

# --hints--

`getFinalOpenedDoors` має бути функцією.

```js
assert(typeof getFinalOpenedDoors === 'function');
```

`getFinalOpenedDoors` має повертати масив.

```js
assert(Array.isArray(getFinalOpenedDoors(100)));
```

`getFinalOpenedDoors` має давати правильний результат.

```js
const solution = [1, 4, 9, 16, 25, 36, 49, 64, 81, 100];
assert.deepEqual(getFinalOpenedDoors(100), solution);
```

# --seed--

## --seed-contents--

```js
function getFinalOpenedDoors(numDoors) {

}
```

# --solutions--

```js
function getFinalOpenedDoors(numDoors) {
  // this is the final pattern (always squares).
  // thus, the most efficient solution simply returns an array of squares up to numDoors).
  const finalState = [];
  let i = 1;
  while (Math.pow(i, 2) <= numDoors) {
    finalState.push(Math.pow(i, 2));
    i++;
  }
  return finalState;
}
```
