---
id: 594dc6c729e5700999302b45
title: Збалансовані дужки
challengeType: 1
forumTopicId: 302230
dashedName: balanced-brackets
---

# --description--

Визначте, чи є згенерований рядок дужок збалансованим; тобто чи складається він повністю з пар відкриваючих/закриваючих дужок (у такому порядку), жодна з яких не вкладана неправильно.

**Приклади:**
| Вхідні дані               | Вихідні дані |
| ------------------------- | ------------ |
| <code>\[]</code>          | true         |
| <code>]\[</code>          | false        |
| <code>[][]</code>         | true         |
| <code>]\[]</code>         | false        |
| <code>\[]]\[\[]</code>    | false        |
| <code>\[\[\[\[]]]]</code> | true         |

# --before-each--

```js
const testCases = [
  '[]',
  ']][[[][][][]][',
  '[][[[[][][[[]]]]]]',
  '][',
  '[[[]]]][[]',
  '][[]',
  '][[][]][[[]]',
  '[[][]]][',
  '[[[]]][[]]]][][[',
  '[]][[]]][[[[][]]',
  '][]][[][',
  '[[]][[][]]',
  '[[]]',
  ']][]][[]][[[',
  '][]][][[',
  '][][',
  '[]]]',
  ''
];
```

# --hints--

`isBalanced` має бути функцією.

```js
assert(typeof isBalanced === 'function');
```

`isBalanced("[]")` має повертати true.

```js
assert(isBalanced(testCases[0]));
```

`isBalanced("]][[[][][][]][")` має повертати false.

```js
assert(!isBalanced(testCases[1]));
```

`isBalanced("[][[[[][][[[]]]]]]")` має повертати true.

```js
assert(isBalanced(testCases[2]));
```

`isBalanced("][")` має повертати false.

```js
assert(!isBalanced(testCases[3]));
```

`isBalanced("[[[]]]][[]")` має повертати false.

```js
assert(!isBalanced(testCases[4]));
```

`isBalanced("][[]")` має повертати false.

```js
assert(!isBalanced(testCases[5]));
```

`isBalanced("][[][]][[[]]")` має повертати false.

```js
assert(!isBalanced(testCases[6]));
```

`isBalanced("[[][]]][")` має повертати false.

```js
assert(!isBalanced(testCases[7]));
```

`isBalanced("[[[]]][[]]]][][[")` має повертати false.

```js
assert(!isBalanced(testCases[8]));
```

`isBalanced("[]][[]]][[[[][]]")` має повертати false.

```js
assert(!isBalanced(testCases[9]));
```

`isBalanced("][]][[][")` має повертати false.

```js
assert(!isBalanced(testCases[10]));
```

`isBalanced("[[]][[][]]")` має повертати true.

```js
assert(isBalanced(testCases[11]));
```

`isBalanced("[[]]")` має повертати true.

```js
assert(isBalanced(testCases[12]));
```

`isBalanced("]][]][[]][[[")` має повертати false.

```js
assert(!isBalanced(testCases[13]));
```

`isBalanced("][]][][[")` має повертати false.

```js
assert(!isBalanced(testCases[14]));
```

`isBalanced("][][")` має повертати false.

```js
assert(!isBalanced(testCases[15]));
```

`isBalanced("[]]]")` має повертати false.

```js
assert(!isBalanced(testCases[16]));
```

`isBalanced("")` має повертати true.

```js
assert(isBalanced(testCases[17]));
```

# --seed--

## --seed-contents--

```js
function isBalanced(str) {

  return true;
}
```

# --solutions--

```js
function isBalanced(str) {
  if (str === '') return true;
  let a = str;
  let b;
  do {
    b = a;
    a = a.replace(/\[\]/g, '');
  } while (a !== b);
  return !a;
}
```
