---
id: 594faaab4e2a8626833e9c3d
title: Розбити рядок з екрануванням
challengeType: 1
forumTopicId: 302338
dashedName: tokenize-a-string-with-escaping
---

# --description--

Напишіть функцію або програму, яка може розбити рядок на кожному неекранованому входженні символу-роздільника.

Вона має приймати три вхідні параметри:

<ul>
  <li><strong>рядок</strong></li>
  <li><strong>символ-роздільник</strong></li>
  <li><strong>символ екранування</strong></li>
</ul>

Вона має повертати список рядків.

Правила розбиття:

<ul>
  <li>Поля, розділені роздільниками, стають елементами вихідного списку.</li>
  <li>Порожні поля слід зберігати, навіть на початку та в кінці.</li>
</ul>

Правила екранування:

<ul>
  <li>"Екранований" означає, що перед символом стоїть символ екранування, який сам не є екранованим.</li>
  <li>Якщо символ екранування стоїть перед символом, який не має спеціального значення, це все одно вважається екрануванням (але не виконує жодної особливої дії).</li>
  <li>Кожне входження символу екранування, яке використовується для екранування, не повинно потрапляти у вихідні дані.</li>
</ul>

Покажіть, що ваша функція задовольняє наступний тест:

Дано рядок

<pre>one^|uno||three^^^^|four^^^|^cuatro|</pre>

і використовуючи `|` як роздільник та `^` як символ екранування, ваша функція має повернути наступний масив:

<pre>  ['one|uno', '', 'three^^', 'four^|cuatro', '']
</pre>

# --before-each--

```js
const testStr1 = 'one^|uno||three^^^^|four^^^|^cuatro|';
const res1 = ['one|uno', '', 'three^^', 'four^|cuatro', ''];

const testStr2 = 'a@&bcd&ef&&@@hi';
const res2 = ['a&bcd', 'ef', '', '@hi'];

const testStr3 = 'hello^|world|how^are^you^|';
const res3 = ['hello|world', 'howareyou|'];
```

# --hints--

`tokenize` має бути функцією.

```js
assert(typeof tokenize === 'function');
```

`tokenize` має повертати масив.

```js
assert(typeof tokenize('a', 'b', 'c') === 'object');
```

`tokenize('one^|uno||three^^^^|four^^^|^cuatro|', '|', '^')` має повертати `['one|uno', '', 'three^^', 'four^|cuatro', '']`

```js
assert.deepEqual(tokenize(testStr1, '|', '^'), res1);
```

`tokenize('a@&bcd&ef&&@@hi', '&', '@')` має повертати `['a&bcd', 'ef', '', '@hi']`

```js
assert.deepEqual(tokenize(testStr2, '&', '@'), res2);
```

`tokenize('hello^|world|how^are^you^|', '|', '^')` має повертати `['hello|world', 'howareyou|']`

```js
assert.deepEqual(tokenize(testStr3, '|', '^'), res3);
```

# --seed--

## --seed-contents--

```js
function tokenize(str, sep, esc) {
  return true;
}
```

# --solutions--

```js
// tokenize :: String -> Character -> Character -> [String]
function tokenize(str, charDelim, charEsc) {
  const dctParse = str.split('')
    .reduce((a, x) => {
      const blnEsc = a.esc;
      const blnBreak = !blnEsc && x === charDelim;
      const blnEscChar = !blnEsc && x === charEsc;

      return {
        esc: blnEscChar,
        token: blnBreak ? '' : (
          a.token + (blnEscChar ? '' : x)
        ),
        list: a.list.concat(blnBreak ? a.token : [])
      };
    }, {
      esc: false,
      token: '',
      list: []
    });

  return dctParse.list.concat(
    dctParse.token
  );
}
```
