---
id: 5951a53863c8a34f02bf1bdc
title: Завдання пошуку найближчої пари
challengeType: 1
forumTopicId: 302232
dashedName: closest-pair-problem
---

# --description--

Напишіть функцію, яка знаходить дві найближчі точки серед заданої множини точок у двовимірному просторі.

Просте рішення — це алгоритм $O(n^2)$ (його можна назвати *алгоритмом грубої сили*); псевдокод (з використанням індексів) може виглядати так:

<pre><strong>bruteForceClosestPair</strong> для P(1), P(2), ... P(N)
<strong>якщо</strong> N &#x3C; 2 <strong>то</strong>
  <strong>повернути</strong> ∞
<strong>інакше</strong>
  minDistance ← |P(1) - P(2)|
  minPoints ← { P(1), P(2) }
  <strong>для кожного</strong> i ∈ [1, N-1]
    <strong>для кожного</strong> j ∈ [i+1, N]
      <strong>якщо</strong> |P(i) - P(j)| &#x3C; minDistance <strong>то</strong>
        minDistance ← |P(i) - P(j)|
        minPoints ← { P(i), P(j) }
      <strong>кінець якщо</strong>
    <strong>кінець для</strong>
  <strong>кінець для</strong>
  <strong>повернути</strong> minDistance, minPoints
<strong>кінець якщо</strong>
</pre>

Кращий алгоритм базується на рекурсивному підході "розділяй і володарюй", який працює за $O(n\log n)$; псевдокод може бути таким:

<pre><strong>closestPair</strong> для (xP, yP)
  де xP — це P(1) .. P(N), відсортовані за координатою x, а
  yP — це P(1) .. P(N), відсортовані за координатою y (за зростанням)
<strong>якщо</strong> N ≤ 3 <strong>то</strong>
  <strong>повернути</strong> найближчі точки з xP за допомогою алгоритму грубої сили
<strong>інакше</strong>
  xL ← точки з xP від 1 до ⌈N/2⌉
  xR ← точки з xP від ⌈N/2⌉+1 до N
  xm ← xP(⌈N/2⌉)<sub>x</sub>
  yL ← { p ∈ yP : p<sub>x</sub> ≤ xm }
  yR ← { p ∈ yP : p<sub>x</sub> > xm }
  (dL, pairL) ← closestPair для (xL, yL)
  (dR, pairR) ← closestPair для (xR, yR)
  (dmin, pairMin) ← (dR, pairR)
  <strong>якщо</strong> dL &#x3C; dR <strong>то</strong>
    (dmin, pairMin) ← (dL, pairL)
  <strong>кінець якщо</strong>
  yS ← { p ∈ yP : |xm - p<sub>x</sub>| &#x3C; dmin }
  nS ← кількість точок у yS
  (closest, closestPair) ← (dmin, pairMin)
  <strong>для</strong> i <strong>від</strong> 1 <strong>до</strong> nS - 1
    k ← i + 1
    <strong>поки</strong> k ≤ nS <strong>та</strong> yS(k)<sub>y</sub> - yS(i)<sub>y</sub> &#x3C; dmin
      <strong>якщо</strong> |yS(k) - yS(i)| &#x3C; closest <strong>то</strong>
        (closest, closestPair) ← (|yS(k) - yS(i)|, {yS(k), yS(i)})
      <strong>кінець якщо</strong>
      k ← k + 1
    <strong>кінець поки</strong>
  <strong>кінець для</strong>
  <strong>повернути</strong> closest, closestPair
<strong>кінець якщо</strong>
</pre>

Для вхідних даних очікуйте аргумент у вигляді масиву `Point` об’єктів із членами `x` і `y`, встановленими у числа. Поверніть об’єкт, що містить пари ключ:значення для `distance` і `pair` (пару двох найближчих точок).

Наприклад, `getClosestPair` з вхідним масивом `points`:

```js
const points = [
  new Point(1, 2),
  new Point(3, 3),
  new Point(2, 2)
];
```

Поверне:

```js
{
  distance: 1,
  pair: [
    {
      x: 1,
      y: 2
    },
    {
      x: 2,
      y: 2
    }
  ]
}
```

**Примітка:** Відсортуйте масив `pair` за їхніми значеннями `x` у зростаючому порядку.


# --before-each--

```js
function getPoints1_() {
  return [
    new Point(0.748501, 4.09624),
    new Point(3.00302, 5.26164),
    new Point(3.61878, 9.52232),
    new Point(7.46911, 4.71611),
    new Point(5.7819, 2.69367),
    new Point(2.34709, 8.74782),
    new Point(2.87169, 5.97774),
    new Point(6.33101, 0.463131),
    new Point(7.46489, 4.6268),
    new Point(1.45428, 0.087596)
  ];
}

const answer1 = {
  distance: 0.0894096443343775,
  pair: [
    { x: 7.46489, y: 4.6268 },
    { x: 7.46911, y: 4.71611 }
  ]
};

function getPoints2_() {
  return [
    new Point(37100, 13118),
    new Point(37134, 1963),
    new Point(37181, 2008),
    new Point(37276, 21611),
    new Point(37307, 9320)
  ];
}

const answer2 = {
  distance: 65.06919393998976,
  pair: [
    { x: 37134, y: 1963 },
    { x: 37181, y: 2008 }
  ]
};

function getPoints3_() {
  return [
    new Point(16910, 54699),
    new Point(14773, 61107),
    new Point(95547, 45344),
    new Point(95951, 17573),
    new Point(5824, 41072),
    new Point(8769, 52562),
    new Point(21182, 41881),
    new Point(53226, 45749),
    new Point(68180, 887),
    new Point(29322, 44017),
    new Point(46817, 64975),
    new Point(10501, 483),
    new Point(57094, 60703),
    new Point(23318, 35472),
    new Point(72452, 88070),
    new Point(67775, 28659),
    new Point(19450, 20518),
    new Point(17314, 26927),
    new Point(98088, 11164),
    new Point(25050, 56835),
    new Point(8364, 6892),
    new Point(37868, 18382),
    new Point(23723, 7701),
    new Point(55767, 11569),
    new Point(70721, 66707),
    new Point(31863, 9837),
    new Point(49358, 30795),
    new Point(13041, 39744),
    new Point(59635, 26523),
    new Point(25859, 1292),
    new Point(1551, 53890),
    new Point(70316, 94479),
    new Point(48549, 86338),
    new Point(46413, 92747),
    new Point(27186, 50426),
    new Point(27591, 22655),
    new Point(10905, 46153),
    new Point(40408, 84202),
    new Point(52821, 73520),
    new Point(84865, 77388),
    new Point(99819, 32527),
    new Point(34404, 75657),
    new Point(78457, 96615),
    new Point(42140, 5564),
    new Point(62175, 92342),
    new Point(54958, 67112),
    new Point(4092, 19709),
    new Point(99415, 60298),
    new Point(51090, 52158),
    new Point(48953, 58567)
  ];
}

const answer3 = {
  distance: 6754.625082119658,
  pair: [
    { x: 46817, y: 64975 },
    { x: 48953, y: 58567 }
  ]
};

function getClosestPair1_() {
  return getClosestPair(getPoints1_());
}

function getClosestPair2_() {
  return getClosestPair(getPoints2_());
}

function getClosestPair3_() {
  return getClosestPair(getPoints3_());
}
```

# --hints--

`getClosestPair` має бути функцією.

```js
assert(typeof getClosestPair === 'function');
```

`getClosestPair(points1).distance` має бути `0.0894096443343775`.

```js
assert.equal(getClosestPair1_().distance, answer1.distance);
```

`getClosestPair(points1).pair` має бути `[ { x: 7.46489, y: 4.6268 }, { x: 7.46911, y: 4.71611 } ]`.

```js
assert.deepEqual(
  JSON.parse(JSON.stringify(getClosestPair1_())).pair,
  answer1.pair
);
```

`getClosestPair(points2).distance` має бути `65.06919393998976`.

```js
assert.equal(getClosestPair2_().distance, answer2.distance);
```

`getClosestPair(points2).pair` має бути `[ { x: 37134, y: 1963 }, { x: 37181, y: 2008 } ]`.

```js
assert.deepEqual(
  JSON.parse(JSON.stringify(getClosestPair2_())).pair,
  answer2.pair
);
```

`getClosestPair(points3).distance` має бути `6754.625082119658`.

```js
assert.equal(getClosestPair3_().distance, answer3.distance);
```

`getClosestPair(points3).pair` має бути `[ { x: 46817, y: 64975 }, { x: 48953, y: 58567 } ]`.

```js
assert.deepEqual(
  JSON.parse(JSON.stringify(getClosestPair3_())).pair,
  answer3.pair
);
```

# --seed--

## --seed-contents--

```js
const Point = function(x, y) {
  this.x = x;
  this.y = y;
};
Point.prototype.getX = function() {
  return this.x;
};
Point.prototype.getY = function() {
  return this.y;
};

function getClosestPair(pointsArr) {

  return true;
}
```

# --solutions--

```js
const Point = function(x, y) {
  this.x = x;
  this.y = y;
};
Point.prototype.getX = function() {
  return this.x;
};
Point.prototype.getY = function() {
  return this.y;
};

const mergeSort = function mergeSort(points, comp) {
    if(points.length < 2) return points;

    var n = points.length,
        i = 0,
        j = 0,
        leftN = Math.floor(n / 2),
        rightN = leftN;

    var leftPart = mergeSort( points.slice(0, leftN), comp),
        rightPart = mergeSort( points.slice(rightN), comp );

    var sortedPart = [];

    while((i < leftPart.length) && (j < rightPart.length)) {
        if(comp(leftPart[i], rightPart[j]) < 0) {
            sortedPart.push(leftPart[i]);
            i += 1;
        }
        else {
            sortedPart.push(rightPart[j]);
            j += 1;
        }
    }
    while(i < leftPart.length) {
        sortedPart.push(leftPart[i]);
        i += 1;
    }
    while(j < rightPart.length) {
        sortedPart.push(rightPart[j]);
        j += 1;
    }
    return sortedPart;
};

const closestPair = function _closestPair(Px, Py) {
    if(Px.length < 2) return { distance: Infinity, pair: [ new Point(0, 0), new Point(0, 0) ] };
    if(Px.length < 3) {
        //find euclid distance
        var d = Math.sqrt( Math.pow(Math.abs(Px[1].x - Px[0].x), 2) + Math.pow(Math.abs(Px[1].y - Px[0].y), 2) );
        return {
            distance: d,
            pair: [ Px[0], Px[1] ]
        };
    }

    var n = Px.length,
        leftN = Math.floor(n / 2),
        rightN = leftN;

    var Xl = Px.slice(0, leftN),
        Xr = Px.slice(rightN),
        Xm = Xl[leftN - 1],
        Yl = [],
        Yr = [];
    //separate Py
    for(var i = 0; i < Py.length; i += 1) {
        if(Py[i].x <= Xm.x)
            Yl.push(Py[i]);
        else
            Yr.push(Py[i]);
    }

    var dLeft = _closestPair(Xl, Yl),
        dRight = _closestPair(Xr, Yr);

    var minDelta = dLeft.distance,
        closestPair = dLeft.pair;
    if(dLeft.distance > dRight.distance) {
        minDelta = dRight.distance;
        closestPair = dRight.pair;
    }

    //filter points around Xm within delta (minDelta)
    var closeY = [];
    for(i = 0; i < Py.length; i += 1) {
        if(Math.abs(Py[i].x - Xm.x) < minDelta) closeY.push(Py[i]);
    }
    //find min within delta. 8 steps max
    for(i = 0; i < closeY.length; i += 1) {
        for(var j = i + 1; j < Math.min( (i + 8), closeY.length ); j += 1) {
            var d = Math.sqrt( Math.pow(Math.abs(closeY[j].x - closeY[i].x), 2) + Math.pow(Math.abs(closeY[j].y - closeY[i].y), 2) );
            if(d < minDelta) {
                minDelta = d;
                closestPair = [ closeY[i], closeY[j] ]
            }
        }
    }

    return {
        distance: minDelta,
        pair: closestPair.sort((pointA, pointB) => pointA.x - pointB.x)
    };
};

function getClosestPair(points) {
  const sortX = function(a, b) { return (a.x < b.x) ? -1 : ((a.x > b.x) ? 1 : 0); }
  const sortY = function(a, b) { return (a.y < b.y) ? -1 : ((a.y > b.y) ? 1 : 0); }

  const Px = mergeSort(points, sortX);
  const Py = mergeSort(points, sortY);

  return closestPair(Px, Py);
}
```
