---
id: 5951ed8945deab770972ae56
title: Вежі Ханоя
challengeType: 1
forumTopicId: 302341
dashedName: towers-of-hanoi
---

# --description--

Розв’яжіть задачу веж Ханоя. Кількість об’єктів буде задана першим параметром, за яким слідують рядки, що ідентифікують кожен стос об’єктів. Створіть вкладений масив, що містить список ходів, `["source", "destination"]`.

Наприклад, параметри `(4, 'A', 'B', 'C')` призведуть до вкладеного масиву ходів `[['A', 'C'], ['B', 'A']]`, що вказує, що перший хід полягав у переміщенні об’єкта зі стосу `A` до `C`, а другий хід — у переміщенні об’єкта зі стосу `B` до `A`.

# --instructions--

Напишіть функцію, яка повертає ходи для складання об’єктів у вкладеному масиві.

# --before-each--

```js
const res3Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B']];
const res7First10Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B'], ['A', 'C'], ['B', 'C'], ['B', 'A']];

function getTowerResults_() {
  return {
    res3: towerOfHanoi(3, 'A', 'B', 'C'),
    res5: towerOfHanoi(5, 'X', 'Y', 'Z'),
    res7: towerOfHanoi(7, 'A', 'B', 'C')
  };
}
```

# --hints--

`towerOfHanoi` має бути функцією.

```js
assert(typeof towerOfHanoi === 'function');
```

`towerOfHanoi(3, ...)` має повертати 7 ходів.

```js
assert(getTowerResults_().res3.length === 7);
```

`towerOfHanoi(3, 'A', 'B', 'C')` має повертати `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B']]`.

```js
assert.deepEqual(getTowerResults_().res3, res3Moves);
```

`towerOfHanoi(5, "X", "Y", "Z")` десятий хід має бути Y -> X.

```js
const { res5 } = getTowerResults_();
assert.deepEqual(res5[9], ['Y', 'X']);
```

`towerOfHanoi(7, 'A', 'B', 'C')` перші десять ходів мають бути `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B'], ['A','C'], ['B','C'], ['B','A']]`

```js
const { res7 } = getTowerResults_();
assert.deepEqual(res7.slice(0, 10), res7First10Moves);
```

# --seed--

## --seed-contents--

```js
function towerOfHanoi(n, a, b, c) {

  return [[]];
}
```

# --solutions--

```js
function towerOfHanoi(n, a, b, c) {
  const res = [];
  towerOfHanoiHelper(n, a, c, b, res);
  return res;
}

function towerOfHanoiHelper(n, a, b, c, res) {
  if (n > 0) {
    towerOfHanoiHelper(n - 1, a, c, b, res);
    res.push([a, c]);
    towerOfHanoiHelper(n - 1, b, a, c, res);
  }
}
```
