---
id: 595608ff8bcd7a50bd490181
title: послідовність hailstone
challengeType: 1
forumTopicId: 302279
dashedName: hailstone-sequence
---

# --description--

Послідовність hailstone чисел можна згенерувати з початкового додатного цілого числа `n` за допомогою:

- Якщо `n` дорівнює `1`, тоді послідовність завершується
- Якщо `n` є `even`, тоді наступне `n` послідовності є `= n/2`
- Якщо `n` є `odd`, тоді наступне `n` послідовності є `= (3 * n) + 1`

(Недоведена) гіпотеза Коллатца стверджує, що послідовність hailstone для будь-якого початкового числа завжди завершується.

Послідовність hailstone також відома як hailstone числа (тому що значення зазвичай піддаються кільком спадам і підйомам, як градини в хмарі), або як послідовність Коллатца.

# --instructions--

1. Створіть процедуру для генерації послідовності hailstone для числа
2. Ваша функція має повертати масив з числом меншим за `limit`, яке має найдовшу послідовність hailstone, та довжиною цієї послідовності. (Але не показуйте саму послідовність!)

# --hints--

`hailstoneSequence` має бути функцією.

```js
assert(typeof hailstoneSequence === 'function');
```

`hailstoneSequence(30)` має повертати масив.

```js
assert(Array.isArray(hailstoneSequence(30)));
```

`hailstoneSequence(30)` має повертати `[27, 112]`.

```js
assert.deepEqual(hailstoneSequence(30), [27, 112]);
```

`hailstoneSequence(50000)` має повертати `[35655, 324]`.

```js
assert.deepEqual(hailstoneSequence(50000), [35655, 324]);
```

`hailstoneSequence(100000)` має повертати `[77031, 351]`.

```js
assert.deepEqual(hailstoneSequence(100000), [77031, 351]);
```

# --seed--

## --seed-contents--

```js
function hailstoneSequence(limit) {
  const res = [];


  return res;
}
```

# --solutions--

```js
function hailstoneSequence (limit) {
  function hailstone(n) {
    const seq = [n];
    while (n > 1) {
      n = n % 2 ? 3 * n + 1 : n / 2;
      seq.push(n);
    }
    return seq;
  }

  let n = 0;
  let max = 0;
  for (let i = limit; --i;) {
    const seq = hailstone(i);
    const sLen = seq.length;

    if (sLen > max) {
      n = i;
      max = sLen;
    }
  }

  return [n, max];
}
```
