---
id: 59637c4d89f6786115efd814
title: послідовність Хофштадтера Q
challengeType: 1
forumTopicId: 302287
dashedName: hofstadter-q-sequence
---

# --description--

Послідовність Хофштадтера Q визначається так:

$Q(1)=Q(2)=1, \\\\ Q(n)=Q\\big(n-Q(n-1)\\big)+Q\\big(n-Q(n-2)), \\quad n>2.$

Вона визначена подібно до послідовності Фібоначчі, але якщо наступний член послідовності Фібоначчі — це сума двох попередніх членів, то в послідовності Q два попередні члени вказують, на скільки кроків назад у послідовності Q потрібно піти, щоб знайти два числа для суми, яка утворює наступний член послідовності.

# --instructions--

Реалізуйте рівняння послідовності Хофштадтера Q у вигляді функції. Функція має приймати число `n` і повертати ціле число.

# --before-each--

```js
const testCase = [1000, 1500, 2000, 2500];
const res = [502, 755, 1005, 1261];
```

# --hints--

`hofstadterQ` має бути функцією.

```js
assert(typeof hofstadterQ === 'function');
```

`hofstadterQ()` має повертати `integer`

```js
assert(Number.isInteger(hofstadterQ(1000)));
```

`hofstadterQ(1000)` має повертати `502`

```js
assert.equal(hofstadterQ(testCase[0]), res[0]);
```

`hofstadterQ(1500)` має повертати `755`

```js
assert.equal(hofstadterQ(testCase[1]), res[1]);
```

`hofstadterQ(2000)` має повертати `1005`

```js
assert.equal(hofstadterQ(testCase[2]), res[2]);
```

`hofstadterQ(2500)` має повертати `1261`

```js
assert.equal(hofstadterQ(testCase[3]), res[3]);
```

# --seed--

## --seed-contents--

```js
function hofstadterQ(n) {

  return n;
}
```

# --solutions--

```js
function hofstadterQ (n) {
  const memo = [1, 1, 1];
  const Q = function (i) {
    let result = memo[i];
    if (typeof result !== 'number') {
      result = Q(i - Q(i - 1)) + Q(i - Q(i - 2));
      memo[i] = result;
    }
    return result;
  };
  return Q(n);
}
```
