---
id: 59667989bf71cf555dd5d2ff
title: s-вирази
challengeType: 1
forumTopicId: 302303
dashedName: s-expressions
---

# --description--

<a href="https://rosettacode.org/wiki/S-expressions" target="_blank" rel="noopener noreferrer nofollow">S-вирази</a> — це зручний спосіб аналізувати та зберігати дані.

# --instructions--

Напишіть простий рідер/парсер для S-виразів, який обробляє рядки в лапках і без, цілі числа та числа з плаваючою крапкою.

Функція має прочитати один, але вкладений S-вираз із рядка і повернути його як (вкладений) масив.

Переноси рядків та інші пробіли можна ігнорувати, якщо вони не містяться в рядку в лапках.

"`()`" всередині рядків у лапках не інтерпретуються, а розглядаються як частина рядка.

Обробка екранування лапок всередині рядка є необов’язковою; отже, "`(foo"bar)`" можна розглядати як рядок "`foo"bar`", або як помилку.

Для цього рідеру не потрібно розпізнавати `\` для екранування, але він має також розпізнавати числа, якщо мова має відповідні типи даних.

Зверніть увагу, що за винятком `()"` (`\`, якщо підтримується екранування) і пробілів, спеціальних символів немає. Все інше дозволено без лапок.

Рідер має вміти прочитати такий вхідний рядок

<pre>((data "quoted data" 123 4.5)
(data (!@# (4.5) "(more" "data)")))
</pre>

і перетворити його на рідну структуру даних.

# --before-each--

```js
const simpleSExpr = '(data1 data2 data3)';
const simpleSolution = ['data1', 'data2', 'data3'];

const basicSExpr = '((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))';
const basicSolution = [["data","\"quoted data\"",123,4.5],["data",["!@#",[4.5],"\"(more\"","\"data)\""]]];
```

# --hints--

`parseSexpr` має бути функцією.

```js
assert(typeof parseSexpr === 'function');
```

`parseSexpr('(data1 data2 data3)')` має повертати `['data1', 'data2', 'data3']`

```js
assert.deepEqual(parseSexpr(simpleSExpr), simpleSolution);
```

`parseSexpr('((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))')` має повертати `[['data', '"quoted data"', 123, 4.5], ['data', ['!@#', [4.5], '"(more"', '"data)"']]]`.

```js
assert.deepEqual(parseSexpr(basicSExpr), basicSolution);
```

# --seed--

## --seed-contents--

```js
function parseSexpr(str) {

  return true;
}
```

# --solutions--

```js
function parseSexpr(str) {
  const t = str.match(/\s*("[^"]*"|\(|\)|"|[^\s()"]+)/g);
  for (var o, c = 0, i = t.length - 1; i >= 0; i--) {
    var n,
      ti = t[i].trim();
    if (ti == '"') return;
    else if (ti == '(') t[i] = '[', c += 1;
    else if (ti == ')') t[i] = ']', c -= 1;
    else if ((n = +ti) == ti) t[i] = n;
    else t[i] = `'${ti.replace('\'', '\\\'')}'`;
    if (i > 0 && ti != ']' && t[i - 1].trim() != '(') t.splice(i, 0, ',');
    if (!c) if (!o) o = true; else return;
  }
  return c ? undefined : eval(t.join(''));
}
```
