---
id: 59669d08d75b60482359409f
title: формат дати
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

Поверніть масив із двох рядків дати поточного дня з такими вимогами:

- Порядок дати в першому рядку має бути: число року, число місяця та число дня, розділені дефісами (`-`).
- Рік у першому рядку має містити чотири цифри.
- Місяць і день у першому рядку не повинні містити провідних нулів.
- Назви днів тижня та місяців у другому рядку не повинні бути скороченими.
- День у другому рядку не повинен містити провідних нулів.

Приклади вихідних даних:

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

`getDateFormats` має бути функцією.

```js
assert(typeof getDateFormats === 'function');
```

`getDateFormats` має повертати об’єкт.

```js
assert(typeof getDateFormats() === 'object');
```

`getDateFormats` має повертати масив із 2 елементів.

```js
assert(getDateFormats().length === 2);
```

`getDateFormats` має повертати правильну дату у потрібному форматі.

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
