---
id: 596a8888ab7c01048de257d5
title: Глибока копія
challengeType: 1
forumTopicId: 302247
dashedName: deepcopy
---

# --description--

Напишіть функцію, яка повертає глибоку копію заданого об’єкта. Копія не повинна бути тим самим об’єктом, що було передано.

Це завдання не перевірятиме:

<ul>
  <li>Об’єкти з пропами, які є функціями</li>
  <li>Об’єкти Date або об’єкти з пропами, які є Date</li>
  <li>Регулярні вирази або об’єкти з пропами, які є регулярними виразами</li>
  <li>Прототип копіювання</li>
</ul>

# --before-each--

```js
const obj1 = { test: 'test' };
const obj2 = {
  t: 'test',
  a: ['an', 'array']
};
const obj3 = {
  t: 'try',
  o: obj2
};
```

# --hints--

`deepcopy` має бути функцією.

```js
assert(typeof deepcopy === 'function');
```

`deepcopy({test: "test"})` має повертати об’єкт.

```js
assert(typeof deepcopy(obj1) === 'object');
``` 

`deepcopy` не має повертати той самий об’єкт, що було передано.

```js
assert(deepcopy(obj2) != obj2);
```

Якщо переданий об’єкт містить масив, `deepcopy` має повертати глибоку копію об’єкта.

```js
assert.deepEqual(deepcopy(obj2), obj2);
```

Якщо переданий об’єкт містить інший об’єкт, `deepcopy` має повертати глибоку копію об’єкта.

```js
assert.deepEqual(deepcopy(obj3), obj3);
```

# --seed--

## --seed-contents--

```js
function deepcopy(obj) {

  return true;
}
```

# --solutions--

```js
function deepcopy(obj) {
  return JSON.parse(JSON.stringify(obj));
}
```
