---
id: 596e414344c3b2872167f0fe
title: Підступність коми
challengeType: 1
forumTopicId: 302234
dashedName: comma-quibbling
---

# --description--

<a href="https://rosettacode.org/wiki/Comma_quibbling" target="_blank" rel="noopener noreferrer nofollow">Підступність коми</a> — це завдання, яке спочатку запропонував Ерік Ліпперт у своєму блозі.

# --instructions--

Напишіть функцію, яка генерує рядок-вихід, що є конкатенацією вхідних слів зі списку/послідовності, де:

<ol>
  <li>Якщо вхідні дані не містять слів, вихідний рядок має бути лише з двох фігурних дужок (<code>"{}"</code>)</li>
  <li>Якщо вхідні дані містять лише одне слово, наприклад <code>["ABC"]</code>, вихідний рядок має містити це слово всередині двох дужок, наприклад <code>"{ABC}"</code></li>
  <li>Якщо вхідні дані містять два слова, наприклад <code>["ABC", "DEF"]</code>, вихідний рядок має містити обидва слова всередині дужок, розділені рядком <code>" and "</code>, наприклад <code>"{ABC and DEF}"</code></li>
  <li>Якщо вхідні дані містять три або більше слів, наприклад <code>["ABC", "DEF", "G", "H"]</code>, вихідний рядок має містити всі слова, крім останнього, розділені <code>", "</code>, а останнє слово — через <code>" and "</code>, і все це всередині дужок; наприклад <code>"{ABC, DEF, G and H}"</code></li>
</ol>

Перевірте свою функцію на наступних вхідних даних і покажіть вихідні дані тут, на цій сторінці:

<ul>
  <li>[] # (немає вхідних слів).</li>
  <li>["ABC"]</li>
  <li>["ABC", "DEF"]</li>
  <li>["ABC", "DEF", "G", "H"]</li>
</ul>

**Примітка:** Для цього завдання припускайте, що слова — це непорожні рядки з великими літерами.

# --before-each--

```js
const testCases = [[], ["ABC"], ["ABC", "DEF"], ["ABC", "DEF", "G", "H"]];
const results = ["{}", "{ABC}", "{ABC and DEF}", "{ABC, DEF, G and H}"];
```

# --hints--

`quibble` має бути функцією.

```js
assert(typeof quibble === 'function');
```

`quibble(["ABC"])` має повертати рядок.

```js
assert(typeof quibble(['ABC']) === 'string');
```

`quibble([])` має повертати "{}".

```js
assert.equal(quibble(testCases[0]), results[0]);
```

`quibble(["ABC"])` має повертати `"{ABC}"`.

```js
assert.equal(quibble(testCases[1]), results[1]);
```

`quibble(["ABC", "DEF"])` має повертати `"{ABC and DEF}"`.

```js
assert.equal(quibble(testCases[2]), results[2]);
```

`quibble(["ABC", "DEF", "G", "H"])` має повертати `"{ABC, DEF, G and H}"`.

```js
assert.equal(quibble(testCases[3]), results[3]);
```

# --seed--

## --seed-contents--

```js
function quibble(words) {

  return true;
}
```

# --solutions--

```js
function quibble(words) {
  return "{" +
    words.slice(0, words.length - 1).join(", ") +
   (words.length > 1 ? " and " : "") +
   (words[words.length - 1] || '') +
  "}";
}
```
