---
id: 596e457071c35c882915b3e4
title: Порівняти список рядків
challengeType: 1
forumTopicId: 302235
dashedName: compare-a-list-of-strings
---

# --description--

Список — це впорядкована множина значень, яка може містити повтори. Ось приклад:

```js
const list = [['AA',  'BB', 'CC'], ['AA', 'ACB', 'AA'], [], ['AA']];
````

Дано список довільної кількості рядків, реалізуйте функцію для кожної з наступних умов:

<ul>
  <li>перевірити, чи всі вони лексично однакові</li>
  <li>перевірити, чи кожен рядок лексично менший за наступний (тобто чи список строго впорядкований за зростанням)</li>
</ul>

# --before-each--

```js
const testCases = [['AA', 'AA', 'AA', 'AA'], ['AA', 'ACB', 'BB', 'CC'], [], ['AA'], ['BB', 'AA']];
```

# --hints--

`allEqual` має бути функцією.

```js
assert(typeof allEqual === 'function');
```

`azSorted` має бути функцією.

```js
assert(typeof azSorted === 'function');
```

`allEqual(["AA", "AA", "AA", "AA"])` має повертати true.

```js
assert(allEqual(testCases[0]));
```

`azSorted(["AA", "AA", "AA", "AA"])` має повертати false.

```js
assert(!azSorted(testCases[0]));
```

`allEqual(["AA", "ACB", "BB", "CC"])` має повертати false.

```js
assert(!allEqual(testCases[1]));
```

`azSorted(["AA", "ACB", "BB", "CC"])` має повертати true.

```js
assert(azSorted(testCases[1]));
```

`allEqual([])` має повертати true.

```js
assert(allEqual(testCases[2]));
```

`azSorted([])` має повертати true.

```js
assert(azSorted(testCases[2]));
```

`allEqual(["AA"])` має повертати true.

```js
assert(allEqual(testCases[3]));
```

`azSorted(["AA"])` має повертати true.

```js
assert(azSorted(testCases[3]));
```

`allEqual(["BB", "AA"])` має повертати false.

```js
assert(!allEqual(testCases[4]));
```

`azSorted(["BB", "AA"])` має повертати false.

```js
assert(!azSorted(testCases[4]));
```

# --seed--

## --seed-contents--

```js
function allEqual(arr) {

  return true;
}

function azSorted(arr) {

  return true;
}
```

# --solutions--

```js
function allEqual(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] === a[i]);
  } return out;
}

function azSorted(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] < a[i]);
  } return out;
}
```
