---
id: 596fda99c69f779975a1b67d
title: Підрахунок входжень підрядка
challengeType: 1
forumTopicId: 302237
dashedName: count-occurrences-of-a-substring
---

# --description--

Створіть функцію або покажіть вбудовану функцію, яка рахує кількість неперекривних входжень підрядка всередині рядка.

Функція має приймати два аргументи:

<ul>
  <li>перший аргумент — рядок для пошуку, та</li>
  <li>другий — підрядок, який потрібно знайти.</li>
</ul>

Вона має повертати ціле число.

Пошук має давати найбільшу кількість неперекривних співпадінь.

Загалом це означає пошук зліва направо або справа наліво.

# --before-each--

```js
const testCases = ['the three truths', 'ababababab', 'abaabba*bbaba*bbab'];
const searchString = ['th', 'abab', 'a*b'];
const results = [3, 2, 2];
```

# --hints--

`countSubstring` має бути функцією.

```js
assert(typeof countSubstring === 'function');
```

`countSubstring("the three truths", "th")` має повертати `3`.

```js
assert.equal(countSubstring(testCases[0], searchString[0]), results[0]);
```

`countSubstring("ababababab", "abab")` має повертати `2`.

```js
assert.equal(countSubstring(testCases[1], searchString[1]), results[1]);
```

`countSubstring("abaabba*bbaba*bbab", "a*b")` має повертати `2`.

```js
assert.equal(countSubstring(testCases[2], searchString[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function countSubstring(str, subStr) {

  return true;
}
```

# --solutions--

```js
function countSubstring(str, subStr) {
  const escapedSubStr = subStr.replace(/[.+*?^$[\]{}()|/]/g, '\\$&');
  const matches = str.match(new RegExp(escapedSubStr, 'g'));
  return matches ? matches.length : 0;
}
```
