---
id: 597089c87eec450c68aa1643
title: Визначте примітивний тип даних
challengeType: 1
forumTopicId: 302248
dashedName: define-a-primitive-data-type
---

# --description--

Визначте тип, який поводиться як ціле число, але має найнижче допустиме значення 1 і найвище допустиме значення 10.

Обробка помилок:

<ul>
  <li>Якщо ви спробуєте створити екземпляр <code>Num</code> зі значенням поза межами від 1 до 10, має бути викинута <code>TypeError</code> з повідомленням про помилку <code>'Out of range'</code>.</li>
  <li>Якщо ви спробуєте створити екземпляр <code>Num</code> зі значенням, яке не є числом, має бути викинута <code>TypeError</code> з повідомленням про помилку <code>'Not a Number'</code>.</li>
</ul>

# --hints--

`Num` має бути функцією.

```js
assert(typeof Num === 'function');
```

`new Num(4)` має повертати об’єкт.

```js
assert(typeof new Num(4) === 'object');
```

`new Num('test')` має викидати TypeError з повідомленням 'Not a Number'.

```js
assert.throws(() => new Num('test'), TypeError);
```

`new Num(0)` має викидати TypeError з повідомленням 'Out of range'.

```js
assert.throws(() => new Num(0), TypeError);
```

`new Num(-5)` має викидати TypeError з повідомленням 'Out of range'.

```js
assert.throws(() => new Num(-5), TypeError);
```

`new Num(10)` має викидати TypeError з повідомленням 'Out of range'.

```js
assert.throws(() => new Num(11), TypeError);
```

`new Num(20)` має викидати TypeError з повідомленням 'Out of range'.

```js
assert.throws(() => new Num(20), TypeError);
```

`new Num(3) + new Num(4)` має дорівнювати 7.

```js
assert.equal(new Num(3) + new Num(4), 7);
```

`new Num(3) - new Num(4)` має дорівнювати -1.

```js
assert.equal(new Num(3) - new Num(4), -1);
```

`new Num(3) * new Num(4)` має дорівнювати 12.

```js
assert.equal(new Num(3) * new Num(4), 12);
```

`new Num(3) / new Num(4)` має дорівнювати 0.75.

```js
assert.equal(new Num(3) / new Num(4), 0.75);
```

`new Num(3) < new Num(4)` має бути true.

```js
assert(new Num(3) < new Num(4));
```

`new Num(3) > new Num(4)` має бути false.

```js
assert(!(new Num(3) > new Num(4)));
```

`(new Num(5)).toString()` має повертати '5'

```js
assert.equal(new Num(5).toString(), '5');
```

# --seed--

## --seed-contents--

```js
function Num(n) {

  return n;
}
```

# --solutions--

```js
function Num(n) {
  if (isNaN(n)) {
    throw new TypeError('Not a Number');
  }
  if (n < 1 || n > 10) {
    throw new TypeError('Out of range');
  }

  this._value = +n;
}
Num.prototype.valueOf = function() { return this._value; };
Num.prototype.toString = function () { return this._value.toString(); };
```
