---
id: 59713bd26bdeb8a594fb9413
title: порахуйте монети
challengeType: 1
forumTopicId: 302238
dashedName: count-the-coins
---

# --description--

Існує чотири типи загальних монет у валюті США:

<ul>
  <li>квартали (25 центів)</li>
  <li>десятки (10 центів)</li>
  <li>нікелі (5 центів), та</li>
  <li>пенні (1 цент)</li>
</ul>

<p>Існує шість способів дати здачу на 15 центів:</p>

<ul>
  <li>десятка і нікель</li>
  <li>десятка і 5 пенні</li>
  <li>3 нікелі</li>
  <li>2 нікелі і 5 пенні</li>
  <li>нікель і 10 пенні</li>
  <li>15 пенні</li>
</ul>

# --instructions--

Реалізуйте функцію, щоб визначити, скільки існує способів дати здачу для заданого вхідного значення `cents`, яке представляє суму в американських пенні, використовуючи ці загальні монети.

# --hints--

`countCoins` має бути функцією.

```js
assert(typeof countCoins === 'function');
```

`countCoins(15)` має повертати `6`.

```js
assert.equal(countCoins(15), 6);
```

`countCoins(85)` має повертати `163`.

```js
assert.equal(countCoins(85), 163);
```

`countCoins(100)` має повертати `242`.

```js
assert.equal(countCoins(100), 242);
```

# --seed--

## --seed-contents--

```js
function countCoins(cents) {

  return true;
}
```

# --solutions--

```js
function countCoins(cents) {
  const operands = [1, 5, 10, 25];
  const targetsLength = cents + 1;
  const operandsLength = operands.length;
  const t = [1];

  for (let a = 0; a < operandsLength; a++) {
    for (let b = 1; b < targetsLength; b++) {
      // initialise undefined target
      t[b] = t[b] ? t[b] : 0;

      // accumulate target + operand ways
      t[b] += (b < operands[a]) ? 0 : t[b - operands[a]];
    }
  }

  return t[targetsLength - 1];
}
```
