---
id: 597f24c1dda4e70f53c79c81
title: послідовність Фібоначчі
challengeType: 1
forumTopicId: 302268
dashedName: fibonacci-sequence
---

# --description--

Напишіть функцію, яка генерує <code>n<sup>те</sup></code> число послідовності Фібоначчі.

<code>n<sup>те</sup></code> число Фібоначчі визначається так:

<code>F<sub>n</sub> = F<sub>n-1</sub> + F<sub>n-2</sub></code>

Перші два члени послідовності — 0 і 1.

Отже, послідовність виглядає так: 0, 1, 1, 2, 3, 5, 8, 13...

# --hints--

`fibonacci` має бути функцією.

```js
assert(typeof fibonacci === 'function');
```

`fibonacci(2)` має повертати число.

```js
assert(typeof fibonacci(2) == 'number');
```

`fibonacci(3)` має повертати 2.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` має повертати 5.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` має повертати 55.

```js
assert.equal(fibonacci(10), 55);
```

# --seed--

## --seed-contents--

```js
function fibonacci(n) {

}
```

# --solutions--

```js
function fibonacci(n) {
  let a = 0, b = 1, t;
  while (--n >= 0) {
    t = a;
    a = b;
    b += t;
  }
  return a;
}
```
