---
id: 59880443fb36441083c6c20e
title: Метод Ейлера
challengeType: 1
forumTopicId: 302258
dashedName: euler-method
---

# --description--

Метод Ейлера чисельно апроксимує розв’язки звичайних диференціальних рівнянь першого порядку (ОДР) з заданим ініціальним значенням. Це явний метод розв’язання задач початкових значень (ЗПЗ), як описано в <a href="https://www.freecodecamp.org/news/eulers-method-explained-with-examples/" title="Euler's Method Explained with Examples" target="_blank" rel="noopener noreferrer nofollow">цьому матеріалі</a>.

ОДР має бути задане у такій формі:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt} = f(t,y(t))$</big></li>
</ul>

з ініціальним значенням

<ul style='list-style: none;'>
  <li><big>$y(t_0) = y_0$</big></li>
</ul>

Щоб отримати чисельне розв’язання, замінимо похідну з лівої частини на скінченне різницеве наближення:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt}  \approx \frac{y(t+h)-y(t)}{h}$</big></li>
</ul>

потім розв’яжемо відносно $y(t+h)$:

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, \frac{dy(t)}{dt}$</big></li>
</ul>

що еквівалентно

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, f(t,y(t))$</big></li>
</ul>

Правило ітеративного розв’язання тоді таке:

<ul style='list-style: none;'>
  <li><big>$y_{n+1} = y_n + h \, f(t_n, y_n)$</big></li>
</ul>

де $h$ — розмір кроку, найважливіший параметр для точності розв’язку. Менший розмір кроку підвищує точність, але також збільшує обчислювальні витрати, тому його завжди потрібно підбирати вручну залежно від конкретної задачі.

**Приклад: закон охолодження Ньютона**

Закон охолодження Ньютона описує, як об’єкт з ініціальною температурою $T(t_0) = T_0$ охолоджується в середовищі з температурою $T_R$:

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, \Delta T$</big></li>
</ul>

або

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, (T(t) - T_R)$</big></li>
</ul>

Це означає, що швидкість охолодження $\\frac{dT(t)}{dt}$ об’єкта пропорційна поточній різниці температур $\\Delta T = (T(t) - T_R)$ із навколишнім середовищем.

Аналітичне розв’язання, з яким ми порівнюватимемо чисельне наближення, має вигляд

<ul style='list-style: none;'>
  <li><big>$T(t) = T_R + (T_0 - T_R) \; e^{-k t}$</big></li>
</ul>

# --instructions--

Реалізуйте процедуру методу Ейлера, а потім використайте її для розв’язання наведеного прикладу закону охолодження Ньютона для трьох різних розмірів кроку:

<ul>
  <li><code>2 s</code></li>
  <li><code>5 s</code> та</li>
  <li><code>10 s</code></li>
</ul>

і порівняйте з аналітичним розв’язком.

**Ініціальні значення:**

<ul>
  <li>ініціальна температура <big>$T_0$</big> має бути <code>100 °C</code></li>
  <li>температура навколишнього середовища <big>$T_R$</big> має бути <code>20 °C</code></li>
  <li>константа охолодження <big>$k$</big> має бути <code>0.07</code></li>
  <li>інтервал часу для обчислення має бути від <code>0 s</code> до <code>100 s</code></li>
</ul>  

Перший параметр функції — ініціальний час, другий параметр — ініціальна температура, третій параметр — пройдений час, четвертий параметр — розмір кроку.

# --hints--

`eulersMethod` має бути функцією.

```js
assert(typeof eulersMethod === 'function');
```

`eulersMethod(0, 100, 100, 2)` має повертати число.

```js
assert(typeof eulersMethod(0, 100, 100, 2) === 'number');
```

`eulersMethod(0, 100, 100, 2)` має повертати 20.0424631833732.

```js
assert.equal(eulersMethod(0, 100, 100, 2), 20.0424631833732);
```

`eulersMethod(0, 100, 100, 5)` має повертати 20.01449963666907.

```js
assert.equal(eulersMethod(0, 100, 100, 5), 20.01449963666907);
```

`eulersMethod(0, 100, 100, 10)` має повертати 20.000472392.

```js
assert.equal(eulersMethod(0, 100, 100, 10), 20.000472392);
```

# --seed--

## --seed-contents--

```js
function eulersMethod(x1, y1, x2, h) {

}
```

# --solutions--

```js
function eulersMethod(x1, y1, x2, h) {
  let x = x1;
  let y = y1;

  while ((x < x2 && x1 < x2) || (x > x2 && x1 > x2)) {
    y += h * (-0.07 * (y - 20));
    x += h;
  }

  return y;
}
```
