---
id: 598eea87e5cf4b116c3ff81a
title: Фактори числа Мерсенна
challengeType: 1
forumTopicId: 302264
dashedName: factors-of-a-mersenne-number
---

# --description--

Число Мерсенна — це число у формі <code>2<sup>P</sup>-1</code>.

Якщо `P` є простим числом, число Мерсенна може бути простим числом Мерсенна. (Якщо `P` не є простим, число Мерсенна також не є простим.)

Під час пошуку простих чисел Мерсенна вигідно виключати показники степеня, знаходячи малий фактор перед початком потенційно тривалого <a href="https://rosettacode.org/wiki/Lucas-Lehmer test" target="_blank" rel="noopener noreferrer nofollow">тесту Лукаса-Лемера</a>.

Існують дуже ефективні алгоритми для визначення, чи ділиться число на <code>2<sup>P</sup>-1</code> (або еквівалентно, чи <code>2<sup>P</sup> mod (число) = 1</code>).

Деякі мови вже мають вбудовані реалізації цієї операції степеня з модулем (називаються modPow або подібно).

Ось як реалізувати цю modPow самостійно:

Наприклад, обчислимо <code>2<sup>23</sup> mod 47</code>.

Перетворіть показник степеня 23 у двійкову форму, отримаємо 10111. Починаючи з <code><tt>square</tt> = 1</code>, повторно підносьте його до квадрату.

Видаліть найстарший біт показника, і якщо він дорівнює 1, помножте `square` на основу степеня (2), потім обчисліть <code><tt>square</tt> modulo 47</code>.

Використайте результат операції modulo з останнього кроку як ініціальне значення `square` у наступному кроці:

<pre>Видалити   Опціонально
square        найстарший біт  помножити на 2  mod 47
------------  --------------  --------------  ------
1*1 = 1       1  0111         1*2 = 2           2
2*2 = 4       0  111          ні                4
4*4 = 16      1  11           16*2 = 32         32
32*32 = 1024  1  1            1024*2 = 2048     27
27*27 = 729   1               729*2 = 1458      1
</pre>

Оскільки <code>2<sup>23</sup> mod 47 = 1</code>, 47 є фактором <code>2<sup>P</sup>-1</code>.

(Щоб це побачити, відніміть 1 з обох сторін: <code>2<sup>23</sup>-1 = 0 mod 47</code>.)

Оскільки ми показали, що 47 є фактором, <code>2<sup>23</sup>-1</code> не є простим числом.

Подальші властивості чисел Мерсенна дозволяють ще більше уточнити процес.

Будь-який фактор `q` числа <code>2<sup>P</sup>-1</code> повинен мати форму `2kP+1`, де `k` — додатне ціле число або нуль. Крім того, `q` має бути `1` або `7 mod 8`.

Нарешті, будь-який потенційний фактор `q` повинен бути <a href="https://rosettacode.org/wiki/Primality_by_trial_division" target="_blank" rel="noopener noreferrer nofollow">простим</a> числом.

Як і в інших алгоритмах пробного ділення, алгоритм зупиняється, коли `2kP+1 > sqrt(N)`. Ці основні тести працюють лише для чисел Мерсенна, де `P` є простим. Наприклад, <code>M<sub>4</sub>=15</code> не дає факторів за допомогою цих технік, але розкладається на 3 і 5, жоден з яких не відповідає `2kP+1`.

# --instructions--

Використовуючи описаний вище метод, знайдіть фактор числа <code>2<sup>p</sup>-1</code>.

# --hints--

`check_mersenne` має бути функцією.

```js
assert(typeof check_mersenne === 'function');
```

`check_mersenne(3)` має повертати рядок.

```js
assert(typeof check_mersenne(3) == 'string');
```

`check_mersenne(3)` має повертати рядок `M3 = 2^3-1 is prime`.

```js
assert.equal(check_mersenne(3), 'M3 = 2^3-1 is prime');
```

`check_mersenne(23)` має повертати рядок `M23 = 2^23-1 is composite with factor 47`.

```js
assert.equal(check_mersenne(23), 'M23 = 2^23-1 is composite with factor 47');
```

`check_mersenne(929)` має повертати рядок `M929 = 2^929-1 is composite with factor 13007`.

```js
assert.equal(
  check_mersenne(929),
  'M929 = 2^929-1 is composite with factor 13007'
);
```

# --seed--

## --seed-contents--

```js
function check_mersenne(p) {

}
```

# --solutions--

```js
function check_mersenne(p){
    function isPrime(value){
      for (let i=2; i < value; i++){
        if (value % i == 0){
          return false;
        }
        if (value % i != 0){
          return true;
         }
      }
    }

    function trial_factor(base, exp, mod){
      let square, bits;
      square = 1;
      bits = exp.toString(2).split('');
      for (let i=0,ln=bits.length; i<ln; i++){
        square = Math.pow(square, 2) * (bits[i] == 1 ? base : 1) % mod;
      }
      return (square == 1);
    }

    function mersenne_factor(p){
      let limit, k, q;
      limit = Math.sqrt(Math.pow(2,p) - 1);
      k = 1;
      while ((2*k*p - 1) < limit){
        q = 2*k*p + 1;
        if (isPrime(q) && (q % 8 == 1 || q % 8 == 7) && trial_factor(2,p,q)){
          return q; // q is a factor of 2**p-1
        }
        k++;
      }
      return null;
    }
  let f, result;
  result="M"+p+" = 2^"+p+"-1 is ";
  f = mersenne_factor(p);
  result+=f == null ? "prime" : "composite with factor "+f;
  return result;
}
```
