---
id: 599d1566a02b571412643b84
title: Ефіопське множення
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

Ефіопське множення — це метод множення цілих чисел, який використовує лише додавання, подвоєння та ділення навпіл.

**Метод:**

<ol>
  <li>Візьміть два числа для множення і запишіть їх у верхній частині двох стовпців</li>
  <li>У лівому стовпці послідовно діліть останнє число навпіл, відкидаючи будь-які остачі, і записуйте результат під останнім у тому ж стовпці, доки не отримаєте значення <code>1</code></li>
  <li>У правому стовпці послідовно подвоюйте останнє число і записуйте результат унизу. Зупиніться, коли додасте результат у тому ж рядку, де в лівому стовпці стоїть <code>1</code></li>
  <li>Перегляньте отриману таблицю і відкиньте будь-який рядок, де значення в лівому стовпці парне</li>
  <li>Підсумуйте значення в правому стовпці, що залишилися, щоб отримати результат множення двох початкових чисел</li>
</ol>

**Наприклад:** `17 × 34`

<pre>17   34
</pre>

Ділення навпіл першого стовпця:

<pre>17   34
8
4
2
1
</pre>

Подвоєння другого стовпця:

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

Закресліть рядки, у яких перша клітинка парна:

<pre>17   34
8    <strike>68</strike>
4   <strike>136</strike>
2   <strike>272</strike>
1   544
</pre>

Підсумуйте залишені числа у правому стовпці:

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

Отже, `17` помножене на `34` за ефіопським методом дорівнює `578`.

# --instructions--

Завдання — визначити три іменовані функції/методи/процедури/підпрограми:

<ol>
  <li>одну для ділення цілого числа навпіл,</li>
  <li>одну для подвоєння цілого числа, і</li>
  <li>одну для визначення, чи є ціле число парним</li>
</ol>

Використайте ці функції, щоб створити функцію, яка виконує ефіопське множення.

<!-- markdownlint-disable MD046-->

# --hints--

`eth_mult` має бути функцією.

```js
assert(typeof eth_mult === 'function');
```

`eth_mult(17,34)` має повертати `578`.

```js
assert.equal(eth_mult(17, 34), 578);
```

`eth_mult(23,46)` має повертати `1058`.

```js
assert.equal(eth_mult(23, 46), 1058);
```

`eth_mult(12,27)` має повертати `324`.

```js
assert.equal(eth_mult(12, 27), 324);
```

`eth_mult(56,98)` має повертати `5488`.

```js
assert.equal(eth_mult(56, 98), 5488);
```

`eth_mult(63,74)` має повертати `4662`.

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
