---
id: 59da22823d04c95919d46269
title: 'Моряки, кокоси та проблема з мавпою'
challengeType: 1
forumTopicId: 302304
dashedName: sailors-coconuts-and-a-monkey-problem
---

# --description--

П’ятеро моряків зазнали корабельної аварії на острові і вдень зібрали велику купу кокосів.

Вночі перший моряк прокидається і вирішує взяти свою першу частку раніше, тому намагається поділити купу кокосів порівну на п’ять купок, але виявляє, що залишається один кокос, тож він кидає його мавпі, а потім ховає «свою» одну з п’яти рівних купок кокосів і зсуває інші чотири купки разом, щоб знову утворити одну видиму купу кокосів, після чого лягає спати.

Щоб не затягувати, кожен із моряків по черзі прокидається вночі і виконує ті самі дії: ділить купу кокосів на п’ять, знаходить, що залишається один кокос, віддає цей один кокос мавпі, а потім зсуває інші чотири купки разом, щоб утворити одну купу.

Вранці (після таємних і окремих дій кожного з п’яти моряків вночі) залишок кокосів ділиться на п’ять рівних купок для кожного моряка, і виявляється, що купа кокосів ділиться порівну між моряками без остачі. (Вранці мавпі нічого не дістається.)

# --instructions--

Створіть функцію, яка повертає мінімальний можливий розмір ініціальної купи кокосів, зібраних вдень для `N` моряків. **Примітка:** Звісно, ця історія відбувається у світі, де можна зібрати будь-яку кількість кокосів за день і зробити кілька поділів купи, тощо, у часі, що відповідає сюжетній лінії, щоб це не впливало на математику.

# --hints--

`splitCoconuts` має бути функцією.

```js
assert(typeof splitCoconuts === 'function');
```

`splitCoconuts(5)` має повертати 3121.

```js
assert(splitCoconuts(5) === 3121);
```

`splitCoconuts(6)` має повертати 233275.

```js
assert(splitCoconuts(6) === 233275);
```

`splitCoconuts(7)` має повертати 823537.

```js
assert(splitCoconuts(7) === 823537);
```

# --seed--

## --seed-contents--

```js
function splitCoconuts(intSailors) {

  return true;
}
```

# --solutions--

```js
function splitCoconuts(intSailors) {
  let intNuts = intSailors;
  let result = splitCoconutsHelper(intNuts, intSailors);
  while (!result) {
    intNuts += 1;
    result = splitCoconutsHelper(intNuts, intSailors);
  }

  return intNuts;
}

function splitCoconutsHelper(intNuts, intSailors, intDepth) {
  const nDepth = intDepth !== undefined ? intDepth : intSailors;
  const portion = Math.floor(intNuts / intSailors);
  const remain = intNuts % intSailors;

  if (portion <= 0 || remain !== (nDepth ? 1 : 0)) {
    return null;
  }

  if (nDepth) {
    return splitCoconutsHelper(
      intNuts - portion - remain, intSailors, nDepth - 1
    );
  }

  return intNuts;
}
```
