---
id: 59e0a8df964e4540d5abe599
title: Реалізувати Brain****
challengeType: 1
forumTopicId: 302261
dashedName: execute-brain
---

# --description--

Напишіть функцію для реалізації інтерпретатора Brain\*\*\*\*. Функція прийматиме рядок як параметр і має повертати рядок як вихідні дані. Детальніше описано нижче:

RCBF — це множина <a href="https://rosettacode.org/wiki/Brainf***" target="_blank" rel="noopener noreferrer nofollow">компіляторів і інтерпретаторів Brainf\*\*\*</a>, написаних для Rosetta Code на різних мовах.

Нижче наведені посилання на кожну з версій RCBF.

Реалізація має правильно підтримувати лише такі інструкції:

| Команда           | Опис                                                                              |
| ----------------- | --------------------------------------------------------------------------------- |
| <code>></code>    | Перемістити вказівник вправо                                                     |
| <code>&lt;</code> | Перемістити вказівник вліво                                                      |
| <code>+</code>    | Збільшити значення комірки пам’яті під вказівником                               |
| <code>-</code>    | Зменшити значення комірки пам’яті під вказівником                               |
| <code>.</code>    | Вивести символ, який позначає значення комірки під вказівником                   |
| <code>,</code>    | Ввести символ і зберегти його в комірці під вказівником                          |
| <code>\[</code>   | Перейти за відповідну <code>]</code>, якщо значення комірки під вказівником 0    |
| <code>]</code>    | Повернутися до відповідної <code>\[</code>, якщо значення комірки під вказівником не 0 |

Дозволяється будь-який розмір комірки, підтримка EOF (*E*nd-*O*-*F*ile) є необов’язковою, як і обмежена чи необмежена пам’ять.

# --hints--

`brain(globalThis.hello)` має повертати рядок

```js
assert(typeof brain(globalThis.hello) === 'string');
```

`brain("++++++[>++++++++++<-]>+++++.")` має повертати "A"

```js
assert.equal(brain('++++++[>++++++++++<-]>+++++.'), 'A');
```

`brain(globalThis.hello)` має повертати `Hello World!\n`

```js
assert.equal(brain(globalThis.hello), 'Hello World!\n');
```

`brain(globalThis.fib)` має повертати `1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89`

```js
assert.equal(brain(globalThis.fib), '1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89');
```

# --before-all--

```js
globalThis.fib=`+

++

+++

++++

+>+>>

>>++++

+++++++

++++++++

+++++++++

++++++++++

++++++>++++

++++++++++++

+++++++++++++

+++<<<<<<[>[>>

>>>>+>+<<<<<<<-

]>>>>>>>[<<<<<<<

+>>>>>>>-]<[>++++

++++++[-<-[>>+>+<<

<-]>>>[<<<+>>>-]+<[

>[-]<[-]]>[<<[>>>+<<

<-]>>[-]]<<]>>>[>>+>+

<<<-]>>>[<<<+>>>-]+<[>

[-]<[-]]>[<<+>>[-]]<<<<

<<<]>>>>>[++++++++++++++

+++++++++++++++++++++++++

+++++++++.[-]]++++++++++<[

->-<]>+++++++++++++++++++++

+++++++++++++++++++++++++++.

[-]<<<<<<<<<<<<[>>>+>+<<<<-]>

>>>[<<<<+>>>>-]<-[>>.>.<<<[-]]

<<[>>+>+<<<-]>>>[<<<+>>>-]<<[<+

>-]>[<+>-]<<<-]`;
globalThis.hello='++++++++[>++++[>++>+++>+++>+<<<<-]>+>+>->>+[<]<-]>>.>---.+++++++..+++.>>.<-.<.+++.------.--------.>>+.>++.'
```

# --seed--

## --seed-contents--

```js
function brain(prog) {

}
```

# --solutions--

```js
function brain(prog){
  var output="";
    var code; // formatted code
  var ip = 0; // current instruction within code
  var nest = 0; // current bracket nesting (for Out button)
  var ahead = []; // locations of matching brackets

  var data = [0]; // data array (mod by +, -)
  var dp = 0; // index into data (mod by <, >)

  var inp = 0; // current input character (fetch with ,)
  var quit = 0;
    var commands = {
    '>':function() { if (++dp >= data.length) data[dp]=0 },
    '<':function() { if (--dp < 0) quit++ },
    '+':function() { ++data[dp] },
    '-':function() { --data[dp] },
    '[':function() { if (!data[dp]) ip = ahead[ip]; else ++nest },
    ']':function() { if ( data[dp]) ip = ahead[ip]; else --nest },
    ',':function() {
        var c = document.getElementById("input").value.charCodeAt(inp++);
        data[dp] = isNaN(c) ? 0 : c; // EOF: other options are -1 or no change
    },
    '.':function() {
            output+=String.fromCharCode(data[dp]);
            /*var s = document.getElementById("output").innerHTML)
             + String.fromCharCode(data[dp]);
            s = s.replace(/\n/g,"<br>").replace(/ /g,"&amp;nbsp;");
            document.getElementById("output").innerHTML = s;*/
        },
    };

    let ar=prog.split('');
    var st = [], back, error = -1;
    for (ip=0; ip<ar.length; ip++) {
        switch(ar[ip]) {
        case '[':
            st.push(ip);
            break;
        case ']':
            if (st.length == 0) error = ip;
            back = st.pop();
            ahead[ip] = back;
            ahead[back] = ip;
            break;
        }
    }

    for(ip=0;ip<ar.length;ip++){
    if(commands.hasOwnProperty(ar[ip]))
          commands[ar[ip]]();
    }

    return output;
}
```
