---
id: 59f40b17e79dbf1ab720ed7a
title: Номери відділів
challengeType: 1
forumTopicId: 302249
dashedName: department-numbers
---

# --description--

Існує дуже організоване місто, яке вирішило призначити номер кожному зі своїх відділів:

<ul>
  <li>Поліцейський відділ</li>
  <li>Відділ санітарії</li>
  <li>Пожежний відділ</li>
</ul>

Кожен відділ може мати номер від 1 до 7 (включно).

Три номери відділів мають бути унікальними (різними один від одного) і їх сума має дорівнювати числу 12.

Начальник поліції не любить непарні числа і хоче, щоб у його відділу був парний номер.

# --instructions--

Напишіть програму, яка виводить усі дійсні комбінації у вигляді масиву.

```js
[2, 3, 7] [2, 4, 6] [2, 6, 4]
[2, 7, 3] [4, 1, 7] [4, 2, 6]
[4, 3, 5] [4, 5, 3] [4, 6, 2]
[4, 7, 1] [6, 1, 5] [6, 2, 4]
[6, 4, 2] [6, 5, 1]
```

# --before-each--

```js
const nums = [1, 2, 3, 4, 5, 6, 7];
const total = 12;
const len = 14;
const result = [
  [2, 3, 7],
  [2, 4, 6],
  [2, 6, 4],
  [2, 7, 3],
  [4, 1, 7],
  [4, 2, 6],
  [4, 3, 5],
  [4, 5, 3],
  [4, 6, 2],
  [4, 7, 1],
  [6, 1, 5],
  [6, 2, 4],
  [6, 4, 2],
  [6, 5, 1]
];
```

# --hints--

`combinations` має бути функцією.

```js
assert(typeof combinations === 'function');
```

`combinations([1, 2, 3], 6)` має повертати масив.

```js
assert(Array.isArray(combinations([1, 2, 3], 6)));
```

`combinations([1, 2, 3, 4, 5, 6, 7], 12)` має повертати масив довжиною 14.

```js
assert(combinations(nums, total).length === len);
```

`combinations([1, 2, 3, 4, 5, 6, 7], 12)` має повертати всі дійсні комбінації.

```js
assert.deepEqual(combinations(nums, total), result);
```

# --seed--

## --seed-contents--

```js
function combinations(possibleNumbers, total) {

  return true;
}
```

# --solutions--

```js
function combinations(possibleNumbers, total) {
  let firstNumber;
  let secondNumber;
  let thirdNumber;
  const allCombinations = [];

  for (let i = 0; i < possibleNumbers.length; i += 1) {
    firstNumber = possibleNumbers[i];

    if (firstNumber % 2 === 0) {
      for (let j = 0; j < possibleNumbers.length; j += 1) {
        secondNumber = possibleNumbers[j];

        if (j !== i && firstNumber + secondNumber <= total) {
          thirdNumber = total - firstNumber - secondNumber;

          if (thirdNumber !== firstNumber && thirdNumber !== secondNumber && possibleNumbers.includes(thirdNumber)) {
            allCombinations.push([firstNumber, secondNumber, thirdNumber]);
          }
        }
      }
    }
  }
  return allCombinations;
}
```
