---
id: 5a23c84252665b21eecc803f
title: Сума цифр цілого числа
challengeType: 1
forumTopicId: 302331
dashedName: sum-digits-of-an-integer
---

# --description--

Напишіть функцію, яка приймає рядок як параметр. Цей рядок представляє число в будь-якій системі числення (меншій за 37) і повертає суму його цифр.

<ul>
  <li><b>1</b><sub>10</sub> дає суму <b>1</b></li>
  <li><b>1234</b><sub>10</sub> дає суму <b>10</b></li>
  <li><b>fe</b><sub>16</sub> дає суму <b>29</b></li>
  <li><b>f0e</b><sub>16</sub> дає суму <b>29</b></li>
</ul>

# --hints--

`sumDigits` має бути функцією.

```js
assert(typeof sumDigits == 'function');
```

`sumDigits("1")` має повертати число.

```js
assert(typeof sumDigits('1') == 'number');
```

`sumDigits("1")` має повертати `1`.

```js
assert.equal(sumDigits('1'), 1);
```

`sumDigits("12345")` має повертати `15`.

```js
assert.equal(sumDigits('12345'), 15);
```

`sumDigits("254")` має повертати `11`.

```js
assert.equal(sumDigits('254'), 11);
```

`sumDigits("fe")` має повертати `29`.

```js
assert.equal(sumDigits('fe'), 29);
```

`sumDigits("f0e")` має повертати `29`.

```js
assert.equal(sumDigits('f0e'), 29);
```

`sumDigits("999ABCXYZ")` має повертати `162`.

```js
assert.equal(sumDigits('999ABCXYZ'), 162);
```

# --seed--

## --seed-contents--

```js
function sumDigits(n) {

}
```

# --solutions--

```js
function sumDigits(n) {
  n += '';
  for (var s = 0, i = 0, e = n.length; i < e; i += 1)
    s += parseInt(n.charAt(i), 36);
  return s;
}
```
