---
id: 5a23c84252665b21eecc8028
title: послідовність Стерна-Броко
challengeType: 1
forumTopicId: 302324
dashedName: stern-brocot-sequence
---

# --description--

Для цього завдання послідовність Стерна-Броко потрібно згенерувати за алгоритмом, схожим на той, що використовується для генерації <a href="https://rosettacode.org/wiki/Fibonacci_sequence" target="_blank" rel="noopener noreferrer nofollow">послідовності Фібоначчі</a>.

<ol>
  <li>Перший і другий члени послідовності обидва дорівнюють 1:</li>
    <ul><li>1, 1</li></ul>
  <li>Почніть з розгляду другого члена послідовності</li>
  <li>Складіть розглянутий член послідовності та його попередній, (1 + 1) = 2, і додайте його в кінець послідовності:</li>
    <ul><li>1, 1, 2</li></ul>
  <li>Додайте розглянутий член послідовності в кінець послідовності:</li>
    <ul><li>1, 1, 2, 1</li></ul>
  <li>Розгляньте наступний член послідовності (третій член, тобто 2)</li>
  <li>ПОВЕРНІТЬСЯ ДО 3 </li>
    <ul>
      <li></li>
      <li> ─── Розгортаючи ще один цикл, отримуємо: ───</li>
      <li></li>
    </ul>
  <li>Складіть розглянутий член послідовності та його попередній, (2 + 1) = 3, і додайте його в кінець послідовності:</li>
    <ul><li>1, 1, 2, 1, 3</li></ul>
  <li>Додайте розглянутий член послідовності в кінець послідовності:</li>
    <ul><li>1, 1, 2, 1, 3, 2</li></ul>
  <li>Розгляньте наступний член послідовності (четвертий член, тобто 1)</li>
</ol>

# --instructions--

Створіть функцію, яка повертає позицію у послідовності Стерна-Броко, на якій вперше зустрічається $ n $, де послідовність генерується за описаним вище методом. Зверніть увагу, що ця послідовність використовує індексацію, що починається з 1.

# --hints--

`sternBrocot` має бути функцією.

```js
assert(typeof sternBrocot == 'function');
```

`sternBrocot(2)` має повертати число.

```js
assert(typeof sternBrocot(2) == 'number');
```

`sternBrocot(2)` має повертати `3`.

```js
assert.equal(sternBrocot(2), 3);
```

`sternBrocot(3)` має повертати `5`.

```js
assert.equal(sternBrocot(3), 5);
```

`sternBrocot(5)` має повертати `11`.

```js
assert.equal(sternBrocot(5), 11);
```

`sternBrocot(7)` має повертати `19`.

```js
assert.equal(sternBrocot(7), 19);
```

`sternBrocot(10)` має повертати `39`.

```js
assert.equal(sternBrocot(10), 39);
```

# --seed--

## --seed-contents--

```js
function sternBrocot(num) {

}
```

# --solutions--

```js
function sternBrocot(num) {
  function f(n) {
    return n < 2
      ? n
      : n & 1
      ? f(Math.floor(n / 2)) + f(Math.floor(n / 2 + 1))
      : f(Math.floor(n / 2));
  }

  function gcd(a, b) {
    return a ? (a < b ? gcd(b % a, a) : gcd(a % b, b)) : b;
  }
  var n;
  for (n = 1; f(n) != num; n++);
  return n;
}
```
