---
id: 5a23c84252665b21eecc7edf
title: Найменше спільне кратне
challengeType: 1
forumTopicId: 302301
dashedName: least-common-multiple
---

# --description--

Найменше спільне кратне для 12 і 18 дорівнює 36, тому що 12 є дільником (12 × 3 = 36), і 18 є дільником (18 × 2 = 36), і немає жодного додатного цілого числа менше за 36, яке має обидва ці дільники. Особливий випадок: якщо $m$ або $n$ дорівнює нулю, тоді найменше спільне кратне також дорівнює нулю. Один зі способів обчислити найменше спільне кратне — ітерувати всі кратні числа $m$, доки не знайдете таке, що також є кратним $n$. Якщо у вас уже є $gcd$ для <a href="https://rosettacode.org/wiki/Greatest_common_divisor" target="_blank" rel="noopener noreferrer nofollow">найбільшого спільного дільника</a>, тоді ця формула обчислює $lcm$. 

$$
\\operatorname{lcm}(m, n) = \\frac{|m \\times n|}{\\operatorname{gcd}(m, n)}
$$

# --instructions--

Обчисліть найменше спільне кратне масиву цілих чисел. Для заданих *m* і *n* найменше спільне кратне — це найменше додатне ціле число, яке має обидва числа *m* і *n* як дільники.

# --hints--

`LCM` має бути функцією.

```js
assert(typeof LCM == 'function');
```

`LCM([2, 4, 8])` має повертати число.

```js
assert(typeof LCM([2, 4, 8]) == 'number');
```

`LCM([2, 4, 8])` має повертати `8`.

```js
assert.equal(LCM([2, 4, 8]), 8);
```

`LCM([4, 8, 12])` має повертати `24`.

```js
assert.equal(LCM([4, 8, 12]), 24);
```

`LCM([3, 4, 5, 12, 40])` має повертати `120`.

```js
assert.equal(LCM([3, 4, 5, 12, 40]), 120);
```

`LCM([11, 33, 90])` має повертати `990`.

```js
assert.equal(LCM([11, 33, 90]), 990);
```

`LCM([-50, 25, -45, -18, 90, 447])` має повертати `67050`.

```js
assert.equal(LCM([-50, 25, -45, -18, 90, 447]), 67050);
```

# --seed--

## --seed-contents--

```js
function LCM(A) {

}
```

# --solutions--

```js
function LCM(A) {
  var n = A.length,
    a = Math.abs(A[0]);
  for (var i = 1; i < n; i++) {
    var b = Math.abs(A[i]),
      c = a;
    while (a && b) {
      a > b ? (a %= b) : (b %= a);
    }
    a = Math.abs(c * A[i]) / (a + b);
  }
  return a;
}
```
