---
id: 5a23c84252665b21eecc8005
title: Алгоритми сортування/Сортування гребінцем
challengeType: 1
forumTopicId: 302313
dashedName: sorting-algorithmscomb-sort
---

# --description--

Реалізуйте *сортування гребінцем*.

**Сортування гребінцем** — це варіант сортування бульбашкою.

Як і сортування Шелла, сортування гребінцем збільшує горизонтальний розрив, який використовується для порівнянь і обмінів.

Найкраще працює ділення розриву на $(1-e^{-\\varphi})^{-1} \\approx 1.247330950103979$, але 1.3 може бути більш практичним.

Деякі реалізації використовують сортування вставками, коли розрив стає меншим за певне значення.

Варіанти:

<ul>
  <li>Combsort11 гарантує, що розрив закінчується на (11, 8, 6, 4, 3, 2, 1), що значно швидше за інші два можливі закінчення.</li>
  <li>Combsort з різними закінченнями переходить до більш ефективного сортування, коли дані майже відсортовані (коли розрив малий). Сортування гребінцем з малим розривом не набагато кращє за сортування бульбашкою.</li>
</ul>

Псевдокод:

<pre><b>function</b> combsort(<b>array</b> input)
  gap := input<b>.size</b> <i>//ініціалізувати розмір розриву</i>
  <b>loop until</b> gap = 1 <b>and</b> swaps = 0
    <i>//оновити значення розриву для наступного проходу. Нижче наведено приклад</i>
    gap := int(gap / 1.25)
    <b>if</b> gap &#x3C; 1 
      <i>//мінімальний розрив — 1</i>
      gap := 1
    <b>end if</b>
    i := 0
    swaps := 0 <i>//див. <a href='https://rosettacode.org/wiki/Sorting_algorithms/Bubble_sort' target='_blank'>сортування бульбашкою</a> для пояснення</i>
    <i>//один "гребінець" по вхідному списку</i>
    <b>loop until</b> i + gap >= input<b>.size</b> <i>//див. <a href='https://rosettacode.org/wiki/Sorting_algorithms/Shell_sort' target='_blank'>сортування Шелла</a> для подібної ідеї</i>
      <b>if</b> input[i] > input[i+gap]
        <b>swap</b>(input[i], input[i+gap])
        swaps := 1 <i>//прапорець, що відбувся обмін, тому</i>
            <i>//список не гарантовано відсортований</i>
      <b>end if</b>
      i := i + 1
    <b>end loop</b>
  <b>end loop</b>
<b>end function</b>
</pre>

# --instructions--

Напишіть функцію, яка сортує заданий масив за допомогою сортування гребінцем.

# --hints--

`combSort` має бути функцією.

```js
assert(typeof combSort == 'function');
```

`combSort([25, 32, 12, 7, 20])` має повертати масив.

```js
assert(Array.isArray(combSort([25, 32, 12, 7, 20])));
```

`combSort([25, 32, 12, 7, 20])` має повертати `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(combSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`combSort([38, 45, 35, 8, 13])` має повертати `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(combSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`combSort([43, 36, 20, 34, 24])` має повертати `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(combSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`combSort([12, 33, 26, 18, 1, 16, 38])` має повертати `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(combSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`combSort([3, 39, 48, 16, 1, 4, 29])` має повертати `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(combSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function combSort(arr) {

}
```

# --solutions--

```js
function combSort(arr) {
  function is_array_sorted(arr) {
    var sorted = true;
    for (var i = 0; i < arr.length - 1; i++) {
      if (arr[i] > arr[i + 1]) {
        sorted = false;
        break;
      }
    }
    return sorted;
  }
  var iteration_count = 0;
  var gap = arr.length - 2;
  var decrease_factor = 1.25;

  // Until array is not sorted, repeat iterations
  while (!is_array_sorted(arr)) {
    // If not first gap
    if (iteration_count > 0)
      // Calculate gap
      gap = gap == 1 ? gap : Math.floor(gap / decrease_factor);

    // Set front and back elements and increment to a gap
    var front = 0;
    var back = gap;
    while (back <= arr.length - 1) {
      // If elements are not ordered swap them
      if (arr[front] > arr[back]) {
        var temp = arr[front];
        arr[front] = arr[back];
        arr[back] = temp;
      }

      // Increment and re-run swapping
      front += 1;
      back += 1;
    }
    iteration_count += 1;
  }

  return arr;
}
```
