---
id: 5a23c84252665b21eecc7ec5
title: задача Йосипа
challengeType: 1
forumTopicId: 302294
dashedName: josephus-problem
---

# --description--

Задача Йосипа — це математична головоломка з похмурим описом: $n$ ув’язнених стоять у колі, послідовно пронумеровані від $0$ до $n-1$.

Кат проходить колом, починаючи з ув’язненого $0$, видаляючи кожного $k$-го ув’язненого та вбиваючи його.

У міру того, як процес триває, коло стає все меншим і меншим, доки не залишається лише один ув’язнений, якого потім звільняють.

Наприклад, якщо є $n=5$ ув’язнених і $k=2$, порядок, у якому вбивають ув’язнених (назвемо це "послідовністю вбивств"), буде 1, 3, 0 і 4, а той, хто виживе, матиме номер #2.

За будь-яких $n, k > 0$ знайдіть, який ув’язнений буде фінальним, хто виживе.

В одному випадку було 41 ув’язнений, і кожного 3<sup>го</sup> вбивали ($k=3$).

Серед них був кмітливий хлопець на ім’я Йосип, який розв’язав цю задачу, став на позицію того, хто вижив, і продовжив жити, щоб розповісти цю історію.

Який номер був у нього?

# --instructions--

Напишіть функцію, яка приймає ініціальне число ув’язнених і `k` як параметри та повертає номер ув’язненого, який виживе.

# --hints--

`josephus` має бути функцією.

```js
assert(typeof josephus == 'function');
```

`josephus(30,3)` має повертати число.

```js
assert(typeof josephus(30, 3) == 'number');
```

`josephus(30,3)` має повертати `28`.

```js
assert.equal(josephus(30, 3), 28);
```

`josephus(30,5)` має повертати `2`.

```js
assert.equal(josephus(30, 5), 2);
```

`josephus(20,2)` має повертати `8`.

```js
assert.equal(josephus(20, 2), 8);
```

`josephus(17,6)` має повертати `1`.

```js
assert.equal(josephus(17, 6), 1);
```

`josephus(29,4)` має повертати `1`.

```js
assert.equal(josephus(29, 4), 1);
```

# --seed--

## --seed-contents--

```js
function josephus(init, kill) {

}
```

# --solutions--

```js
function josephus(init, kill) {
  const arr = Array.from(Array(init).keys());
  let curr = -1
  while (arr.length > 1) {
    curr = (curr + kill) % arr.length;
    arr.splice(curr, 1);
    curr--;
  }
  return arr[0];
}
```
