---
id: 5a23c84252665b21eecc8036
title: Видалити керуючі коди та розширені символи з рядка
challengeType: 1
forumTopicId: 302327
dashedName: strip-control-codes-and-extended-characters-from-a-string
---

# --description--

Завдання полягає у видаленні керуючих кодів та розширених символів з рядка. Рішення має показати, як досягти кожного з наступних результатів: рядок без керуючих кодів та розширених символів. В ASCII керуючі коди мають десяткові коди від 0 до 31 та 127. На системі, що базується на ASCII, якщо видалити керуючі коди, отриманий рядок міститиме всі символи в діапазоні від 32 до 126 десяткових у ASCII-таблиці. На системі, що не базується на ASCII, для цілей цього завдання вважаємо розширеним символом будь-який символ, який не має відповідного гліфа в ASCII-таблиці (в межах діапазону ASCII від 32 до 126 десяткових).

# --hints--

`strip` має бути функцією.

```js
assert(typeof strip == 'function');
```

`strip("abc")` має повертати рядок.

```js
assert(typeof strip('abc') == 'string');
```

`strip("\ba\x00b\n\rc\fd\xc3")` має повертати `"abcd"`.

```js
assert.equal(strip('\ba\x00b\n\rc\fd\xc3'), 'abcd');
```

`strip("\u0000\n abc\u00E9def\u007F")` має повертати `" abcdef"`.

```js
assert.equal(strip('\u0000\n abc\u00E9def\u007F'), ' abcdef');
```

`strip("a\n\tb\u2102d\u2147f")` має повертати `"abdf"`.

```js
assert.equal(strip('a\n\tb\u2102d\u2147f'), 'abdf');
```

`strip("Français.")` має повертати `"Franais."`.

```js
assert.equal(strip('Français.'), 'Franais.');
```

`strip("123\tabc\u0007DEF\u007F+-*/€æŧðłþ")` має повертати `"123abcDEF+-*/"`.

```js
assert.equal(strip('123\tabc\u0007DEF\u007F+-*/€æŧðłþ'), '123abcDEF+-*/');
```

# --seed--

## --seed-contents--

```js
function strip(s) {

}
```

# --solutions--

```js
function strip(s) {
  return s
    .split('')
    .filter(function(x) {
      var n = x.charCodeAt(0);

      return 31 < n && 127 > n;
    })
    .join('');
}
```
