---
id: 5a23c84252665b21eecc8017
title: Soundex
challengeType: 1
forumTopicId: 302320
dashedName: soundex
---

# --description--

**Алгоритм Soundex** працює з *наміром* слів. Він створює представлення для слів, які звучать схоже.

Він використовується для пошуку <em>імен</em> та <em>адрес</em>. Це означає, що людина, яка заповнює <em>ім’я</em>, може зосередитися на тому, як воно звучить, замість того, щоб виправляти написання <em>імен</em>.

Наприклад:

Якщо ви вперше чуєте ім’я `Quenci` і неправильно його написали, ви отримаєте **Soundex** код `Q520`.

Коли ви наступного разу правильно напишете ім’я `Quincy`, ви все одно отримаєте той самий код `Q520`, що означає, що ви можете пов’язати кілька варіантів вимови імені з однією <em>особою</em> без необхідності додавати кожне написання.

Ось правила:

<ul>
  <li>Якщо голосна (A, E, I, O, U) розділяє два приголосні з однаковим кодом soundex, приголосний праворуч від голосної кодується. Tymczak кодується як T-522 (T, 5 для M, 2 для C, Z ігнорується (див. правило "Side-by-Side" вище), 2 для K). Оскільки голосна "A" розділяє Z і K, K кодується.</li>
  <li>Якщо "H" або "W" розділяють два приголосні з однаковим кодом soundex, приголосний праворуч від голосної не кодується. Приклад: Ashcraft кодується як A-261 (A, 2 для S, C ігнорується, 6 для R, 1 для F). Він не кодується як A-226.</li>
</ul>

# --instructions--

Напишіть функцію, яка приймає рядок як параметр і повертає закодований рядок.

# --hints--

`soundex` має бути функцією.

```js
assert(typeof soundex == 'function');
```

`soundex("Soundex")` має повертати рядок.

```js
assert(typeof soundex('Soundex') == 'string');
```

`soundex("Soundex")` має повертати `"S532"`.

```js
assert.equal(soundex('Soundex'), 'S532');
```

`soundex("Example")` має повертати `"E251"`.

```js
assert.equal(soundex('Example'), 'E251');
```

`soundex("Sownteks")` має повертати `"S532"`.

```js
assert.equal(soundex('Sownteks'), 'S532');
```

`soundex("Ekzampul")` має повертати `"E251"`.

```js
assert.equal(soundex('Ekzampul'), 'E251');
```

`soundex("Euler")` має повертати `"E460"`.

```js
assert.equal(soundex('Euler'), 'E460');
```

`soundex("Gauss")` має повертати `"G200"`.

```js
assert.equal(soundex('Gauss'), 'G200');
```

`soundex("Hilbert")` має повертати `"H416"`.

```js
assert.equal(soundex('Hilbert'), 'H416');
```

`soundex("Knuth")` має повертати `"K530"`.

```js
assert.equal(soundex('Knuth'), 'K530');
```

`soundex("Lloyd")` має повертати `"L300"`.

```js
assert.equal(soundex('Lloyd'), 'L300');
```

`soundex("Lukasiewicz")` має повертати `"L222"`.

```js
assert.equal(soundex('Lukasiewicz'), 'L222');
```

# --seed--

## --seed-contents--

```js
function soundex(s) {

}
```

# --solutions--

```js
function soundex(s) {
  var a = s.toLowerCase().split('');
  var f = a.shift(),
    r = '',
    codes = {
      a: '',
      e: '',
      i: '',
      o: '',
      u: '',
      b: 1,
      f: 1,
      p: 1,
      v: 1,
      c: 2,
      g: 2,
      j: 2,
      k: 2,
      q: 2,
      s: 2,
      x: 2,
      z: 2,
      d: 3,
      t: 3,
      l: 4,
      m: 5,
      n: 5,
      r: 6
    };
  r =
    f +
    a
      .map(function(v, i, a) {
        return codes[v];
      })
      .filter(function(v, i, a) {
        return i === 0 ? v !== codes[f] : v !== a[i - 1];
      })
      .join('');

  return (r + '000').slice(0, 4).toUpperCase();
}
```
