---
id: 5a23c84252665b21eecc7eb0
title: "I перед E, крім випадків після C"
challengeType: 1
forumTopicId: 302288
dashedName: i-before-e-except-after-c
---

# --description--

**<a href="http://www.rosettacode.org/wiki/I_before_E_except_after_C" target="_blank" rel="noopener noreferrer nofollow">"I перед E, крім випадків після C"</a>** — це загальне правило англійської орфографії. Якщо ви не впевнені, чи слово пишеться з диграфом `ei` чи `ie`, то рима підказує, що правильний порядок — `ie`, якщо перед ним немає літери `c`, у такому разі це може бути `ei`.

Використовуючи наведені слова, перевірте, чи є дві підумови фрази правдоподібними окремо:

<ol>
  <li>
    <i>"I перед E, коли перед ними немає C".</i>
  </li>
  <li>
    <i>"E перед I, коли перед ними є C".</i>
  </li>
</ol>

Якщо обидві підфрази правдоподібні, тоді можна сказати, що оригінальна фраза є правдоподібною.

# --instructions--

Напишіть функцію, яка приймає слово і перевіряє, чи дотримується воно цього правила. Функція має повертати true, якщо слово відповідає правилу, і false, якщо ні.

# --hints--

`IBeforeExceptC` має бути функцією.

```js
assert(typeof IBeforeExceptC == 'function');
```

`IBeforeExceptC("receive")` має повертати булеве значення.

```js
assert(typeof IBeforeExceptC('receive') == 'boolean');
```

`IBeforeExceptC("receive")` має повертати `true`.

```js
assert.equal(IBeforeExceptC('receive'), true);
```

`IBeforeExceptC("science")` має повертати `false`.

```js
assert.equal(IBeforeExceptC('science'), false);
```

`IBeforeExceptC("imperceivable")` має повертати `true`.

```js
assert.equal(IBeforeExceptC('imperceivable'), true);
```

`IBeforeExceptC("inconceivable")` має повертати `true`.

```js
assert.equal(IBeforeExceptC('inconceivable'), true);
```

`IBeforeExceptC("insufficient")` має повертати `false`.

```js
assert.equal(IBeforeExceptC('insufficient'), false);
```

`IBeforeExceptC("omniscient")` має повертати `false`.

```js
assert.equal(IBeforeExceptC('omniscient'), false);
```

# --seed--

## --seed-contents--

```js
function IBeforeExceptC(word) {

}
```

# --solutions--

```js
function IBeforeExceptC(word)
{
    if(word.indexOf("c")==-1 && word.indexOf("ie")!=-1)
        return true;
    else if(word.indexOf("cei")!=-1)
        return true;
    return false;
}
```
