---
id: 5a23c84252665b21eecc7ec4
title: jortSort
challengeType: 1
forumTopicId: 302293
dashedName: jortsort
---

# --description--

jortSort — це набір інструментів для сортування, який змушує користувача виконувати роботу і гарантує ефективність, бо вам більше ніколи не доведеться сортувати. Спочатку його представила Дженн «Moneydollars» Шиффер на престижній конференції JSConf2014.

jortSort має бути функцією, яка приймає один масив порівнюваних об’єктів як аргумент. Потім вона сортує масив за зростанням і порівнює відсортований масив із початково наданим масивом. Якщо масиви співпадають (тобто початковий масив уже був відсортований), функція повертає true. Якщо масиви не співпадають (тобто початковий масив не був відсортований), функція повертає false.

# --hints--

`jortsort` має бути функцією.

```js
assert(typeof jortsort == 'function');
```

`jortsort([1,2,3,4,5])` має повертати булевий тип.

```js
assert(typeof jortsort([1, 2, 3, 4, 5]) == 'boolean');
```

`jortsort([1,2,3,4,5])` має повертати `true`.

```js
assert.equal(jortsort([1, 2, 3, 4, 5]), true);
```

`jortsort([1,2,13,4,5])` має повертати `false`.

```js
assert.equal(jortsort([1, 2, 13, 4, 5]), false);
```

`jortsort([12,4,51,2,4])` має повертати `false`.

```js
assert.equal(jortsort([12, 4, 51, 2, 4]), false);
```

`jortsort([1,2])` має повертати `true`.

```js
assert.equal(jortsort([1, 2]), true);
```

`jortsort([5,4,3,2,1])` має повертати `false`.

```js
assert.equal(jortsort([5, 4, 3, 2, 1]), false);
```

`jortsort([1,1,1,1,1])` має повертати `true`.

```js
assert.equal(jortsort([1, 1, 1, 1, 1]), true);
```

# --seed--

## --seed-contents--

```js
function jortsort(array) {

}
```

# --solutions--

```js
function jortsort(array) {
  // sort the array
  var originalArray = array.slice(0);
  array.sort( function(a,b){return a - b} );

  // compare to see if it was originally sorted
  for (var i = 0; i < originalArray.length; ++i) {
    if (originalArray[i] !== array[i]) return false;
  }

  return true;
};
```
