---
id: 5a23c84252665b21eecc7e7b
title: Генератор/експоненціальний
challengeType: 1
forumTopicId: 302275
dashedName: generatorexponential
---

# --description--

Генератор — це виконуваний об’єкт (як функція або процедура), який містить код, що повертає послідовність значень по одному за раз, так що кожного разу, коли ви викликаєте генератор, надається наступне значення послідовності.

Генератори часто будуються на основі корутин або об’єктів, щоб внутрішній стан об’єкта оброблявся «природно».

Генератори часто використовують у ситуаціях, коли послідовність потенційно нескінченна, і де можливо побудувати наступне значення послідовності, маючи лише мінімальний стан.

# --instructions--

Напишіть функцію, яка використовує генератори для створення квадратів і кубів. Створіть новий генератор, який фільтрує всі куби з генератора квадратів.

Функція має повертати \\( n^{th} \\) значення відфільтрованого генератора.

Наприклад, для \\(n=7\\) функція має повернути 81, оскільки послідовність буде 4, 9, 16, 25, 36, 49, 81. Тут 64 відфільтровано, бо це куб.

# --hints--

`exponentialGenerator` має бути функцією.

```js
assert(typeof exponentialGenerator == 'function');
```

`exponentialGenerator()` має повертати число.

```js
assert(typeof exponentialGenerator(10) == 'number');
```

`exponentialGenerator(10)` має повертати `144`.

```js
assert.equal(exponentialGenerator(10), 144);
```

`exponentialGenerator(12)` має повертати `196`.

```js
assert.equal(exponentialGenerator(12), 196);
```

`exponentialGenerator(14)` має повертати `256`.

```js
assert.equal(exponentialGenerator(14), 256);
```

`exponentialGenerator(20)` має повертати `484`.

```js
assert.equal(exponentialGenerator(20), 484);
```

`exponentialGenerator(25)` має повертати `784`.

```js
assert.equal(exponentialGenerator(25), 784);
```

# --seed--

## --seed-contents--

```js
function exponentialGenerator(n) {

}
```

# --solutions--

```js
function exponentialGenerator(n){
  function* PowersGenerator(m) {
    var n=0;
    while(1) {
        yield Math.pow(n, m);
        n += 1;
    }
  }

  function* FilteredGenerator(g, f){
    var value = g.next().value;
    var filter = f.next().value;
    while(1) {
        if( value < filter ) {
            yield value;
            value = g.next().value;
        } else if ( value > filter ) {
            filter = f.next().value;
        } else {
            value = g.next().value;
            filter = f.next().value;
        }
    }
  }

  var squares = PowersGenerator(2);
  var cubes = PowersGenerator(3);

  var filtered = FilteredGenerator(squares, cubes);

  var curr=0;
  for(var i=0;i<n;i++) curr=filtered.next();

  return curr.value;
}
```
