---
id: 5a23c84252665b21eecc7ede
title: Визначення високосного року
challengeType: 1
forumTopicId: 302300
dashedName: leap-year
---

# --description--

Визначте, чи є заданий рік високосним за григоріанським календарем.

# --hints--

`isLeapYear` має бути функцією.

```js
assert(typeof isLeapYear == 'function');
```

`isLeapYear()` має повертати булеве значення.

```js
assert(typeof isLeapYear(2018) == 'boolean');
```

`isLeapYear(2018)` має повертати `false`.

```js
assert.equal(isLeapYear(2018), false);
```

`isLeapYear(2016)` має повертати `true`.

```js
assert.equal(isLeapYear(2016), true);
```

`isLeapYear(2000)` має повертати `true`.

```js
assert.equal(isLeapYear(2000), true);
```

`isLeapYear(1900)` має повертати `false`.

```js
assert.equal(isLeapYear(1900), false);
```

`isLeapYear(1996)` має повертати `true`.

```js
assert.equal(isLeapYear(1996), true);
```

`isLeapYear(1800)` має повертати `false`.

```js
assert.equal(isLeapYear(1800), false);
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

}
```

# --solutions--

```js
function isLeapYear(year) {
  return year % 100 === 0 ? year % 400 === 0 : year % 4 === 0;
}
```
