---
id: 5a23c84252665b21eecc7ee0
title: Ліві факторіали
challengeType: 1
forumTopicId: 302302
dashedName: left-factorials
---

# --description--

**Ліві факторіали**, $ !n $, можуть означати або *субфакторіали*, або *суми факторіалів*. Та сама нотація може збивати з пантелику, оскільки використовується для двох різних визначень. Іноді *субфакторіали* (також відомі як *деренжменти*) можуть позначатися будь-якою з таких нотацій:

<ul>
  <li>$!n`$</li>
  <li>$!n$</li>
  <li>$n¡$</li>
</ul>

(Візуально це може бути неочевидно, але останній приклад використовує перевернутий знак оклику.) У цьому завданні буде використано таку формулу для **лівого факторіала**:

$ !n = \\sum\_{k=0}^{n-1} k! $

де $!0 = 0$

# --instructions--

Напишіть функцію для обчислення лівого факторіала заданого числа.

# --hints--

`leftFactorial` має бути функцією.

```js
assert(typeof leftFactorial == 'function');
```

`leftFactorial(0)` має повертати число.

```js
assert(typeof leftFactorial(0) == 'number');
```

`leftFactorial(0)` має повертати `0`.

```js
assert.equal(leftFactorial(0), 0);
```

`leftFactorial(1)` має повертати `1`.

```js
assert.equal(leftFactorial(1), 1);
```

`leftFactorial(2)` має повертати `2`.

```js
assert.equal(leftFactorial(2), 2);
```

`leftFactorial(3)` має повертати `4`.

```js
assert.equal(leftFactorial(3), 4);
```

`leftFactorial(10)` має повертати `409114`.

```js
assert.equal(leftFactorial(10), 409114);
```

`leftFactorial(17)` має повертати `22324392524314`.

```js
assert.equal(leftFactorial(17), 22324392524314);
```

`leftFactorial(19)` має повертати `6780385526348314`.

```js
assert.equal(leftFactorial(19), 6780385526348314);
```

# --seed--

## --seed-contents--

```js
function leftFactorial(n) {

}
```

# --solutions--

```js
function leftFactorial(n) {
  if (n == 0) return 0;
  if (n == 1) return 1;

  // Note: for n>=20, the result may not be correct.
  // This is because JavaScript uses 53 bit integers and
  // for n>=20 result becomes too large.

  let res = 2,
    fact = 2;
  for (var i = 2; i < n; i++) {
    res += fact;
    fact *= i + 1;
  }

  return res;
}
```
