---
id: 5a23c84252665b21eecc7eca
title: Числа Капрекара
challengeType: 1
forumTopicId: 302296
dashedName: kaprekar-numbers
---

# --description--

Позитивне ціле число є числом Капрекара, якщо:

<ul>
  <li>воно дорівнює 1, або</li>
  <li>десяткове представлення його квадрата можна розбити один раз на дві частини, які складаються з натуральних чисел і сума яких дорівнює початковому числу.</li>
</ul>

Зверніть увагу, що розбиття, яке дає частину, що складається виключно з нулів, не є дійсним, оскільки 0 не вважається позитивним числом.

Приклади чисел Капрекара:

<ul>
  <li><code>2223</code> є числом Капрекара, оскільки <code>2223 * 2223 = 4941729</code>, <code>4941729</code> можна розбити на <code>494</code> і <code>1729</code>, а <code>494 + 1729 = 2223</code></li>
  <li>Послідовність чисел Капрекара відома як A006886 і починається з <code>1, 9, 45, 55, ...</code></li>
</ul>

# --instructions--

Напишіть функцію, яка приймає число $n$, основу $bs$ і повертає true, якщо число є числом Капрекара для заданої основи. Інакше функція повертає false.

# --hints--

`isKaprekar` має бути функцією.

```js
assert(typeof isKaprekar == 'function');
```

`isKaprekar(1, 10)` має повертати булеве значення.

```js
assert(typeof isKaprekar(1, 10) == 'boolean');
```

`isKaprekar(1, 10)` має повертати `true`.

```js
assert.equal(isKaprekar(1, 10), true);
```

`isKaprekar(9, 10)` має повертати `true`.

```js
assert.equal(isKaprekar(9, 10), true);
```

`isKaprekar(2223, 10)` має повертати `true`.

```js
assert.equal(isKaprekar(2223, 10), true);
```

`isKaprekar(22823, 10)` має повертати `false`.

```js
assert.equal(isKaprekar(22823, 10), false);
```

`isKaprekar(9, 17)` має повертати `false`.

```js
assert.equal(isKaprekar(9, 17), false);
```

`isKaprekar(225, 17)` має повертати `true`.

```js
assert.equal(isKaprekar(225, 17), true);
```

`isKaprekar(999, 17)` має повертати `false`.

```js
assert.equal(isKaprekar(999, 17), false);
```

# --seed--

## --seed-contents--

```js
function isKaprekar(n, bs) {

}
```

# --solutions--

```js
function isKaprekar(n, bs) {
  if (n < 1) return false;
  if (n == 1) return true;
  for (var a = n * n, b = 0, s = 1; a; s *= bs) {
    b += (a % bs) * s;
    a = Math.floor(a / bs);
    if (b && a + b == n) return true;
  }
  return false;
}
```
