---
id: 5e4ce2eaac708cc68c1df260
title: Відстань Левенштейна
challengeType: 1
forumTopicId: 385264
dashedName: levenshtein-distance
---

# --description--

У теорії інформації та інформатиці **відстань Левенштейна** — це метрика для вимірювання різниці між двома послідовностями (тобто редагувальна відстань). Відстань Левенштейна між двома рядками визначається як мінімальна кількість редагувань, необхідних для перетворення одного рядка в інший, при цьому дозволеними операціями редагування є вставка, видалення або заміна одного символу.

Приклад:

Відстань Левенштейна між "**kitten**" і "**sitting**" дорівнює 3, оскільки наступні три редагування змінюють один рядок на інший, і неможливо зробити це менш ніж трьома редагуваннями:

<ul>
  <li><strong>k</strong>itten   <strong>s</strong>itten    (заміна 'k' на 's')</li>
  <li>sitt<strong>e</strong>n   sitt<strong>i</strong>n    (заміна 'e' на 'i')</li>
  <li>sittin   sittin<strong>g</strong>    (вставка 'g' в кінець).</li>
</ul>

*Відстань Левенштейна між "**rosettacode**" і "**raisethysword**" дорівнює **8**.*

*Відстань між двома рядками однакова, якщо обидва рядки перевернути.*

# --instructions--

Напишіть функцію, яка повертає відстань Левенштейна між двома рядками, переданими як параметри.

# --hints--

`levenshtein` має бути функцією.

```js
assert(typeof levenshtein == 'function');
```

`levenshtein("mist", "dist")` має повертати число.

```js
assert(typeof levenshtein('mist', 'dist') == 'number');
```

`levenshtein("mist", "dist")` має повертати `1`.

```js
assert.equal(levenshtein('mist', 'dist'), 1);
```

`levenshtein("tier", "tor")` має повертати `2`.

```js
assert.equal(levenshtein('tier', 'tor'), 2);
```

`levenshtein("kitten", "sitting")` має повертати `3`.

```js
assert.equal(levenshtein('kitten', 'sitting'), 3);
```

`levenshtein("stop", "tops")` має повертати `2`.

```js
assert.equal(levenshtein('stop', 'tops'), 2);
```

`levenshtein("rosettacode", "raisethysword")` має повертати `8`.

```js
assert.equal(levenshtein('rosettacode', 'raisethysword'), 8);
```

`levenshtein("mississippi", "swiss miss")` має повертати `8`.

```js
assert.equal(levenshtein('mississippi', 'swiss miss'), 8);
```

# --seed--

## --seed-contents--

```js
function levenshtein(a, b) {

}
```

# --solutions--

```js
function levenshtein(a, b) {
  var t = [], u, i, j, m = a.length, n = b.length;
  if (!m) { return n; }
  if (!n) { return m; }
  for (j = 0; j <= n; j++) { t[j] = j; }
  for (i = 1; i <= m; i++) {
    for (u = [i], j = 1; j <= n; j++) {
      u[j] = a[i - 1] === b[j - 1] ? t[j - 1] : Math.min(t[j - 1], t[j], u[j - 1]) + 1;
    } t = u;
  } return u[n];
}
```
