---
id: 5e4ce2f5ac708cc68c1df261
title: Лінійний конгруентний генератор
challengeType: 1
forumTopicId: 385266
dashedName: linear-congruential-generator
---

# --description--

Лінійний конгруентний генератор (LCG) — це <em>алгоритм</em>, який створює послідовність псевдовипадкових чисел, обчислених за допомогою розривної кусочної лінійної формули. Всі лінійні конгруентні генератори використовують таку формулу:

$$r_{n + 1} = (a \times r_n + c) \bmod m$$

Де:

<ul>
<li>$ r_0 $ — початкове значення (насіння).</li>
<li>$r_1$, $r_2$, $r_3$, ..., — випадкові числа.</li>
<li>$a$, $c$, $m$ — константи.</li>
</ul>

Якщо обрати значення $a$, $c$ та $m$ уважно, генератор створює рівномірний розподіл цілих чисел від $0$ до $m - 1$.

Числа <abbr title="лінійний конгруентний генератор">LCG</abbr> мають низьку якість. $r_n$ та $r\_{n + 1}$ не є незалежними, як справжні випадкові числа. Той, хто знає $r_n$, може передбачити $r\_{n + 1}$, тому <abbr title="лінійний конгруентний генератор">LCG</abbr> не є криптографічно безпечним. Проте <abbr title="лінійний конгруентний генератор">LCG</abbr> достатньо добрий для простих завдань, як-от тест простоти Міллера-Рабіна або роздачі у FreeCell. Серед переваг <abbr title="лінійний конгруентний генератор">LCG</abbr> — легкість відтворення послідовності чисел з того самого $r_0$. Також таку послідовність можна відтворити в іншій мові програмування, бо формула дуже проста.

# --instructions--

Напишіть функцію, яка приймає $r_0,a,c,m,n$ як параметри і повертає $r_n$.

# --hints--

`linearCongGenerator` має бути функцією.

```js
assert(typeof linearCongGenerator == 'function');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` має повертати число.

```js
assert(typeof linearCongGenerator(324, 1145, 177, 2148, 3) == 'number');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` має повертати `855`.

```js
assert.equal(linearCongGenerator(324, 1145, 177, 2148, 3), 855);
```

`linearCongGenerator(234, 11245, 145, 83648, 4)` має повертати `1110`.

```js
assert.equal(linearCongGenerator(234, 11245, 145, 83648, 4), 1110);
```

`linearCongGenerator(85, 11, 1234, 214748, 5)` має повертати `62217`.

```js
assert.equal(linearCongGenerator(85, 11, 1234, 214748, 5), 62217);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 1)` має повертати `12345`.

```js
assert.equal(linearCongGenerator(0, 1103515245, 12345, 2147483648, 1), 12345);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 2)` має повертати `1406932606`.

```js
assert.equal(
  linearCongGenerator(0, 1103515245, 12345, 2147483648, 2),
  1406932606
);
```

# --seed--

## --seed-contents--

```js
function linearCongGenerator(r0, a, c, m, n) {

}
```

# --solutions--

```js
function linearCongGenerator(r0, a, c, m, n) {
    for (let i = 0; i < n; i++) {
        r0 = (a * r0 + c) % m;
    }
    return r0;
}
```
