---
id: 5e6dd14797f5ce267c2f19d0
title: послідовність look-and-say
challengeType: 1
forumTopicId: 385277
dashedName: look-and-say-sequence
---

# --description--

Послідовність Look and say — це рекурсивно визначена послідовність чисел.

Визначення послідовності

<ul><li>Візьміть десяткове число</li>
<li><span>Подивіться</span> на число, візуально згрупувавши послідовні повторення однакової цифри.</li>
<li><span>Скажіть</span> число, зліва направо, група за групою; скільки разів зустрічається ця цифра — а потім сама згрупована цифра.</li></ul><span> Це стає наступним числом послідовності.</span>

Приклад:

<ul><li>Починаючи з числа 1, у вас є <span>одна</span> 1, що дає 11</li>
<li>Починаючи з 11, у вас є <span>дві</span> 1. Тобто: 21</li>
<li>Починаючи з 21, у вас є <span>одна</span> 2, потім <span>одна</span> 1. Тобто: (12)(11), що стає 1211</li>
<li>Починаючи з 1211, у вас є <span>одна</span> 1, <span>одна</span> 2, потім <span>дві</span> 1. Тобто: (11)(12)(21), що стає 111221</li></ul>

# --instructions--

Напишіть функцію, яка приймає рядок як параметр, обробляє його та повертає отриманий рядок.

# --hints--

`lookAndSay` має бути функцією.

```js
assert(typeof lookAndSay == 'function');
```

`lookAndSay("1")` має повертати рядок.

```js
assert(typeof lookAndSay('1') == 'string');
```

`lookAndSay("1")` має повертати `"11"`.

```js
assert.equal(lookAndSay('1'), '11');
```

`lookAndSay("11")` має повертати `"21"`.

```js
assert.equal(lookAndSay('11'), '21');
```

`lookAndSay("21")` має повертати `"1211"`.

```js
assert.equal(lookAndSay('21'), '1211');
```

`lookAndSay("1211")` має повертати `"111221"`.

```js
assert.equal(lookAndSay('1211'), '111221');
```

`lookAndSay("3542")` має повертати `"13151412"`.

```js
assert.equal(lookAndSay('3542'), '13151412');
```

# --seed--

## --seed-contents--

```js
function lookAndSay(str) {

}
```

# --solutions--

```js
function lookAndSay(str) {
    return str.replace(/(.)\1*/g, function(seq, p1) {
      return seq.length.toString() + p1;
    });
}
```
