---
id: 5e6dd1278e6ca105cde40ea9
title: найдовша загальна підпослідовність
challengeType: 1
forumTopicId: 385271
dashedName: longest-common-subsequence
---

# --description--

**Найдовша загальна підпослідовність** (або **LCS**) груп A і B — це найдовша група елементів із A і B, які є спільними для обох груп і йдуть у тому ж порядку в кожній групі. Наприклад, послідовності `1234` і `1224533324` мають LCS `1234`:

<u>1234</u>

<u>12</u>245<u>3</u>332<u>4</u>

Для прикладу з рядками розглянемо послідовності `thisisatest` і `testing123testing`. LCS буде `tsitest`:

<u>t</u>hi<u>si</u>sa<u>test</u>

<u>t</u>e<u>s</u>t<u>i</u>ng123<u>test</u>ing.

Ваш код має працювати лише з рядками.

# --instructions--

Напишіть функцію з урахуванням регістру, яка повертає LCS двох рядків. Не потрібно показувати кілька LCS.

# --hints--

`lcs` має бути функцією.

```js
assert(typeof lcs == 'function');
```

`lcs("thisisatest", "testing123testing")` має повертати рядок.

```js
assert(typeof lcs('thisisatest', 'testing123testing') == 'string');
```

`lcs("thisisatest", "testing123testing")` має повертати `"tsitest"`.

```js
assert.equal(lcs('thisisatest', 'testing123testing'), 'tsitest');
```

`lcs("ABCDGH", "AEDFHR")` має повертати `"ADH"`.

```js
assert.equal(lcs('ABCDGH', 'AEDFHR'), 'ADH');
```

`lcs("AGGTAB", "GXTXAYB")` має повертати `"GTAB"`.

```js
assert.equal(lcs('AGGTAB', 'GXTXAYB'), 'GTAB');
```

`lcs("BDACDB", "BDCB")` має повертати `"BDCB"`.

```js
assert.equal(lcs('BDACDB', 'BDCB'), 'BDCB');
```

`lcs("ABAZDC", "BACBAD")` має повертати `"ABAD"`.

```js
assert.equal(lcs('ABAZDC', 'BACBAD'), 'ABAD');
```

# --seed--

## --seed-contents--

```js
function lcs(a, b) {

}
```

# --solutions--

```js
function lcs(a, b) {
  var aSub = a.substring(0, a.length - 1);
  var bSub = b.substring(0, b.length - 1);

  if (a.length === 0 || b.length === 0) {
    return '';
  } else if (a.charAt(a.length - 1) === b.charAt(b.length - 1)) {
    return lcs(aSub, bSub) + a.charAt(a.length - 1);
  } else {
    var x = lcs(a, bSub);
    var y = lcs(aSub, b);
    return (x.length > y.length) ? x : y;
  }
}
```
