---
id: 5e6decd8ec8d7db960950d1c
title: Розклад LU
challengeType: 1
forumTopicId: 385280
dashedName: lu-decomposition
---

# --description--

Кожна квадратна матриця $A$ може бути розкладена у добуток нижньої трикутної матриці $L$ та верхньої трикутної матриці $U$. Це відоме як розклад LU.

$A = LU$

Це модифікована форма методу Гауса.

Хоча розклад Холецкого працює лише для симетричних, додатньо визначених матриць, більш загальний розклад LU підходить для будь-якої квадратної матриці.

Існує кілька алгоритмів для обчислення $L$ та $U$.

Щоб вивести *алгоритм Кру*, для прикладу 3x3, потрібно розв’язати таку систему:

\\begin{align}A = \\begin{pmatrix} a\_{11} & a\_{12} & a\_{13}\\\\ a\_{21} & a\_{22} & a\_{23}\\\\ a\_{31} & a\_{32} & a\_{33}\\\\ \\end{pmatrix}= \\begin{pmatrix} l\_{11} & 0 & 0 \\\\ l\_{21} & l\_{22} & 0 \\\\ l\_{31} & l\_{32} & l\_{33}\\\\ \\end{pmatrix} \\begin{pmatrix} u\_{11} & u\_{12} & u\_{13} \\\\ 0 & u\_{22} & u\_{23} \\\\ 0 & 0 & u\_{33} \\end{pmatrix} = LU\\end{align}

Тепер нам потрібно розв’язати 9 рівнянь з 12 невідомими. Щоб система була однозначно розв’язною, зазвичай діагональні елементи $L$ встановлюють рівними 1

$l\_{11}=1$

$l\_{22}=1$

$l\_{33}=1$

тоді отримуємо розв’язну систему з 9 невідомими та 9 рівняннями.

\\begin{align}A = \\begin{pmatrix} a\_{11} & a\_{12} & a\_{13}\\\\ a\_{21} & a\_{22} & a\_{23}\\\\ a\_{31} & a\_{32} & a\_{33}\\\\ \\end{pmatrix} = \\begin{pmatrix} 1 & 0 & 0 \\\\ l\_{21} & 1 & 0 \\\\ l\_{31} & l\_{32} & 1\\\\ \\end{pmatrix} \\begin{pmatrix} u\_{11} & u\_{12} & u\_{13} \\\\ 0 & u\_{22} & u\_{23} \\\\ 0 & 0 & u\_{33} \\end{pmatrix} = \\begin{pmatrix} u\_{11} & u\_{12} & u\_{13} \\\\ u\_{11}l\_{21} & u\_{12}l\_{21}+u\_{22} & u\_{13}l\_{21}+u\_{23} \\\\ u\_{11}l\_{31} & u\_{12}l\_{31}+u\_{22}l\_{32} & u\_{13}l\_{31} + u\_{23}l\_{32}+u\_{33} \\end{pmatrix} = LU\\end{align}

Розв’язуючи для інших $l$ та $u$, отримуємо такі рівняння:

$u\_{11}=a\_{11}$

$u\_{12}=a\_{12}$

$u\_{13}=a\_{13}$

$u\_{22}=a\_{22} - u\_{12}l\_{21}$

$u\_{23}=a\_{23} - u\_{13}l\_{21}$

$u\_{33}=a\_{33} - (u\_{13}l\_{31} + u\_{23}l\_{32})$

а для $l$:

$l\_{21}=\\frac{1}{u\_{11}} a\_{21}$

$l\_{31}=\\frac{1}{u\_{11}} a\_{31}$

$l\_{32}=\\frac{1}{u\_{22}} (a\_{32} - u\_{12}l\_{31})$

Ми бачимо, що існує шаблон обчислень, який можна виразити такими формулами, спочатку для $U$

$u\_{ij} = a\_{ij} - \\sum\_{k=1}^{i-1} u\_{kj}l\_{ik}$

а потім для $L$

$l\_{ij} = \\frac{1}{u\_{jj}} (a\_{ij} - \\sum\_{k=1}^{j-1} u\_{kj}l\_{ik})$

У другій формулі видно, що щоб отримати $l\_{ij}$ під діагоналлю, потрібно поділити на діагональний елемент (півот) $u\_{jj}$, тому виникають проблеми, коли $u\_{jj}$ дорівнює 0 або дуже мале, що призводить до числової нестабільності.

Рішенням цієї проблеми є *півотування* $A$, що означає перестановку рядків $A$ перед розкладом $LU$ так, щоб найбільший елемент кожного стовпця опинився на діагоналі $A$. Перестановка рядків означає множення $A$ на матрицю перестановки $P$:

$PA \\Rightarrow A'$

Приклад:

\\begin{align} \\begin{pmatrix} 0 & 1 \\\\ 1 & 0 \\end{pmatrix} \\begin{pmatrix} 1 & 4 \\\\ 2 & 3 \\end{pmatrix} \\Rightarrow \\begin{pmatrix} 2 & 3 \\\\ 1 & 4 \\end{pmatrix} \\end{align}

Алгоритм розкладу застосовується до переставленої матриці так, що

$PA = LU$

# --instructions--

Завдання — реалізувати рутину, яка прийматиме квадратну матрицю розміру nxn $A$ і повертатиме нижню трикутну матрицю $L$, верхню трикутну матрицю $U$ та матрицю перестановки $P$, так щоб виконувалась наведена вище рівність. Повернене значення має бути у формі `[L, U, P]`.

# --before-each--

```js
const matrix1 = [
  [1, 3, 5],
  [2, 4, 7],
  [1, 1, 0]
];
const expected1 = [
  [
    [1, 0, 0],
    [0.5, 1, 0],
    [0.5, -1, 1]
  ],
  [
    [2, 4, 7],
    [0, 1, 1.5],
    [0, 0, -2]
  ],
  [
    [0, 1, 0],
    [1, 0, 0],
    [0, 0, 1]
  ]
];

const matrix2 = [
  [11, 9, 24, 2],
  [1, 5, 2, 6],
  [3, 17, 18, 1],
  [2, 5, 7, 1]
];
const expected2 = [
  [
    [1, 0, 0, 0],
    [0.2727272727272727, 1, 0, 0],
    [0.09090909090909091, 0.2875, 1, 0],
    [0.18181818181818182, 0.23124999999999996, 0.0035971223021580693, 1]
  ],
  [
    [11, 9, 24, 2],
    [0, 14.545454545454547, 11.454545454545455, 0.4545454545454546],
    [0, 0, -3.4749999999999996, 5.6875],
    [0, 0, 0, 0.510791366906476]
  ],
  [
    [1, 0, 0, 0],
    [0, 0, 1, 0],
    [0, 1, 0, 0],
    [0, 0, 0, 1]
  ]
];

const matrix3 = [
  [1, 1, 1],
  [4, 3, -1],
  [3, 5, 3]
];
const expected3 = [
  [
    [1, 0, 0],
    [0.75, 1, 0],
    [0.25, 0.09090909090909091, 1]
  ],
  [
    [4, 3, -1],
    [0, 2.75, 3.75],
    [0, 0, 0.9090909090909091]
  ],
  [
    [0, 1, 0],
    [0, 0, 1],
    [1, 0, 0]
  ]
];

const matrix4 = [
  [1, -2, 3],
  [2, -5, 12],
  [0, 2, -10]
];
const expected4 = [
  [
    [1, 0, 0],
    [0, 1, 0],
    [0.5, 0.25, 1]
  ],
  [
    [2, -5, 12],
    [0, 2, -10],
    [0, 0, -0.5]
  ],
  [
    [0, 1, 0],
    [0, 0, 1],
    [1, 0, 0]
  ]
];
```

# --hints--

`luDecomposition` має бути функцією.

```js
assert(typeof luDecomposition == 'function');
```

`luDecomposition([[1, 3, 5], [2, 4, 7], [1, 1, 0]])` має повертати масив.

```js
assert(Array.isArray(luDecomposition(matrix1)));
```

`luDecomposition([[1, 3, 5], [2, 4, 7], [1, 1, 0]])` має повертати `[[[1, 0, 0], [0.5, 1, 0], [0.5, -1, 1]], [[2, 4, 7], [0, 1, 1.5], [0, 0, -2]], [[0, 1, 0], [1, 0, 0], [0, 0, 1]]]`.

```js
assert.deepEqual(luDecomposition(matrix1), expected1);
```

`luDecomposition([[11, 9, 24, 2], [1, 5, 2, 6], [3, 17, 18, 1], [2, 5, 7, 1]])` має повертати `[[[1, 0, 0, 0], [0.2727272727272727, 1, 0, 0], [0.09090909090909091, 0.2875, 1, 0], [0.18181818181818182, 0.23124999999999996, 0.0035971223021580693, 1]], [[11, 9, 24, 2], [0, 14.545454545454547, 11.454545454545455, 0.4545454545454546], [0, 0, -3.4749999999999996, 5.6875], [0, 0, 0, 0.510791366906476]], [[1, 0, 0, 0], [0, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 1]]]`.

```js
assert.deepEqual(luDecomposition(matrix2), expected2);
```

`luDecomposition([[1, 1, 1], [4, 3, -1], [3, 5, 3]])` має повертати `[[[1, 0, 0], [0.75, 1, 0], [0.25, 0.09090909090909091, 1]], [[4, 3, -1], [0, 2.75, 3.75], [0, 0, 0.9090909090909091]], [[0, 1, 0], [0, 0, 1], [1, 0, 0]]]`.

```js
assert.deepEqual(luDecomposition(matrix3), expected3);
```

`luDecomposition([[1, -2, 3], [2, -5, 12], [0, 2, -10]])` має повертати `[[[1, 0, 0], [0, 1, 0], [0.5, 0.25, 1]], [[2, -5, 12], [0, 2, -10], [0, 0, -0.5]], [[0, 1, 0], [0, 0, 1], [1, 0, 0]]]`.

```js
assert.deepEqual(luDecomposition(matrix4), expected4);
```

# --seed--

## --seed-contents--

```js
function luDecomposition(A) {

}
```

# --solutions--

```js
function luDecomposition(A) {

    function dotProduct(a, b) {
        var sum = 0;
        for (var i = 0; i < a.length; i++)
            sum += a[i] * b[i]
        return sum;
    }

    function matrixMul(A, B) {
        var result = new Array(A.length);
        for (var i = 0; i < A.length; i++)
            result[i] = new Array(B[0].length)
        var aux = new Array(B.length);

        for (var j = 0; j < B[0].length; j++) {

            for (var k = 0; k < B.length; k++)
                aux[k] = B[k][j];

            for (var i = 0; i < A.length; i++)
                result[i][j] = dotProduct(A[i], aux);
        }
        return result;
    }

    function pivotize(m) {
        var n = m.length;
        var id = new Array(n);
        for (var i = 0; i < n; i++) {
            id[i] = new Array(n);
            id[i].fill(0)
            id[i][i] = 1;
        }

        for (var i = 0; i < n; i++) {
            var maxm = m[i][i];
            var row = i;
            for (var j = i; j < n; j++)
                if (m[j][i] > maxm) {
                    maxm = m[j][i];
                    row = j;
                }

            if (i != row) {
                var tmp = id[i];
                id[i] = id[row];
                id[row] = tmp;
            }
        }
        return id;
    }

    var n = A.length;
    var L = new Array(n);
    for (var i = 0; i < n; i++) { L[i] = new Array(n); L[i].fill(0) }
    var U = new Array(n);
    for (var i = 0; i < n; i++) { U[i] = new Array(n); U[i].fill(0) }
    var P = pivotize(A);
    var A2 = matrixMul(P, A);

    for (var j = 0; j < n; j++) {
        L[j][j] = 1;
        for (var i = 0; i < j + 1; i++) {
            var s1 = 0;
            for (var k = 0; k < i; k++)
                s1 += U[k][j] * L[i][k];
            U[i][j] = A2[i][j] - s1;
        }
        for (var i = j; i < n; i++) {
            var s2 = 0;
            for (var k = 0; k < j; k++)
                s2 += U[k][j] * L[i][k];
            L[i][j] = (A2[i][j] - s2) / U[j][j];
        }
    }
    return [L, U, P];
}
```
