---
id: 5e94a54cc7b022105bf0fd2c
title: Частота слів
challengeType: 1
forumTopicId: 393913
dashedName: word-frequency
---

# --description--

Дано рядок тексту та ціле число n. Поверніть n найпоширеніших слів у файлі (та кількість їхніх входжень) у порядку спадання частоти.

# --instructions--

Напишіть функцію, яка рахує кількість входжень кожного слова і повертає n найпоширеніших слів разом із кількістю їхніх входжень у порядку спадання частоти.

Функція має повертати двовимірний масив, де кожен елемент має такий вигляд: `[word, freq]`. `word` має бути версією слова в нижньому регістрі, а `freq` — числом, що позначає кількість.

Функція має повертати порожній масив, якщо рядок не надано.

Функція має бути нечутливою до регістру, наприклад, рядки "Hello" і "hello" слід розглядати як однакові.

Ви можете вважати слова зі спеціальними символами, такими як підкреслення, дефіси, апострофи, коми тощо, окремими словами.

Наприклад, для рядка "Hello hello goodbye" ваша функція має повернути `[['hello', 2], ['goodbye', 1]]`.

# --hints--

`wordFrequency` має бути функцією.

```js
assert.isFunction(wordFrequency);
```

`wordFrequency` має повертати масив.

```js
assert.isArray(wordFrequency('test'));
```

`wordFrequency("Hello hello world", 2)` має повертати `[['hello', 2], ['world', 1]]`

```js
const example_1 = 'Hello hello world';
const example_1_solution = [['hello', 2], ['world', 1]];

assert.deepEqual(wordFrequency(example_1, 2), example_1_solution);
```

`wordFrequency("The quick brown fox jumped over the lazy dog", 1)` має повертати `[['the', 2]]`

```js
const example_2 = 'The quick brown fox jumped over the lazy dog';
const example_2_solution = [['the', 2]];

assert.deepEqual(wordFrequency(example_2, 1), example_2_solution);
```

`wordFrequency("Opensource opensource open-source open source", 1)` має повертати `[['opensource', 2]]`

```js
const example_3 = 'Opensource opensource open-source open source';
const example_3_solution = [['opensource', 2]];

assert.deepEqual(wordFrequency(example_3, 1), example_3_solution);
```

`wordFrequency("Apple App apply aPP aPPlE", 3)` має повертати `[['app', 2], ['apple', 2], ['apply', 1]]` або `[['apple', 2], ['app', 2], ['apply', 1]]`

```js
const example_4 = 'Apple App apply aPP aPPlE';
const example_4_solution = [["app", 2],["apple", 2],["apply", 1]];

const arr = wordFrequency(example_4, 3);
assert.includeDeepMembers(arr, example_4_solution);
```

`wordFrequency("c d a d c a b d d c", 4)` має повертати `[['d', 4], ['c', 3], ['a', 2], ['b', 1]]`

```js
const example_5 = 'c d a d c a b d d c';
const example_5_solution = [['d', 4], ['c', 3], ['a', 2], ['b', 1]];

assert.deepEqual(wordFrequency(example_5, 4), example_5_solution);
```

`wordFrequency("", 5)` має повертати `[]`

```js
const example_6 = '';
const example_6_solution = [];

assert.deepEqual(wordFrequency(example_6, 5), example_6_solution);
```

# --seed--

## --seed-contents--

```js
function wordFrequency(txt, n) {

}
```

# --solutions--

```js
function wordFrequency(txt, n) {
  var words = txt.split(/\s+/);
  var wordCount = {};
  words.forEach(word => {
    if (word == '') {
      return;
    }
    const lowerWord = word.toLowerCase();
    wordCount[lowerWord] =
      lowerWord in wordCount ? wordCount[lowerWord] + 1 : 1;
  });

  var wordsArray = [];
  for (let [word, count] of Object.entries(wordCount)) {
    wordsArray.push([word, count]);
  }

  wordsArray.sort((a, b) => {
    if (a[1] !== b[1]) {
      return b[1] - a[1];
    } else if (a[0] !== b[0]) {
      return a[0] < b[0] ? -1 : 1;
    }
    return 0;
  });
  return wordsArray.slice(0, n);
}
```
