---
id: 5ea28156e79528a9ab248f27
title: Тест Луна для номерів кредитних карток
challengeType: 1
forumTopicId: 385284
dashedName: luhn-test-of-credit-card-numbers
---

# --description--

Тест Луна використовується деякими компаніями кредитних карток, щоб відрізнити дійсні номери кредитних карток від випадкового набору цифр.

Компанії, які використовують номери кредитних карток, що можна перевірити за допомогою тесту Луна, мають номери, які проходять наступний тест:

<ol>
  <li> Поміняйте порядок цифр у номері на зворотний.</li>
  <li> Візьміть першу, третю, ... і кожну іншу непарну цифру у зворотному порядку та підсумуйте їх, щоб отримати часткову суму s1.</li>
  <li> Візьміть другу, четверту ... і кожну іншу парну цифру у зворотному порядку:</li>
    <ol>
      <li>Помножте кожну цифру на два і, якщо відповідь більша за дев’ять, підсумуйте цифри, щоб отримати часткові суми для парних цифр.</li>
      <li>Підсумуйте часткові суми парних цифр, щоб отримати s2.</li>
    </ol>
  <li>Якщо s1 + s2 закінчується на нуль, тоді початковий номер є дійсним номером кредитної картки, підтвердженим тестом Луна.</li>
</ol>

Наприклад, якщо тестове число 49927398716:

```md
Reverse the digits:
  61789372994
Sum the odd digits:
  6 + 7 + 9 + 7 + 9 + 4 = 42 = s1
The even digits:
    1, 8, 3, 2, 9
  Two times each even digit:
    2, 16, 6, 4, 18
  Sum the digits of each multiplication:
    2, 7, 6, 4, 9
  Sum the last:
    2 + 7 + 6 + 4 + 9 = 28 = s2

s1 + s2 = 70 which ends in zero which means that 49927398716 passes the Luhn test.
```

# --instructions--

Напишіть функцію, яка перевірятиме число за тестом Луна. Поверніть true, якщо це дійсне число. Інакше поверніть false.

# --hints--

`luhnTest` має бути функцією.

```js
assert(typeof luhnTest === 'function');
```

`luhnTest("4111111111111111")` має повертати булеве значення.

```js
assert(typeof luhnTest('4111111111111111') === 'boolean');
```

`luhnTest("4111111111111111")` має повертати `true`.

```js
assert.equal(luhnTest('4111111111111111'), true);
```

`luhnTest("4111111111111112")` має повертати `false`.

```js
assert.equal(luhnTest('4111111111111112'), false);
```

`luhnTest("49927398716")` має повертати `true`.

```js
assert.equal(luhnTest('49927398716'), true);
```

`luhnTest("49927398717")` має повертати `false`.

```js
assert.equal(luhnTest('49927398717'), false);
```

`luhnTest("1234567812345678")` має повертати `false`.

```js
assert.equal(luhnTest('1234567812345678'), false);
```

`luhnTest("1234567812345670")` має повертати `true`.

```js
assert.equal(luhnTest('1234567812345670'), true);
```

# --seed--

## --seed-contents--

```js
function luhnTest(str) {

}
```

# --solutions--

```js
function luhnTest(str) {
  var luhnArr = [0, 2, 4, 6, 8, 1, 3, 5, 7, 9];
  var counter = 0;
  var incNum;
  var odd = false;
  var temp = String(str).replace(/[^\d]/g, '');
  if (temp.length == 0) return false;
  for (var i = temp.length - 1; i >= 0; --i) {
    incNum = parseInt(temp.charAt(i), 10);
    counter += (odd = !odd) ? incNum : luhnArr[incNum];
  }
  return counter % 10 == 0;
}
```
