---
id: 5eaf48389ee512d4d103684b
title: Самоописові числа
challengeType: 1
forumTopicId: 385289
dashedName: self-describing-numbers
---

# --description--

Існує кілька так званих "самоописових" цілих чисел.

Ціле число вважається "самоописовим", якщо воно має властивість, що, коли позиції цифр позначені від 0 до N-1, цифра на кожній позиції дорівнює кількості разів, скільки ця цифра зустрічається у числі.

Наприклад, **2020** — це чотиризначне самоописове число:

<ul>
    <li> позиція 0 має значення 2, і у числі є дві 0; </li>
    <li> позиція 1 має значення 0, і у числі немає 1; </li>
    <li> позиція 2 має значення 2, і у числі є дві 2; </li>
    <li> позиція 3 має значення 0, і у числі немає 3; </li>
</ul>

Самоописові числа менші за 100 000 000: 1210, 2020, 21200, 3211000, 42101000.

# --instructions--

Напишіть функцію, яка приймає позитивне ціле число як параметр. Якщо воно є самоописовим, поверніть true. Інакше поверніть false.

# --hints--

`isSelfDescribing` має бути функцією.

```js
assert(typeof isSelfDescribing == 'function');
```

`isSelfDescribing()` має повертати булеве значення.

```js
assert(typeof isSelfDescribing(2020) == 'boolean');
```

`isSelfDescribing(2020)` має повертати `true`.

```js
assert.equal(isSelfDescribing(2020), true);
```

`isSelfDescribing(3021)` має повертати `false`.

```js
assert.equal(isSelfDescribing(3021), false);
```

`isSelfDescribing(3211000)` має повертати `true`.

```js
assert.equal(isSelfDescribing(3211000), true);
```

# --seed--

## --seed-contents--

```js
function isSelfDescribing(n) {

}
```

# --solutions--

```js
function isSelfDescribing(n) {
    let digits = String(n).split("");
    digits = digits.map(function(e) {return parseInt(e)});
    let count = digits.map((x) => {return 0})
    digits.forEach((d) =>{
        if (d >= count.length) {
            return false
        }
        count[d] += 1;
    });

     if (digits === count) {
        return true;
    }
    if (digits.length != count.length) {
        return false;
    }
    
    for (let i=0; i< digits.length; i++){
      if (digits[i] !== count[i]) {
        return false;
      }
    }
    return true;
}
```
