---
id: 69aaafeaa4037197c75587cd
title: キューを実装する
challengeType: 26
dashedName: implement-a-queue
---

# --description--

このラボでは、関数を使ってキューのデータ構造を実装します。キューは先入れ先出し（FIFO）のデータ構造で、要素はキューの後ろに追加され、前から取り出されます。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. キューの後ろに要素を追加する`enqueue`関数が必要です。
1. キューの前の要素を取り出して返す`dequeue`関数が必要です。
1. キューの前の要素を取り出さずに返す`front`関数が必要です。
1. キュー内の要素数を返す`size`関数が必要です。
1. キューが空の場合は`true`を返し、そうでなければ`false`を返す`isEmpty`関数が必要です。

**注意:** ほとんどのテストは`enqueue`関数に依存しています。`dequeue`、`front`、`size`、`isEmpty`のテストはキューに要素を追加する必要があるため、まず`enqueue`を実装してください。

# --hints--

`enqueue`関数が必要です。

```js
assert.isFunction(enqueue);
```

`enqueue`関数はキューの後ろに要素を追加する必要があります。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`dequeue`関数が必要です。

```js
assert.isFunction(dequeue);
```

`dequeue`関数はキューの前の要素を取り出して返す必要があります。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

`dequeue`は空のキューに対して呼び出されたときに`undefined`を返す必要があります。

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue`は取り出し後もキューの後ろに要素を追加し続ける必要があります。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

異なるキューのインスタンスは状態を共有してはいけません。

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

`front`関数が必要です。

```js
assert.isFunction(front);
```

`front`関数はキューの前の要素を取り出さずに返す必要があります。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`front`は空のキューに対して呼び出されたときに`undefined`を返す必要があります。

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

`size`関数が必要です。

```js
assert.isFunction(size);
```

`size`関数はキュー内の正しい要素数を返す必要があります。

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

`isEmpty`関数が必要です。

```js
assert.isFunction(isEmpty);
```

`isEmpty`関数は空のキューに対して`true`を返す必要があります。

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

`isEmpty`関数は空でないキューに対して`false`を返す必要があります。

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

`isEmpty`関数はキューからすべての要素が取り除かれた後に`true`を返す必要があります。

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
