---
id: 69708d59833ceded7d00541f
title: Реалізувати операції зв’язаного списку
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

У цій лабораторній роботі ви реалізуєте додаткові операції для структури даних зв’язаний список, базуючись на базовій функціональності додавання та видалення вузлів.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача**

1. Ви повинні мати `contains` функцію, яка приймає зв’язаний список і елемент. Вона має повертати `true`, якщо вказаний елемент існує у зв’язаному списку, інакше — `false`.
1. Ви повинні мати `getAt` функцію, яка приймає зв’язаний список і індекс. Вона має повертати елемент за заданим індексом у зв’язаному списку. Якщо індекс виходить за межі, має повертати `undefined`.
1. Ви повинні мати `insertAt` функцію, яка приймає зв’язаний список, індекс і елемент. Вона має вставити заданий елемент на вказану позицію у зв’язаному списку. Якщо індекс виходить за межі, список не повинен змінюватися.
1. Ви повинні мати `removeAt` функцію, яка приймає зв’язаний список і індекс. Вона має видалити вузол за заданим індексом у зв’язаному списку. Якщо індекс виходить за межі, список не повинен змінюватися.
1. Ви повинні мати `clear` функцію, яка приймає зв’язаний список. Вона має видалити всі елементи зі зв’язаного списку, ефективно скидаючи його до порожнього стану.

**Примітка:** Деякі пізніші тести залежать від раніше реалізованих методів. Наприклад, якщо `getAt` реалізовано неправильно, тести для функцій на кшталт `insertAt` і `removeAt` можуть не пройти, навіть якщо ці функції майже правильні.

# --hints--

`initList` має повертати об’єкт із `head`, встановленим у `null`, і `length`, встановленим у `0`.

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

`add` має збільшувати довжину списку на один.

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

`remove` має зменшувати довжину списку на один.

```js
const list = initList();  
add(list, 'apple');  
add(list, 'banana');  
remove(list, 'apple');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(list.head.element, 'banana');  
assert.isNull(list.head.next);
```

Ви повинні мати `contains` функцію.

```js
assert.isFunction(contains);
```

`contains` має повертати `true`, якщо елемент існує у списку.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains` має повертати `false`, якщо елемент не існує у списку.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains` має повертати `false` для порожнього списку.

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

Ви повинні мати `getAt` функцію.

```js
assert.isFunction(getAt);
```

`getAt` має повертати елемент за заданим індексом.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt` має повертати `undefined` для недійсних індексів.

```js
const list = initList();  
add(list, 'apple');  
assert.isUndefined(getAt(list, -1));  
assert.isUndefined(getAt(list, 1));  
assert.isUndefined(getAt(list, 5));
```

`getAt` має повертати `undefined` для порожнього списку.

```js
const list = initList();  
assert.isUndefined(getAt(list, 0));
```

Ви повинні мати `insertAt` функцію.

```js
assert.isFunction(insertAt);
```

`insertAt` має вставляти елемент на початок списку.

```js
const list = initList();  
add(list, 'second');  
add(list, 'third');  
insertAt(list, 0, 'first');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` має вставляти елемент у середину списку.

```js
const list = initList();  
add(list, 'first');  
add(list, 'third');  
insertAt(list, 1, 'second');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` має вставляти елемент у кінець списку.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` має вставляти у порожній список за індексом `0`.  

```js  
const list = initList();  
insertAt(list, 0, 'first');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.isNull(list.head.next); 
```

`insertAt` не має змінювати список для недійсних індексів.

```js
const list = initList();  
add(list, 'apple');  
insertAt(list, -1, 'invalid');  
insertAt(list, 10, 'invalid');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
assert.isUndefined(getAt(list, 1));
```

Ви повинні мати `removeAt` функцію.

```js
assert.isFunction(removeAt);
```

`removeAt` має видаляти елемент за вказаним індексом.

```js
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'third');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` має правильно видаляти перший елемент.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

`removeAt` має правильно видаляти останній елемент.  

```js  
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 2);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` має скинути список при видаленні єдиного елемента.

```js
const list = initList();  
add(list, 'only');  
removeAt(list, 0);  
assert.strictEqual(list.length, 0);  
assert.isNull(list.head);  
assert.isTrue(isEmpty(list));
```

`removeAt` не має змінювати список для недійсних індексів.

```js
const list = initList();  
add(list, 'apple');  
removeAt(list, -1);  
removeAt(list, 10);  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
```

Ви повинні мати `clear` функцію.

```js
assert.isFunction(clear);
```

`clear` має видаляти всі елементи зі списку.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`clear` має дозволяти повторно використовувати список.  

```js  
const list = initList();  
add(list, 'first');  
clear(list);  
clear(list);  
add(list, 'second');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'second');  
assert.isNull(list.head.next); 
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  
}

function getAt(list, index) {
 
}

function insertAt(list, index, element) {
  
}

function removeAt(list, index) {
 
}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
