---
id: 66b102ee0bdbad34a9f85ba0
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 문자열에 대해 계속 배우면서 선생님 챗봇을 만들어 보세요.

먼저, `console` 문장을 추가하고 메시지로 `"Hi there!"`을 입력하세요.

# --hints--

`console` 문장이 있어야 합니다.

```js
assert.match(code, /console\.log(.*)/);
```

콘솔 문장은 메시지 `"Hi there!"`을 출력해야 합니다.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
