---
id: 66b59be6ab830800c4df9146
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

이제 이전 단계에서 만든 변수를 사용할 차례입니다.

이전 수업에서 다음과 같이 템플릿 리터럴을 다루는 방법을 배웠습니다:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

`sentence` 변수를 먼저 만드세요.

템플릿 리터럴 구문을 사용하여 `Today, you will learn about [topic variable goes here] in [subject variable goes here].` 변수에 문자열 `sentence`을 할당하세요.

`[topic variable goes here]`와 `[subject variable goes here]` 자리 표시자를 `topic`와 `subject` 변수 및 `${}` 구문으로 바꿀 것입니다.

마지막으로 `sentence` 변수를 콘솔에 출력하세요.

# --hints--

`sentence` 변수가 있어야 합니다.

```js
assert.isNotNull(sentence);
```

`sentence` 변수는 문자열이어야 합니다.

```js
assert.isString(sentence);
```

템플릿 리터럴 구문을 사용하여 `Today, you will learn about [topic variable goes here] in [subject variable goes here].` 변수에 문자열 `sentence`을 할당해야 합니다. 도움이 필요하면 예제를 참조하세요.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

`sentence` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
