---
id: 66b59829dba144ff1351220f
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이제 `botName` 변수를 사용해 인사말을 만들어 보세요.

이전 수업에서 템플릿 리터럴을 사용해 문자열을 연결하는 방법을 배웠습니다. 예를 들어 다음과 같습니다:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

먼저 `greeting`라는 변수를 만드세요.

다음으로 템플릿 리터럴 구문을 사용해 `My name is`라는 문자열을 `botName` 변수 뒤에 붙이고, 마침표(`.`)로 끝나는 문자열을 할당하세요.

마지막으로 `greeting` 변수를 콘솔에 출력하세요.

# --hints--

`greeting`라는 변수가 있어야 합니다.

```js
assert.isNotNull(greeting);
```

`greeting` 변수에 문자열을 할당해야 합니다.

```js
assert.isString(greeting);
```

템플릿 리터럴을 사용해 `My name is` 문자열과 `botName` 변수를 마침표(`.`)와 함께 연결해야 합니다.

```js
assert.equal(greeting, "My name is teacherBot.");
```

`console` 문장으로 `greeting` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
