---
id: 66b5928b86e907fdfbf34e56
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이제 봇의 이름을 설정할 차례입니다.

`botName`라는 변수를 만들고 `"teacherBot"` 문자열 값을 할당하세요.

# --hints--

`botName`라는 변수가 있어야 합니다.

```js
assert.isNotNull(botName);
```

`botName` 변수는 문자열이어야 합니다.

```js
assert.isString(botName);
```

`botName` 변수는 `"teacherBot"` 값을 가져야 합니다.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
