---
id: 66b59b12a745e10011158f7b
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

다음 단계는 앞으로 봇 메시지에서 사용할 몇 가지 변수를 더 만드는 것입니다.

`subject`라는 변수를 만들고 문자열 값 `"JavaScript"`을 할당하세요.

그다음 `topic`라는 변수를 만들고 문자열 값 `"strings"`을 할당하세요.

# --hints--

`subject`라는 변수가 있어야 합니다.

```js
assert.isNotNull(subject);
```

`subject` 변수에 문자열을 할당해야 합니다.

```js
assert.isString(subject);
```

`"JavaScript"` 변수에 문자열 `subject`을 할당해야 합니다.

```js
assert.equal(subject, "JavaScript");
```

`topic`라는 변수가 있어야 합니다.

```js
assert.isNotNull(topic);
```

`topic` 변수에 문자열을 할당해야 합니다.

```js
assert.isString(topic);
```

`"strings"` 변수에 문자열 `topic`을 할당해야 합니다.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
