---
id: 66b6d482bbb9e12f2e5ee1ae
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

프로젝트의 다음 부분에서는 봇이 문자열 `length` 속성을 다루는 방법을 가르칩니다.

먼저 `strLengthIntro`라는 새 변수를 만드세요.

그다음 템플릿 리터럴 구문을 사용해 `Here is an example of using the length property on the word [subject].` 변수에 문자열 `strLengthIntro`를 할당하세요.

이전에 했던 것처럼 `[subject]`를 `subject` 변수로 바꾸세요.

마지막으로 `strLengthIntro` 변수를 콘솔에 출력하세요.

# --hints--

`strLengthIntro`라는 변수가 있어야 합니다.

```js
assert.isNotNull(strLengthIntro);
```

`strLengthIntro` 변수는 문자열이어야 합니다.

```js
assert.isString(strLengthIntro);
```

템플릿 리터럴 구문을 사용해 `Here is an example of using the length property on the word [subject goes here].` 변수에 문자열 `strLengthIntro`를 할당해야 합니다.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

`strLengthIntro` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
