---
id: 66b6e80d6c3f0b329c360283
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

이 워크숍의 다음 부분은 문자열에서 문자를 접근하는 것을 검토하는 것입니다.

먼저 `Here is an example of accessing the first letter in the word [subject].` 메시지를 콘솔에 출력하세요.

이전 단계에서 했던 것처럼 `[subject]`을 `subject` 변수로 바꾸고 올바른 템플릿 리터럴 구문을 사용하세요.

# --hints--

`console` 문장은 `Here is an example of accessing the first letter in the word [subject].` 메시지를 출력해야 합니다.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+first\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

--fcc-editable-region--

--fcc-editable-region--
```
