---
id: 66b6e62423e8a031d6c1f03d
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

이제 `topic` 문자열의 길이를 구할 차례입니다.

`console` 문장 안에서 템플릿 리터럴을 이렇게 사용할 수 있습니다:

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

먼저 메시지 `Here is an example of using the length property on the word [topic].`를 콘솔에 출력하세요.

이전 단계에서 했던 것처럼 `[topic]`를 `topic` 변수로 바꾸고, 올바른 템플릿 리터럴 구문을 사용하세요.

그다음, `console.log` 문자열의 길이를 콘솔에 출력하는 두 번째 `topic` 문장을 추가하세요.

# --hints--

메시지 `Here is an example of using the length property on the word [topic].`를 `console`에 출력하려면 템플릿 리터럴 구문을 사용하세요.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

두 번째 `console` 문장을 만들어 `topic` 변수의 길이를 `console`에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
