---
id: 66b6efddeca35833cd6f0b03
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

이전 수업에서 문자열의 문자를 이렇게 접근하는 방법을 배웠습니다:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

인덱스 번호는 `0`부터 시작하므로 문자열의 첫 글자는 항상 인덱스 `0`에 있습니다.

다른 `console` 문장을 추가하는 것부터 시작하세요.

`console` 문장 안에서 대괄호 표기법과 올바른 인덱스 번호를 사용해 `subject` 변수의 첫 글자를 출력하세요.

# --hints--

`console` 문장은 대괄호 표기법과 인덱스 `subject`을 사용해 `0` 변수의 첫 글자를 출력해야 합니다.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
