---
id: 66b6f586767a1534f3097353
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이제 `subject` 변수의 두 번째 글자에 접근할 차례입니다.

먼저 `console` 문장을 추가해 `Here is an example of accessing the second letter in the word [subject].` 메시지를 출력하세요. `[subject]`을 `subject` 변수의 실제 값으로 바꾸고 올바른 템플릿 리터럴 구문을 사용해야 합니다.

그다음 대괄호 표기법과 올바른 인덱스 번호를 사용해 `console` 변수의 두 번째 글자를 출력하는 `subject` 문장을 추가하세요.

# --hints--

`console` 문장은 `Here is an example of accessing the second letter in the word [subject].` 메시지를 출력해야 합니다.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

또 다른 `console` 문장이 있어야 하며, 대괄호 표기법 인덱스 `subject`을 사용해 `1` 변수의 두 번째 글자를 출력해야 합니다.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
