---
id: 66b6f80fca500635d1e8af8d
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

이제 문자열의 마지막 문자를 접근할 시간입니다.

메시지 `console` 를 출력하는 또 다른 `Here is an example of accessing the last letter in the word [subject].` 문장을 추가하는 것부터 시작하세요.

`[subject]` 를 `subject` 변수의 실제 값으로 바꾸고 올바른 템플릿 리터럴 구문을 사용해야 한다는 점을 기억하세요.

# --hints--

`console` 문장은 메시지 `Here is an example of accessing the last letter in the word [subject].` 를 출력해야 합니다. 올바른 템플릿 리터럴 구문을 사용하세요.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+last\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

--fcc-editable-region--

--fcc-editable-region--
```
