---
id: 66b6fdb76441c738719039fa
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

이전 수업에서 문자열의 마지막 문자를 이렇게 접근하는 방법을 배웠습니다:

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1`는 항상 문자열의 마지막 인덱스 번호를 알려줍니다.

`lastCharacter`이라는 새 변수를 만들고 `subject` 변수의 마지막 문자 값을 할당하세요.

그런 다음 `lastCharacter` 변수의 값을 콘솔에 출력하세요.

# --hints--

`lastCharacter`라는 변수가 있어야 합니다.

```js
assert.isNotNull(lastCharacter);
```

`subject` 변수의 마지막 문자 값을 `lastCharacter` 변수에 할당해야 합니다. 도움이 필요하면 예제를 참조하세요.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

`lastCharacter` 변수의 값을 `console`에 출력해야 합니다.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
