---
id: 66b6ffb42b88e33943788abf
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

워크숍의 마지막 부분에서는 문자열에서 부분 문자열의 인덱스 위치를 찾는 방법을 검토합니다. 부분 문자열은 문자열의 일부임을 기억하세요.

먼저 `learningIsFunSentence`라는 변수를 만들고 `"Learning is fun."` 문자열 값을 할당하세요.

# --hints--

`learningIsFunSentence`라는 변수가 있어야 합니다.

```js
assert.isNotNull(learningIsFunSentence);
```

`learningIsFunSentence`는 문자열 값을 가져야 합니다.

```js
assert.isString(learningIsFunSentence);
```

`"Learning is fun."` 값을 `learningIsFunSentence` 변수에 할당해야 합니다.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
