---
id: 66b6e39031393e30f2c48d0f
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

문자열의 길이를 구하려면 <dfn>length</dfn> 속성을 사용할 수 있습니다. 이 속성은 문자열에 포함된 문자 수를 반환합니다.

예시는 다음과 같습니다:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

`length` 문자열의 길이를 구하기 위해 `subject` 속성을 사용하고, 그 값을 콘솔에 출력하는 것부터 시작하세요.

# --hints--


`length` 문자열의 길이를 구하려면 `subject` 속성을 사용하세요.

```js
assert.match(code, /subject\.length/);
```

`subject` 문자열의 길이를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
