---
id: 66b70b8a611cbf3bcc5c6c5f
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

이전 수업에서 `indexOf` 메서드를 사용하는 방법을 배웠습니다:

```js
const sentence = "I love to learn.";

// returns index 2
console.log(sentence.indexOf("love"));

// returns -1
console.log(sentence.indexOf("hate"));
```

`indexOf` 메서드는 문자열에서 부분 문자열이 처음 나타나는 인덱스 위치를 반환한다는 점을 기억하세요. 부분 문자열을 찾지 못하면 `-1`를 반환합니다.

`console` 문장을 새로 추가하여 `indexOf` 변수에서 `learningIsFunSentence` 부분 문자열의 인덱스 위치를 찾기 위해 `"Learning"` 메서드를 사용한 결과를 출력하세요.

# --hints--

`console`은 `indexOf` 변수에서 부분 문자열 `learningIsFunSentence`의 인덱스 위치를 찾기 위해 `"Learning"` 메서드를 사용해야 합니다.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(Learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

--fcc-editable-region--

--fcc-editable-region--
```
