---
id: 66b7137348cfb53fd3ec6c73
title: 17단계
challengeType: 1
dashedName: step-17
---

# --description--

다음으로, `console` 변수에서 `indexOf` 부분 문자열의 위치를 찾기 위해 `learningIsFunSentence` 메서드를 사용한 결과를 출력하는 새로운 `"fun"` 문장을 추가하세요.

그 아래 `console` 문장 아래에, `console` 메서드를 사용해 `indexOf` 부분 문자열의 위치를 찾은 결과를 출력하는 새로운 `"learning"` 문장을 추가하세요.

마지막 `console` 문장이 출력하는 내용을 주목하세요.

# --hints--

`console`은 `"fun"` 변수에서 `learningIsFunSentence` 부분 문자열의 인덱스 위치를 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(fun)\1\s*\)\s*\)/);
```

두 번째 `console` 문장은 `"learning"` 변수에서 `learningIsFunSentence` 부분 문자열의 인덱스 위치를 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

console.log(learningIsFunSentence.indexOf("Learning"));

--fcc-editable-region--

--fcc-editable-region--
```
