---
id: 66e2edb23013ca596f8d9568
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

いよいよ`getPunctuationCount`関数をテストする時間です。

`punctuationCount`変数を作成し、`getPunctuationCount`関数を`"WHAT?!?!?!?!?"`という引数で呼び出した結果を代入してください。

その後、次の内容をコンソールに出力してください：`"Punctuation Count: [Punctuation count goes here]"`。`[Punctuation count goes here]`は実際の変数名に置き換えてください。ここではテンプレート文字列か`+`演算子を使った文字列連結のどちらかを選べます。

# --hints--

`punctuationCount`変数を作成してください。

```js
assert.isNotNull(punctuationCount)
```

`punctuationCount`変数には`getPunctuationCount("WHAT?!?!?!?!?")`の結果を代入してください。

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

次の内容をコンソールに出力してください：`"Punctuation Count: [Punctuation count goes here]"`。`[Punctuation count goes here]`は実際の変数名に置き換えてください。文字列連結の構文を正しく使うことを確認してください。

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
