---
id: 66e2f06b191e305c00574e4d
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

最後に、`getWordCount` 関数を作成して `sentence` パラメータを受け取り、単語数を数えてください。この関数は呼び出されたときに渡された文章の単語の合計数を返す必要があります。

# --hints--

`getWordCount` 関数を作成してください。

```js
assert.isFunction(getWordCount);
```

`getWordCount` 関数は `sentence` パラメータを持つ必要があります。

```js
assert.match(getWordCount.toString(), /sentence/);
```

`getWordCount` 関数は数値を返す必要があります。

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

文章が `"When are you gonna start learning to code?"` の場合、`getWordCount` 関数は `8` を返すべきです。

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

文章が `"What's going on?"` の場合、`getWordCount` 関数は `3` を返すべきです。

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

単語数のカウントは大文字・小文字を区別しないようにしてください。

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

`getWordCount` 関数はどんな文章でも正しい単語数を返す必要があります。

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both front-end and back-end development"), 7);
```

`getWordCount` 関数は空文字列や空白だけの文字列に対しても正しい単語数を返す必要があります。

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
