---
id: 66e2df8c5676fe51cead8ec8
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

いよいよ `getVowelCount` 関数をテストする時間です。

`vowelCount` 変数を作成し、`getVowelCount` 関数を `"Apples are tasty fruits"` を引数として呼び出した結果を代入してください。

その後、次の内容をコンソールに出力してください：`"Vowel Count: [vowel count goes here]"`。`[vowel count goes here]` は実際の変数名に置き換えてください。ここではテンプレート文字列か `+` 演算子を使った文字列連結のどちらかを選べます。

# --hints--

`vowelCount` 変数を作成してください。

```js
assert.isNotNull(vowelCount)
```

`vowelCount` 変数は `getVowelCount("Apples are tasty fruits")` の結果に設定してください。

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

`vowelCount` 変数を `"Vowel Count: ${vowelCount}"` をパラメータにしてコンソールに出力してください。

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
