---
id: 66e2e99caa0cb3570f4822b4
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

いよいよ`getConsonantCount`関数をテストする時間です。

`consonantCount`変数を作成し、`getConsonantCount`関数を`"Coding is fun"`を引数として呼び出した結果を代入してください。

その後、次の内容をコンソールに出力してください：`"Consonant Count: [Consonant count goes here]"`。`[Consonant count goes here]`は実際の変数名に置き換えてください。ここではテンプレート文字列か`+`演算子を使った文字列連結のどちらかを選べます。

# --hints--

`consonantCount`変数を作成してください。

```js
assert.isNotNull(consonantCount)
```

`consonantCount`変数は`getConsonantCount("Coding is fun")`の結果に設定する必要があります。

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

`consonantCount`変数を`"Consonant Count: ${consonantCount}"`をパラメータにしてコンソールに出力してください。

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
