---
id: 66e2e5d890ac8a54c9d703bd
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

子音の数を数える時です。`getConsonantCount` 関数を `sentence` パラメータで作成してください。

関数の中でループを使い、呼び出されたときに関数に渡される `sentence` の中の子音の数を数えてください。子音とは、次の文字以外のすべての文字です：`"aeiou"`。

あなたの `getConsonantCount` 関数は数値を返す必要があります。

# --hints--

`getConsonantCount` 関数を作成してください。

```js
assert.isFunction(getConsonantCount);
```

あなたの `getConsonantCount` 関数は `sentence` パラメータを持つべきです。

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

あなたの `getConsonantCount` 関数は数値を返すべきです。

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

文が `"Coding is fun"` のとき、`getConsonantCount` 関数は `7` を返すべきです。

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

文が `"Hello, World!"` のとき、`getConsonantCount` 関数は `7` を返すべきです。

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

子音の数え方は大文字・小文字を区別しないでください。

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

あなたの `getConsonantCount` 関数はどんな文でも正しい子音の数を返すべきです。

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
