---
id: 66e2eab8a5638f57b637b7cc
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

今度は句読点の数を数えましょう。

`getPunctuationCount` 関数を `sentence` パラメータで作成してください。

関数の中で、呼び出されたときに渡される `sentence` の中の句読点文字の数を数えるループを作成してください。ここでの句読点文字とは、スペース（`" "`）や文字以外のすべての文字を指します。

`getPunctuationCount` 関数は数値を返す必要があります。

# --hints--

`getPunctuationCount` 関数を作成してください。

```js
assert.isFunction(getPunctuationCount);
```

`getPunctuationCount` 関数は `sentence` パラメータを持つべきです。

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

`getPunctuationCount` 関数は数値を返すべきです。

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

文が `"What's going on here?"` のとき、`getPunctuationCount` 関数は `2` を返すべきです。

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

文が `"What????!"` のとき、`getPunctuationCount` 関数は `5` を返すべきです。

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

`getPunctuationCount` 関数はどんな文に対しても正しい句読点の数を返すべきです。

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
