---
id: 66e2d680e129e1423116a541
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、文章を受け取り、単語数、母音、子音、句読点の数を数える文章解析ツールを作成します。

まず、`getVowelCount` 関数を作成し、`sentence` というパラメータを持たせてください。関数は文章内の母音の合計数を返すようにしてください。

# --hints--

`getVowelCount` 関数を作成してください。

```js
assert.isFunction(getVowelCount);
```

`getVowelCount` 関数は `sentence` パラメータを持つ必要があります。

```js
assert.match(getVowelCount.toString(), /sentence/);
```

`getVowelCount` 関数は数値を返す必要があります。

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

文章が `"Apples are tasty fruits"` の場合、`getVowelCount` 関数は `7` を返すべきです。

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

文章が `"Hello, World!"` の場合、関数は `3` を返すべきです。

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

母音のカウントは大文字・小文字を区別しないようにしてください。

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

`getVowelCount` 関数はどんな文章でも正しい母音の数を返す必要があります。

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
