---
id: 67ef8f695db52583424bd118
title: Progetta una pagina di recensioni di film
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `main`.
2. All'interno dell'elemento `main`, dovresti avere un elemento `h1` per il titolo del film.
3. Sotto l'elemento `h1`, dovresti avere un elemento `img` che mostra la copertina del film. Il tuo elemento `img` dovrebbe avere un testo descrittivo `alt` che descrive l'immagine. Sei libero di usare l'immagine seguente se vuoi: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. Dovresti avere un elemento `p` contenente una breve descrizione del film.
5. Dovresti avere un altro elemento `p` per mostrare la valutazione del film. Al suo interno, dovresti avere questi elementi nell'ordine indicato:
   - Un elemento `strong` con il testo `Movie Rating:`.
   - Un elemento `span` con un attributo `aria-hidden` impostato su `true` contenente una rappresentazione visiva della valutazione usando stelle `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`.
   - Un valore numerico, che rappresenta la valutazione, tra parentesi tonde (ad esempio `9.2/10`) dopo lo span.
6. Dovresti avere un elemento `h2` con il testo `Cast Members`.
7. Dovresti avere un elemento `ul`.
8. All'interno dell'elemento `ul`, dovresti avere più elementi `li`, ciascuno contenente un elemento `strong` per il nome dell'attore seguito dal nome del personaggio corrispondente preceduto dal testo `as`. (ad esempio, `James Holloway as Ethan Carter`).

# --hints--

Dovresti avere un elemento `main`.

```js
assert.isNotNull(document.querySelector("main"));
```

All'interno dell'elemento `main`, dovresti avere un elemento `h1` contenente il titolo del film.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

Sotto l'elemento `h1`, dovresti avere un elemento `img` con un attributo `alt` descrittivo.

```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

Dovresti avere un elemento `p` dopo l'immagine che contiene la descrizione del film.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

Dovresti avere un altro elemento `p` che contiene la valutazione.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

All'interno del secondo elemento `p`, dovresti avere un elemento `strong` con il testo `Movie Rating: `.

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

Dovrebbe esserci un elemento `span` all'interno del paragrafo della valutazione.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

L'elemento `span` all'interno del paragrafo della valutazione dovrebbe avere dieci stelle, piene (⭐) o vuote (☆), seguite da un valore numerico tra parentesi tonde dopo `span`.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

L'elemento `span` all'interno del paragrafo della valutazione dovrebbe avere un attributo `aria-hidden` impostato su `true`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

Dovresti avere un elemento `h2`.

```js
assert.exists(document.querySelector("h2"));
```

Il tuo `h2` dovrebbe venire dopo il tuo secondo elemento `p`.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

Il tuo `h2` dovrebbe avere il testo `Cast Members`.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

Dovresti avere un elemento `ul` dopo `h2`.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

Il `ul` dovrebbe contenere più elementi `li`.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

Ciascun `li` dovrebbe contenere un elemento `strong` per il nome dell'attore seguito dal nome del personaggio corrispondente preceduto dal testo `as`. (ad esempio, `James Holloway as Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
