---
id: 673b3d6b7ef7318eef926d5a
title: 建置情緒板
challengeType: 25
dashedName: build-a-mood-board
demoType: onClick
---

# --description--

情緒板是一個由影像和文字組成的拼貼，傳達關於特定主題的一般概念、目標或感覺。

在本實驗中，你將使用可重複使用的 React 組件來創建情緒板。CSS 已經為你提供。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建並匯出一個 `MoodBoardItem` 組件，該組件接受三個 props：`color`、`image` 和 `description`。
2. 你的 `MoodBoardItem` 組件應該傳回一個帶有類別 `mood-board-item` 的 `div` 作為其頂層元素。
3. 你應該使用行內樣式將 `.mood-board-item` 元素的背景顏色設為 `color` 屬性的值。
4. 你應該在 `.mood-board-item` 元素內，算繪一個 `img` 元素，該元素具有 `mood-board-image` 類別，且其 `src` 屬性設為 `image` prop 的值。
5. 你應該在 `.mood-board-item` 項目內，算繪一個 `h3` 元素，該元素具有 `mood-board-text` 類別，且其文字為 `description` prop 的值。
6. 你應該創建並匯出一個 `MoodBoard`。
7. 你的 `MoodBoard` 組件應該傳回一個 `div` 作為其頂層元素。
8. 你的 `MoodBoard` 組件應該呈現一個帶有 `mood-board-heading` 類別和文字 `Destination Mood Board` 的 `h1` 元素。
9. 你的 `MoodBoard` 組件應該呈現一個具有 `mood-board` 類別的 `div`。
10. 你的 `MoodBoard` 組件應該在 `.mood-board` 元素內至少呈現三個 `MoodBoardItem` 組件，每個都應該傳遞具有有效值的 `color`、`image` 和 `description` props。

如果你願意，可以在你的 Mood Board 中使用以下影像。

- `https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/shore.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/grass.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/ship.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg`

# --hints--

你應該匯出一個 `MoodBoardItem` 組件。

```js
assert.isFunction(window.index.MoodBoardItem, 1);
```

你的 `MoodBoardItem` 組件應該在其頂層元素傳回一個具有 `mood-board-item` 類別的 `div`。

```js
const item = document.getElementsByClassName('mood-board-item')[0];
assert.exists(item);
assert.equal(item.tagName, 'DIV');
```

`.mood-board-item` 元素的背景顏色應該使用行內樣式設定為 `color` 屬性的值。

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { color: "red" });
  const moodBoardItem = container.querySelector(".mood-board-item");
  assert.equal(moodBoardItem.style.backgroundColor, "red");
```

你的 `MoodBoardItem` 組件應該呈現一個帶有 `mood-board-image` 類別的 `img` 元素，並將其 `src` 設定為 `image` prop 的值。

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { image: "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg" });
  const img = container.querySelector(".mood-board-image");
  assert.exists(img);
  assert.equal(img.tagName, "IMG");
  assert.equal(img.src, "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg");
```

你的 `MoodBoardItem` 組件應該呈現一個帶有 `mood-board-text` 類別的 `h3` 元素，並將其文字設為 `description` prop 的值。

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { description: "Caribbean" });
  const text = container.querySelector(".mood-board-text");
  assert.exists(text);
  assert.equal(text.tagName, "H3");
  assert.equal(text.textContent, "Caribbean");
```

你應該匯出一個 `MoodBoard` 組件。

```js
assert.isFunction(window.index.MoodBoard, 1);
```

你的 `MoodBoard` 組件應該傳回一個 `div` 作為其頂層元素。

```js
const root = document.getElementById('root');
assert.equal(root.children[0].tagName, 'DIV');
```

你的 `MoodBoard` 組件應該呈現一個帶有 `mood-board-heading` 類別和文字 `Destination Mood Board` 的 `h1` 元素。

```js
const heading = document.getElementsByClassName('mood-board-heading')[0];
assert.exists(heading);
assert.equal(heading.tagName, 'H1');
assert.equal(heading.textContent, 'Destination Mood Board');
```

你的 `MoodBoard` 組件應該至少呈現三個 `MoodBoardItem` 組件，每個都應該傳遞具有有效值的 `color`、`image` 和 `description` 屬性。

```js
const moodboard = document.getElementsByClassName('mood-board')[0];
const items = moodboard.querySelectorAll(".mood-board-item");
assert.isAtLeast(items.length, 3);     

items.forEach((item) => {
  const color = item.style.backgroundColor;
  const img = item.querySelector("img");
  const text = item.querySelector("h3");
  
  assert.isAtLeast(color.length, 1);
  assert.isAtLeast(img.src.length, 1);
  assert.isAtLeast(text.textContent.length, 1);
});
```

你的 `MoodBoard` 組件應該被呈現在頁面的 `#root` 元素。

```js
const root = document.getElementById('root');
const moodBoard = root.getElementsByClassName('mood-board')[0];
assert.exists(moodBoard);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx

```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx
export function MoodBoardItem(props) {
  return (
    <div className='mood-board-item' style={{ backgroundColor: props.color }}>
      <img src={props.image} alt='Mood' className='mood-board-image' />
      <h3 className='mood-board-text'>{props.description}</h3>
    </div>
  );
}

export function MoodBoard() {
  return (
    <div>
      <h1 className='mood-board-heading'>Destination Mood Board</h1>
      <div className='mood-board'>
        <MoodBoardItem
          color='#2da64f'
          image='https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg'
          description='Caribbean'
        />
        <MoodBoardItem
          color='#8e44ad'
          image='https://cdn.freecodecamp.org/curriculum/labs/shore.jpg'
          description='Gawadar Beach'
        />
        <MoodBoardItem
          color='#3498db'
          image='https://cdn.freecodecamp.org/curriculum/labs/grass.jpg'
          description='Cape Town'
        />
        <MoodBoardItem
          color='#bf3d7e'
          image='https://cdn.freecodecamp.org/curriculum/labs/ship.jpg'
          description='Suez Canal'
        />
        <MoodBoardItem
          color='#e74c3c'
          image='https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg'
          description='Santorini'
        />
        <MoodBoardItem
          color='#95a5a6'
          image='https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg'
          description='Istanbul'
        />
      </div>
    </div>
  );
}
```
