---
id: 69899123889dcf5b19be1847
title: 實作深度優先搜尋演算法
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

在本實驗中，你將實作一種稱為 <dfn>深度優先搜尋</dfn>（DFS）的圖形遍歷演算法。

當 <dfn>breadth-first search</dfn>（BFS）逐層探索鄰居時，<dfn>depth-first search</dfn> 則會盡可能深入單一路徑的邊，然後再回溯。

當演算法到達路徑的終點（沒有未造訪鄰居的節點）時，它會回溯到最近仍有未探索邊的節點，並從那裡繼續進行搜尋。

要反覆地實作這個演算法，你會想使用一個**堆疊**。在 JavaScript 中，標準的 `Array` 使用 `.push()` 和 `.pop()` 方法作為堆疊非常合適，遵循**後進先出（LIFO）**原則。這確保最近被添加到堆疊的鄰居是下一個被探索的對象。

這個演算法的一個簡單輸出是從給定節點可達的節點列表。因此，你也會想要追蹤你造訪過的節點。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該定義一個名為 `dfs` 的函式（程式）。
1. `dfs` 函式應該接受兩個引數：

    - `graph`：一個鄰接矩陣（一個陣列的陣列）。
    - `root`：一個數字節點標籤（起始索引），其數值介於 `0` 和 `n - 1` 之間，其中 `n` 是圖中節點的總數。

1. 該函式應傳回一個包含從起始節點可達所有節點標籤的陣列。

# --hints--

你應該有一個名為 `dfs` 的函式（程式），它接受兩個引數。

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` 應該傳回一個包含 `1`、`2`、`3` 和 `0` 的列表。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` 應該傳回一個包含 `1`、`2`、`3` 和 `0` 的列表。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` 應該傳回 `[3]`。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` 應該傳回一個包含 `3` 和 `2` 的列表。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` 應該傳回一個包含 `0` 和 `1` 的列表。

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
