---
id: 587d7b8d367417b2b2512b5b
title: 学习函数式编程
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

函数式编程是一种编程风格，其中方案是简单、独立的函数，且函数作用域外没有任何副作用：`INPUT -> PROCESS -> OUTPUT`

函数式编程：

独立函数——不依赖于程序状态，包括可能变化的全局变量

纯函数——相同的 `input` 总是产生相同的输出

具有有限副作用的函数——对函数外部程序状态的任何更改或变异都受到严格控制

# --instructions--

freeCodeCamp 的成员们爱喝茶。

在代码编辑器中，已经为你定义好了`prepareTea`和`getTea`函数。 调用 `getTea` 函数为团队准备 40 杯茶，并将它们存储在 `tea4TeamFCC` 变量里。

# --hints--

`tea4TeamFCC` 变量里应有 40 杯为团队准备的茶。

```js
assert(tea4TeamFCC.length === 40);
```

`tea4TeamFCC` 变量里应有几杯 greenTea（绿茶）。

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40); 
```
