---
id: 67ed03ac474c48692f41749e
title: Побудова хеш-таблиці
challengeType: 27
dashedName: build-a-hash-table
saveSubmissionToDB: true
---

# --description--

У цій лабораторній роботі ви побудуєте хеш-таблицю з нуля. Хеш-таблиця — це структура даних, яка зберігає пари ключ-значення. Хеш-таблиця працює, приймаючи ключ як вхідні дані, а потім хешуючи цей ключ за допомогою певної хеш-функції.

Для цієї лабораторної роботи хеш-функція буде простою: вона сумуватиме Unicode-значення кожного символу в ключі. Отримане хеш-значення потім використовуватиметься як фактичний ключ для зберігання пов’язаного значення. Те саме хеш-значення також використовуватиметься для отримання та видалення значення, пов’язаного з ключем.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні визначити клас з ім’ям `HashTable`, який матиме атрибут `collection`, ініціалізований порожнім словником під час створення нового екземпляра `HashTable`. Словник `collection` повинен зберігати пари ключ-значення на основі хешованого значення ключа.

2. Клас `HashTable` повинен мати чотири методи екземпляра: `hash`, `add`, `remove` і `lookup`.

3. Метод `hash` повинен:

    - Приймати рядок як параметр.
    - Повернути хешоване значення, обчислене як суму Unicode (ASCII) значень кожного символу в рядку. Для цього обчислення можна використовувати функцію `ord`.

4. Метод `add` повинен:
    - Приймати два аргументи, що представляють пару ключ-значення, і обчислювати хеш ключа.
    - Використовувати обчислене хеш-значення як ключ для зберігання словника, що містить пару ключ-значення, всередині словника `collection`.
    - Якщо кілька ключів дають однакове хеш-значення, їхні пари ключ-значення повинні зберігатися у вже існуючому вкладеному словнику під тим самим хеш-значенням.

5. Метод `remove` повинен:

    - Приймати ключ як аргумент і обчислювати його хеш.
    - Перевіряти, чи існує ключ у колекції.
    - Видаляти відповідну пару ключ-значення з хеш-таблиці.
    - Якщо ключ не існує в колекції, не повинен викликати помилку або видаляти щось.

6. Метод `lookup` повинен:

    - Приймати ключ як аргумент.
    - Обчислювати хеш ключа і повертати відповідне значення, збережене в хеш-таблиці.
    - Якщо ключ не існує в колекції, повинен повертати `None`.

# --hints--

Ви повинні визначити клас `HashTable`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_class("HashTable")
    `)) 
})
```

Під час створення нового екземпляра класу `HashTable` його атрибут `collection` повинен бути ініціалізований порожнім словником.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.collection == {}
    `) 
})
```

Клас `HashTable` повинен мати метод `hash`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('hash')`))}})
```

Метод `hash` повинен приймати рядок як параметр.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.hash)
    len(sig.parameters) == 1
  `))
})
```

Метод `hash` повинен приймати рядок як аргумент і повертати суму Unicode-значень кожного символу в рядку.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("fcc")
        assert hash_result == 300

        # prevent hardcoding
        assert ht.hash("golf") == 424
        assert ht.hash("read") == 412
`) })
```

Клас `HashTable` повинен мати метод `add`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('add')`))}})
```

Метод `add` повинен приймати ключ і значення як параметри.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.add)
    len(sig.parameters) == 2
  `))
})
```

Клас `HashTable` повинен мати метод `remove`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('remove')`))}})
```

Метод `remove` повинен приймати ключ як параметр.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.remove)
    len(sig.parameters) == 1
  `))
})
```

Коли ви намагаєтеся видалити ключ, якого немає в колекції, не повинно виникати помилки або видалення.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("rose", "flower")
        index = ht.hash("rose")
        original = ht.collection.copy()

        ht.remove("tulip")
        ht.remove("sore")

        assert ht.collection == original
        assert "rose" in ht.collection[index]
        assert "tulip" not in ht.collection.get(index, {})
    `) 
})
```

Якщо кілька ключів хешуються в один індекс, метод `remove` повинен видаляти лише конкретну пару ключ-значення, а не весь словник за цим індексом.

```js
({
  test: () => runPython(`
    ht = HashTable()
    ht.add("rose", "flower")
    ht.add("sore", "pain")  # "rose" and "sore" both hash to the same index

    index = ht.hash("rose")
    ht.remove("rose")

    assert index in ht.collection
    assert "rose" not in ht.collection[index]
    assert "sore" in ht.collection[index]
    assert ht.collection[index]["sore"] == "pain"
  `)
})
```

Клас `HashTable` повинен мати метод `lookup`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('lookup')`))}})
```

Метод `lookup` повинен приймати ключ як параметр.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.lookup)
    len(sig.parameters) == 1
  `))
})
```

`HashTable().hash('golf')` повинен повертати `424`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("golf")
        assert hash_result == 424

        # prevent hardcoding
        assert ht.hash("dear") == 412
        assert ht.hash("cat") == 312
    `) 
})
```

`HashTable().add('golf', 'sport')` повинен додавати пару ключ-значення до колекції за ключем `424`.

```js

({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")
    expected_value = {424: {'golf': 'sport'}}

    assert ht.collection == expected_value
    `) 
})
```

`HashTable().add('dear', 'friend')` і `HashTable().add('read', 'book')` повинні додавати обидві пари ключ-значення до колекції за індексом `412` як вкладений словник.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("dear", "friend")
    ht.add("read", "book")

    expected_value = {
        "dear": "friend",
        "read": "book"
    }

    assert ht.collection.get(412) == expected_value
    `) 
})
```

Коли ключ існує в хеш-таблиці, метод `remove()` повинен видаляти заданий ключ і відповідне значення з колекції.

```js
({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")

    expected_value_before_removal = {
        "golf": "sport"
    }

    index = ht.hash("golf")

    assert ht.collection.get(index) == expected_value_before_removal

    ht.remove("golf")

    assert "golf" not in ht.collection.get(index, {})
    `) 
})
```

Коли пара ключ-значення `'golf', 'sport'` існує в хеш-таблиці, `HashTable().lookup('golf')` повинен повертати `sport`.

```js
({ 
    test: () => runPython(`
    ht = HashTable()

    ht.add("golf", "sport")

    expected_value = "sport"
    assert ht.lookup("golf") == expected_value
    `) 
})

```

Коли пара ключ-значення `'golf', 'sport'` не існує в колекції, `HashTable().lookup('golf')` повинен повертати `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.lookup("golf") is None
`) })

```

Коли ключ `'fcc'` існує в колекції, `HashTable().lookup('cfc')` повинен повертати `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("fcc", "coding")
        assert ht.lookup("cfc") is None
`) })
```

Коли ви додаєте `('rose', 'flower')` до хеш-таблиці, його атрибут `collection` повинен виглядати так: `{ 441: { 'rose': 'flower' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("rose", "flower")

    expected_value = {
        "rose": "flower"
    }

    assert ht.collection.get(441) == expected_value

    # prevent hardcoding

    ht2 = HashTable()

    ht2.add("kebab", "food")
    expected_value = {
        "kebab": "food"
    }

    assert ht2.collection.get(501) == expected_value
    `)
})
```

Коли ви додаєте ключ, який хешується в те саме значення, що й існуючий ключ, наприклад `fcc` і `cfc`, `collection` повинен виглядати так: `{ 300: { 'fcc': 'coding', 'cfc':  'chemical' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("fcc", "coding")
    ht.add("cfc", "chemical")

    expected_value = {
        "fcc": "coding",
        "cfc": "chemical"
    }

    assert ht.collection.get(300) == expected_value

    # prevent hardcoding

    ht2 = HashTable()
    ht2.add("cat", "animal")
    ht2.add("act", "verb")

    expected_value = {
        "cat": "animal",
        "act": "verb"
    }

    assert ht2.collection.get(312) == expected_value
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class HashTable:
    def __init__(self):
        self.collection = {}

    def hash(self, string):
        hashed = 0
        for char in string:
            hashed += ord(char)  
        return hashed

    def add(self, key, val):
        the_hash = self.hash(key)
        if the_hash not in self.collection:
            self.collection[the_hash] = {}
        self.collection[the_hash][key] = val

    def remove(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            del self.collection[the_hash][key]
            if not self.collection[the_hash]: 
                del self.collection[the_hash]

    def lookup(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            return self.collection[the_hash][key]
        return None
```
