---
id: ac6993d51946422351508a41
title: "implementare l'algoritmo per troncare una stringa"
challengeType: 26
dashedName: truncate-a-string
---

# --description--

In questo laboratorio, eserciterai il troncamento di una stringa a una certa lunghezza.

**Obiettivo:** Soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione `truncateString` che accetta due argomenti, il primo una stringa, il secondo un numero.
1. Se la lunghezza della stringa è maggiore del numero dato, la stringa dovrebbe essere troncata per ridurre la lunghezza in modo che sia uguale al numero dato, e `...` dovrebbe essere aggiunto in coda alla stringa troncata.
1. Se la lunghezza della stringa è uguale o inferiore al numero dato, la stringa dovrebbe essere restituita invariata.

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` dovrebbe restituire la stringa `A-tisket...`.

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` dovrebbe restituire la stringa `Peter Piper...`.

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` dovrebbe restituire la stringa `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` dovrebbe restituire la stringa `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` dovrebbe restituire la stringa `A...`.

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` dovrebbe restituire la stringa `Ab...`.

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
