---
id: 6814d8e1516e86b171929de4
title: "التحدي 1: توازن الحروف المتحركة"
challengeType: 28
dashedName: challenge-1
---

# --description--

بالنظر إلى سلسلة نصية، حدّد ما إذا كان عدد الحروف المتحركة في النصف الأول من السلسلة يساوي عدد الحروف المتحركة في النصف الثاني.

- يمكن أن تحتوي السلسلة النصية على أي حروف.
- تُعتبر الحروف `a`، `e`، `i`، `o`، و`u`، سواء كانت بأحرف كبيرة أو أحرف صغيرة، حروفًا متحركة.
- إذا كان عدد الحروف في السلسلة فرديًا، تجاهل الحرف الأوسط.

# --hints--

يجب أن تُرجع الدالة `isBalanced("racecar")` القيمة `true`.

```js
assert.isTrue(isBalanced("racecar"));
```

يجب أن تُرجع الدالة `isBalanced("Lorem Ipsum")` القيمة `true`.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

يجب أن تُرجع الدالة `isBalanced("Kitty Ipsum")` القيمة `false`.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

يجب أن تُرجع الدالة `isBalanced("string")` القيمة `false`.

```js
assert.isFalse(isBalanced("string"));
```

يجب أن تُرجع الدالة `isBalanced(" ")` القيمة `true`.

```js
assert.isTrue(isBalanced(" "));
```

يجب أن تُرجع الدالة `isBalanced("abcdefghijklmnopqrstuvwxyz")` القيمة `false`.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

يجب أن تُرجع الدالة `isBalanced("123A#b!E&*456-o.U")` القيمة `true`.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const half = Math.floor(s.length / 2);

  let firstHalf = s.slice(0, half);
  let secondHalf = s.length % 2 === 0 ? s.slice(half) : s.slice(half + 1);

  const countVowels = str =>
    str
      .toLowerCase()
      .split('')
      .filter(c => vowels.includes(c))
      .length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
