---
id: 6814d8e1516e86b171929de4
title: "تحدٍّ 1: توازن الحروف المتحركة"
challengeType: 28
dashedName: challenge-1
---

# --description--

اعطِ سلسلة نصية، حدّد ما إذا كان عدد الحروف المتحركة في النصف الأول من السلسلة يساوي عدد الحروف المتحركة في النصف الثاني.

- يمكن أن تحتوي السلسلة النصية على أي حروف.
- تُعتبر الحروف `a`، `e`، `i`، `o`، و `u`، سواء كانت بأحرف كبيرة أو أحرف صغيرة، حروفًا متحركة.
- إذا كان عدد الحروف في السلسلة فرديًا، تجاهل الحرف الأوسط.

# --hints--

يجب أن تُرجع الدالة `isBalanced("racecar")` القيمة `true`.

```js
assert.isTrue(isBalanced("racecar"));
```

يجب أن تُرجع الدالة `isBalanced("Lorem Ipsum")` القيمة `true`.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

يجب أن تُرجع الدالة `isBalanced("Kitty Ipsum")` القيمة `false`.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

يجب أن تُرجع الدالة `isBalanced("string")` القيمة `false`.

```js
assert.isFalse(isBalanced("string"));
```

يجب أن تُرجع الدالة `isBalanced(" ")` القيمة `true`.

```js
assert.isTrue(isBalanced(" "));
```

يجب أن تُرجع الدالة `isBalanced("abcdefghijklmnopqrstuvwxyz")` القيمة `false`.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

يجب أن تُرجع الدالة `isBalanced("123A#b!E&*456-o.U")` القيمة `true`.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

يجب أن ترجع الدالة `isBalanced("a")` القيمة `true`.

```js
assert.isTrue(isBalanced("a"));
```

يجب أن ترجع الدالة `isBalanced("E")` القيمة `true`.

```js
assert.isTrue(isBalanced("E"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const chars = [...s];
  const half = Math.floor(chars.length / 2);

  const firstHalf = chars.slice(0, half);
  const secondHalf =
    chars.length % 2 === 0 ? chars.slice(half) : chars.slice(half + 1);

  const countVowels = arr =>
    arr.filter(c => vowels.includes(c.toLowerCase())).length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
