---
id: 681cb1b0dab50c87ddb2e518
title: "تحدٍّ 7: المجموع المستهدف"
challengeType: 28
dashedName: challenge-7
---

# --description--

معطاة مصفوفة من الأرقام وعدد صحيح مستهدف، ابحث عن رقمين فريدين في المصفوفة مجموعهما يساوي القيمة المستهدفة. أرجع مصفوفة تحتوي على مؤشرات هذين الرقمين، أو `"Target not found"` إذا لم يوجد رقمان مجموعهما يساوي الهدف.

- يجب أن تحتوي المصفوفة المرجعة على المؤشرات بترتيب تصاعدي.

# --hints--

`findTarget([2, 7, 11, 15], 9)` يجب أن تُرجع `[0, 1]`.

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)` يجب أن تُرجع `[1, 2]`.

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)` يجب أن تُرجع `[4, 5]`.

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)` يجب أن تُرجع `"Target not found"`.

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
