---
id: 681cb1afdab50c87ddb2e517
title: "تحدٍّ 6: مدقق الكلمات المتشابهة"
challengeType: 28
dashedName: challenge-6
---

# --description--

معطيتان سلسلتا نصية، حدّد إذا كانتا كلمات متشابهة (تحتويان على نفس الأحرف بأي ترتيب).

- تجاهل حالة الأحرف والمسافات البيضاء.

# --hints--

`areAnagrams("listen", "silent")` يجب أن تُرجع `true`.

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` يجب أن تُرجع `true`.

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` يجب أن تُرجع `true`.

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` يجب أن تُرجع `false`.

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` يجب أن تُرجع `false`.

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` يجب أن تُرجع `false`.

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
