---
id: 681cb1afdab50c87ddb2e515
title: "تحدٍّ 4: S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--

أعطِ سلسلة نصية، احذف كل الفراغات منها، أدخل فراغين بين كل حرف وآخر، حوّل كل الحروف الأبجدية إلى حروف كبيرة، وأرجع النتيجة.

- يجب أن تبقى الأحرف غير الأبجدية كما هي (باستثناء الفراغات).

# --hints--

`spaceJam("freeCodeCamp")` يجب أن تُرجع `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` يجب أن تُرجع `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")` يجب أن تُرجع `"H  E  L  L  O  W  O  R  L  D  ?  !"`.

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` يجب أن تُرجع `"C  @  T  $  &  D  0  G  $"`.

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` يجب أن تُرجع `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
