---
id: 681cb1b0dab50c87ddb2e51a
title: "تحدٍّ 9: مجموع المربعات"
challengeType: 28
dashedName: challenge-9
---

# --description--

أعطِ عددًا صحيحًا موجبًا يصل إلى 1000، أرجع مجموع المربعات لجميع الأعداد الصحيحة من 1 حتى ذلك العدد.

# --hints--

`sumOfSquares(5)` يجب أن تُرجع `55`.

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` يجب أن تُرجع `385`.

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` يجب أن تُرجع `5525`.

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` يجب أن تُرجع `41791750`.

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` يجب أن تُرجع `333833500`.

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
