---
id: 681cb1a2dab50c87ddb2e514
title: "التحدي 3: متتالية فيبوناتشي"
challengeType: 28
dashedName: challenge-3
---

# --description--

متتالية فيبوناتشي هي سلسلة من الأعداد حيث يكون كل عدد هو مجموع العددين السابقين له. عند البدء بـ `0` و `1`، تكون أول 10 أعداد في المتتالية هي `0`، `1`، `1`، `2`، `3`، `5`، `8`، `13`، `21`، `34`.

معطى مصفوفة تحتوي على أول رقمين من متتالية فيبوناتشي، وعدد صحيح يمثل طول المتتالية، أرجع مصفوفة تحتوي على المتتالية بالطول المعطى.

- يجب أن تتعامل دالتك مع متتاليات بأي طول أكبر من أو يساوي صفر.
- إذا كان الطول صفرًا، أرجع مصفوفة فارغة.
- لاحظ أن الأعداد الابتدائية جزء من المتتالية.

# --hints--

`fibonacciSequence([0, 1], 20)` يجب أن تُرجع `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)` يجب أن تُرجع `[21]`.

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)` يجب أن تُرجع `[]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)` يجب أن تُرجع `[10, 20]`.

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)` يجب أن تُرجع `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
