---
id: 681cb1b0dab50c87ddb2e519
title: "تحدٍّ 8: مضروب"
challengeType: 28
dashedName: challenge-8
---

# --description--

أعطِ عددًا صحيحًا من صفر إلى 20، أرجع مضروب ذلك العدد. المضروب هو حاصل ضرب كل الأعداد بين 1 والعدد المعطى.

- مضروب الصفر هو 1.

# --hints--

`factorial(0)` يجب أن تُرجع `1`.

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` يجب أن تُرجع `120`.

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` يجب أن تُرجع `2432902008176640000`.

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
