---
id: 681cb05adab50c87ddb2e513
title: "تحدٍّ 2: التحقق من الأساس"
challengeType: 28
dashedName: challenge-2
---

# --description--

معطاة سلسلة نصية تمثل رقمًا، وقاعدة عدد صحيحة من 2 إلى 36، حدّد ما إذا كان الرقم صالحًا في تلك القاعدة.

- قد تحتوي السلسلة النصية على أعداد صحيحة، وأحرف كبيرة أو أحرف صغيرة.
- يجب أن يكون التحقق غير حساس لحالة الأحرف.
- يمكن أن تكون القاعدة أي رقم من 2 إلى 36.
- الرقم صالح إذا كان كل حرف فيه رقمًا صالحًا في القاعدة المعطاة.
- أمثلة على الأرقام الصالحة للقيَم:
  - القاعدة 2: 0-1
  - القاعدة 8: 0-7
  - القاعدة 10: 0-9
  - القاعدة 16: 0-9 و A-F
  - القاعدة 36: 0-9 و A-Z

# --hints--

`isValidNumber("10101", 2)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` يجب أن تُعيد `false`.

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` يجب أن تُعيد `false`.

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` يجب أن تُعيد `false`.

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` يجب أن تُعيد `false`.

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` يجب أن تُعيد `false`.

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` يجب أن تُعيد `true`.

```js
assert.isTrue(isValidNumber("z", 36))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
