---
id: 681cb1afdab50c87ddb2e516
title: "تحدٍّ 5: نص Jbelmud"
challengeType: 28
dashedName: challenge-5
---

# --description--

بالنظر إلى سلسلة نصية، أعد نسخة مشوشة من تلك السلسلة حيث تُحوّل كل كلمة باستخدام القيود التالية:

- تبقى الحروف الأولى والأخيرة من الكلمات في مكانها
- تُرتب جميع الحروف بين الحرف الأول والأخير أبجديًا
- لن تحتوي السلاسل النصية المُدخلة على علامات ترقيم، وستكون كلها أحرف صغيرة.

# --hints--

`jbelmu("hello world")` يجب أن تُعيد `"hello wlord"`.

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` يجب أن تُعيد `"i love jbelmud text"`.

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` يجب أن تُعيد `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` يجب أن تُعيد `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
