---
id: 681cb1b0dab50c87ddb2e51b
title: "تحدٍّ 10: 3 ضربات"
challengeType: 28
dashedName: challenge-10
---

# --description--

معطى عدد صحيح بين 1 و 10,000، أرجع عدد الأعداد من 1 حتى ذلك العدد التي يحتوي مربعها على الرقم 3 على الأقل مرة واحدة.

# --hints--

`squaresWithThree(1)` يجب أن تُرجع `0`.

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` يجب أن تُرجع `1`.

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` يجب أن تُرجع `19`.

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` يجب أن تُرجع `326`.

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` يجب أن تُرجع `4531`.

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
