---
id: 6821ebd4237de8297eaee791
title: "تحدٍّ 15: تنسيق سنام الجمل"
challengeType: 28
dashedName: challenge-15
---

# --description--

أعطِ سلسلة نصية، أعد إصدارها بتنسيق سنام الجمل باستخدام القواعد التالية:

- الكلمات في سلسلة المعلمة مفصولة بحرف واحد أو أكثر من المجموعة التالية: مسافة (` `)، شرطة (`-`)، أو شرطة سفلية (`_`). اعتبر أي متتالية من هذه كفاصل للكلمات.
- يجب أن تكون الكلمة الأولى كلها بأحرف صغيرة.
- يجب أن تبدأ كل كلمة لاحقة بحرف كبير، مع بقية الحروف بأحرف صغيرة.
- يجب إزالة كل المسافات والفواصل.

# --hints--

`toCamelCase("hello world")` يجب أن تُعيد `"helloWorld"`.

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` يجب أن تُعيد `"helloWorld"`.

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")` يجب أن تُعيد `"secretAgentX"`.

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` يجب أن تُعيد `"freeCodeCamp"`.

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` يجب أن تُعيد `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
