---
id: 6821ebce237de8297eaee790
title: "تحدٍّ 14: معركة الشخصيات"
challengeType: 28
dashedName: challenge-14
---

# --description--

معطاة سلسلتا نص تمثلان جيشك وجيش الخصم، يقاتل كل حرف من جيشك الحرف في نفس الموضع من جيش الخصم وفق القواعد التالية:

- الأحرف `a-z` لها قوة من `1-26`، على التوالي.
- الأحرف `A-Z` لها قوة من `27-52`، على التوالي.
- الأرقام `0-9` لها قوة تساوي قيمتها الظاهرة.
- جميع الأحرف الأخرى قيمتها صفر.
- يمكن لكل حرف أن يخوض معركة واحدة فقط.

في كل معركة، ينتصر الحرف الأقوى. الجيش الذي يحقق انتصارات أكثر يفوز بالحرب. أرجع القيم التالية:

- `"Opponent retreated"` إذا كان جيشك يحتوي على أحرف أكثر من جيش الخصم.
- `"We retreated"` إذا كان جيش الخصم يحتوي على أحرف أكثر من جيشك.
- `"We won"` إذا فاز جيشك بمعارك أكثر.
- `"We lost"` إذا فاز جيش الخصم بمعارك أكثر.
- `"It was a tie"` إذا فاز الجيشان بنفس عدد المعارك.

# --hints--

`battle("Hello", "World")` يجب أن تُرجع `"We lost"`.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` يجب أن تُرجع `"We won"`.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` يجب أن تُرجع `"We won"`.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` يجب أن تُرجع `"Opponent retreated"`.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` يجب أن تُرجع `"We retreated"`.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` يجب أن تُرجع `"It was a tie"`.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` يجب أن تُرجع `"It was a tie"`.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
