---
id: 6821ebf8237de8297eaee798
title: "تحدٍّ 22: متتالية Tribonacci"
challengeType: 28
dashedName: challenge-22
---

# --description--

متتالية Tribonacci هي سلسلة من الأعداد حيث يكون كل عدد هو مجموع الأعداد الثلاثة السابقة له. عند البدء بـ `0`، `0` و `1`، تكون أول 10 أعداد في المتتالية هي `0`، `0`، `1`، `1`، `2`، `4`، `7`، `13`، `24`، `44`.

معطى مصفوفة تحتوي على أول ثلاثة أعداد من متتالية Tribonacci، وعدد صحيح يمثل طول المتتالية، أرجع مصفوفة تحتوي على المتتالية بالطول المعطى.

- يجب أن تتعامل دالتك مع متتاليات بأي طول أكبر من أو يساوي صفر.
- إذا كان الطول صفرًا، أرجع مصفوفة فارغة.
- لاحظ أن الأعداد الابتدائية هي جزء من المتتالية.

# --hints--

`tribonacciSequence([0, 0, 1], 20)` يجب أن تُرجع `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)` يجب أن تُرجع `[21]`.


```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)` يجب أن تُرجع `[]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)` يجب أن تُرجع `[10, 20]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)` يجب أن تُرجع `[10, 20, 30]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)` يجب أن تُرجع `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
