---
id: 6821ebf3237de8297eaee797
title: "تحدٍّ 21: مولد الألوان الست عشرية"
challengeType: 28
dashedName: challenge-21
---

# --description--

بالنظر إلى سلسلة نصية لاسم لون في CSS، أنشئ شفرة لون سداسية عشرية (hex) عشوائية تكون اللون السائد فيها هو اللون المعطى.

- يجب أن تتعامل الدالة مع `"red"`، `"green"`، أو `"blue"` كمعلمة مُدخَل.
- إذا لم يكن المُدخَل واحدًا من هذه، يجب أن تُرجع الدالة `"Invalid color"`.
- يجب أن تُرجع الدالة شفرة لون سداسية عشرية عشوائية مكونة من ستة أحرف حيث تكون قيمة اللون المُدخَل أكبر من أي من الألوان الأخرى.
- مثال على مخرجات صحيحة لمُدخَل معين:

| المُدخَل   | المخرجات   |
|---------|----------|
| `"red"`   | `"FF0000"` |
| `"green"` | `"00FF00"` |
| `"blue"`  | `"0000FF"` |

# --hints--

يجب أن تُرجع الدالة `generateHex("yellow")` القيمة `"Invalid color"`.

```js
assert.equal(generateHex("yellow"), "Invalid color");
```

يجب أن تُرجع الدالة `generateHex("red")` سلسلة نصية مكونة من ستة أحرف.

```js
assert.lengthOf(generateHex("red"), 6);
```

يجب أن تُرجع الدالة `generateHex("red")` شفرة لون سداسية عشرية صحيحة مكونة من ستة أحرف.

```js
const hex = generateHex("red").toUpperCase();
const isValidHex = /^[0-9A-F]{6}$/.test(hex);
assert.isTrue(isValidHex);
```

يجب أن تُرجع الدالة `generateHex("red")` شفرة لون سداسية عشرية صحيحة تكون قيمة اللون الأحمر فيها أعلى من بقية الألوان.

```js
const hex = generateHex("red").toUpperCase();
const isValidHex = /^[0-9A-F]{6}$/.test(hex);
assert.isTrue(isValidHex);

const r = parseInt(hex.slice(0, 2), 16);
const g = parseInt(hex.slice(2, 4), 16);
const b = parseInt(hex.slice(4, 6), 16);

assert.isAbove(r, g);
assert.isAbove(r, b);
```

يجب أن تُرجع استدعاء الدالة `generateHex("red")` مرتين قيمتين مختلفتين للون السداسي حيث يكون اللون الأحمر هو السائد.

```js
const hex1 = generateHex("red").toUpperCase();
const isValidHex1 = /^[0-9A-F]{6}$/.test(hex1)
assert.isTrue(isValidHex1);

const r1 = parseInt(hex1.slice(0, 2), 16);
const g1 = parseInt(hex1.slice(2, 4), 16);
const b1 = parseInt(hex1.slice(4, 6), 16);

assert.isAbove(r1, g1);
assert.isAbove(r1, b1);

const hex2 = generateHex("red").toUpperCase();
const isValidHex2 = /^[0-9A-F]{6}$/.test(hex2);
assert.isTrue(isValidHex2);

const r2 = parseInt(hex2.slice(0, 2), 16);
const g2 = parseInt(hex2.slice(2, 4), 16);
const b2 = parseInt(hex2.slice(4, 6), 16);

assert.isAbove(r2, g2);
assert.isAbove(r2, b2);
assert.notEqual(hex1, hex2);
```

يجب أن تُرجع استدعاء الدالة `generateHex("green")` مرتين قيمتين مختلفتين للون السداسي حيث يكون اللون الأخضر هو السائد.

```js
const hex1 = generateHex("green").toUpperCase();
const isValidHex1 = /^[0-9A-F]{6}$/.test(hex1)
assert.isTrue(isValidHex1);

const r1 = parseInt(hex1.slice(0, 2), 16);
const g1 = parseInt(hex1.slice(2, 4), 16);
const b1 = parseInt(hex1.slice(4, 6), 16);

assert.isAbove(g1, r1);
assert.isAbove(g1, b1);

const hex2 = generateHex("green").toUpperCase();
const isValidHex2 = /^[0-9A-F]{6}$/.test(hex2);
assert.isTrue(isValidHex2);

const r2 = parseInt(hex2.slice(0, 2), 16);
const g2 = parseInt(hex2.slice(2, 4), 16);
const b2 = parseInt(hex2.slice(4, 6), 16);

assert.isAbove(g2, r2);
assert.isAbove(g2, b2);
assert.notEqual(hex1, hex2);
```

يجب أن تُرجع استدعاء الدالة `generateHex("blue")` مرتين قيمتين مختلفتين للون السداسي حيث يكون اللون الأزرق هو السائد.

```js
const hex1 = generateHex("blue").toUpperCase();
const isValidHex1 = /^[0-9A-F]{6}$/.test(hex1)
assert.isTrue(isValidHex1);

const r1 = parseInt(hex1.slice(0, 2), 16);
const g1 = parseInt(hex1.slice(2, 4), 16);
const b1 = parseInt(hex1.slice(4, 6), 16);

assert.isAbove(b1, r1);
assert.isAbove(b1, g1);

const hex2 = generateHex("blue").toUpperCase();
const isValidHex2 = /^[0-9A-F]{6}$/.test(hex2);
assert.isTrue(isValidHex2);

const r2 = parseInt(hex2.slice(0, 2), 16);
const g2 = parseInt(hex2.slice(2, 4), 16);
const b2 = parseInt(hex2.slice(4, 6), 16);

assert.isAbove(b2, r2);
assert.isAbove(b2, g2);
assert.notEqual(hex1, hex2);
```

# --seed--

## --seed-contents--

```js
function generateHex(color) {

  return color;
}
```

# --solutions--

```js
function generateHex(color) {
  const toHex = n => n.toString(16).padStart(2, "0").toUpperCase();

  const dominant = Math.floor(Math.random() * 86) + 170;
  const weak1 = Math.floor(Math.random() * 170);
  const weak2 = Math.floor(Math.random() * 170);

  let r, g, b;

  switch (color) {
    case "red":
      r = dominant;
      g = weak1;
      b = weak2;
      break;
    case "green":
      r = weak1;
      g = dominant;
      b = weak2;
      break;
    case "blue":
      r = weak1;
      g = weak2;
      b = dominant;
      break;
    default:
      return "Invalid color";
  }

  return `${toHex(r)}${toHex(g)}${toHex(b)}`;
}
```
