---
id: 6821ec02237de8297eaee79a
title: "تحدٍّ 24: جملة تحتوي كل الأحرف"
challengeType: 28
dashedName: challenge-24
---

# --description--

معطى كلمة أو جملة وسلسلة من الأحرف الصغيرة، حدّد إذا كانت الكلمة أو الجملة تستخدم كل الأحرف من المجموعة المعطاة مرة واحدة على الأقل ولا تستخدم أحرفًا أخرى.

- تجاهل الأحرف غير الأبجدية في الكلمة أو الجملة.
- تجاهل حالة الأحرف في الكلمة أو الجملة.

# --hints--

`isPangram("hello", "helo")` يجب أن تُعيد `true`

```js
assert.isTrue(isPangram("hello", "helo"));
```

`isPangram("hello", "hel")` يجب أن تُعيد `false`

```js
assert.isFalse(isPangram("hello", "hel"));
```

`isPangram("hello", "helow")` يجب أن تُعيد `false`

```js
assert.isFalse(isPangram("hello", "helow"));
```

`isPangram("hello world", "helowrd")` يجب أن تُعيد `true`

```js
assert.isTrue(isPangram("hello world", "helowrd"));
```

`isPangram("Hello World!", "helowrd")` يجب أن تُعيد `true`

```js
assert.isTrue(isPangram("Hello World!", "helowrd"));
```

`isPangram("Hello World!", "heliowrd")` يجب أن تُعيد `false`

```js
assert.isFalse(isPangram("Hello World!", "heliowrd"));
```

`isPangram("freeCodeCamp", "frcdmp")` يجب أن تُعيد `false`

```js
assert.isFalse(isPangram("freeCodeCamp", "frcdmp"));
```

`isPangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz")` يجب أن تُعيد `true`

```js
assert.isTrue(isPangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz"));
```

# --seed--

## --seed-contents--

```js
function isPangram(sentence, letters) {

  return sentence;
}
```

# --solutions--

```js
function isPangram(sentence, letters) {
  const usedLetters = [];
  for (let i = 0; i < sentence.length; i++) {
    const letter = sentence[i].toLowerCase();
    if (!usedLetters.includes(letter) && /[a-z]/.test(letter)) {
      usedLetters.push(letter);
    }
  }

  const sortedLetters = letters.split('').sort().join('');
  const sortedUsedLetters = usedLetters.sort().join('');

  return sortedLetters === sortedUsedLetters;
}
```
