---
id: 6821ebda237de8297eaee792
title: "تحدٍّ 16: عكس الأقواس العادية"
challengeType: 28
dashedName: challenge-16
---

# --description--

أعطِ سلسلة نصية تحتوي على أقواس عادية متسكّنة بشكل صحيح، أرجع النسخة المفككة من السلسلة باستخدام القواعد التالية:

- يجب عكس كل الأحرف داخل كل زوج من الأقواس العادية.
- يجب إزالة الأقواس العادية من النتيجة النهائية.
- إذا كانت الأقواس العادية مسكنة، يجب عكس الزوج الداخلي أولًا، ثم تضمين نتيجته في عكس الزوج الخارجي.
- افترض أن كل الأقواس العادية متوازنة ومتسكّنة بشكل صحيح.

# --hints--

`decode("(f(b(dc)e)a)")` يجب أن تُرجع `"abcdef"`.

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` يجب أن تُرجع `"Can you read this?"`.

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")` يجب أن تُرجع `"freeCodeCamp"`.

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```
