---
id: 68216ef80f957572e7c340c5
title: "تحدٍّ 12: فك شفرة الرسالة"
challengeType: 28
dashedName: challenge-12
---

# --description--

أعطِ سلسلة نصية لرسالة سرية، ورقم صحيح يمثل عدد الحروف التي استُخدمت لتحويل الرسالة وترميزها، أعد السلسلة النصية المفكوكة.

- الرقم الموجب يعني أن الرسالة تحولت للأمام في الأبجدية.
- الرقم السالب يعني أن الرسالة تحولت للخلف في الأبجدية.
- حالة الحروف مهمة، يجب أن تحتفظ الأحرف المفكوكة بحالة الأحرف الأصلية المشفرة.
- لا تُفكك الأحرف غير الأبجدية.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` يجب أن تُعيد `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` يجب أن تُعيد `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` يجب أن تُعيد `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` يجب أن تُعيد `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
