---
id: 6821ebea237de8297eaee795
title: "تحدٍّ 19: ضوء الشموع"
challengeType: 28
dashedName: challenge-19
---

# --description--

معطى عدد صحيح يمثل عدد الشموع التي تبدأ بها، وعدد صحيح يمثل عدد الشموع المحترقة اللازمة لصنع شمعة جديدة، أرجع عدد الشموع التي ستستخدمها بعد صنع وحرق أكبر عدد ممكن.

على سبيل المثال، إذا كان لديك 7 شموع ويحتاج صنع شمعة جديدة إلى شمعتين محترقتين:

1. احرق 7 شموع لتحصل على 7 بقايا،
2. أعد تدوير 6 بقايا لصنع 3 شموع جديدة (تبقى 1 بقايا)،
3. احرق 3 شموع لتحصل على 3 بقايا إضافية (المجموع 4)،
4. أعد تدوير 4 بقايا لصنع 2 شمعة جديدة،
5. احرق 2 شمعة لتحصل على 2 بقايا،
6. أعد تدوير 2 بقايا لصنع شمعة جديدة واحدة،
7. احرق شمعة واحدة.

ستكون قد حرقت 13 شمعة في المجموع في هذا المثال.

# --hints--

`burnCandles(7, 2)` يجب أن تُرجع `13`

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` يجب أن تُرجع `12`

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` يجب أن تُرجع `29`

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` يجب أن تُرجع `22`

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` يجب أن تُرجع `3517`

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```
