---
id: 6821ebdf237de8297eaee793
title: "تحدٍّ 17: عدم ترتيب العمليات"
challengeType: 28
dashedName: challenge-17
---

# --description--

معطى مصفوفة من الأعداد الصحيحة ومصفوفة من العوامل النصية، طبّق العمليات على الأعداد بالتتابع من اليسار إلى اليمين. كرر العمليات حسب الحاجة حتى تُستخدم كل الأعداد. أرجع النتيجة النهائية.

على سبيل المثال، معطى `[1, 2, 3, 4, 5]` و `['+', '*']`، أرجع نتيجة تقييم `1 + 2 * 3 + 4 * 5` من اليسار إلى اليمين متجاهلًا الترتيب القياسي للعمليات.

- العوامل الصالحة هي `+`، `-`، `*`، `/`، و `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` يجب أن تُرجع `3`

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` يجب أن تُرجع `38`

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` يجب أن تُرجع `60`

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` يجب أن تُرجع `30`

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` يجب أن تُرجع `-2`

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
