---
id: 6821ebfd237de8297eaee799
title: "تحدٍّ 23: تحويل RGB إلى Hex"
challengeType: 28
dashedName: challenge-23
---

# --description--

معطى سلسلة لون CSS بصيغة `rgb(r, g, b)`، أرجع ما يعادلها بالنظام الست عشري.

إليك بعض الأمثلة على المخرجات مقابل مُدخَل معين:

| المُدخَل   | المخرجات   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- اجعل كل الأحرف أحرف صغيرة.
- أرجع `#` متبوعًا بستة أحرف. لا تستخدم أي قيم مختصرة.

# --hints--

`rgbToHex("rgb(255, 255, 255)")` يجب أن تُرجع `"#ffffff"`.

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` يجب أن تُرجع `"#010b6f"`.

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` يجب أن تُرجع `"#add8e6"`.

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` يجب أن تُرجع `"#4f7bc9"`.

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
