---
id: 6821ebe4237de8297eaee794
title: "تحدٍّ 18: الثاني الأفضل"
challengeType: 28
dashedName: challenge-18
---

# --description--

معطاة مصفوفة من الأعداد الصحيحة تمثل سعر حواسيب محمولة مختلفة، وعدد صحيح يمثل ميزانيتك، أرجع:

1. الحاسوب المحمول الثاني الأغلى إذا كان ضمن ميزانيتك، أو
2. الحاسوب المحمول الأغلى الذي يقع ضمن ميزانيتك، أو
3. `0` إذا لم يكن هناك حواسيب محمولة ضمن ميزانيتك.

- يجب تجاهل الأسعار المكررة.

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)` يجب أن تُرجع `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)` يجب أن تُرجع `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)` يجب أن تُرجع `1899`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)` يجب أن تُرجع `0`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` يجب أن تُرجع `1400`

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
