---
id: 6821ebc9237de8297eaee78f
title: "تحدٍّ 13: عدد أولي غير طبيعي"
challengeType: 28
dashedName: challenge-13
---

# --description--

معطى عدد صحيح، حدّد إذا كان هذا العدد عددًا أوليًا أو عددًا أوليًا سالبًا.

- العدد الأولي هو عدد صحيح موجب أكبر من 1 لا يقبل القسمة إلا على 1 وعلى نفسه فقط.
- العدد الأولي السالب هو النسخة السالبة من عدد أولي موجب.
- `1` و `0` لا يُعتبران أعدادًا أولية.

# --hints--

يجب أن تُرجع الدالة `isUnnaturalPrime(1)` القيمة `false`.

```js
assert.isFalse(isUnnaturalPrime(1));
```

يجب أن تُرجع الدالة `isUnnaturalPrime(-1)` القيمة `false`.

```js
assert.isFalse(isUnnaturalPrime(-1));
```

يجب أن تُرجع الدالة `isUnnaturalPrime(19)` القيمة `true`.

```js
assert.isTrue(isUnnaturalPrime(19));
```

يجب أن تُرجع الدالة `isUnnaturalPrime(-23)` القيمة `true`.

```js
assert.isTrue(isUnnaturalPrime(-23));
```

يجب أن تُرجع الدالة `isUnnaturalPrime(0)` القيمة `false`.

```js
assert.isFalse(isUnnaturalPrime(0));
```

يجب أن تُرجع الدالة `isUnnaturalPrime(97)` القيمة `true`.

```js
assert.isTrue(isUnnaturalPrime(97));
```

يجب أن تُرجع الدالة `isUnnaturalPrime(-61)` القيمة `true`.

```js
assert.isTrue(isUnnaturalPrime(-61));
```

يجب أن تُرجع الدالة `isUnnaturalPrime(99)` القيمة `false`.

```js
assert.isFalse(isUnnaturalPrime(99));
```

يجب أن تُرجع الدالة `isUnnaturalPrime(-44)` القيمة `false`.

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
