---
id: 6821ebee237de8297eaee796
title: "تحدٍّ 20: العناصر المكررة في المصفوفة"
challengeType: 28
dashedName: challenge-20
---

# --description--

بالنظر إلى مصفوفة من الأعداد الصحيحة، أعد مصفوفة من الأعداد الصحيحة التي تظهر أكثر من مرة في المصفوفة الأصلية، مرتبة تصاعديًا. إذا لم تظهر أي قيمة أكثر من مرة، أعد مصفوفة فارغة.

- تضمّن نسخة واحدة فقط من كل قيمة في المصفوفة المعادة.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` يجب أن تُعيد `[]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` يجب أن تُعيد `[1, 2]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` يجب أن تُعيد `[-6, 0, 2, 4, 5, 23]`.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
