---
id: 68adce01c0e1144d0a90295e
title: "تحدٍّ 29: منشئ الاختصارات"
challengeType: 28
dashedName: challenge-29
---

# --description--

معطاة سلسلة نصية تحتوي على كلمة واحدة أو أكثر، أرجع اختصارًا للكلمات باستخدام القيود التالية:

- يجب أن يتكون الاختصار من الحرف الأول من كل كلمة بحروف كبيرة، ما لم يُذكر خلاف ذلك.
- يجب تجاهل الحرف الأول من هذه الكلمات ما لم تكن الكلمة الأولى في السلسلة النصية المعطاة: `a`، `for`، `an`، `and`، `by`، و`of`.
- يجب إرجاع حروف الاختصار بالترتيب الذي تظهر به.
- يجب ألا يحتوي الاختصار على أي فراغات.

# --hints--

`buildAcronym("Search Engine Optimization")` يجب أن تُرجع `"SEO"`.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` يجب أن تُرجع `"FAQ"`.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` يجب أن تُرجع `"NASA"`.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` يجب أن تُرجع `"FBI"`.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` يجب أن تُرجع `"FYI"`.

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")` يجب أن تُرجع `"BTW"`.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` يجب أن تُرجع `"AUHWPOTIMSH"`.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
