---
id: 68adce01c0e1144d0a902956
title: "تحدٍّ 25: مكرر الحروف المتحركة"
challengeType: 28
dashedName: challenge-25
---

# --description--

أعطِ سلسلة نصية، أرجع إصدارًا جديدًا من السلسلة حيث تتكرر كل حرف متحرك مرة أكثر من الحرف المتحرك السابق الذي واجهته. على سبيل المثال، يجب أن يبقى الحرف المتحرك الأول في الجملة كما هو دون تغيير. يجب أن يظهر الحرف المتحرك الثاني مرتين متتاليتين. يجب أن يظهر الحرف المتحرك الثالث ثلاث مرات متتالية، وهكذا.

- الأحرف `a`، `e`، `i`، `o`، و `u`، سواء كانت بأحرف كبيرة أو أحرف صغيرة، تُعتبر حروفًا متحركة.
- يجب أن يحتفظ الحرف المتحرك الأصلي بحالته.
- يجب أن تكون الحروف المتحركة المكررة بأحرف صغيرة.
- يجب أن تحتفظ جميع الأحرف غير المتحركة بحالتها الأصلية.

# --hints--

`repeatVowels("hello world")` يجب أن تُرجع `"helloo wooorld"`.

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` يجب أن تُرجع `"freeeCooodeeeeCaaaaamp"`.

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` يجب أن تُرجع `"AEeIiiOoooUuuuu"`.

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` يجب أن تُرجع `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
