---
id: 68adce01c0e1144d0a902958
title: "تحدٍّ 26: مدقق عنوان IPv4"
challengeType: 28
dashedName: challenge-26
---

# --description--

معطاة سلسلة نصية، حدّد إذا كانت عنوان IPv4 صالحًا. يتكون عنوان IPv4 الصالح من أربعة أعداد صحيحة مفصولة بنقاط (`.`). يجب أن يحقق كل عدد الشروط التالية:

- أن يكون بين 0 و255 شاملة.
- ألا يحتوي على أصفار بادئة (مثلاً 0 مسموح، لكن 01 غير مسموح).
- يُسمح فقط بالأحرف الرقمية.

# --hints--

يجب أن تُرجع الدالة `isValidIPv4("192.168.1.1")` القيمة `true`.

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

يجب أن تُرجع الدالة `isValidIPv4("0.0.0.0")` القيمة `true`.

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

يجب أن تُرجع الدالة `isValidIPv4("255.01.50.111")` القيمة `false`.

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

يجب أن تُرجع الدالة `isValidIPv4("255.00.50.111")` القيمة `false`.

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

يجب أن تُرجع الدالة `isValidIPv4("256.101.50.115")` القيمة `false`.

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

يجب أن تُرجع الدالة `isValidIPv4("192.168.101.")` القيمة `false`.

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

يجب أن تُرجع الدالة `isValidIPv4("192168145213")` القيمة `false`.

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
