---
id: 68adce01c0e1144d0a90295c
title: "تحدٍّ 28: محلل الأعداد الرومانية"
challengeType: 28
dashedName: challenge-28
---

# --description--

أعطِ سلسلة نصية تمثل عددًا رومانيًا، أرجع قيمته الصحيحة.

تتكون الأعداد الرومانية من الرموز والقيم التالية:

| الرمز | القيمة |
|--------|-------|
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

- تُقرأ الأعداد من اليسار إلى اليمين. إذا ظهر عدد أصغر قبل عدد أكبر، تُطرح قيمته. وإلا، تُجمع القيم.

# --hints--

يجب أن تُرجع الدالة `parseRomanNumeral("III")` القيمة `3`.

```js
assert.equal(parseRomanNumeral("III"), 3);
```

يجب أن تُرجع الدالة `parseRomanNumeral("IV")` القيمة `4`.

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

يجب أن تُرجع الدالة `parseRomanNumeral("XXVI")` القيمة `26`.

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

يجب أن تُرجع الدالة `parseRomanNumeral("XCIX")` القيمة `99`.

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

يجب أن تُرجع الدالة `parseRomanNumeral("CDLX")` القيمة `460`.

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

يجب أن تُرجع الدالة `parseRomanNumeral("DIV")` القيمة `504`.

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

يجب أن تُرجع الدالة `parseRomanNumeral("MMXXV")` القيمة `2025`.

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
