---
id: 68af0687ef34c76c28ffa54f
title: "تحدٍّ 34: الأعداد المفقودة"
challengeType: 28
dashedName: challenge-34
---

# --description--

معطاة مصفوفة من الأعداد الصحيحة من 1 إلى `n`، شاملة، أرجع مصفوفة بكل الأعداد الصحيحة المفقودة بين 1 و`n` (حيث `n` هو أكبر رقم في المصفوفة المعطاة).

- قد تكون المصفوفة المعطاة غير مرتبة وقد تحتوي على أعداد مكررة.
- يجب أن تكون المصفوفة المرجعة مرتبة تصاعديًا.
- إذا لم تكن هناك أعداد مفقودة، أرجع مصفوفة فارغة.

# --hints--

`findMissingNumbers([1, 3, 5])` يجب أن تُرجع `[2, 4]`.

```js
assert.deepEqual(findMissingNumbers([1, 3, 5]), [2, 4]);
```

`findMissingNumbers([1, 2, 3, 4, 5])` يجب أن تُرجع `[]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 10])` يجب أن تُرجع `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([10, 1, 10, 1, 10, 1])` يجب أن تُرجع `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([3, 1, 4, 1, 5, 9])` يجب أن تُرجع `[2, 6, 7, 8]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` يجب أن تُرجع `[11]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11]);
```

# --seed--

## --seed-contents--

```js
function findMissingNumbers(arr) {

  return arr;
}
```

# --solutions--

```js
function findMissingNumbers(arr) {
  if (arr.length === 0) return [];
  
  const n = Math.max(...arr);
  const set = new Set(arr);
  const result = [];

  for (let i = 1; i <= n; i++) {
    if (!set.has(i)) {
      result.push(i);
    }
  }
  return result;
}
```
