---
id: 68af0687ef34c76c28ffa547
title: "تحدٍّ 30: أحرف فريدة"
challengeType: 28
dashedName: challenge-30
---

# --description--

معطاة سلسلة نصية، حدّد إذا كانت كل الأحرف في السلسلة فريدة.

- يجب اعتبار الأحرف الكبيرة والأحرف الصغيرة أحرفًا مختلفة.

# --hints--

`allUnique("abc")` يجب أن تُرجع `true`.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` يجب أن تُرجع `true`.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` يجب أن تُرجع `true`.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` يجب أن تُرجع `true`.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` يجب أن تُرجع `false`.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` يجب أن تُرجع `false`.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` يجب أن تُرجع `false`.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
