---
id: 68af0687ef34c76c28ffa54b
title: "تحدٍّ 32: عكس الجملة"
challengeType: 28
dashedName: challenge-32
---

# --description--

عند إعطائك سلسلة نصية من الكلمات، أعد سلسلة جديدة بالكلمات بترتيب عكسي. على سبيل المثال، يجب أن تكون الكلمة الأولى في نهاية السلسلة المعادة، والكلمة الأخيرة في بداية السلسلة المعادة.

- في السلسلة المعطاة، يمكن أن تفصل الكلمات بمسافة واحدة أو أكثر.
- يجب أن تحتوي السلسلة المعادة على مسافة واحدة فقط بين الكلمات.

# --hints--

`reverseSentence("world hello")` يجب أن تعيد `"hello world"`.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` يجب أن تعيد `"git commit push"`.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` يجب أن تعيد `"sudo install npm"`.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` يجب أن تعيد `"export function default import"`.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
