---
id: 68af0687ef34c76c28ffa54d
title: "تحدٍّ 33: وقت الشاشة"
challengeType: 28
dashedName: challenge-33
---

# --description--

معطى مصفوفة من سبعة أعداد صحيحة، تمثل وقت أسبوع، حيث كل عدد صحيح هو عدد الساعات التي قضيتها على هاتفك في ذلك اليوم، حدّد إذا كان وقت الشاشة كثيرًا بناءً على هذه القيود:

- إذا كان أي يوم واحد يحتوي على 10 ساعات أو أكثر، فهذا كثير.
- إذا كان متوسط أي ثلاثة أيام متتالية أكبر من أو يساوي 8 ساعات، فهذا كثير.
- إذا كان متوسط الأيام السبعة أكبر من أو يساوي 6 ساعات، فهذا كثير.

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` يجب أن تُرجع `false`.

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` يجب أن تُرجع `false`.

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` يجب أن تُرجع `false`.

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` يجب أن تُرجع `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` يجب أن تُرجع `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` يجب أن تُرجع `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` يجب أن تُرجع `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```
