---
id: 68b06e589bf2273243814777
title: "تحدٍّ 39: ملء الخزان"
challengeType: 28
dashedName: challenge-39
---

# --description--

بما أن حجم خزان الوقود ومستوى الوقود الحالي وسعر الجالون معطاة، أرجع تكلفة ملء الخزان بالكامل.

- `tankSize` هو السعة الكلية للخزان بالجالونات.
- `fuelLevel` هو كمية الوقود الحالية في الخزان بالجالونات.
- `pricePerGallon` هو تكلفة جالون واحد من الوقود.
- يجب تقريب القيمة المرجعة إلى منزلتين عشريتين بالشكل: `"$d.dd"`.

# --hints--

`costToFill(20, 0, 4.00)` يجب أن تُرجع `"$80.00"`.

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` يجب أن تُرجع `"$17.50"`.

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` يجب أن تُرجع `"$29.25"`.

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` يجب أن تُرجع `"$0.00"`.

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` يجب أن تُرجع `"$21.89"`.

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
