---
id: 68b06e589bf2273243814773
title: "تحدٍّ 37: تحويل أول حرف في الجملة إلى حرف كبير"
challengeType: 28
dashedName: challenge-37
---

# --description--

أعط فقرة، أعد فقرة جديدة حيث يكون أول حرف في كل جملة كبيرًا.

- يجب الحفاظ على جميع الأحرف الأخرى.
- يمكن أن تنتهي الجمل بنقطة (`.`)، أو علامة استفهام واحدة أو أكثر (`?`)، أو علامة تعجب واحدة أو أكثر (`!`).

# --hints--

`capitalize("this is a simple sentence.")` يجب أن تُعيد `"This is a simple sentence."`.

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")` يجب أن تُعيد `"Hello world. How are you?"`.

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")` يجب أن تُعيد `"I did today's coding challenge... It was fun!!"`.

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")` يجب أن تُعيد `"Crazy!!!Strange???Unconventional...Sentences."`.

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")` يجب أن تُعيد `"There's a space before this period . Why is there a space before that period ?"`.

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```
