---
id: 68b06e589bf2273243814771
title: "تحدٍّ 36: مُعدِّل منظم الحرارة"
challengeType: 28
dashedName: challenge-36
---

# --description--

بالنظر إلى درجة حرارة الغرفة الحالية ودرجة الحرارة المستهدفة، أعد سلسلة نصية تشير إلى كيفية تعديل درجة حرارة الغرفة بناءً على هذه القيود:

- أعد `"heat"` إذا كانت درجة الحرارة الحالية أقل من الهدف.
- أعد `"cool"` إذا كانت درجة الحرارة الحالية أعلى من الهدف.
- أعد `"hold"` إذا كانت درجة الحرارة الحالية مساوية للهدف.

# --hints--

`adjustThermostat(68, 72)` يجب أن تُعيد `"heat"`.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` يجب أن تُعيد `"cool"`.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` يجب أن تُعيد `"hold"`.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` يجب أن تُعيد `"heat"`.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` يجب أن تُعيد `"cool"`.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` يجب أن تُعيد `"hold"`.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
