---
id: 68b06e589bf2273243814775
title: "تحدٍّ 38: مولد العناوين الصديقة للروابط"
challengeType: 28
dashedName: challenge-38
---

# --description--

بالنظر إلى سلسلة نصية، أعد نسخة صديقة للروابط من السلسلة باستخدام القيود التالية:

- يجب أن تكون كل الأحرف أحرف صغيرة.
- يجب إزالة كل الأحرف التي ليست أحرفًا أو أرقامًا أو فراغات.
- يجب استبدال كل الفراغات برمز الفراغ المشفر في الرابط `%20`.
- يجب استبدال الفراغات المتتالية برمز `%20` واحد فقط.
- يجب ألا تحتوي السلسلة المعادة على `%20` في البداية أو النهاية.

# --hints--

`generateSlug("helloWorld")` يجب أن تُعيد `"helloworld"`.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` يجب أن تُعيد `"hello%20world"`.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` يجب أن تُعيد `"helloworld"`.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` يجب أن تُعيد `"hello%20world"`.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` يجب أن تُعيد `"h3110%20w0r1d"`.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
