---
id: 68b06e589bf227324381476f
title: "تحدٍّ 35: تكرار الكلمات"
challengeType: 28
dashedName: challenge-35
---

# --description--

أعطِ فقرة، أرجع مصفوفة تحتوي على أكثر ثلاث كلمات تكرارًا.

- الكلمات في الفقرة مفصولة بمسافات.
- تجاهل حالة الأحرف في الفقرة المعطاة. على سبيل المثال، اعتبر `Hello` و `hello` نفس الكلمة.
- تجاهل علامات الترقيم في الفقرة المعطاة. علامات الترقيم هي الفواصل (`,`)، النقاط (`.`)، وعلامات التعجب (`!`).
- يجب أن تحتوي المصفوفة المرجعة على كلمات بأحرف صغيرة فقط.
- يجب أن تكون المصفوفة المرجعة مرتبة ترتيبًا تنازليًا بحيث تكون الكلمة الأكثر تكرارًا أولاً.

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` يجب أن تُرجع `["coding", "python", "in"]`.

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` يجب أن تُرجع `["i", "like", "coding"]`.

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` يجب أن تُرجع `["debug", "test", "deploy"]`.

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
