---
id: 68b1f72371a5ac895ac70a06
title: "تحدٍّ 42: تخزين الفيديو"
challengeType: 28
dashedName: challenge-42
---

# --description--

معطى حجم فيديو، وحدة لحجم الفيديو، سعة قرص صلب، ووحدة للقرص الصلب، أرجع عدد مقاطع الفيديو التي يمكن للقرص الصلب تخزينها باستخدام القيود التالية:

- وحدة حجم الفيديو يمكن أن تكون بايت (`"B"`)، كيلوبايت (`"KB"`)، ميغابايت (`"MB"`)، أو غيغابايت (`"GB"`).
- إذا لم تُعطَ إحدى وحدات الفيديو السابقة، أرجع `"Invalid video unit"`.
- وحدة سعة القرص الصلب يمكن أن تكون غيغابايت (`"GB"`) أو تيرابايت (`"TB"`).
- إذا لم تُعطَ إحدى وحدات القرص الصلب السابقة، أرجع `"Invalid drive unit"`.
- أرجع عدد مقاطع الفيديو الكاملة التي يمكن للقرص استيعابها.
- استخدم التحويلات التالية:

| الوحدة | ما يعادلها |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |
| 1 TB |   1000 GB  |

على سبيل المثال، معطى `500`، `"MB"`، `100`، و`"GB"` كمعلمات، حدد كم عدد مقاطع الفيديو بحجم 500 ميغابايت يمكن تخزينها على قرص صلب سعته 100 غيغابايت.

# --hints--

`numberOfVideos(500, "MB", 100, "GB")` يجب أن تُرجع `200`.

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")` يجب أن تُرجع `"Invalid video unit"`.

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")` يجب أن تُرجع `"Invalid drive unit"`.

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")` يجب أن تُرجع `4000`.

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")` يجب أن تُرجع `1466`.

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
