---
id: 68b1f72371a5ac895ac70a08
title: "تحدٍّ 43: الأرقام مقابل الحروف"
challengeType: 28
dashedName: challenge-43
---

# --description--

بالنظر إلى سلسلة نصية، أعد `"digits"` إذا كانت السلسلة تحتوي على أرقام أكثر من الحروف، و`"letters"` إذا كانت تحتوي على حروف أكثر من الأرقام، و`"tie"` إذا كان عدد الأرقام والحروف متساويًا.

- تتكون الأرقام من `0-9`.
- تتكون الحروف من `a-z` بحروف كبيرة أو صغيرة.
- تجاهل أي حروف أخرى.

# --hints--

`digitsOrLetters("abc123")` يجب أن تُعيد `"tie"`.

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` يجب أن تُعيد `"letters"`.

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` يجب أن تُعيد `"digits"`.

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` يجب أن تُعيد `"letters"`.

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` يجب أن تُعيد `"digits"`.

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` يجب أن تُعيد `"tie"`.

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
