---
id: 68b1f72371a5ac895ac70a02
title: "تحدٍّ 40: تخزين الصور"
challengeType: 28
dashedName: challenge-40
---

# --description--

معطى حجم الصورة بالميغابايت (MB) وسعة القرص الصلب بالجيجابايت (GB)، أرجع عدد الصور التي يمكن للقرص الصلب تخزينها باستخدام القيود التالية:

- 1 جيجابايت يساوي 1000 ميغابايت.
- أرجع عدد الصور الكاملة التي يمكن للقرص تخزينها.

# --hints--

يجب أن تُرجع `numberOfPhotos(1, 1)` القيمة `1000`.

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

يجب أن تُرجع `numberOfPhotos(2, 1)` القيمة `500`.

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

يجب أن تُرجع `numberOfPhotos(4, 256)` القيمة `64000`.

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

يجب أن تُرجع `numberOfPhotos(3.5, 750)` القيمة `214285`.

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

يجب أن تُرجع `numberOfPhotos(3.5, 5.5)` القيمة `1571`.

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
