---
id: 68b1f72371a5ac895ac70a04
title: "تحدٍّ 41: تخزين الملفات"
challengeType: 28
dashedName: challenge-41
---

# --description--

معطى حجم ملف، وحدة حجم الملف، وسعة القرص الصلب بالجيجابايت (GB)، أرجع عدد الملفات التي يمكن للقرص الصلب تخزينها باستخدام القيود التالية:

- وحدة حجم الملف يمكن أن تكون بايت (`"B"`)، كيلوبايت (`"KB"`)، أو ميغابايت (`"MB"`).
- أرجع عدد الملفات الكاملة التي يمكن للقرص استيعابها.
- استخدم التحويلات التالية:

| الوحدة | ما يعادلها |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

على سبيل المثال، معطى `500`، `"KB"`، و `1` كمعلمات، حدد كم عدد ملفات 500 KB التي يمكن أن تتسع على قرص صلب بسعة 1 GB.

# --hints--

`numberOfFiles(500, "KB", 1)` يجب أن تُرجع `2000`.

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` يجب أن تُرجع `20000`.

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` يجب أن تُرجع `200`.

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` يجب أن تُرجع `366210`.

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` يجب أن تُرجع `453514`.

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` يجب أن تُرجع `166666`.

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
