---
id: 68b1f72371a5ac895ac70a0a
title: "تحدٍّ 44: مرآة السلسلة النصية"
challengeType: 28
dashedName: challenge-44
---

# --description--

معطيتان سلسلتان نصيتان، حدّد إذا كانت السلسلة النصية الثانية مرآة للأولى.

- تُعتبر السلسلة النصية مرآة إذا احتوت على نفس الأحرف بترتيب عكسي.
- اعتبر الأحرف الكبيرة والأحرف الصغيرة مميزة عن بعضها.
- تجاهل كل الأحرف غير الأبجدية.

# --hints--

`isMirror("helloworld", "helloworld")` يجب أن تُرجع `false`.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` يجب أن تُرجع `false`.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` يجب أن تُرجع `false`.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
