---
id: 68b7687dded630607aceccb1
title: "تحدٍّ 48: كاشف الرسائل المزعجة"
challengeType: 28
dashedName: challenge-48
---

# --description--

معطى رقم هاتف بالصيغة `"+A (BBB) CCC-DDDD"`، حيث تمثل كل حرف رقمًا كما يلي:

- `A` تمثل رمز الدولة ويمكن أن تكون أي عدد من الأرقام.
- `BBB` تمثل رمز المنطقة وستكون دائمًا ثلاثة أرقام.
- `CCC` و `DDDD` تمثلان الرقم المحلي وسيكون طولهما دائمًا ثلاثة وأربعة أرقام على التوالي.

حدد ما إذا كان الرقم مزعجًا بناءً على المعايير التالية:

- رمز الدولة أطول من رقمين أو لا يبدأ بصفر (`0`).
- رمز المنطقة أكبر من 900 أو أقل من 200.
- مجموع أول ثلاثة أرقام من الرقم المحلي يظهر ضمن آخر أربعة أرقام من الرقم المحلي.
- الرقم يحتوي على نفس الرقم مكرر أربع مرات أو أكثر على التوالي (مع تجاهل أحرف التنسيق).

# --hints--

`isSpam("+0 (200) 234-0182")` يجب أن تُرجع `false`.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` يجب أن تُرجع `true`.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` يجب أن تُرجع `true`.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` يجب أن تُرجع `true`.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` يجب أن تُرجع `true`.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` يجب أن تُرجع `true`.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` يجب أن تُرجع `true`.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` يجب أن تُرجع `false`.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
