---
id: 68b7cadffed0e75a517da677
title: "تحدٍّ 54: قوة كلمة المرور P@ssw0rd Str3ngth!"
challengeType: 28
dashedName: challenge-54
---

# --description--

بالنظر إلى سلسلة نصية لكلمة مرور، أعد `"weak"` أو `"medium"` أو `"strong"` بناءً على قوة كلمة المرور.

تُقيَّم كلمة المرور وفقًا للقواعد التالية:

- يجب أن تكون على الأقل 8 أحرف.
- تحتوي على أحرف كبيرة وأحرف صغيرة.
- تحتوي على رقم واحد على الأقل.
- تحتوي على حرف خاص واحد على الأقل من هذه المجموعة: `!`، `@`، `#`، `$`، `%`، `^`، `&`، أو `*`.

أعد `"weak"` إذا استوفت كلمة المرور أقل من شرطين.
أعد `"medium"` إذا استوفت كلمة المرور شرطين أو ثلاثة.
أعد `"strong"` إذا استوفت كلمة المرور جميع الشروط الأربعة.

# --hints--

يجب أن تُعيد `checkStrength("123456")` القيمة `"weak"`.

```js
assert.equal(checkStrength("123456"), "weak");
```

يجب أن تُعيد `checkStrength("pass!!!")` القيمة `"weak"`.

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

يجب أن تُعيد `checkStrength("Qwerty")` القيمة `"weak"`.

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

يجب أن تُعيد `checkStrength("PASSWORD")` القيمة `"weak"`.

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

يجب أن تُعيد `checkStrength("PASSWORD!")` القيمة `"medium"`.

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

يجب أن تُعيد `checkStrength("PassWord%^!")` القيمة `"medium"`.

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

يجب أن تُعيد `checkStrength("qwerty12345")` القيمة `"medium"`.

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

يجب أن تُعيد `checkStrength("S3cur3P@ssw0rd")` القيمة `"strong"`.

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

يجب أن تُعيد `checkStrength("C0d3&Fun!")` القيمة `"strong"`.

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
