---
id: 68b7cadffed0e75a517da675
title: "تحدٍّ 53: التحويل من عشري إلى ثنائي"
challengeType: 28
dashedName: challenge-53
---

# --description--

أعطِ عددًا صحيحًا غير سالب، أرجع تمثيله الثنائي كسلسلة نصية.

العدد الثنائي يستخدم فقط الرقمين `0` و `1` لتمثيل أي عدد. لتحويل عدد عشري إلى ثنائي، اقسم العدد على `2` بشكل متكرر وسجل الباقي. كرر حتى يصبح العدد صفرًا. اقرأ البواقي من الأخير إلى الأول. على سبيل المثال، لتحويل `12` إلى ثنائي:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` في النظام الثنائي هو `1100`.

# --hints--

`toBinary(5)` يجب أن تُرجع `"101"`.

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)` يجب أن تُرجع `"1100"`.

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)` يجب أن تُرجع `"110010"`.

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)` يجب أن تُرجع `"1100011"`.

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
