---
id: 68b7687dded630607aceccb3
title: "تحدٍّ 49: محلل القسم العلوي لملف CSV"
challengeType: 28
dashedName: challenge-49
---

# --description--

بالنظر إلى السطر الأول من ملف القيم المفصولة بفواصل (CSV)، أرجع مصفوفة تحتوي على العناوين.

- يحتوي السطر الأول من ملف CSV على عناوين مفصولة بفواصل.
- أزل أي مسافات بيضاء بادئة أو لاحقة من كل عنوان.

# --hints--

`getHeadings("name,age,city")` يجب أن تُرجع `["name", "age", "city"]`.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` يجب أن تُرجع `["first name", "last name", "phone"]`.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` يجب أن تُرجع `["username", "email", "signup date"]`.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
