---
id: 68b7687dded630607aceccad
title: "تحدٍّ 46: ثاني أكبر قيمة"
challengeType: 28
dashedName: challenge-46
---

# --description--

بالنظر إلى مصفوفة، أعد ثاني أكبر رقم مميز.

# --hints--

`secondLargest([1, 2, 3, 4])` يجب أن تُعيد `3`.

```js
assert.equal(secondLargest([1, 2, 3, 4]), 3);
```

`secondLargest([20, 139, 94, 67, 31])` يجب أن تُعيد `94`.

```js
assert.equal(secondLargest([20, 139, 94, 67, 31]), 94);
```

`secondLargest([2, 3, 4, 6, 6])` يجب أن تُعيد `4`.

```js
assert.equal(secondLargest([2, 3, 4, 6, 6]), 4);
```

`secondLargest([10, -17, 55.5, 44, 91, 0])` يجب أن تُعيد `55.5`.

```js
assert.equal(secondLargest([10, -17, 55.5, 44, 91, 0]), 55.5);
```

`secondLargest([1, 0, -1, 0, 1, 0, -1, 1, 0])` يجب أن تُعيد `0`.

```js
assert.equal(secondLargest([1, 0, -1, 0, 1, 0, -1, 1, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function secondLargest(arr) {

  return arr;
}
```

# --solutions--

```js
function secondLargest(arr) {
  const unique = [...new Set(arr)];
  unique.sort((a, b) => b - a);
  return unique[1];
}
```
