---
id: 68b7687dded630607aceccab
title: "تحدٍّ 45: المربع الكامل"
challengeType: 28
dashedName: challenge-45
---

# --description--

أعطِ عددًا صحيحًا، حدّد إذا كان مربعًا كاملاً.

- الرقم مربع كامل إذا استطعت الضرب بين عدد صحيح ونفسه لتحصل على هذا الرقم. على سبيل المثال، 9 هو مربع كامل لأنك تستطيع ضرب 3 بنفسه لتحصل عليه.

# --hints--

`isPerfectSquare(9)` يجب أن تُرجع `true`.

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` يجب أن تُرجع `true`.

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` يجب أن تُرجع `true`.

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` يجب أن تُرجع `false`.

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` يجب أن تُرجع `false`.

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` يجب أن تُرجع `false`.

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` يجب أن تُرجع `true`.

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` يجب أن تُرجع `true`.

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
