---
id: 68b7cadffed0e75a517da673
title: "تحدٍّ 52: التحويل من الثنائي إلى العشري"
challengeType: 28
dashedName: challenge-52
---

# --description--

أعطِ سلسلة نصية تمثّل رقمًا ثنائيًا، أرجع ما يعادله في النظام العشري كرقم.

يستخدم الرقم الثنائي فقط الرقمين `0` و `1` لتمثيل أي رقم. لتحويل الرقم الثنائي إلى عشري، اضرب كل رقم في قوة من `2` واجمع النتائج معًا. ابدأ بضرب الرقم الأيمن في `2^0`، والرقم الذي يليه إلى اليسار في `2^1`، وهكذا. بعد ضرب كل الأرقام في قوة من `2`، اجمع النتيجة معًا.

على سبيل المثال، الرقم الثنائي `101` يساوي `5` في النظام العشري لأن:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

يجب أن تُرجع الدالة `toDecimal("101")` القيمة `5`.

```js
assert.equal(toDecimal("101"), 5);
```

يجب أن تُرجع الدالة `toDecimal("1010")` القيمة `10`.

```js
assert.equal(toDecimal("1010"), 10);
```

يجب أن تُرجع الدالة `toDecimal("10010")` القيمة `18`.

```js
assert.equal(toDecimal("10010"), 18);
```

يجب أن تُرجع الدالة `toDecimal("1010101")` القيمة `85`.

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
