---
id: 68c1a929005bf54d342aa8d6
title: "تحدٍّ 59: أسبوع الفضاء اليوم 5: منطقة جولدي لوكس"
challengeType: 28
dashedName: challenge-59
---

# --description--

في اليوم الخامس من أسبوع الفضاء، ستحسب "منطقة جولدي لوكس" لنجم - المنطقة حول النجم حيث تكون الظروف "مناسبة تمامًا" لوجود الماء السائل.

معطى كتلة نجم، أرجع مصفوفة تحتوي على مسافات بداية ونهاية منطقة جولدي لوكس بوحدات فلكية.

لحساب منطقة جولدي لوكس:

1. احسب لمعان النجم برفع كتلته إلى القوة 3.5.
2. بداية المنطقة تساوي 0.95 مضروبة في الجذر التربيعي للمعان النجم.
3. نهاية المنطقة تساوي 1.37 مضروبة في الجذر التربيعي للمعان النجم.

- أرجع المسافات بعد تقريبها إلى منزلتين عشريتين.

على سبيل المثال، معطى `1` ككتلة، أرجع `[0.95, 1.37]`.

# --hints--

`goldilocksZone(1)` يجب أن تُرجع `[0.95, 1.37]`.

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` يجب أن تُرجع `[0.28, 0.41]`.

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` يجب أن تُرجع `[21.85, 31.51]`.

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` يجب أن تُرجع `[9.38, 13.52]`.

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` يجب أن تُرجع `[179.69, 259.13]`.

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
