---
id: 68c1a929005bf54d342aa8d2
title: "تحدٍّ 55: أسبوع الفضاء اليوم 1: تصنيف النجوم"
challengeType: 28
dashedName: challenge-55
---

# --description--

يوافق 4 أكتوبر بداية أسبوع الفضاء العالمي. الأيام السبعة القادمة ستقدم لك تحديات برمجية ذات طابع فلكي.

في تحدي اليوم، يُعطى لك درجة حرارة سطح نجم بالكلفن (K) وتحتاج إلى تحديد تصنيفه النجمي بناءً على النطاقات التالية:

- `"O"`: 30,000 K أو أعلى
- `"B"`: من 10,000 K إلى 29,999 K
- `"A"`: من 7,500 K إلى 9,999 K
- `"F"`: من 6,000 K إلى 7,499 K
- `"G"`: من 5,200 K إلى 5,999 K
- `"K"`: من 3,700 K إلى 5,199 K
- `"M"`: من 0 K إلى 3,699 K

- أعد تصنيف النجم المعطى.

# --hints--

`classification(5778)` يجب أن تُعيد `"G"`.

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` يجب أن تُعيد `"M"`.

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` يجب أن تُعيد `"A"`.

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` يجب أن تُعيد `"K"`.

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` يجب أن تُعيد `"M"`.

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` يجب أن تُعيد `"O"`.

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` يجب أن تُعيد `"F"`.

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` يجب أن تُعيد `"B"`.

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
