---
id: 68c1a929005bf54d342aa8d4
title: "تحدٍّ 57: أسبوع الفضاء اليوم 3: اتصل بالمنزل"
challengeType: 28
dashedName: challenge-57
---

# --description--

في اليوم الثالث من أسبوع الفضاء، تحصل على مصفوفة من الأرقام تمثل المسافات (بالكيلومترات) بينك وبين الأقمار الصناعية وكوكبك الأصلي في مسار الاتصال. حدّد كم من الوقت سيستغرق وصول رسالة تُرسل عبر المسار إلى كوكب الوجهة باستخدام القيود التالية:

- القيمة الأولى في المصفوفة هي المسافة من موقعك إلى القمر الصناعي الأول.
- كل قيمة تالية، ما عدا الأخيرة، هي المسافة إلى القمر الصناعي التالي.
- القيمة الأخيرة في المصفوفة هي المسافة من القمر الصناعي السابق إلى كوكبك الأصلي.
- تنتقل الرسالة بسرعة 300,000 كم/ثانية.
- كل قمر صناعي تمر الرسالة **من خلاله** يضيف تأخيرًا في الإرسال قدره 0.5 ثانية.
- أرجع رقمًا مقربًا إلى 4 منازل عشرية، مع إزالة الأصفار الزائدة في النهاية.

# --hints--

`sendMessage([300000, 300000])` يجب أن تُرجع `2.5`.

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` يجب أن تُرجع `3.0627`.

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` يجب أن تُرجع `364.5`.

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` يجب أن تُرجع `1674.3333`.

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` يجب أن تُرجع `2.4086`.

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` يجب أن تُرجع `21.1597`.

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
