---
id: 68c1a929005bf54d342aa8d3
title: "تحدٍّ 56: أسبوع الفضاء اليوم الثاني: البحث عن كواكب خارج المجموعة الشمسية"
challengeType: 28
dashedName: challenge-56
---

# --description--

في اليوم الثاني من أسبوع الفضاء، تُعطى سلسلة نصية حيث يمثل كل حرف قراءة سطوع نجم. حدّد إذا كانت القراءات قد اكتشفت كوكبًا خارج المجموعة الشمسية باستخدام طريقة العبور. طريقة العبور هي عندما يمر كوكب أمام نجم، مما يقلل السطوع المرصود له.

- قراءات السطوع تتكون فقط من الأحرف `0-9` و `A-Z` حيث تمثل كل قراءة القيم الرقمية التالية:
- الأحرف `0-9` تمثل مستويات السطوع من `0-9`.
- الأحرف `A-Z` تمثل مستويات السطوع من `10-35`.

يُعتبر النجم أنه يحتوي على كوكب خارج المجموعة الشمسية إذا كانت أي قراءة واحدة أقل من أو تساوي 80% من متوسط جميع القراءات. على سبيل المثال، إذا كان متوسط سطوع النجم 10، يُعتبر أنه يحتوي على كوكب إذا كانت أي قراءة واحدة 8 أو أقل.

# --hints--

يجب أن تُعيد الدالة `hasExoplanet("665544554")` القيمة `false`.

```js
assert.isFalse(hasExoplanet("665544554"));
```

يجب أن تُعيد الدالة `hasExoplanet("FGFFCFFGG")` القيمة `true`.

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

يجب أن تُعيد الدالة `hasExoplanet("MONOPLONOMONPLNOMPNOMP")` القيمة `false`.

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

يجب أن تُعيد الدالة `hasExoplanet("FREECODECAMP")` القيمة `true`.

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

يجب أن تُعيد الدالة `hasExoplanet("9AB98AB9BC98A")` القيمة `false`.

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

يجب أن تُعيد الدالة `hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` القيمة `true`.

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
