---
id: 68c1a929005bf54d342aa8d5
title: "تحدٍّ 58: أسبوع الفضاء اليوم الرابع: موقع الهبوط"
challengeType: 28
dashedName: challenge-58
---

# --description--

في اليوم الرابع من أسبوع الفضاء، تحصل على مصفوفة من الأرقام (مصفوفة من المصفوفات)، تمثل مواقع هبوط محتملة للمركبة الجوالة. ابحث عن موقع الهبوط الأكثر أمانًا بناءً على القواعد التالية:

- كل موقع في المصفوفة يحتوي على رقم من `0-9`، شاملة.
- أي `0` يمثل موقع هبوط محتمل.
- أي رقم غير `0` خطير جدًا للهبوط. كلما زاد الرقم، زادت الخطورة.
- الموقع الأكثر أمانًا هو الخانة التي تحتوي على `0` والتي تحيط بها خلايا (حتى 4 جيران، تجاهل الأقطار) بأقل مجموع خطر.
- تجاهل الجيران خارج الحدود (الزوايا والحواف لها جيران أقل فقط).
- أعد مؤشرات موقع الهبوط الأكثر أمانًا. سيكون هناك دائمًا موقع واحد فقط هو الأكثر أمانًا.

على سبيل المثال، إذا كانت المصفوفة:

```js
[
  [1, 0],
  [2, 0]
]
```

أعد `[0, 1]`، وهي مؤشرات الـ `0` في المصفوفة الأولى.

# --hints--

`findLandingSpot([[1, 0], [2, 0]])` يجب أن تعيد `[0, 1]`.

```js
assert.deepEqual(findLandingSpot([[1, 0], [2, 0]]), [0, 1]);
```

`findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` يجب أن تعيد `[1, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1]);
```

`findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` يجب أن تعيد `[2, 2]`.

```js
assert.deepEqual(findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2]);
```

`findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` يجب أن تعيد `[2, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1]);
```

# --seed--

## --seed-contents--

```js
function findLandingSpot(matrix) {

  return matrix;
}
```

# --solutions--

```js
function findLandingSpot(matrix) {
  let bestSpot = null;
  let lowestNeighborSum = Infinity;

  for (let i = 0; i < matrix.length; i++) {
    for (let j = 0; j < matrix[i].length; j++) {

      if (matrix[i][j] === 0) {
        let currentNeighborSum = 0;

        if (i > 0) currentNeighborSum += matrix[i - 1][j];
        if (j < matrix[i].length - 1) currentNeighborSum += matrix[i][j + 1];
        if (i < matrix.length - 1) currentNeighborSum += matrix[i + 1][j];
        if (j > 0) currentNeighborSum += matrix[i][j - 1];

        if (currentNeighborSum < lowestNeighborSum) {
          lowestNeighborSum = currentNeighborSum;
          bestSpot = [i, j];
        }
      }
    }
  }

  return bestSpot;
}
```
