---
id: 68c497f3aaefc9fd9f1b0e25
title: "تحدٍّ 61: أسبوع الفضاء اليوم 7: وقود الإطلاق"
challengeType: 28
dashedName: challenge-61
---

# --description--

في اليوم الأخير من أسبوع الفضاء، ستُعطى كتلة الحمولة بالكيلوغرام (kg) التي تريد إرسالها إلى المدار. حدّد كمية الوقود اللازمة لإرسال الحمولة إلى المدار باستخدام القواعد التالية:

- الصواريخ تحتاج إلى 1 كغ من الوقود لكل 5 كغ من الكتلة التي يجب رفعها.
- الوقود نفسه له كتلة. لذلك عند إضافة الوقود، تزداد الكتلة التي يجب رفعها، مما يتطلب مزيدًا من الوقود، وهذا يزيد الكتلة، وهكذا.
- لحساب إجمالي الوقود اللازم: ابدأ بكتلة الحمولة، احسب الوقود اللازم لها، أضف هذا الوقود إلى الكتلة الإجمالية، وأعد الحساب. كرر هذه العملية حتى تصبح كمية الوقود الإضافي المطلوبة أقل من 1 كغ، ثم توقف.
- تجاهل كتلة الصاروخ نفسه. احسب فقط الوقود اللازم لرفع الحمولة ووقودها الخاص.

على سبيل المثال، إذا كانت كتلة الحمولة 50 كغ، ستحتاج إلى 10 كغ من الوقود لرفعها (الحمولة / 5)، مما يزيد الكتلة الإجمالية إلى 60 كغ، والتي تحتاج إلى 12 كغ للرفع (2 كغ إضافية)، مما يزيد الكتلة الإجمالية إلى 62 كغ، والتي تحتاج إلى 12.4 كغ للرفع - 0.4 كغ إضافية - وهي أقل من 1 كغ إضافية، لذا نتوقف هنا. الكتلة الإجمالية للرفع هي 62.4 كغ، منها 50 كغ هي الحمولة الأصلية و12.4 كغ وقود.

- أرجع كمية الوقود اللازمة مقربة إلى منزلة عشرية واحدة.

# --hints--

`launchFuel(50)` يجب أن تُرجع `12.4`.

```js
assert.equal(launchFuel(50), 12.4);
```

`launchFuel(500)` يجب أن تُرجع `124.8`.

```js
assert.equal(launchFuel(500), 124.8);
```

`launchFuel(243)` يجب أن تُرجع `60.7`.

```js
assert.equal(launchFuel(243), 60.7);
```

`launchFuel(11000)` يجب أن تُرجع `2749.8`.

```js
assert.equal(launchFuel(11000), 2749.8);
```

`launchFuel(6214)` يجب أن تُرجع `1553.4`.

```js
assert.equal(launchFuel(6214), 1553.4);
```

# --seed--

## --seed-contents--

```js
function launchFuel(payload) {

  return payload;
}
```

# --solutions--

```js
function launchFuel(payload) {
  let totalMass = payload;
  let additionalFuel = totalMass / 5;
  let totalFuel = additionalFuel;

  while (additionalFuel >= 1) {
    totalMass += additionalFuel;
    additionalFuel = (totalMass / 5) - totalFuel;
    totalFuel += additionalFuel;
  }

  return Number(totalFuel.toFixed(1));
}
```
