---
id: 68c497f3aaefc9fd9f1b0e24
title: "تحدٍّ 60: أسبوع الفضاء اليوم 6: مرحلة القمر"
challengeType: 28
dashedName: challenge-60
---

# --description--

في اليوم السادس من أسبوع الفضاء، ستُعطى تاريخًا بصيغة `"YYYY-MM-DD"` وتحتاج إلى تحديد مرحلة القمر لذلك اليوم باستخدام القواعد التالية:

استخدم دورة قمرية مبسطة مدتها 28 يومًا، مقسمة إلى أربع مراحل متساوية:

- `"New"`: الأيام 1 - 7
- `"Waxing"`: الأيام 8 - 14
- `"Full"`: الأيام 15 - 21
- `"Waning"`: الأيام 22 - 28

بعد اليوم 28، تتكرر الدورة مع اليوم 1، وهو القمر الجديد.

- استخدم `"2000-01-06"` كمرجع للقمر الجديد (اليوم 1 من الدورة) لتحديد مرحلة اليوم المعطى.
- لن تُعطى تواريخ قبل تاريخ المرجع.
- أرجع المرحلة الصحيحة كسلسلة نصية.

**ملاحظة:** اليوم 1 يمثل يوم القمر الجديد، أي لم تمر أي أيام منذ آخر قمر جديد.

# --hints--

`moonPhase("2000-01-12")` يجب أن تُرجع `"New"`.

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` يجب أن تُرجع `"Waxing"`.

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` يجب أن تُرجع `"Full"`.

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` يجب أن تُرجع `"Waning"`.

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` يجب أن تُرجع `"New"`.

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
