---
id: 68cae5b538ff798bbd4da008
title: "تحدٍّ 70: مستخرج خاصيات HTML"
challengeType: 28
dashedName: challenge-70
---

# --description--

بالنظر إلى سلسلة نصية لعنصر HTML صالح، أعد خاصيات العنصر باستخدام المعايير التالية:

- سيُعطى عنصر واحد فقط.
- تكون الخاصيات بالشكل: `attribute="value"`.
- أعد مصفوفة من السلاسل النصية تحتوي كل خاصية وقيمتها، مفصولة بفاصلة، بالشكل التالي: `["attribute1, value1", "attribute2, value2"]`.
- أعد الخاصيات بنفس الترتيب الذي تظهر به.
- إذا لم توجد خاصيات، أعد مصفوفة فارغة.

# --hints--

`extractAttributes('<span class="red"></span>')` يجب أن تُعيد `["class, red"]`.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` يجب أن تُعيد `["charset, UTF-8"]`.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` يجب أن تُعيد `[]`.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` يجب أن تُعيد `["name, email", "type, email", "required, true"]`.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` يجب أن تُعيد `["id, submit", "class, btn btn-primary"]`.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
