---
id: 68cae5b538ff798bbd4da00a
title: "تحدٍّ 72: مُعدِّل منظم الحرارة 2"
challengeType: 28
dashedName: challenge-72
---

# --description--

معطى درجة حرارة الغرفة الحالية بالفهرنهايت ودرجة الحرارة المستهدفة بالسلسلة النصية، أعد سلسلة نصية توضح كيفية تعديل درجة حرارة الغرفة بناءً على هذه القيود:

- أعد `"Heat: X degrees Fahrenheit"` إذا كانت درجة الحرارة الحالية أقل من الهدف. حيث `X` هو عدد الدرجات بالفهرنهايت اللازمة لتسخين الغرفة للوصول إلى الهدف، مقربًا إلى منزلة عشرية واحدة.
- أعد `"Cool: X degrees Fahrenheit"` إذا كانت درجة الحرارة الحالية أعلى من الهدف. حيث `X` هو عدد الدرجات بالفهرنهايت اللازمة لتبريد الغرفة للوصول إلى الهدف، مقربًا إلى منزلة عشرية واحدة.
- أعد `"Hold"` إذا كانت درجة الحرارة الحالية مساوية للهدف.

لتحويل درجة الحرارة من سلسة نصية إلى فهرنهايت، اضرب درجة الحرارة بالسلسلة النصية في 1.8 وأضف 32 إلى النتيجة (`F = (C * 1.8) + 32`).

# --hints--

`adjustThermostat(32, 0)` يجب أن يعيد `"Hold"`.

```js
assert.equal(adjustThermostat(32, 0), "Hold");
```

`adjustThermostat(70, 25)` يجب أن يعيد `"Heat: 7.0 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(70, 25), "Heat: 7.0 degrees Fahrenheit");
```

`adjustThermostat(72, 18)` يجب أن يعيد `"Cool: 7.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(72, 18), "Cool: 7.6 degrees Fahrenheit");
```

`adjustThermostat(212, 100)` يجب أن يعيد `"Hold"`.

```js
assert.equal(adjustThermostat(212, 100), "Hold");
```

`adjustThermostat(59, 22)` يجب أن يعيد `"Heat: 12.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(59, 22), "Heat: 12.6 degrees Fahrenheit");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(currentF, targetC) {

  return currentF;
}
```

# --solutions--

```js
function adjustThermostat(currentF, targetC) {
  const targetF = targetC * 1.8 + 32;
  const diff = Math.abs(targetF - currentF).toFixed(1);

  if (currentF < targetF) {
    return `Heat: ${diff} degrees Fahrenheit`;
  } else if (currentF > targetF) {
    return `Cool: ${diff} degrees Fahrenheit`;
  } else {
    return "Hold"
  }
}
```
