---
id: 68cae5b538ff798bbd4da004
title: "تحدٍّ 66: مزيل وسوم HTML"
challengeType: 28
dashedName: challenge-66
---

# --description--

معطى سلسلة نصية من شفرة HTML، احذف الوسوم وأعد المحتوى النصي العادي.

- ستحتوي السلسلة النصية المُدخلة على HTML صالح فقط.
- قد تكون وسوم HTML مسكنة.
- احذف الوسوم وأي خاصيات.

على سبيل المثال، يجب أن تُعيد `'<a href="#">Click here</a>'` القيمة `"Click here"`.

# --hints--

يجب أن تُعيد الدالة `stripTags('<a href="#">Click here</a>')` القيمة `"Click here"`.

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

يجب أن تُعيد الدالة `stripTags('<p class="center">Hello <b>World</b>!</p>')` القيمة `"Hello World!"`.

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

يجب أن تُعيد الدالة `stripTags('<img src="cat.jpg" alt="Cat">')` سلسلة فارغة (`""`).

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

يجب أن تُعيد الدالة `stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` القيمة `sectionsection`.

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
