---
id: 68cae5b538ff798bbd4da005
title: "تحدٍّ 67: مدقق البريد الإلكتروني"
challengeType: 28
dashedName: challenge-67
---

# --description--

معطاة سلسلة نصية، حدّد إذا كانت عنوان بريد إلكتروني صالح باستخدام القيود التالية:

- يجب أن تحتوي على رمز `@` واحد فقط.
- الجزء المحلي (قبل `@`):
  - يمكن أن يحتوي فقط على حروف (`a-z`، `A-Z`)، أرقام (`0-9`)، نقاط (`.`)، شرطات سفلية (`_`)، أو شرطات (`-`).
  - لا يمكن أن يبدأ أو ينتهي بنقطة.
- الجزء الخاص بالنطاق (بعد `@`):
  - يجب أن يحتوي على نقطة واحدة على الأقل.
  - يجب أن ينتهي بنقطة تليها حرفان على الأقل.
- لا يمكن أن يحتوي الجزء المحلي أو جزء النطاق على نقطتين متتاليتين.

# --hints--

يجب أن تُرجع الدالة `validate("a@b.cd")` القيمة `true`.

```js
assert.isTrue(validate("a@b.cd"));
```

يجب أن تُرجع الدالة `validate("hell.-w.rld@example.com")` القيمة `true`.

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

يجب أن تُرجع الدالة `validate(".b@sh.rc")` القيمة `false`.

```js
assert.isFalse(validate(".b@sh.rc"));
```

يجب أن تُرجع الدالة `validate("example@test.c0")` القيمة `false`.

```js
assert.isFalse(validate("example@test.c0"));
```

يجب أن تُرجع الدالة `validate("freecodecamp.org")` القيمة `false`.

```js
assert.isFalse(validate("freecodecamp.org"));
```

يجب أن تُرجع الدالة `validate("develop.ment_user@c0D!NG.R.CKS")` القيمة `true`.

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

يجب أن تُرجع الدالة `validate("hello.@wo.rld")` القيمة `false`.

```js
assert.isFalse(validate("hello.@wo.rld"));
```

يجب أن تُرجع الدالة `validate("hello@world..com")` القيمة `false`.

```js
assert.isFalse(validate("hello@world..com"));
```

يجب أن تُرجع الدالة `validate("develop..ment_user@c0D!NG.R.CKS")` القيمة `false`.

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

يجب أن تُرجع الدالة `validate("git@commit@push.io")` القيمة `false`.

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
