---
id: 68cae5b538ff798bbd4da002
title: "تحدٍّ 64: من 24 إلى 12"
challengeType: 28
dashedName: challenge-64
---

# --description--

أعطِ سلسلة نصية تمثل وقتًا من اليوم بصيغة 24 ساعة `"HHMM"`، أرجع الوقت بصيغته المكافئة في 12 ساعة `"H:MM AM"` أو `"H:MM PM"`.

- سيكون المُدخَل دائمًا سلسلة نصية مكونة من أربعة أرقام بصيغة 24 ساعة، من `"0000"` إلى `"2359"`.

# --hints--

`to12("1124")` يجب أن تُرجع `"11:24 AM"`.

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")` يجب أن تُرجع `"9:00 AM"`.

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")` يجب أن تُرجع `"2:55 PM"`.

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")` يجب أن تُرجع `"11:46 PM"`.

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")` يجب أن تُرجع `"12:30 AM"`.

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
