---
id: 68cae5b538ff798bbd4da009
title: "تحدٍّ 71: حاسبة البقشيش"
challengeType: 28
dashedName: challenge-71
---

# --description--

بما أن لديك سعر وجبتك ونسبة بقشيش مخصصة، أرجع مصفوفة تحتوي على ثلاثة قيم للبقشيش؛ 15%، 20%، والمبلغ المخصص.

- ستُعطى الأسعار بالتنسيق: `"$N.NN"`.
- ستُعطى نسب البقشيش المخصصة بهذا التنسيق: `"25%"`.
- أرجع المبالغ بنفس تنسيق `"$N.NN"`، مع تقريبها إلى منزلتين عشريتين.

على سبيل المثال، إذا كان سعر الوجبة `"$10.00"`، وقيمة البقشيش المخصصة `"25%"`، أرجع `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculateTips("$10.00", "25%")` يجب أن تُرجع `["$1.50", "$2.00", "$2.50"]`.

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` يجب أن تُرجع `["$13.45", "$17.93", "$23.31"]`.

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` يجب أن تُرجع `["$2.98", "$3.97", "$1.79"]`.

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
