---
id: 68cae5b538ff798bbd4da003
title: "تحدٍّ 65: عد السلاسل النصية"
challengeType: 28
dashedName: challenge-65
---

# --description--

معطيتان سلسلتا نص، حدّد كم مرة تظهر السلسلة النصية الثانية في الأولى.

- يمكن أن تتداخل سلسلة النمط داخل السلسلة الأولى. على سبيل المثال، تحتوي `"aaa"` على `"aa"` مرتين. الحرفان الأولان `a` والحرفان الثانيان.

# --hints--

`count('abcdefg', 'def')` يجب أن تُرجع `1`.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` يجب أن تُرجع `0`.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` يجب أن تُرجع `2`.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` يجب أن تُرجع `3`.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` يجب أن تُرجع `10`.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
