---
id: 68cae5b538ff798bbd4da001
title: "تحدٍّ 63: معركة الكلمات"
challengeType: 28
dashedName: challenge-63
---

# --description--

معطى جملتان تمثلان فريقك وفريقًا معارضًا، حيث تتقاتل كل كلمة من فريقك مع الكلمة المقابلة لها من الفريق المعارض، حدّد الفريق الفائز باستخدام القواعد التالية:

- الجمل المعطاة تحتوي دائمًا على نفس عدد الكلمات.
- الكلمات مفصولة بمسافة واحدة فقط وتحتوي على حروف فقط.
- قيمة كل كلمة هي مجموع قيم حروفها.
- الحروف من `a` إلى `z` تمثل القيم من `1` إلى `26`. على سبيل المثال، `a` قيمتها `1`، و`z` قيمتها `26`.
- الحرف الكبير يضاعف قيمة الحرف. على سبيل المثال، `A` قيمته `2`، و`Z` قيمته `52`.
- الكلمات تتقاتل بالترتيب: الكلمة الأولى من فريقك تتقاتل مع الكلمة الأولى من الفريق المعارض، وهكذا.
- تفوز الكلمة إذا كانت قيمتها أكبر من قيمة الكلمة المقابلة.
- الفريق الذي يملك عددًا أكبر من الكلمات الفائزة هو الفائز.

أرجع `"We win"` إذا كان فريقك هو الفائز، و`"We lose"` إذا خسر فريقك، و`"Draw"` إذا كان عدد الكلمات الفائزة متساويًا بين الفريقين.

# --hints--

`battle("hello world", "hello word")` يجب أن تُرجع `"We win"`.

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")` يجب أن تُرجع `"We win"`.

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")` يجب أن تُرجع `"We lose"`.

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")` يجب أن تُرجع `"Draw"`.

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")` يجب أن تُرجع `"We win"`.

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` يجب أن تُرجع `"We lose"`.

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")` يجب أن تُرجع `"Draw"`.

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
