---
id: 68d2ba1468508398389487cf
title: "تحدٍّ 74: الأغاني المفضلة"
challengeType: 28
dashedName: challenge-74
---

# --description--

هل تتذكر أجهزة iPod؟ صدر أول نموذج قبل 24 سنة اليوم، في 23 أكتوبر 2001.

معطى مصفوفة من كائنات الأغاني التي تمثل قائمة تشغيل iPod الخاصة بك، أرجع مصفوفة بعناوين أكثر أغنيتين تشغيلًا، مع وضع الأغنية الأكثر تشغيلًا أولًا.

- كل كائن يحتوي على خاصية `"title"` (سلسلة نصية)، وخاصية `"plays"` (عدد صحيح).

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` يجب أن تُرجع `["Sync or Swim", "Earbud Blues"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` يجب أن تُرجع `["Clickwheel Love", "99 Downloads"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` يجب أن تُرجع `["Song B", "Song C"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```
