---
id: 68d2ba1468508398389487ce
title: "تحدٍّ 73: تحدث بحكمة، يجب عليك"
challengeType: 28
dashedName: challenge-73
---

# --description--

بالنظر إلى جملة، أعد إصدارًا منها يبدو كنصيحة من معلم حكيم باستخدام القواعد التالية:

- تُفصل الكلمات بمسافة واحدة.
- ابحث عن أول ظهور لإحدى الكلمات التالية في الجملة: `"have"`، `"must"`، `"are"`، `"will"`، `"can"`.
- انقل كل الكلمات التي تسبق هذه الكلمة وتشملها إلى نهاية الجملة و:
  - حافظ على ترتيب الكلمات عند نقلها.
  - اجعلها كلها أحرف صغيرة.
  - وأضف فاصلة ومسافة قبلها.
- اجعل الحرف الأول من الكلمة الجديدة الأولى في الجملة كبيرًا.
- جميع الجمل المعطاة تنتهي بعلامة ترقيم واحدة. احتفظ بعلامة الترقيم الأصلية وانقلها إلى نهاية الجملة الجديدة.
- أعد الجملة الجديدة، مع التأكد من وجود مسافة واحدة بين كل كلمة وأخرى وعدم وجود مسافات في بداية أو نهاية الجملة.

على سبيل المثال، إذا كانت الجملة `"You must speak wisely."` فأعد `"Speak wisely, you must."`

# --hints--

`wiseSpeak("You must speak wisely.")` يجب أن تُعيد `"Speak wisely, you must."`

```js
assert.equal(wiseSpeak("You must speak wisely."), "Speak wisely, you must.");
```

`wiseSpeak("You can do it!")` يجب أن تُعيد `"Do it, you can!"`

```js
assert.equal(wiseSpeak("You can do it!"), "Do it, you can!");
```

`wiseSpeak("Do you think you will complete this?")` يجب أن تُعيد `"Complete this, do you think you will?"`

```js
assert.equal(wiseSpeak("Do you think you will complete this?"), "Complete this, do you think you will?");
```

`wiseSpeak("All your base are belong to us.")` يجب أن تُعيد `"Belong to us, all your base are."`

```js
assert.equal(wiseSpeak("All your base are belong to us."), "Belong to us, all your base are.");
```

`wiseSpeak("You have much to learn.")` يجب أن تُعيد `"Much to learn, you have."`

```js
assert.equal(wiseSpeak("You have much to learn."), "Much to learn, you have.");
```

# --seed--

## --seed-contents--

```js
function wiseSpeak(sentence) {

  return sentence;
}
```

# --solutions--

```js
function wiseSpeak(sentence) {
  const triggers = ["have", "must", "are", "will", "can"];
  const punctuation = sentence[sentence.length - 1];
  const words = sentence.split(" ");

  const triggerIndex = words.findIndex(w => triggers.includes(w));

  const toMove = words.slice(0, triggerIndex + 1).map(w => w.toLowerCase());
  const remaining = words.slice(triggerIndex + 1);

  const newStart = remaining.map((w, i) => {
    return i === 0 ? w[0].toUpperCase() + w.slice(1) :
      i === remaining.length - 1 ? w.slice(0, w.length - 1) : w;
  })

  return newStart.join(" ") + ", " + toMove.join(" ") + punctuation;
}
```
