---
id: 68d2ba1468508398389487d0
title: "تحدٍّ 75: الكنز المخفي"
challengeType: 28
dashedName: challenge-75
---

# --description--

معطى مصفوفة ثنائية الأبعاد تمثل خريطة لقاع المحيط تحتوي على كنز مخفي، ومصفوفة بالإحداثيات (`[row, column]`) للغوص التالي في بحثك عن الكنز، أرجع `"Empty"` أو `"Found"` أو `"Recovered"` باستخدام القواعد التالية:

- ستحتوي المصفوفة الثنائية الأبعاد المعطاة على كنز واحد غير مسترد بالضبط، يشغل عدة خلايا.
- كل خلية في المصفوفة تحتوي على واحدة من القيم التالية:
  - `"-"`: لا يوجد كنز.
  - `"O"`: جزء من الكنز لم يُعثر عليه بعد.
  - `"X"`: جزء من الكنز تم العثور عليه بالفعل.

- إذا لم يكن في موقع الغوص كنز، أرجع `"Empty"`.
- إذا وجد موقع الغوص كنزًا، لكن لا يزال جزء آخر من الكنز غير مكتشف، أرجع `"Found"`.
- إذا وجد موقع الغوص آخر جزء غير مكتشف من الكنز، أرجع `"Recovered"`.

على سبيل المثال، معطى:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

و `[2, 1]` كإحداثيات موقع الغوص، أرجع `"Recovered"` لأن الغوص وجد آخر جزء غير مكتشف من الكنز.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` يجب أن تُرجع `"Recovered"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` يجب أن تُرجع `"Empty"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` يجب أن تُرجع `"Found"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` يجب أن تُرجع `"Found"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` يجب أن تُرجع `"Recovered"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` يجب أن تُرجع `"Empty"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
