---
id: 68d30845cc08266018fc46bc
title: "تحدٍّ 76: الحمض النووي التكميلي"
challengeType: 28
dashedName: challenge-76
---

# --description--

أعطِ سلسلة نصية تمثل متتالية DNA، أعد السلسلة التكاملية لها باستخدام القواعد التالية:

- يتكون DNA من الحروف `"A"`، `"C"`، `"G"`، و`"T"`.
- الحرفان `"A"` و`"T"` متكاملان.
- الحرفان `"C"` و`"G"` متكاملان.

على سبيل المثال، إذا كانت السلسلة `"ACGT"`، أعد `"TGCA"`.

# --hints--

يجب أن تُعيد الدالة `complementaryDNA("ACGT")` القيمة `"TGCA"`.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

يجب أن تُعيد الدالة `complementaryDNA("ATGCGTACGTTAGC")` القيمة `"TACGCATGCAATCG"`.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

يجب أن تُعيد الدالة `complementaryDNA("GGCTTACGATCGAAG")` القيمة `"CCGAATGCTAGCTTC"`.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

يجب أن تُعيد الدالة `complementaryDNA("GATCTAGCTAGGCTAGCTAG")` القيمة `"CTAGATCGATCCGATCGATC"`.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
