---
id: 68d30845cc08266018fc46be
title: "تحدٍّ 78: متتالية أعداد صحيحة"
challengeType: 28
dashedName: challenge-78
---

# --description--

أعطِ عددًا صحيحًا موجبًا، أرجع سلسلة نصية تحتوي كل الأعداد الصحيحة من `1` حتى العدد المعطى، بما في ذلك هذا العدد، بترتيب رقمي.

على سبيل المثال، إذا كان العدد `5`، أرجع `"12345"`.

# --hints--

يجب أن تُرجع الدالة `sequence(5)` القيمة `"12345"`.

```js
assert.equal(sequence(5), "12345");
```

يجب أن تُرجع الدالة `sequence(10)` القيمة `"12345678910"`.

```js
assert.equal(sequence(10), "12345678910");
```

يجب أن تُرجع الدالة `sequence(1)` القيمة `"1"`.

```js
assert.strictEqual(sequence(1), "1");
```

يجب أن تُرجع الدالة `sequence(27)` القيمة `"123456789101112131415161718192021222324252627"`.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
