---
id: 68d30fc57588d97fd3027b30
title: "تحدٍّ 79: Navigator"
challengeType: 28
dashedName: challenge-79
---

# --description--

في 28 أكتوبر 1994، أُصدر Netscape Navigator، مما ساعد ملايين الأشخاص على استكشاف الويب المبكر.

معطى مصفوفة من أوامر المتصفح التي نفذتها على Netscape Navigator، أعد الصفحة الحالية التي تتصفحها بعد تنفيذ كل الأوامر باستخدام القواعد التالية:

- تبدأ دائمًا على صفحة `"Home"`، التي لن تكون مدرجة في مصفوفة الأوامر.
- الأوامر الصالحة هي:
  - `"Visit Page"`: حيث `"Page"` هو اسم الصفحة التي تزورها. على سبيل المثال، يأخذك الأمر `"Visit About"` إلى صفحة `"About"`. عند زيارة صفحة جديدة، تأكد من حذف أي سجل للتقدم الأمامي لديك.
  - `"Back"`: يعيدك إلى الصفحة السابقة في سجل التصفح أو يبقيك على الصفحة الحالية إذا لم تكن هناك صفحة سابقة.
  - `"Forward"`: ينقلك إلى الأمام في السجل إلى الصفحة التي جئت منها أو يبقيك على الصفحة الحالية إذا لم تكن هناك صفحة للأمام.

على سبيل المثال، معطى `["Visit About Us", "Back", "Forward"]`، أعد `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` يجب أن يعيد `"About Us"`.

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` يجب أن يعيد `"Home"`.

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` يجب أن يعيد `"Home"`.

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` يجب أن يعيد `"Gallery"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` يجب أن يعيد `"Home"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` يجب أن يعيد `"Contact"`.

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` يجب أن يعيد `"Visit Us"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```
