---
id: 68d30845cc08266018fc46bd
title: "تحدٍّ 77: منسق المدة الزمنية"
challengeType: 28
dashedName: challenge-77
---

# --description--

بالنظر إلى عدد صحيح من الثواني، أعد سلسلة نصية تمثل نفس المدة بالتنسيق `"H:MM:SS"`، حيث `"H"` هو عدد الساعات، و`"MM"` هو عدد الدقائق، و`"SS"` هو عدد الثواني. أعد الوقت باستخدام القواعد التالية:

- الثواني: يجب أن تكون دائمًا رقمين.
- الدقائق: يجب حذف الأصفار البادئة عندما لا تكون ضرورية. استخدم `"0"` إذا كانت المدة أقل من دقيقة واحدة.
- الساعات: يجب تضمينها فقط إذا كانت أكبر من صفر.

# --hints--

`format(500)` يجب أن يعيد `"8:20"`.

```js
assert.equal(format(500), "8:20");
```

`format(4000)` يجب أن يعيد `"1:06:40"`.

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` يجب أن يعيد `"0:01"`.

```js
assert.equal(format(1), "0:01");
```

`format(5555)` يجب أن يعيد `"1:32:35"`.

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` يجب أن يعيد `"27:46:39"`.

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
