---
id: 68e39ed6106dac2f0a98fd63
title: "تحدٍّ 81: العدد الأولي الـ nth"
challengeType: 28
dashedName: challenge-81
---

# --description--

العدد الأولي هو عدد صحيح موجب أكبر من 1 لا يقبل القسمة إلا على 1 وعلى نفسه فقط. أول خمسة أعداد أولية هي `2`، `3`، `5`، `7`، و`11`.

معطى عدد صحيح موجب `n`، أرجع العدد الأولي الـ `n`. على سبيل المثال، إذا كان `5` فأرجع العدد الأولي الخامس: `11`.

# --hints--

يجب أن تُرجع الدالة `nthPrime(5)` القيمة `11`.

```js
assert.equal(nthPrime(5), 11);
```

يجب أن تُرجع الدالة `nthPrime(10)` القيمة `29`.

```js
assert.equal(nthPrime(10), 29);
```

يجب أن تُرجع الدالة `nthPrime(16)` القيمة `53`.

```js
assert.equal(nthPrime(16), 53);
```

يجب أن تُرجع الدالة `nthPrime(99)` القيمة `523`.

```js
assert.equal(nthPrime(99), 523);
```

يجب أن تُرجع الدالة `nthPrime(1000)` القيمة `7919`.

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
