---
id: 68e39ed6106dac2f0a98fd64
title: "تحدٍّ 82: SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

معطاة سلسلة نصية تمثل اسم متغير، حوّلها إلى "spooky case" باستخدام القيود التالية:

- استبدل كل الشرطات السفلية (`_`) والواصلات (`-`) بمد (`~`).
- اجعل الحرف الأول من السلسلة وكل حرف ثاني بعده بحروف كبيرة. تجاهل حرف المد عند العد. اجعل بقية الحروف أحرفًا صغيرة.

على سبيل المثال، إذا كانت السلسلة `hello_world`، فأرجع `HeLlO~wOrLd`.

# --hints--

`spookify("hello_world")` يجب أن تُرجع `"HeLlO~wOrLd"`.

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")` يجب أن تُرجع `"SpOoKy~CaSe"`.

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")` يجب أن تُرجع `"TrIcK~oR~tReAt"`.

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")` يجب أن تُرجع `"C~a~N~d~Y~~b~O~w~L"`.

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` يجب أن تُرجع `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
