---
id: 68e39ed6106dac2f0a98fd65
title: "تحدٍّ 83: التحقق من التوقيع"
challengeType: 28
dashedName: challenge-83
---

# --description--

معطاة سلسلة نصية للرسالة، وسلسلة نصية للمفتاح السري، ورقم التوقيع، حدّد إذا كان التوقيع صحيحًا باستخدام هذه الطريقة في الترميز:

- الحروف في الرسالة والمفتاح السري لها القيم التالية:
  - الحروف من `a` إلى `z` لها القيم من `1` إلى `26` على التوالي.
  - الحروف من `A` إلى `Z` لها القيم من `27` إلى `52` على التوالي.
- جميع الأحرف الأخرى لا قيمة لها.
- احسب التوقيع بجمع مجموع قيم الرسالة مع مجموع قيم المفتاح السري.

على سبيل المثال، معطاة الرسالة `"foo"` والمفتاح السري `"bar"`، سيكون التوقيع `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

تحقق إذا كان التوقيع المحسوب يطابق التوقيع المقدم.

# --hints--

`verify("foo", "bar", 57)` يجب أن تُرجع `true`.

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` يجب أن تُرجع `false`.

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` يجب أن تُرجع `true`.

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` يجب أن تُرجع `false`.

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` يجب أن تُرجع `true`.

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` يجب أن تُرجع `true`.

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
