---
id: 68e39ed6106dac2f0a98fd62
title: "تحدٍّ 80: مُرتّب البريد الإلكتروني"
challengeType: 28
dashedName: challenge-80
---

# --description--

في 29 أكتوبر 1971، أُرسل أول بريد إلكتروني على الإطلاق، مُقدّمًا صيغة `username@domain` التي نستخدمها حتى الآن. الآن، هناك مليارات من عناوين البريد الإلكتروني.

في هذا التحدٍّ، تحصل على قائمة بعناوين البريد الإلكتروني وتحتاج إلى فرزها أبجديًا حسب اسم النطاق أولًا (الجزء بعد `@`)، ثم حسب اسم المستخدم ثانيًا (الجزء قبل `@`).

- يجب أن يكون الفرز غير حساس لحالة الأحرف.
- إذا كان هناك أكثر من بريد إلكتروني بنفس النطاق، فرزهم حسب اسم المستخدم.
- أرجع مصفوفة بالعناوين المرتبة.
- يجب أن تحتفظ العناوين المرجعة بحالتها الأصلية.

على سبيل المثال، إذا أعطيت `["jill@mail.com", "john@example.com", "jane@example.com"]`، أرجع `["jane@example.com", "john@example.com", "jill@mail.com"]`.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` يجب أن تُرجع `["jane@example.com", "john@example.com", "jill@mail.com"]`.

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` يجب أن تُرجع `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`.

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` يجب أن تُرجع `["user@x.com", "user@y.com", "user@z.com"]`.

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` يجب أن تُرجع `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`.

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` يجب أن تُرجع `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`.

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```
