---
id: 68ee9e3066cfd4eb2328e8a5
title: "تحدٍّ 86: بحث الصور"
challengeType: 28
dashedName: challenge-86
---

# --description--

في 4 نوفمبر 2001، أطلقت Google خدمة بحث الصور، مما سمح للناس بالعثور على الصور باستخدام مصطلحات البحث. في هذا التحدٍّ، ستقلد وظيفة بحث الصور.

معطى مصفوفة من أسماء الصور ومصطلح بحث، أعد مصفوفة بأسماء الصور التي تحتوي على مصطلح البحث.

- تجاهل حالة الأحرف عند مطابقة مصطلحات البحث.
- أعد الصور بنفس الترتيب الذي تظهر به في مصفوفة المُدخَل.

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` يجب أن تُعيد `["dog.png"]`.

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` يجب أن تُعيد `["Sunset.jpg", "sunflower.jpeg"]`.

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` يجب أن تُعيد `["Moon.png", "stars.png"]`.

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` يجب أن تُعيد `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`.

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```
