---
id: 68ee9e3066cfd4eb2328e8a7
title: "تحدٍّ 88: محدد يوم الأسبوع"
challengeType: 28
dashedName: challenge-88
---

# --description--

أعطِ تاريخًا كسلسلة نصية بالصيغة `YYYY-MM-DD`، وأرجع يوم الأسبوع.

أيام الإرجاع الصحيحة هي:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

تأكد من تجاهل المناطق الزمنية.

# --hints--

`getWeekday("2025-11-06")` يجب أن تُرجع `Thursday`.

```js
assert.equal(getWeekday("2025-11-06"), "Thursday");
```

`getWeekday("1999-12-31")` يجب أن تُرجع `Friday`.

```js
assert.equal(getWeekday("1999-12-31"), "Friday");
```

`getWeekday("1111-11-11")` يجب أن تُرجع `Saturday`.

```js
assert.equal(getWeekday("1111-11-11"), "Saturday");
```

`getWeekday("2112-12-21")` يجب أن تُرجع `Wednesday`.

```js
assert.equal(getWeekday("2112-12-21"), "Wednesday");
```

`getWeekday("2345-10-01")` يجب أن تُرجع `Monday`.

```js
assert.equal(getWeekday("2345-10-01"), "Monday");
```

# --seed--

## --seed-contents--

```js
function getWeekday(dateString) {

  return dateString;
}
```

# --solutions--

```js
function getWeekday(dateString) {
  const days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"];
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  return days[date.getUTCDay()];
}
```
