---
id: 68ee9e3066cfd4eb2328e8a4
title: "تحدٍّ 85: عداد الكلمات"
challengeType: 28
dashedName: challenge-85
---

# --description--

أعطِ سلسلة نصية تمثل جملة، وأرجع عدد الكلمات الموجودة في الجملة.

- الكلمات هي أي متتالية من الأحرف غير الفراغية وتُفصل بمسافة واحدة.

# --hints--

`countWords("Hello world")` يجب أن تُرجع `2`.

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` يجب أن تُرجع `9`.

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` يجب أن تُرجع `4`.

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` يجب أن تُرجع `7`.

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` يجب أن تُرجع `7`.

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
