---
id: 68ee9e3066cfd4eb2328e8a8
title: "تحدٍّ 89: عدّ الأوراق"
challengeType: 28
dashedName: challenge-89
---

# --description--

تحتوي مجموعة أوراق اللعب القياسية على 13 ورقة فريدة في كل نوع. معطى عدد صحيح يمثل عدد الأوراق التي ستختارها من المجموعة، أرجع عدد التراكيب الفريدة للأوراق التي يمكنك اختيارها.

- الترتيب لا يهم. اختيار الورقة A ثم الورقة B هو نفسه اختيار الورقة B ثم الورقة A.

على سبيل المثال، إذا كان المعطى `52`، أرجع `1`. هناك تركيبة واحدة فقط من 52 ورقة يمكن اختيارها من مجموعة أوراق مكونة من 52 ورقة. وإذا كان المعطى `2`، أرجع `1326`، هناك 1326 تركيبة من الأوراق يمكنك الحصول عليها عند اختيار ورقتين من المجموعة.

# --hints--

`combinations(52)` يجب أن تُرجع `1`.

```js
assert.equal(combinations(52), 1);
```

`combinations(1)` يجب أن تُرجع `52`.

```js
assert.equal(combinations(1), 52);
```

`combinations(2)` يجب أن تُرجع `1326`.

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)` يجب أن تُرجع `2598960`.

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)` يجب أن تُرجع `15820024220`.

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)` يجب أن تُرجع `1326`.

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
