---
id: 68f6587287ad1f4ad39b0c84
title: "تحدٍّ 98: عد المستطيلات"
challengeType: 28
dashedName: challenge-98
---

# --description--

معطى عددان صحيحان موجبان يمثلان عرض وارتفاع مستطيل، حدّد عدد المستطيلات التي يمكن أن تتسع داخل المستطيل المعطى.

- عد فقط المستطيلات التي يكون عرضها وارتفاعها أعدادًا صحيحة.

على سبيل المثال، إذا كان المعطى `1` و `3`، أرجع `6`. ثلاثة مستطيلات 1x1، مستطيلان 1x2، ومستطيل واحد 1x3.

# --hints--

يجب أن تُرجع الدالة `countRectangles(1, 3)` القيمة `6`.

```js
assert.equal(countRectangles(1, 3), 6);
```

يجب أن تُرجع الدالة `countRectangles(3, 2)` القيمة `18`.

```js
assert.equal(countRectangles(3, 2), 18);
```

يجب أن تُرجع الدالة `countRectangles(1, 2)` القيمة `3`.

```js
assert.equal(countRectangles(1, 2), 3);
```

يجب أن تُرجع الدالة `countRectangles(5, 4)` القيمة `150`.

```js
assert.equal(countRectangles(5, 4), 150);
```

يجب أن تُرجع الدالة `countRectangles(11, 19)` القيمة `12540`.

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
