---
id: 68f6587287ad1f4ad39b0c7d
title: "تحدٍّ 91: البحث عن كلمة"
challengeType: 28
dashedName: challenge-91
---

# --description--

معطى مصفوفة (مصفوفة من المصفوفات) من أحرف مفردة وكلمة للعثور عليها، أرجع مؤشرات البداية والنهاية للكلمة في المصفوفة.

- المصفوفة المعطاة ستكون مملوءة بأحرف صغيرة (`a-z`).
- الكلمة التي تبحث عنها ستكون موجودة في المصفوفة مرة واحدة بالضبط.
- الكلمة التي تبحث عنها ستكون دائمًا على خط مستقيم في أحد الاتجاهات التالية:
  - من اليسار إلى اليمين
  - من اليمين إلى اليسار
  - من الأعلى إلى الأسفل
  - من الأسفل إلى الأعلى

على سبيل المثال، معطى المصفوفة:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

والكلمة `"cat"`، أرجع:

```md
[[0, 1], [2, 1]]
```

حيث أن `[0, 1]` هي مؤشرات الحرف `"c"` (بداية الكلمة)، و`[2, 1]` هي مؤشرات الحرف `"t"` (نهاية الكلمة).

# --hints--

`findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` يجب أن تُرجع `[[0, 1], [2, 1]]`.

```js
assert.deepEqual(findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]]);
```

`findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` يجب أن تُرجع `[[0, 0], [0, 2]]`.

```js
assert.deepEqual(findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]]);
```

`findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` يجب أن تُرجع `[[3, 3], [0, 3]]`.

```js
assert.deepEqual(findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]]);
```

`findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` يجب أن تُرجع `[[1, 3], [1, 1]]`.

```js
assert.deepEqual(findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]]);
```

# --seed--

## --seed-contents--

```js
function findWord(matrix, word) {

  return matrix;
}
```

# --solutions--

```js
function findWord(matrix, word) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const len = word.length;

  const directions = [
    [0, 1],
    [0, -1],
    [1, 0],
    [-1, 0]
  ];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      for (let [dr, dc] of directions) {
        let match = true;
        for (let i = 0; i < len; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols || matrix[nr][nc] !== word[i]) {
            match = false;
            break;
          }
        }
        if (match) {
          return [
            [r, c],
            [r + dr * (len - 1), c + dc * (len - 1)]
          ];
        }
      }
    }
  }
}
```
