---
id: 68f6587287ad1f4ad39b0c82
title: "تحدٍّ 96: هل هو عطلة نهاية الأسبوع؟"
challengeType: 28
dashedName: challenge-96
---

# --description--

بالنظر إلى تاريخ بالصيغة `"YYYY-MM-DD"`، أعد عدد الأيام المتبقية حتى عطلة نهاية الأسبوع.

- يبدأ عطلة نهاية الأسبوع يوم السبت.
- إذا كان التاريخ المعطى يوم السبت أو الأحد، أعد `"It's the weekend!"`.
- وإلا، أعد `"X days until the weekend."`، حيث `X` هو عدد الأيام حتى يوم السبت.
- إذا كان `X` يساوي `1`، استخدم `"day"` (مفرد) بدلًا من `"days"` (جمع).
- تأكد من أن الحساب يتجاهل المنطقة الزمنية المحلية.

# --hints--

`daysUntilWeekend("2025-11-14")` يجب أن يعيد `"1 day until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` يجب أن يعيد `"3 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` يجب أن يعيد `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` يجب أن يعيد `"4 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` يجب أن يعيد `"5 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` يجب أن يعيد `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
