---
id: 68f6587287ad1f4ad39b0c7e
title: "تحدٍّ 92: مستخرج الامتداد"
challengeType: 28
dashedName: challenge-92
---

# --description--

أعطِ سلسلة نصية تمثل اسم ملف، أرجع امتداد الملف.

- الامتداد هو الجزء من اسم الملف الذي يأتي بعد آخر نقطة (`.`).
- إذا لم يحتوي اسم الملف على نقطة أو انتهى بنقطة، أرجع `"none"`.
- يجب إرجاع الامتداد كما هو، مع الحفاظ على حالة الأحرف.

# --hints--

`getExtension("document.txt")` يجب أن تُرجع `"txt"`.

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` يجب أن تُرجع `"none"`.

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` يجب أن تُرجع `"PNG"`.

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` يجب أن تُرجع `"gitignore"`.

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` يجب أن تُرجع `"gz"`.

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` يجب أن تُرجع `"none"`.

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
