---
id: 68f6587287ad1f4ad39b0c80
title: "تحدٍّ 94: مولد توقيع البريد الإلكتروني"
challengeType: 28
dashedName: challenge-94
---

# --description--

بالنظر إلى سلاسل نصية لاسم شخص، واللقب، والشركة، أعد توقيع بريد إلكتروني كسلسلة نصية واحدة باستخدام القواعد التالية:

- يجب أن يظهر الاسم أولًا، مسبوقًا ببادئة تعتمد على الحرف الأول من الاسم. للأسماء التي تبدأ بـ (بغض النظر عن حالة الحروف):
  - `A-I`: استخدم `>>` كبادئة.
  - `J-R`: استخدم `--` كبادئة.
  - `S-Z`: استخدم `::` كبادئة.
- يجب أن يتبع الاسم فاصلة ومسافة (`, `).
- يجب أن يتبع اللقب والشركة الفاصلة والمسافة، مفصولين بـ `" at "` (مع وجود مسافات حولها).

على سبيل المثال، إذا كانت القيم المعطاة هي `"Quinn Waverly"`، و`"Founder and CEO"`، و`"TechCo"`، فأعد `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` يجب أن تعيد `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` يجب أن تعيد `">>Alice Reed, Engineer at TechCo"`.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` يجب أن تعيد `"::Tina Vaughn, Developer at example.com"`.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` يجب أن تعيد `">>B. B., Product Tester at AcmeCorp"`.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` يجب أن تعيد `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
