---
id: 68f6587287ad1f4ad39b0c81
title: "تحدٍّ 95: إزاحة المصفوفة"
challengeType: 28
dashedName: challenge-95
---

# --description--

معطاة مصفوفة وعدد صحيح يمثل عدد المواقع التي يجب إزاحة المصفوفة بها، أرجع المصفوفة بعد الإزاحة.

- العدد الصحيح الموجب يزيح المصفوفة إلى اليسار.
- العدد الصحيح السالب يزيح المصفوفة إلى اليمين.
- الإزاحة تدور حول المصفوفة.

على سبيل المثال، معطاة `[1, 2, 3]` و`1`، ازح المصفوفة بمقدار 1 إلى اليسار، معيدًا `[2, 3, 1]`.

# --hints--

`shiftArray([1, 2, 3], 1)` يجب أن تُرجع `[2, 3, 1]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` يجب أن تُرجع `[3, 1, 2]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` يجب أن تُرجع `["charlie", "alpha", "bravo"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` يجب أن تُرجع `["bravo", "charlie", "alpha"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` يجب أن تُرجع `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
