---
id: 68f6587287ad1f4ad39b0c7f
title: "تحدٍّ 93: الحروف المتحركة والحروف الساكنة"
challengeType: 28
dashedName: challenge-93
---

# --description--

أعطِ سلسلة نصية، أرجع مصفوفة تحتوي على عدد الحروف المتحركة وعدد الحروف الساكنة في السلسلة.

- الحروف المتحركة تتكون من `a`، `e`، `i`، `o`، `u` بأي حالة.
- الحروف الساكنة تتكون من جميع الحروف الأخرى بأي حالة.
- تجاهل أي حروف غير حروف أبجدية.

على سبيل المثال، عند إعطاء `"Hello World"`، أرجع `[3, 7]`.

# --hints--

`count("Hello World")` يجب أن تُرجع `[3, 7]`.

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` يجب أن تُرجع `[3, 7]`.

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` يجب أن تُرجع `[1, 5]`.

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` يجب أن تُرجع `[5, 7]`.

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` يجب أن تُرجع `[3, 7]`.

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` يجب أن تُرجع `[11, 24]`.

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
