---
id: 68f6587287ad1f4ad39b0c7c
title: "تحدٍّ 90: حد الحروف"
challengeType: 28
dashedName: challenge-90
---

# --description--

في هذا التحدٍّ، تُعطى سلسلة نصية وتحتاج إلى تحديد ما إذا كانت تناسب منشورًا على وسائل التواصل الاجتماعي. أرجع السلاسل النصية التالية بناءً على القواعد المعطاة:

- `"short post"` إذا كانت تناسب حد 40 حرفًا.
- `"long post"` إذا كانت أكثر من 40 حرفًا وتناسب حد 80 حرفًا.
- `"invalid post"` إذا كانت طويلة جدًا بحيث لا تناسب أيًا من الحدين.

# --hints--

`canPost("Hello world")` يجب أن تُرجع `"short post"`.

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")` يجب أن تُرجع `"long post"`.

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")` يجب أن تُرجع `"invalid post"`.

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```
