---
id: 68f6587287ad1f4ad39b0c83
title: "تحدٍّ 97: القاسم المشترك الأكبر"
challengeType: 28
dashedName: challenge-97
---

# --description--

أعطِ عددين صحيحين موجبين، أرجع القاسم المشترك الأكبر (GCD) لهما.

- القاسم المشترك الأكبر لعددين صحيحين هو أكبر عدد يقسم العددين بدون أن يترك باقيًا.

على سبيل المثال، القواسم لـ `4` هي `1`، `2`، و`4`. القواسم لـ `6` هي `1`، `2`، `3`، و`6`. إذًا عند إعطاء `4` و`6`، أرجع `2`، وهو أكبر عدد يظهر في مجموعتي القواسم.

# --hints--

`gcd(4, 6)` يجب أن تُرجع `2`.

```js
assert.equal(gcd(4, 6), 2);
```

`gcd(20, 15)` يجب أن تُرجع `5`.

```js
assert.equal(gcd(20, 15), 5);
```

`gcd(13, 17)` يجب أن تُرجع `1`.

```js
assert.equal(gcd(13, 17), 1);
```

`gcd(654, 456)` يجب أن تُرجع `6`.

```js
assert.equal(gcd(654, 456), 6);
```

`gcd(3456, 4320)` يجب أن تُرجع `864`.

```js
assert.equal(gcd(3456, 4320), 864);
```

# --seed--

## --seed-contents--

```js
function gcd(x, y) {

  return x;
}
```

# --solutions--

```js
function gcd(x, y) {
  while (y !== 0) {
    const temp = y;
    y = x % y;
    x = temp;
  }
  return x;
}
```
