---
id: 68f6587287ad1f4ad39b0c85
title: "تحدٍّ 99: اختبار البصمة"
challengeType: 28
dashedName: challenge-99
---

# --description--

معطيتان سلسلتا نصية تمثلان بصمات أصابع، حدّد ما إذا كانتا متطابقتين باستخدام القواعد التالية:

- تتكون كل بصمة من أحرف صغيرة فقط (`a-z`).
- تُعتبر بصمتان متطابقتين إذا:
  - كانتا بنفس الطول.
  - لا يتجاوز عدد الأحرف المختلفة 10% من طول البصمة.

# --hints--

`isMatch("helloworld", "helloworld")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMatch("helloworld", "helloworld"));
```

`isMatch("helloworld", "helloworlds")` يجب أن تُرجع `false`.

```js
assert.isFalse(isMatch("helloworld", "helloworlds"));
```

`isMatch("helloworld", "jelloworld")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMatch("helloworld", "jelloworld"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"));
```

`isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")` يجب أن تُرجع `false`.

```js
assert.isFalse(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"));
```

# --seed--

## --seed-contents--

```js
function isMatch(fingerprintA, fingerprintB) {

  return fingerprintA;
}
```

# --solutions--

```js
function isMatch(fingerprintA, fingerprintB) {
  if (fingerprintA.length !== fingerprintB.length) return false;

  const length = fingerprintA.length;
  let mismatches = 0;

  for (let i = 0; i < length; i++) {
    if (fingerprintA[i] !== fingerprintB[i]) {
      mismatches++;
      if (mismatches > length * 0.1) return false;
    }
  }

  return true;
}
```
