---
id: 68ffb91507a5b645769328cc
title: "تحدٍّ 109: كم عمري مرة أخرى؟"
challengeType: 28
dashedName: challenge-109
---

# --description--

بالنظر إلى تاريخ ميلاد شخص ما بصيغة `YYYY-MM-DD`، أرجع عمر الشخص حتى 27 نوفمبر 2025.

- افترض أن جميع تواريخ الميلاد صالحة وتقع قبل 27 نوفمبر 2025.
- أرجع العمر كعدد صحيح.
- تأكد من احتساب ما إذا كان الشخص قد احتفل بعيد ميلاده في 2025 بالفعل.

# --hints--

`calculateAge("2000-11-20")` يجب أن تُرجع `25`.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` يجب أن تُرجع `24`.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` يجب أن تُرجع `11`.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` يجب أن تُرجع `31`.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` يجب أن تُرجع `30`.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
