---
id: 68ffb91507a5b645769328c5
title: "تحدٍّ 102: أطول كلمة"
challengeType: 28
dashedName: challenge-102
---

# --description--

أعطِ سلسلة نصية تمثل جملة، وأرجع أطول كلمة في الجملة.

- الكلمات مفصولة بمسافة واحدة.
- فقط الحروف (`a-z`، غير حساسة لحالة الحروف) تُحتسب لطول الكلمة.
- إذا وُجدت كلمات متعددة بنفس الطول، أرجع أول كلمة تظهر.
- أرجع الكلمة كما تظهر في السلسلة المعطاة، مع إزالة علامات الترقيم.

# --hints--

`longestWord("The quick red fox")` يجب أن تُرجع `"quick"`.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` يجب أن تُرجع `"challenge"`.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` يجب أن تُرجع `"This"`.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` يجب أن تُرجع `"exclamation"`.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` يجب أن تُرجع `"tie"`.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` يجب أن تُرجع `"Wouldnt"`.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
