---
id: 68ffb91507a5b645769328cb
title: "تحدٍّ 108: BuzzFizz"
challengeType: 28
dashedName: challenge-108
---

# --description--

معطاة مصفوفة، حدّد إذا كانت متتالية FizzBuzz صحيحة من 1 حتى آخر عنصر في المصفوفة. المتتالية صحيحة إذا:

- الأعداد التي هي مضاعفات 3 تُستبدَل بـ "Fizz"
- الأعداد التي هي مضاعفات 5 تُستبدَل بـ "Buzz"
- الأعداد التي هي مضاعفات كل من 3 و5 تُستبدَل بـ "FizzBuzz"
- تبقى جميع الأعداد الأخرى كأعداد صحيحة بترتيب تصاعدي، بدءًا من 1.
- يجب أن تبدأ المصفوفة بالعدد 1 ولا تحتوي على عناصر مفقودة أو زائدة.

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4])` يجب أن تُعيد `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4]));
```

`isFizzBuzz([1, 2, 3, 4])` يجب أن تُعيد `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, 3, 4]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", 7])` يجب أن تُعيد `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", 7]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"])` يجب أن تُعيد `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"])` يجب أن تُعيد `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"])` يجب أن تُعيد `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"])` يجب أن تُعيد `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(sequence) {

  return sequence;
}
```

# --solutions--

```js
function isFizzBuzz(sequence) {
  for (let i = 0; i < sequence.length; i++) {
    const n = i + 1;
    if (n % 15 === 0 && sequence[i] !== "FizzBuzz") return false;
    if (n % 3 === 0 && n % 5 !== 0 && sequence[i] !== "Fizz") return false;
    if (n % 5 === 0 && n % 3 !== 0 && sequence[i] !== "Buzz") return false;
    if (n % 3 !== 0 && n % 5 !== 0 && sequence[i] !== n) return false;
  }
  return true;
}
```
