---
id: 68ffb91507a5b645769328c4
title: "تحدٍّ 101: محول عنوان Markdown"
challengeType: 28
dashedName: challenge-101
---

# --description--

أعطِ سلسلة نصية تمثل عنوان Markdown، وأرجع العنوان المكافئ في HTML.

يجب أن يكون عنوان Markdown صالحًا:

- يبدأ بصفر أو أكثر من الفراغات، يتبعها
- من 1 إلى 6 رموز هاش (`#`) متتالية، ثم
- فراغ واحد على الأقل. وأخيرًا،
- نص العنوان.

عدد رموز الهاش يحدد مستوى العنوان. على سبيل المثال، رمز هاش واحد يقابل وسم `h1`، وستة رموز هاش تقابل وسم `h6`.

إذا لم تكن السلسلة النصية المعطاة بالصيغة الدقيقة أعلاه، أرجع `"Invalid format"`.

على سبيل المثال، إذا كانت السلسلة `"# My level 1 heading"`، أرجع `"<h1>My level 1 heading</h1>"`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل النصية عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`convert("# My level 1 heading")` يجب أن تُرجع `"<h1>My level 1 heading</h1>"`.

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")` يجب أن تُرجع `"Invalid format"`.

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")` يجب أن تُرجع `"<h5>My level 5 heading</h5>"`.

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")` يجب أن تُرجع `"Invalid format"`.

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")` يجب أن تُرجع `"<h3>My level 3 heading</h3>"`.

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")` يجب أن تُرجع `"Invalid format"`.

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")` يجب أن تُرجع `"<h2>My #2 heading</h2>"`.

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```
