---
id: 68ffb91507a5b645769328c9
title: "تحدٍّ 106: مدقق الرسائل"
challengeType: 28
dashedName: challenge-106
---

# --description--

معطاة سلسلة نصية للرسالة وسلسلة نصية للتحقق، حدّد ما إذا كانت الرسالة صحيحة.

- تكون الرسالة صحيحة إذا بدأ كل كلمة في الرسالة بالحرف المقابل في سلسلة التحقق، وبالترتيب.
- الحروف غير حساسة لحالة الأحرف.
- الكلمات في الرسالة مفصولة بمسافات مفردة.

# --hints--

`isValidMessage("hello world", "hw")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
