---
id: 69162d64f96574d9bb629f00
title: "تحدٍّ 115: محول عنصر قائمة مرتبة في Markdown"
challengeType: 28
dashedName: challenge-115
---

# --description--

أعطِ سلسلة نصية تمثل عنصر قائمة مرتبة في Markdown، وأرجع السلسلة المكافئة بصيغة HTML.

يجب أن يكون عنصر القائمة المرتبة الصحيح في Markdown:

- يبدأ بصفر أو أكثر من الفراغات، يتبعها
- رقم (1 أو أكثر) ونقطة (`.`)، يتبعها
- فراغ واحد على الأقل، ثم
- نص عنصر القائمة.

إذا لم تكن السلسلة النصية بالصيغة الدقيقة أعلاه، أرجع `"Invalid format"`. وإلا، لف نص عنصر القائمة داخل وسم `li` وأرجع السلسلة.

على سبيل المثال، إذا كانت السلسلة `"1. My item"`، أرجع `"<li>My item</li>"`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل النصية عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`convertListItem("1. My item")` يجب أن تُرجع `"<li>My item</li>"`.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` يجب أن تُرجع `"<li>Another item</li>"`.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` يجب أن تُرجع `"Invalid format"`.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` يجب أن تُرجع `"<li>list item text</li>"`.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` يجب أن تُرجع `"Invalid format"`.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` يجب أن تُرجع `"Invalid format"`.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
