---
id: 69162d64f96574d9bb629f04
title: "تحدٍّ 119: ضغط السلسلة النصية"
challengeType: 28
dashedName: challenge-119
---

# --description--

بالنظر إلى جملة سلسلة نصية، أعد نسخة مضغوطة من الجملة حيث تُستبدَل الكلمات المتكررة المتتالية بالكلمة متبوعة بعدد مرات تكرارها داخل أقواس عادية.

- تُضغط فقط التكرارات المتتالية.
- تُفصل الكلمات بمسافات مفردة.

على سبيل المثال، إذا كانت القيمة المعطاة هي `"yes yes yes please"`، فأعد `"yes(3) please"`.

# --hints--

يجب أن تُعيد الدالة `compressString("yes yes yes please")` القيمة `"yes(3) please"`.

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

يجب أن تُعيد الدالة `compressString("I have have have apples")` القيمة `"I have(3) apples"`.

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

يجب أن تُعيد الدالة `compressString("one one three and to the the the the")` القيمة `"one(2) three and to the(4)"`.

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

يجب أن تُعيد الدالة `compressString("route route route route route route tee tee tee tee tee tee")` القيمة `"route(6) tee(6)"`.

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
