---
id: 69162d64f96574d9bb629efc
title: "تحدٍّ 111: مسار الكرة"
challengeType: 28
dashedName: challenge-111
---

# --description--

التحدٍّ اليوم مستوحى من لعبة الفيديو Pong، التي صدرت في 29 نوفمبر 1972.

معطى مصفوفة (مصفوفة من المصفوفات) تتضمن موقع الكرة (`2`) والموقع السابق للكرة (`1`)، أرجع مؤشرات المصفوفة للموقع التالي للكرة.

- الكرة تتحرك دائمًا في خط مستقيم.
- اتجاه الحركة يُحدد بناءً على كيفية تحرك الكرة من `1` إلى `2`.
- تُعتبر حواف المصفوفة جدرانًا. إذا اصطدمت الكرة بـ:
  - جدار علوي أو سفلي، ترتد بعكس اتجاهها الرأسي.
  - جدار أيسر أو أيمن، ترتد بعكس اتجاهها الأفقي.
  - زاوية، ترتد بعكس كلا الاتجاهين.

# --hints--

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])` يجب أن تُرجع `[2, 3]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3]);
```

`getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])` يجب أن تُرجع `[3, 0]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0]);
```

`getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])` يجب أن تُرجع `[1, 2]`.

```js
assert.deepEqual(getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])` يجب أن تُرجع `[1, 1]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])` يجب أن تُرجع `[2, 2]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2]);
```

# --seed--

## --seed-contents--

```js
function getNextLocation(matrix) {

  return matrix;
}
```

# --solutions--

```js
function getNextLocation(matrix) {
  let prev = null;
  let curr = null;

  for (let row = 0; row < matrix.length; row++) {
    for (let col = 0; col < matrix[0].length; col++) {
      if (matrix[row][col] === 1) prev = [row, col];
      if (matrix[row][col] === 2) curr = [row, col];
    }
  }

  const [prevRow, prevCol] = prev;
  const [currRow, currCol] = curr;

  let dirX = currCol - prevCol;
  let dirY = currRow - prevRow;

  let nextRow = currRow + dirY;
  let nextCol = currCol + dirX;

  const maxRow = matrix.length - 1;
  const maxCol = matrix[0].length - 1;


  if (nextCol < 0 || nextCol > maxCol) dirX *= -1;
  if (nextRow < 0 || nextRow > maxRow) dirY *= -1;

  nextRow = currRow + dirY;
  nextCol = currCol + dirX;

  return [nextRow, nextCol];
}
```
