---
id: 69162d64f96574d9bb629f01
title: "تحدٍّ 116: عد التباديل"
challengeType: 28
dashedName: challenge-116
---

# --description--

أعطِ سلسلة نصية، أرجع عدد التباديل المميزة التي يمكن تكوينها من حروفها.

- التبديل هو أي إعادة ترتيب لحروف السلسلة النصية.
- لا تحسب التباديل المكررة.
- إذا احتوت السلسلة على حروف مكررة، يجب حساب الترتيبات المكررة مرة واحدة فقط.
- السلسلة ستحتوي على حروف فقط (`A-Z`، `a-z`).

على سبيل المثال، إذا كانت السلسلة `"abb"`، أرجع `3` لأن هناك ثلاث طرق فريدة لترتيب الحروف: `"abb"`، `"bab"`، و `"bba"`.

# --hints--

`countPermutations("abb")` يجب أن تُرجع `3`.

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` يجب أن تُرجع `6`.

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` يجب أن تُرجع `630`.

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` يجب أن تُرجع `39916800`.

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
