---
id: 69162d64f96574d9bb629efd
title: "تحدٍّ 112: كاشف الذكاء الاصطناعي"
challengeType: 28
dashedName: challenge-112
---

# --description--

التحدٍّ اليوم مستوحى من إصدار ChatGPT في 30 نوفمبر 2022.

معطى سلسلة نصية تحتوي على جملة واحدة أو أكثر، حدّد ما إذا كان من المحتمل أن تكون مولدة بواسطة الذكاء الاصطناعي باستخدام القواعد التالية:

- تحتوي على شرطتين أو أكثر (`-`).
- تحتوي على مجموعتين أو أكثر من الأقواس (`()`). يمكن أن يكون هناك نص داخل الأقواس.
- تحتوي على ثلاث كلمات أو أكثر تحتوي كل منها على 7 أحرف أو أكثر.

- الكلمات مفصولة بمسافة واحدة فقط وتتكون من حروف فقط (`A-Z`). لا تشمل علامات الترقيم أو غير الحروف كجزء من الكلمة.

إذا استوفت الجملة المعطاة أيًا من القواعد أعلاه، أرجع `"AI"`، وإلا أرجع `"Human"`.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` يجب أن تُرجع `"Human"`.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` يجب أن تُرجع `"Human"`.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` يجب أن تُرجع `"AI"`.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` يجب أن تُرجع `"AI"`.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` يجب أن تُرجع `"AI"`.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
