---
id: 69162d64f96574d9bb629efe
title: "تحدٍّ 113: من الأميال إلى الكيلومترات"
challengeType: 28
dashedName: challenge-113
---

# --description--

معطى مسافة بالأميال كرقم، أرجع المسافة المكافئة بالكيلومترات.

- سيكون المُدخَل دائمًا رقمًا غير سالب.
- الميل الواحد يساوي 1.60934 كيلومتر.
- قرّب النتيجة إلى منزلتين عشريتين.
- أزل الأصفار الزائدة غير الضرورية من النتيجة المقربة.

# --hints--

يجب أن تُرجع الدالة `convertToKm(1)` القيمة `1.61`.

```js
assert.equal(convertToKm(1), 1.61);
```

يجب أن تُرجع الدالة `convertToKm(21)` القيمة `33.8`.

```js
assert.equal(convertToKm(21), 33.8);
```

يجب أن تُرجع الدالة `convertToKm(3.5)` القيمة `5.63`.

```js
assert.equal(convertToKm(3.5), 5.63);
```

يجب أن تُرجع الدالة `convertToKm(0)` القيمة `0`.

```js
assert.equal(convertToKm(0), 0);
```

يجب أن تُرجع الدالة `convertToKm(0.621371)` القيمة `1`.

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
