---
id: 69162d64f96574d9bb629eff
title: "تحدٍّ 114: من Camel إلى Snake"
challengeType: 28
dashedName: challenge-114
---

# --description--

بالنظر إلى سلسلة نصية بصيغة camel case، أعد النسخة بصيغة snake case باستخدام القواعد التالية:

- ستحتوي السلسلة النصية المُدخَلة على أحرف فقط (`A-Z` و `a-z`) وستبدأ دائمًا بحرف صغير.
- كل حرف كبير في سلسلة camel case يبدأ كلمة جديدة.
- حوّل كل الأحرف إلى أحرف صغيرة.
- افصل الكلمات بشرطة سفلية (`_`).

# --hints--

`toSnake("helloWorld")` يجب أن تُعيد `"hello_world"`.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` يجب أن تُعيد `"my_variable_name"`.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` يجب أن تُعيد `"freecodecamp_daily_challenges"`.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
