---
id: 691b559495c5cb5a37b9b488
title: "تحدٍّ 128: عدّ الحروف الساكنة"
challengeType: 28
dashedName: challenge-128
---

# --description--

معطاة سلسلة نصية ورقم مستهدف، حدّد ما إذا كانت السلسلة تحتوي بالضبط على العدد المستهدف من الحروف الساكنة.

- الحروف الساكنة هي كل الحروف الأبجدية ما عدا `"a"`، `"e"`، `"i"`، `"o"`، و`"u"` بأي حالة.
- تجاهل الأرقام وعلامات الترقيم والمسافات وأي حروف غير أبجدية عند العد.

# --hints--

يجب أن تُرجع الدالة `hasConsonantCount("helloworld", 7)` القيمة `true`.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

يجب أن تُرجع الدالة `hasConsonantCount("eieio", 5)` القيمة `false`.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

يجب أن تُرجع الدالة `hasConsonantCount("freeCodeCamp Rocks!", 11)` القيمة `true`.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

يجب أن تُرجع الدالة `hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` القيمة `false`.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

يجب أن تُرجع الدالة `hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` القيمة `true`.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
