---
id: 691b559495c5cb5a37b9b485
title: "تحدٍّ 125: لعبة الحياة"
challengeType: 28
dashedName: challenge-125
---

# --description--

معطاة مصفوفة (مصفوفة من المصفوفات) تمثل الحالة الحالية في لعبة الحياة لـConway، أعد الحالة التالية للمصفوفة باستخدام هذه القواعد:

- كل خلية إما `1` (حية) أو `0` (ميتة).
- جيران الخلية هم حتى ثمانية خلايا محيطة بها (عموديًا، أفقيًا، وقطريًا).
- الخلايا على الحواف لها عدد أقل من ثمانية جيران.

قواعد تحديث كل خلية:

- أي خلية حية لديها أقل من جارين أحياء تموت (قلة السكان).
- أي خلية حية لديها جاران أو ثلاثة جيران أحياء تبقى على قيد الحياة.
- أي خلية حية لديها أكثر من ثلاثة جيران أحياء تموت (زيادة السكان).
- أي خلية ميتة لديها بالضبط ثلاثة جيران أحياء تصبح حية (تكاثر).

على سبيل المثال، معطاة:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

أعد:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

تُحدّث كل خلية حسب عدد الجيران الأحياء. على سبيل المثال، تبقى `[0][0]` ميتة (لديها 2 جيران أحياء)، وتبقى `[0][1]` حية (لديها 2 جيران أحياء)، وتموت `[0][2]` (لديها 3 جيران أحياء)، وهكذا.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` يجب أن تُعيد `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` يجب أن تُعيد `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` يجب أن تُعيد `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` يجب أن تُعيد `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
