---
id: 691b559495c5cb5a37b9b486
title: "تحدٍّ 126: اجعل الحروف كبيرة"
challengeType: 28
dashedName: challenge-126
---

# --description--

معطى سلسلة نصية عنوان، أعد سلسلة نصية جديدة منسقة بحالة العنوان باستخدام القواعد التالية:

- اجعل الحرف الأول من كل كلمة كبيرًا.
- اجعل جميع الحروف الأخرى في كل كلمة أحرف صغيرة.
- الكلمات مفصولة دائمًا بمسافة واحدة.

# --hints--

`titleCase("hello world")` يجب أن تُعيد `"Hello World"`.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` يجب أن تُعيد `"The Quick Brown Fox"`.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` يجب أن تُعيد `"Javascript And Python"`.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` يجب أن تُعيد `"Avocado Toast For Breakfast"`.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
