---
id: 691b559495c5cb5a37b9b489
title: "تحدٍّ 129: محلل الاقتباس في Markdown"
challengeType: 28
dashedName: challenge-129
---

# --description--

أعطِ سلسلة نصية تتضمن اقتباسًا في Markdown، وأرجع سلسلة HTML المكافئة.

الاقتباس في Markdown هو أي سطر يحقق ما يلي:

- يبدأ بصفر أو أكثر من الفراغات
- يتبعه علامة أكبر من (`>`)
- ثم واحد أو أكثر من الفراغات
- وأخيرًا نص الاقتباس.

أرجع نص الاقتباس محاطًا بوسمي HTML `blockquote` الافتتاحي والختامي.

على سبيل المثال، إذا أعطيت `"> This is a quote"`، أرجع `<blockquote>This is a quote</blockquote>`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML في السلاسل النصية عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parseBlockquote("> This is a quote")` يجب أن تُرجع `"<blockquote>This is a quote</blockquote>"`.

```js
assert.equal(parseBlockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>");
```

`parseBlockquote(" > This is also a quote")` يجب أن تُرجع `"<blockquote>This is also a quote</blockquote>"`.

```js
assert.equal(parseBlockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>");
```

`parseBlockquote("  >    So  Is  This")` يجب أن تُرجع `"<blockquote>So  Is  This</blockquote>"`.

```js
assert.equal(parseBlockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>");
```

# --seed--

## --seed-contents--

```js
function parseBlockquote(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBlockquote(markdown) {
  const blockquoteRegex = /^\s*>\s+(.+)$/;
  const match = markdown.match(blockquoteRegex);

  const text = match[1];
  return `<blockquote>${text}</blockquote>`;
}
```
