---
id: 691b559495c5cb5a37b9b482
title: "تحدٍّ 122: محلل النص العريض في Markdown"
challengeType: 28
dashedName: challenge-122
---

# --description--

أعطِ سلسلة نصية قد تحتوي على بعض النصوص العريضة في Markdown، أعد السلسلة المكافئة بصيغة HTML.

- النص العريض في Markdown هو أي نص يبدأ وينتهي بنجمتين (`**`) أو بشرطتين سفليتين (`__`).
- لا يمكن أن توجد مسافات بين النص والنجمتين أو الشرطتين السفليتين، لكن يمكن أن توجد
  مسافات داخل النص نفسه.
- حوّل كل النصوص العريضة إلى وسم HTML `b` وأرجع السلسلة النصية.

على سبيل المثال، إذا كانت السلسلة `"**This is bold**"`، أعد `"<b>This is bold</b>"`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parseBold("**This is bold**")` يجب أن تعيد `"<b>This is bold</b>"`.

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` يجب أن تعيد `"<b>This is also bold</b>"`.

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` يجب أن تعيد `"**This is not bold **"`.

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` يجب أن تعيد `"__ This is also not bold__"`.

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` يجب أن تعيد `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
