---
id: 691b559495c5cb5a37b9b480
title: "تحدٍّ 120: من باوند إلى كيلوجرام"
challengeType: 28
dashedName: challenge-120
---

# --description--

بالنظر إلى وزن بوحدة الباوند كرقم، أعد السلسلة النصية `"(lbs) pounds equals (kgs) kilograms."`.

- استبدل `"(lbs)"` بالرقم المُدخَل.
- استبدل `"(kgs)"` بالقيمة المحولة إلى كيلوجرامات، مع تقريبها إلى منزلتين عشريتين ودوماً تضمين منزلتين عشريتين في القيمة.
- 1 باوند يساوي 0.453592 كيلوجرام.
- إذا كان المُدخَل هو `1`، استخدم `"pound"` بدلاً من `"pounds"`.
- إذا كانت القيمة المحولة هي `1`، استخدم `"kilogram"` بدلاً من `"kilograms"`.

# --hints--

`convertToKgs(1)` يجب أن تُعيد `"1 pound equals 0.45 kilograms."`.

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` يجب أن تُعيد `"0 pounds equals 0.00 kilograms."`.

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` يجب أن تُعيد `"100 pounds equals 45.36 kilograms."`.

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` يجب أن تُعيد `"2.5 pounds equals 1.13 kilograms."`.

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` يجب أن تُعيد `"2.20462 pounds equals 1.00 kilogram."`.

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
