---
id: 691b559495c5cb5a37b9b483
title: "تحدٍّ 123: منشئ الأعداد الرومانية"
challengeType: 28
dashedName: challenge-123
---

# --description--

أعطِ عددًا صحيحًا، وأرجع قيمته المكافئة بالأعداد الرومانية.

تستخدم الأعداد الرومانية هذه الرموز:

| الرمز | القيمة |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

تُكتب الأعداد الرومانية من الأكبر إلى الأصغر، من اليسار إلى اليمين باستخدام القواعد التالية:

- تُستخدم الإضافة عندما يتبع رمز رمز آخر له نفس القيمة أو قيمة أصغر. على سبيل المثال، يُكتب `18` كـ `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- تُستخدم الطرح عندما يظهر رمز أصغر قبل رمز أكبر، لتمثيل 4 أو 9 في أي منزلة. على سبيل المثال، يُكتب 19 كـ `XIX` (10 + (10 - 1)).
- لا يجوز تكرار أي رمز أكثر من ثلاث مرات متتالية. يُستخدم الطرح عندما تحتاج إلى كتابة رمز أكثر من ثلاث مرات متتالية. لذلك، أكبر رقم يمكنك كتابته هو 3999.

إليك مثالًا آخر: إذا أعطيت `1464`، أرجع `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

يجب أن تُرجع `toRoman(18)` القيمة `"XVIII"`.

```js
assert.equal(toRoman(18), "XVIII");
```

يجب أن تُرجع `toRoman(19)` القيمة `"XIX"`.

```js
assert.equal(toRoman(19), "XIX");
```

يجب أن تُرجع `toRoman(1464)` القيمة `"MCDLXIV"`.

```js
assert.equal(toRoman(1464), "MCDLXIV");
```

يجب أن تُرجع `toRoman(2025)` القيمة `"MMXXV"`.

```js
assert.equal(toRoman(2025), "MMXXV");
```

يجب أن تُرجع `toRoman(3999)` القيمة `"MMMCMXCIX"`.

```js
assert.equal(toRoman(3999), "MMMCMXCIX");
```

# --seed--

## --seed-contents--

```js
function toRoman(num) {
  return num;
}
```

# --solutions--

```js
function toRoman(num) {
  const symbols = [
    ["M", 1000],
    ["CM", 900],
    ["D", 500],
    ["CD", 400],
    ["C", 100],
    ["XC", 90],
    ["L", 50],
    ["XL", 40],
    ["X", 10],
    ["IX", 9],
    ["V", 5],
    ["IV", 4],
    ["I", 1]
  ];

  let result = "";
  for (const [sym, val] of symbols) {
    while (num >= val) {
      result += sym;
      num -= val;
    }
  }
  return result;
}
```
