---
id: 691b559495c5cb5a37b9b487
title: "تحدٍّ 127: فحص السرعة"
challengeType: 28
dashedName: challenge-127
---

# --description--

بما أن السرعة التي تسير بها بالأميال في الساعة (MPH)، وحد السرعة بالكيلومترات في الساعة (KPH)، حدّد ما إذا كنت تتجاوز السرعة وإذا كنت ستحصل على تحذير أو مخالفة.

- الميل الواحد يساوي 1.60934 كيلومتر.
- إذا كنت تسير بسرعة أقل من أو تساوي حد السرعة، أرجع `"Not Speeding"`.
- إذا كنت تسير بسرعة تزيد 5 كيلومترات في الساعة أو أقل فوق حد السرعة، أرجع `"Warning"`.
- إذا كنت تسير بسرعة تزيد أكثر من 5 كيلومترات في الساعة فوق حد السرعة، أرجع `"Ticket"`.

# --hints--

`speedCheck(30, 70)` يجب أن تُرجع `"Not Speeding"`.

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` يجب أن تُرجع `"Warning"`.

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` يجب أن تُرجع `"Not Speeding"`.

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` يجب أن تُرجع `"Ticket"`.

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` يجب أن تُرجع `"Warning"`.

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` يجب أن تُرجع `"Ticket"`.

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
