---
id: 691b559495c5cb5a37b9b484
title: "تحدٍّ 124: تحديث المخزون"
challengeType: 28
dashedName: challenge-124
---

# --description--

معطى مصفوفة ثنائية الأبعاد تمثل مخزون متجرك، ومصفوفة ثنائية الأبعاد أخرى تمثل شحنة استلمتها، أرجع المخزون المحدث.

- كل عنصر في المصفوفات سيكون بالشكل: `[quantity, "item"]`، حيث `quantity` عدد صحيح و`"item"` سلسلة نصية.
- حدّث العناصر في المخزون بإضافة كمية أي عناصر متطابقة من الشحنة.
- إذا كان العنصر المستلم غير موجود في المخزون الحالي، أضفه كإدخال جديد في نهاية المخزون.
- أرجع المخزون بنفس الترتيب الذي أعطي به مع وضع العناصر الجديدة في النهاية حسب ترتيب ظهورها في الشحنة.

على سبيل المثال، مع مخزون `[[2, "apples"], [5, "bananas"]]` وشحنة `[[1, "apples"], [3, "bananas"]]`، أرجع `[[3, "apples"], [8, "bananas"]]`.

# --hints--

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` يجب أن تُرجع `[[3, "apples"], [8, "bananas"]]`.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]]);
```

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` يجب أن تُرجع `[[3, "apples"], [8, "bananas"], [4, "oranges"]]`.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]]);
```

`updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` يجب أن تُرجع `[[10, "apples"], [30, "bananas"], [20, "oranges"]]`.

```js
assert.deepEqual(updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]]);
```

`updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` يجب أن تُرجع `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]`.

```js
assert.deepEqual(updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]);
```

# --seed--

## --seed-contents--

```js
function updateInventory(inventory, shipment) {
  return inventory;
}
```

# --solutions--

```js
function updateInventory(inventory, shipment) {
  const inventoryMap = new Map();
  inventory.forEach(([qty, item], index) => {
    inventoryMap.set(item, index);
  });

  shipment.forEach(([qty, item]) => {
    if (inventoryMap.has(item)) {
      const index = inventoryMap.get(item);
      inventory[index][0] += qty;
    } else {
      inventory.push([qty, item]);
      inventoryMap.set(item, inventory.length - 1);
    }
  });

  return inventory;
}
```
