---
id: 6925e2068081f40f549ced1d
title: "تحدٍّ 139: حجر، ورقة، مقص"
challengeType: 28
dashedName: challenge-139
---

# --description--

معطاة سلسلتا نص، الأولى تمثل اللاعب 1 والثانية تمثل اللاعب 2، حدّد الفائز في مباراة حجر، ورقة، مقص.

- ستكون السلاسل النصية دائمًا `"Rock"`، `"Paper"`، أو `"Scissors"`.
- `"Rock"` تغلب `"Scissors"`.
- `"Paper"` تغلب `"Rock"`.
- `"Scissors"` تغلب `"Paper"`.

أرجع:

- `"Player 1 wins"` إذا فاز اللاعب 1.
- `"Player 2 wins"` إذا فاز اللاعب 2.
- `"Tie"` إذا اختار اللاعبان نفس الخيار.

# --hints--

`rockPaperScissors("Rock", "Rock")` يجب أن تُرجع `"Tie"`.

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` يجب أن تُرجع `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` يجب أن تُرجع `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` يجب أن تُرجع `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` يجب أن تُرجع `"Tie"`.

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` يجب أن تُرجع `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
