---
id: 6925e2068081f40f549ced1c
title: "تحدٍّ 138: مجموع القواسم"
challengeType: 28
dashedName: challenge-138
---

# --description--

أعطِ عددًا صحيحًا موجبًا، أرجع مجموع جميع قواسمه.

- القاسم هو أي عدد صحيح يقسم العدد بدون باقي (الباقي هو `0`).
- احسب كل قاسم مرة واحدة فقط.

على سبيل المثال، إذا كان العدد `6`، أرجع `12` لأن قواسم `6` هي `1` و`2` و`3` و`6`، ومجموعها هو `12`.

# --hints--

يجب أن تُرجع الدالة `sumDivisors(6)` القيمة `12`.

```js
assert.equal(sumDivisors(6), 12);
```

يجب أن تُرجع الدالة `sumDivisors(13)` القيمة `14`.

```js
assert.equal(sumDivisors(13), 14);
```

يجب أن تُرجع الدالة `sumDivisors(28)` القيمة `56`.

```js
assert.equal(sumDivisors(28), 56);
```

يجب أن تُرجع الدالة `sumDivisors(84)` القيمة `224`.

```js
assert.equal(sumDivisors(84), 224);
```

يجب أن تُرجع الدالة `sumDivisors(549)` القيمة `806`.

```js
assert.equal(sumDivisors(549), 806);
```

يجب أن تُرجع الدالة `sumDivisors(9348)` القيمة `23520`.

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
