---
id: 6925e2068081f40f549ced1a
title: "تحدٍّ 136: محلل صور Markdown"
challengeType: 28
dashedName: challenge-136
---

# --description--

عند إعطائك سلسلة نصية لصورة بصيغة Markdown، أعد السلسلة المكافئة بصيغة HTML.

تأخذ صورة Markdown الصيغة التالية: `"![alt text](image_url)"`. حيث:

- `alt text` هو وصف الصورة (قيمة الخاصية `alt`).
- `image_url` هو عنوان URL المصدر للصورة (قيمة الخاصية `src`).

أعد سلسلة نصية لوسم HTML `img` مع تعيين الخاصية `src` إلى عنوان URL للصورة والخاصية `alt` إلى نص الوصف.

على سبيل المثال، عند إعطائك `"![Cute cat](cat.png)"` أعد `'<img src="cat.png" alt="Cute cat">'`;

- تأكد من أن الوسم، وترتيب الخصائص، والتباعد، واستخدام علامات الاقتباس مطابق لما سبق.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل النصية عند تسجيل الرسائل — تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parseImage("![Cute cat](cat.png)")` يجب أن تعيد `'<img src="cat.png" alt="Cute cat">'`.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` يجب أن تعيد `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` يجب أن تعيد `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
