---
id: 6925e2068081f40f549ced1b
title: "تحدٍّ 137: مولد ندفة الثلج"
challengeType: 28
dashedName: challenge-137
---

# --description--

معطى سلسلة نصية متعددة الأسطر تستخدم أحرف السطر الجديد (`\n`) لتمثيل فاصل الأسطر، أعد سلسلة جديدة حيث تُعكس كل سطر أفقيًا وتُلحق بنهاية السطر الأصلي.

- عكس السطر يتم بعكس جميع حروفه، بما في ذلك الفراغات.

على سبيل المثال، معطى `"* \n *\n* "`، والذي يُسجل في وحدة التحكم كالتالي:

```sh
* 
 *
* 
```

أعد `"*  *\n ** \n*  *"`، والذي يُسجل في وحدة التحكم كالتالي:

```sh
*  *
 ** 
*  *
```

انتبه جيدًا إلى الفراغات في السلاسل المعطاة والمعادة. تأكد من عدم حذف أي منها.

# --hints--

`generateSnowflake("* \n *\n* ")` يجب أن تُعيد `"*  *\n ** \n*  *"`.

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` يجب أن تُعيد `"X=~~=X"`.

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` يجب أن تُعيد `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` يجب أن تُعيد `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` يجب أن تُعيد `"*  --  *\n * -- * \n*  --  *"`.

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
