---
id: 69272dcf1c24b44fd79137c5
title: "تحدٍّ 142: مجموع السلسلة النصية"
challengeType: 28
dashedName: challenge-142
---

# --description--

أعطِ سلسلة نصية تحتوي أرقامًا وحروفًا أخرى، أرجع مجموع كل الأرقام في السلسلة.

- اعتبر الأرقام المتتالية رقمًا واحدًا. على سبيل المثال، `"13"` تُحسب كـ 13، وليس 1 + 3.
- تجاهل أي حروف غير أرقام.

# --hints--

`stringSum("3apples2bananas")` يجب أن تُرجع `5`.

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` يجب أن تُرجع `17`.

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` يجب أن تُرجع `125344`.

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` يجب أن تُرجع `321`.

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` يجب أن تُرجع `1653`.

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
