---
id: 69272dcf1c24b44fd79137c6
title: "تحدٍّ 143: محلل النص المائل في Markdown"
challengeType: 28
dashedName: challenge-143
---

# --description--

أعطِ سلسلة نصية قد تحتوي على بعض النص المائل في Markdown، أرجع سلسلة HTML المكافئة.

- النص المائل في Markdown هو أي نص يبدأ وينتهي بنجمة واحدة (`*`) أو شرطة سفلية واحدة (`_`).
- لا يمكن أن توجد مسافات بين النص والنجمة أو الشرطة السفلية، لكن يمكن أن توجد مسافات داخل النص نفسه.
- حوّل كل النصوص المائلة إلى وسم `i` في HTML وأرجع السلسلة.

على سبيل المثال، إذا أعطيت `"*This is italic*"`، أرجع `"<i>This is italic</i>"`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML في السلاسل النصية عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parseItalics("*This is italic*")` يجب أن تُرجع `"<i>This is italic</i>"`.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` يجب أن تُرجع `"<i>This is also italic</i>"`.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` يجب أن تُرجع `"*This is not italic *"`.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` يجب أن تُرجع `"_ This is also not italic_"`.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` يجب أن تُرجع `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
