---
id: 69272dcf1c24b44fd79137c4
title: "تحدٍّ 141: وقود الإقلاع"
challengeType: 28
dashedName: challenge-141
---

# --description--

بالنظر إلى عدد جالونات الوقود الموجودة حاليًا في طائرتك، وعدد اللترات المطلوبة من الوقود للوصول إلى وجهتك، حدّد كم جالونًا إضافيًا من الوقود يجب أن تضيف.

- 1 جالون يساوي 3.78541 لترات.
- إذا كانت الطائرة تحتوي على وقود كافٍ بالفعل، أرجع `0`.
- يمكنك إضافة جالونات كاملة فقط.
- لا تدرج الأعداد العشرية في الرقم المرجع.

# --hints--

يجب أن تُرجع الدالة `fuelToAdd(0, 1)` القيمة `1`.

```js
assert.equal(fuelToAdd(0, 1), 1);
```

يجب أن تُرجع الدالة `fuelToAdd(5, 40)` القيمة `6`.

```js
assert.equal(fuelToAdd(5, 40), 6);
```

يجب أن تُرجع الدالة `fuelToAdd(10, 30)` القيمة `0`.

```js
assert.equal(fuelToAdd(10, 30), 0);
```

يجب أن تُرجع الدالة `fuelToAdd(896, 20500)` القيمة `4520`.

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

يجب أن تُرجع الدالة `fuelToAdd(1000, 50000)` القيمة `12209`.

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
