---
id: 69272dcf1c24b44fd79137c3
title: "تحدٍّ 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

أعطِ سلسلة نصية تمثّل اسم متغير، أرجع اسم المتغير بعد تحويله إلى SCREAMING_SNAKE_CASE.

أسماء المتغيرات المعطاة ستكون مكتوبة بأحد التنسيقات التالية:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

في التنسيقات السابقة، تُفصل الكلمات بشرطة سفلية (`_`) أو شرطة (`-`)، أو تبدأ كلمة جديدة بحرف كبير.

لتحويل إلى SCREAMING_SNAKE_CASE:

- اجعل كل الحروف كبيرة
- افصل الكلمات بشرطة سفلية (`_`)

# --hints--

`toScreamingSnakeCase("userEmail")` يجب أن تُرجع `"USER_EMAIL"`.

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` يجب أن تُرجع `"USER_PASSWORD"`.

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` يجب أن تُرجع `"USER_ID"`.

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` يجب أن تُرجع `"USER_ADDRESS"`.

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` يجب أن تُرجع `"USERNAME"`.

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

`toScreamingSnakeCase("my_variable_name")` يجب أن تُرجع `"MY_VARIABLE_NAME"`.

```js
assert.equal(toScreamingSnakeCase("my_variable_name"), "MY_VARIABLE_NAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();
}
```
