---
id: 69306364df283fcaff2e1ad5
title: "تحدٍّ 144: سلسلة دقة الوضوح"
challengeType: 28
dashedName: challenge-144
---

# --description--

معطاة مصفوفة من مصفوفات، حيث تمثل كل مصفوفة فرعية يومًا من أنشطتك المتعلقة بالدقة وتحتوي على ثلاثة أعداد صحيحة: عدد الخطوات التي مشيتها في ذلك اليوم، ودقائق وقت الشاشة في ذلك اليوم، وعدد الصفحات التي قرأتها في ذلك اليوم؛ حدّد إذا كنت تحافظ على قراراتك.

- المصفوفة الفرعية الأولى هي اليوم 1، والثانية اليوم 2، وهكذا.

يُعتبر اليوم ناجحًا إذا تحققت جميع الأهداف الثلاثة التالية:

- مشيت على الأقل 10,000 خطوة.
- لم تتجاوز وقت الشاشة 120 دقيقة.
- قرأت على الأقل خمس صفحات.

إذا كانت جميع الأيام المعطاة ناجحة، أرجع `"Resolution on track: N day streak."` حيث `N` هو عدد الأيام الناجحة.

إذا كان يوم واحد أو أكثر غير ناجح، أرجع `"Resolution failed on day X: N day streak."` حيث `X` هو رقم أول يوم غير ناجح، و`N` هو عدد الأيام الناجحة قبل أول يوم غير ناجح.

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` يجب أن تُرجع `"Resolution on track: 3 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])` يجب أن تُرجع `"Resolution failed on day 2: 1 day streak."`

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` يجب أن تُرجع `"Resolution failed on day 3: 2 day streak."`

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` يجب أن تُرجع `"Resolution failed on day 4: 3 day streak."`

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` يجب أن تُرجع `"Resolution on track: 10 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```
