---
id: 69306364df283fcaff2e1ad8
title: "تحدٍّ 147: حاسبة السنة الكبيسة"
challengeType: 28
dashedName: challenge-147
---

# --description--

بالنظر إلى سنة صحيحة، حدّد ما إذا كانت سنة كبيسة.

تُعتبر السنة كبيسة إذا استوفت القواعد التالية:

- تكون السنة قابلة للقسمة على 4 بدون باقي، و
- لا تكون السنة قابلة للقسمة على 100 بدون باقي، إلا إذا
- كانت السنة قابلة للقسمة على 400 بدون باقي.

# --hints--

يجب أن تُعيد الدالة `isLeapYear(2024)` القيمة `true`.

```js
assert.isTrue(isLeapYear(2024));
```

يجب أن تُعيد الدالة `isLeapYear(2023)` القيمة `false`.

```js
assert.isFalse(isLeapYear(2023));
```

يجب أن تُعيد الدالة `isLeapYear(2100)` القيمة `false`.

```js
assert.isFalse(isLeapYear(2100));
```

يجب أن تُعيد الدالة `isLeapYear(2000)` القيمة `true`.

```js
assert.isTrue(isLeapYear(2000));
```

يجب أن تُعيد الدالة `isLeapYear(1999)` القيمة `false`.

```js
assert.isFalse(isLeapYear(1999));
```

يجب أن تُعيد الدالة `isLeapYear(2040)` القيمة `true`.

```js
assert.isTrue(isLeapYear(2040));
```

يجب أن تُعيد الدالة `isLeapYear(2026)` القيمة `false`.

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
