---
id: 69306364df283fcaff2e1ad6
title: "تحدٍّ 145: العدد النوني في متتالية فيبوناتشي"
challengeType: 28
dashedName: challenge-145
---

# --description--

بالنظر إلى عدد صحيح `n`، أعد العدد `n` في متتالية فيبوناتشي.

متتالية فيبوناتشي هي سلسلة أعداد يكون كل عدد فيها مجموع العددين السابقين له. أول 10 أعداد في المتتالية هي `0`، `1`، `1`، `2`، `3`، `5`، `8`، `13`، `21`، `34`.

# --hints--

يجب أن تُعيد الدالة `nthFibonacci(4)` القيمة `2`.

```js
assert.equal(nthFibonacci(4), 2);
```

يجب أن تُعيد الدالة `nthFibonacci(10)` القيمة `34`.

```js
assert.equal(nthFibonacci(10), 34);
```

يجب أن تُعيد الدالة `nthFibonacci(15)` القيمة `377`.

```js
assert.equal(nthFibonacci(15), 377);
```

يجب أن تُعيد الدالة `nthFibonacci(40)` القيمة `63245986`.

```js
assert.equal(nthFibonacci(40), 63245986);
```

يجب أن تُعيد الدالة `nthFibonacci(75)` القيمة `1304969544928657`.

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
