---
id: 69306364df283fcaff2e1ad9
title: "تحدٍّ 148: ضغط الإطارات"
challengeType: 28
dashedName: challenge-148
---

# --description--

معطى مصفوفة بأربعة أرقام تمثل ضغوط الإطارات بوحدة psi لأربعة إطارات في سيارتك، ومصفوفة أخرى برقمين تمثلان الحد الأدنى والحد الأقصى للضغط لإطاراتك بوحدة bar، أرجع مصفوفة بأربعة سلاسل نصية تصف حالة كل إطار.

- 1 bar يساوي 14.5038 psi.

أرجع مصفوفة بالقيم التالية لكل إطار:

- `"Low"` إذا كان ضغط الإطار أقل من الحد الأدنى المسموح.
- `"Good"` إذا كان بين الحد الأدنى والحد الأقصى المسموح.
- `"High"` إذا كان أعلى من الحد الأقصى المسموح.

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])` يجب أن تُرجع `["Good", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])` يجب أن تُرجع `["Good", "Low", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])` يجب أن تُرجع `["Low", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])` يجب أن تُرجع `["High", "High", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])` يجب أن تُرجع `["Good", "Good", "Good", "Good"]`.

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
