---
id: 69306364df283fcaff2e1ada
title: "تحدٍّ 149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

أعطِ سلسلة نصية، وأرجع سلسلة جديدة حيث تُحوّل جميع الحروف المتحركة إلى أحرف كبيرة، وجميع الحروف الأبجدية الأخرى إلى أحرف صغيرة.

- الحروف المتحركة هي `"a"`، `"e"`، `"i"`، `"o"`، و`"u"` بأي حالة.
- يجب أن تبقى الأحرف غير الأبجدية دون تغيير.

# --hints--

`vowelCase("vowelcase")` يجب أن تُرجع `"vOwElcAsE"`.

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` يجب أن تُرجع `"cOdIng Is fUn"`.

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` يجب أن تُرجع `"hEllO, wOrld!"`.

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` يجب أن تُرجع `"gIt chErry-pIck"`.

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` يجب أن تُرجع `"hEAd~1"`.

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
