---
id: 69306364df283fcaff2e1ad7
title: "تحدٍّ 146: المقعد لليد اليسرى على الطاولة"
challengeType: 28
dashedName: challenge-146
---

# --description--

معطى مصفوفة 4x2 (مصفوفة من مصفوفات) تمثل ترتيب المقاعد لوجبة، حدّد عدد المقاعد التي يمكن للشخص الأيسر الجلوس عليها.

- لا يمكن للشخص الأيسر الجلوس حيث يجلس شخص أيمن في المقعد المجاور على اليسار مباشرةً.

في المصفوفة المعطاة:

- المقعد الذي يحتوي على `"R"` يشغله شخص أيمن.
- المقعد الذي يحتوي على `"L"` يشغله شخص أيسر.
- المقعد الذي يحتوي على `"U"` هو مقعد غير مشغول.
- المقاعد المتاحة للجلوس هي فقط المقاعد غير المشغولة.
- المقاعد في الصف العلوي تواجه "للأسفل"، والمقاعد في الصف السفلي تواجه "للأعلى" (مثل الطاولة)، لذا اليسار واليمين نسبيان لاتجاه المقعد.
- المقاعد في الزوايا لها مقعد واحد فقط بجانبها.

على سبيل المثال، في المصفوفة المعطاة:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

لا يمكن الجلوس في المقعد العلوي الأيسر لأن هناك شخصًا أيمن على اليسار. يمكن الجلوس في المقعدين المفتوحين الآخرين لأنه لا يوجد شخص أيمن على اليسار.

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` يجب أن تُرجع `2`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` يجب أن تُرجع `8`.

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` يجب أن تُرجع `0`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` يجب أن تُرجع `1`.

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` يجب أن تُرجع `5`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```
