---
id: 69373793f5a867f769cde139
title: "تحدٍّ 154: المعدل الطبيعي للحفرة"
challengeType: 28
dashedName: challenge-154
---

# --description--

معطى عددان صحيحان، المعدل الطبيعي لحفرة الغولف وعدد خطوط الرسم التي استخدمها اللاعب في تلك الحفرة، أرجع نتيجة اللاعب باستخدام مصطلحات الغولف.

أرجع:

- `"Hole in one!"` إذا استغرق خط رسم واحد.
- `"Eagle"` إذا استغرق خطي رسم أقل من المعدل الطبيعي.
- `"Birdie"` إذا استغرق خط رسم أقل من المعدل الطبيعي.
- `"Par"` إذا استغرق نفس عدد خطوط الرسم كالمعدل الطبيعي.
- `"Bogey"` إذا استغرق خط رسم أكثر من المعدل الطبيعي.
- `"Double bogey"` إذا استغرق خطي رسم أكثر من المعدل الطبيعي.

# --hints--

`golfScore(3, 3)` يجب أن تُرجع `"Par"`.

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` يجب أن تُرجع `"Birdie"`.

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` يجب أن تُرجع `"Hole in one!"`.

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` يجب أن تُرجع `"Double bogey"`.

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` يجب أن تُرجع `"Bogey"`.

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` يجب أن تُرجع `"Eagle"`.

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
