---
id: 69373793f5a867f769cde135
title: "تحدٍّ 150: محلل قائمة غير مرتبة في Markdown"
challengeType: 28
dashedName: challenge-150
---

# --description--

معطى سلسلة نصية لقائمة غير مرتبة صحيحة في Markdown، أرجع سلسلة HTML المكافئة.

تتكون القائمة غير المرتبة من عنصر أو أكثر من عناصر القائمة. يظهر عنصر القائمة الصحيح في سطر منفصل و:

- يبدأ بشرطة (`"-"`)، تليها
- مسافة واحدة على الأقل، ثم
- نص عنصر القائمة.

تعطى القائمة كسلسلة واحدة مع فواصل أسطر مفصولة بحرف السطر الجديد (`"\n"`). لا تضمّن أحرف السطر الجديد في نص العنصر.

لف كل عنصر قائمة داخل وسم HTML `li`، ولجميع العناصر داخل وسم `ul`.

على سبيل المثال، معطى `"- Item A\n- Item B"`، أرجع `"<ul><li>Item A</li><li>Item B</li></ul>"`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML في السلاسل عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم المضمنة.

# --hints--

`parseUnorderedList("- Item A\n- Item B")` يجب أن تُرجع `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")` يجب أن تُرجع `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` يجب أن تُرجع `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")` يجب أن تُرجع `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
