---
id: 69373793f5a867f769cde137
title: "تحدٍّ 152: عدد أولي دائري"
challengeType: 28
dashedName: challenge-152
---

# --description--

أعطِ عددًا صحيحًا، حدّد ما إذا كان عددًا أوليًا دائريًا.

العدد الأولي الدائري هو عدد صحيح تكون جميع تدويراته لأرقامه أعدادًا أولية.

على سبيل المثال، `197` هو عدد أولي دائري لأن جميع تدويراته لأرقامه: `197`، `971`، و `719`، أعداد أولية.

# --hints--

يجب أن تُرجع الدالة `isCircularPrime(197)` القيمة `true`.

```js
assert.isTrue(isCircularPrime(197));
```

يجب أن تُرجع الدالة `isCircularPrime(23)` القيمة `false`.

```js
assert.isFalse(isCircularPrime(23));
```

يجب أن تُرجع الدالة `isCircularPrime(13)` القيمة `true`.

```js
assert.isTrue(isCircularPrime(13));
```

يجب أن تُرجع الدالة `isCircularPrime(89)` القيمة `false`.

```js
assert.isFalse(isCircularPrime(89));
```

يجب أن تُرجع الدالة `isCircularPrime(1193)` القيمة `true`.

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
