---
id: 69373793f5a867f769cde138
title: "تحدٍّ 153: لعبة إكس-أو"
challengeType: 28
dashedName: challenge-153
---

# --description--

معطاة مصفوفة 3×3 (مصفوفة من مصفوفات) تمثل لعبة إكس-أو مكتملة، حدّد الفائز.

- كل عنصر في المصفوفة المعطاة هو إما `"X"` أو `"O"`.

يفوز اللاعب إذا كان لديه ثلاثة من علاماته متتالية - أفقيًا أو عموديًا أو قطريًا.

أرجع:

- `"X wins"` إذا كان اللاعب X لديه ثلاثة متتالية.
- `"O wins"` إذا كان اللاعب O لديه ثلاثة متتالية.
- `"Draw"` إذا لم يكن لدى أي لاعب ثلاثة متتالية.

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` يجب أن تُرجع `"X wins"`.

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])` يجب أن تُرجع `"X wins"`.

```js
assert.equal(ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` يجب أن تُرجع `"Draw"`.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]])` يجب أن تُرجع `"O wins"`.

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` يجب أن تُرجع `"X wins"`.

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` يجب أن تُرجع `"Draw"`.

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```
