---
id: 69373793f5a867f769cde13a
title: "تحدٍّ 155: زراعة المحصول"
challengeType: 28
dashedName: challenge-155
---

# --description--

معطى عدد صحيح يُمثّل حجم حقل مزرعتك، و`"acres"` أو `"hectares"` التي تمثّل وحدة قياس حجم الحقل، ونوع المحصول، حدّد عدد النباتات من ذلك النوع التي يمكنك زراعتها في الحقل.

- 1 فدان يساوي 4046.86 مترًا مربعًا.
- 1 هكتار يساوي 10,000 متر مربع.

إليك قائمة بالمحاصيل التي ستُعطى كمُدخَل والمساحة التي يشغلها نبات واحد منها:

|المحصول|المساحة لكل نبات|
|-|-|
|`"corn"`|1 متر مربع|
|`"wheat"`|0.1 متر مربع|
|`"soybeans"`|0.5 متر مربع|
|`"tomatoes"`|0.25 متر مربع|
|`"lettuce"`|0.2 متر مربع|

أرجع عدد النباتات التي تتسع لها المساحة، مقربة إلى أقرب عدد صحيح للأسفل.

# --hints--

`getNumberOfPlants(1, "acres", "corn")` يجب أن تُرجع `4046`.

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` يجب أن تُرجع `100000`.

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` يجب أن تُرجع `161874`.

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` يجب أن تُرجع `150000`.

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` يجب أن تُرجع `271139`.

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
