---
id: 6939b873185d8e00d453563d
title: "تحدٍّ 158: تبديل المصفوفة"
challengeType: 28
dashedName: challenge-158
---

# --description--

أعطِ مصفوفة تحتوي على قيمتين، أرجع مصفوفة مع تبديل القيمتين.

على سبيل المثال، إذا كانت المصفوفة `["A", "B"]` فأرجع `["B", "A"]`.

# --hints--

يجب أن تُرجع الدالة `arraySwap(["A", "B"])` القيمة `["B", "A"]`.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

يجب أن تُرجع الدالة `arraySwap([25, 20])` القيمة `[20, 25]`.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

يجب أن تُرجع الدالة `arraySwap([true, false])` القيمة `[false, true]`.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

يجب أن تُرجع الدالة `arraySwap(["1", 1])` القيمة `[1, "1"]`.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
