---
id: 6939b873185d8e00d453563c
title: "تحدٍّ 157: محلل روابط Markdown"
challengeType: 28
dashedName: challenge-157
---

# --description--

بالنظر إلى سلسلة نصية لرابط في Markdown، أعد السلسلة المكافئة بلغة HTML.

صورة Markdown لها الصيغة التالية: `"[link_text](link_url)"`. أعد سلسلة وسم HTML `a` مع تعيين `href` إلى `link_url` و`link_text` كمحتوى الوسم.

على سبيل المثال، إذا كانت السلسلة `"[freeCodeCamp](https://freecodecamp.org/)"` فأعد `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML داخل السلاسل النصية عند تسجيل الرسائل — تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")` يجب أن تعيد `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` يجب أن تعيد `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` يجب أن تعيد `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
