---
id: 6939b873185d8e00d453563e
title: "تحدٍّ 159: الوتر عدد صحيح"
challengeType: 28
dashedName: challenge-159
---

# --description--

معطى عددان صحيحان موجبان يمثلان أطوال الضلعين القائمين (الجانبين القصيرين) لمثلث قائم، حدّد ما إذا كان الوتر عددًا صحيحًا.

يُحسب طول الوتر بجمع مربعات طولي الضلعين القائمين ثم أخذ الجذر التربيعي للمجموع (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

يجب أن تُرجع الدالة `isIntegerHypotenuse(3, 4)` القيمة `true`.

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

يجب أن تُرجع الدالة `isIntegerHypotenuse(2, 3)` القيمة `false`.

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

يجب أن تُرجع الدالة `isIntegerHypotenuse(5, 12)` القيمة `true`.

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

يجب أن تُرجع الدالة `isIntegerHypotenuse(10, 10)` القيمة `false`.

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

يجب أن تُرجع الدالة `isIntegerHypotenuse(780, 1040)` القيمة `true`.

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

يجب أن تُرجع الدالة `isIntegerHypotenuse(250, 333)` القيمة `false`.

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
