---
id: 6939b873185d8e00d453563b
title: "تحدٍّ 156: فردي أم زوجي؟"
challengeType: 28
dashedName: challenge-156
---

# --description--

أعطِ عددًا صحيحًا موجبًا، أرجع `"Odd"` إذا كان عددًا فرديًا، و`"Even"` إذا كان زوجيًا.

# --hints--

`oddOrEven(1)` يجب أن تُرجع `"Odd"`.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` يجب أن تُرجع `"Even"`.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` يجب أن تُرجع `"Odd"`.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` يجب أن تُرجع `"Even"`.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` يجب أن تُرجع `"Odd"`.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
