---
id: 694596b0585c11170ac7c7f9
title: "تحدٍّ 161: الشحن المجاني"
challengeType: 28
dashedName: challenge-161
---

# --description--

معطاة مصفوفة من السلاسل النصية تمثل العناصر في عربة التسوق الخاصة بك، ورقم يمثل الحد الأدنى لمبلغ الطلب للتأهل للشحن المجاني، حدّد ما إذا كانت العناصر في عربة التسوق الخاصة بك تؤهل للشحن المجاني.

ستحتوي المصفوفة المعطاة على عناصر من القائمة أدناه:

| العنصر | السعر |
| - | - |
| `"shirt"`	| 34.25 |
| `"jeans"`	| 48.50 |
| `"shoes"` |	75.00 |
| `"hat"`	| 19.95 |
| `"socks"` | 15.00 |
| `"jacket"` | 109.95 |

# --hints--

`getsFreeShipping(["shoes"], 50)` يجب أن تُرجع `true`.

```js
assert.isTrue(getsFreeShipping(["shoes"], 50));
```

`getsFreeShipping(["hat", "socks"], 50)` يجب أن تُرجع `false`.

```js
assert.isFalse(getsFreeShipping(["hat", "socks"], 50));
```

`getsFreeShipping(["jeans", "shirt", "jacket"], 75)` يجب أن تُرجع `true`.

```js
assert.isTrue(getsFreeShipping(["jeans", "shirt", "jacket"], 75));
```

`getsFreeShipping(["socks", "socks", "hat"], 75)` يجب أن تُرجع `false`.

```js
assert.isFalse(getsFreeShipping(["socks", "socks", "hat"], 75));
```

`getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100)` يجب أن تُرجع `true`.

```js
assert.isTrue(getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100));
```

`getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200)` يجب أن تُرجع `false`.

```js
assert.isFalse(getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200));
```

# --seed--

## --seed-contents--

```js
function getsFreeShipping(cart, minimum) {

  return cart;
}
```

# --solutions--

```js
function getsFreeShipping(cart, minimum) {
  const prices = {
    shirt: 34.25,
    jeans: 48.50,
    shoes: 75.00,
    hat: 19.95,
    socks: 15.00,
    jacket: 109.95
  };

  let total = 0;

  for (const item of cart) {
    total += prices[item];
  }

  return total >= minimum;
}
```
