---
id: 694596b0585c11170ac7c7fa
title: "تحدٍّ 162: استهلاك الطاقة"
challengeType: 28
dashedName: challenge-162
---

# --description--

بالنظر إلى عدد السعرات الحرارية المحروقة أثناء التمرين، وعدد واط-ساعات التي استخدمتها أجهزتك الإلكترونية خلال ذلك التمرين، حدّد أيهما استهلك طاقة أكثر.

لمقارنتها، حوّل القيمتين إلى جول باستخدام التحويلات التالية:

- 1 سعر حراري يساوي 4184 جول.
- 1 واط-ساعة يساوي 3600 جول.

أرجع:

- `"Workout"` إذا استهلك التمرين طاقة أكثر.
- `"Devices"` إذا استهلك الجهاز طاقة أكثر.
- `"Equal"` إذا استهلك كلاهما نفس كمية الطاقة.

# --hints--

`compareEnergy(250, 50)` يجب أن تُرجع `"Workout"`.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` يجب أن تُرجع `"Devices"`.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` يجب أن تُرجع `"Equal"`.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` يجب أن تُرجع `"Workout"`.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` يجب أن تُرجع `"Devices"`.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` يجب أن تُرجع `"Equal"`.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
