---
id: 694596b0585c11170ac7c7fb
title: "تحدٍّ 163: حالة الحروف الساكنة"
challengeType: 28
dashedName: challenge-163
---

# --description--

معطاة سلسلة نصية تمثل اسم متغير، حوّلها إلى حالة الحروف الساكنة باستخدام القواعد التالية:

- يجب تحويل جميع الحروف الساكنة إلى أحرف كبيرة.
- يجب تحويل جميع الحروف المتحركة (`a`، `e`، `i`، `o`، `u` بأي حالة) إلى أحرف صغيرة.
- يجب تحويل جميع الشرطات (`-`) إلى شرطات سفلية (`_`).

# --hints--

`toConsonantCase("helloworld")` يجب أن تُرجع `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` يجب أن تُرجع `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")` يجب أن تُرجع `"_HeLLo_WoRLD_"`.

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` يجب أن تُرجع `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
