---
id: 694596b0585c11170ac7c7fc
title: "تحدٍّ 164: محلل كود مضمَّن في Markdown"
challengeType: 28
dashedName: challenge-164
---

# --description--

أعطِ سلسلة نصية من Markdown تحتوي على كتلة أو أكثر من كود مضمَّن، أرجع سلسلة HTML المكافئة.

تستخدم كتل الكود المضمَّن في Markdown علامة backtick واحدة (``` ` ```) في بداية ونهاية نص الكود.

أرجع السلسلة المعطاة مع تحويل جميع كتل الكود إلى وسم HTML `code`.

على سبيل المثال، بالنسبة للسلسلة النصية ```"Use `let` to declare the variable."```، أعد `"Use <code>let</code> to declare the variable."`.

ملاحظة: قد لا تعرض وحدة التحكم وسوم HTML في السلاسل عند تسجيل الرسائل. تحقق من وحدة تحكم المتصفح لرؤية السجلات مع الوسوم مضمنة.

# --hints--

```parseInlineCode("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
assert.equal(parseInlineCode("Use `let` to declare the variable."), "Use <code>let</code> to declare the variable.");
```

```parseInlineCode("Use `let` or `const` to declare a variable.")``` should return `"Use <code>let</code> or <code>const</code> to declare a variable."`.

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

```parseInlineCode("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
assert.equal(parseInlineCode("Run `npm install` then `npm start`."), "Run <code>npm install</code> then <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
