---
id: 696655d24b614176d4c9b78d
title: "تحدٍّ 170: يوم فردي أم زوجي"
challengeType: 28
dashedName: challenge-170
---

# --description--

معطى طابع زمني (عدد المللي ثانية منذ بداية عصر Unix)، أرجع:

- `"odd"` إذا كان يوم الشهر لذلك الطابع الزمني فرديًا.
- `"even"` إذا كان يوم الشهر لذلك الطابع الزمني زوجيًا.

على سبيل المثال، معطى `1769472000000`، وهو طابع زمني ليوم 27 يناير 2026، أرجع `"odd"` لأن رقم اليوم (`27`) عدد فردي.

ملاحظة: الطابع الزمني بالمللي ثانية ويجب استخدام التاريخ في توقيت UTC، وليس في توقيتك المحلي.

# --hints--

`oddOrEvenDay(1769472000000)` يجب أن تُرجع `"odd"`.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` يجب أن تُرجع `"even"`.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` يجب أن تُرجع `"odd"`.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` يجب أن تُرجع `"odd"`.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` يجب أن تُرجع `"even"`.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
