---
id: 696655d24b614176d4c9b78a
title: "تحدٍّ 167: حرف بينغو!"
challengeType: 28
dashedName: challenge-167
---

# --description--

عند إعطائك رقمًا، أعد الحرف الخاص بلعبة بينغو المرتبط به (بحروف كبيرة). تُصنّف أرقام بينغو كما يلي:

| الحرف | نطاق الأرقام |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

يجب أن تُعيد الدالة `getBingoLetter(75)` القيمة `"O"`.

```js
assert.equal(getBingoLetter(75), "O");
```

يجب أن تُعيد الدالة `getBingoLetter(54)` القيمة `"G"`.

```js
assert.equal(getBingoLetter(54), "G");
```

يجب أن تُعيد الدالة `getBingoLetter(25)` القيمة `"I"`.

```js
assert.equal(getBingoLetter(25), "I");
```

يجب أن تُعيد الدالة `getBingoLetter(38)` القيمة `"N"`.

```js
assert.equal(getBingoLetter(38), "N");
```

يجب أن تُعيد الدالة `getBingoLetter(11)` القيمة `"B"`.

```js
assert.equal(getBingoLetter(11), "B");
```

# --seed--

## --seed-contents--

```js
function getBingoLetter(n) {

  return n;
}
```

# --solutions--

```js
function getBingoLetter(n) {
  if (n >= 1 && n <= 15) return "B";
  if (n >= 16 && n <= 30) return "I";
  if (n >= 31 && n <= 45) return "N";
  if (n >= 46 && n <= 60) return "G";
  if (n >= 61 && n <= 75) return "O";
  return null;
}
```
