---
id: 696655d24b614176d4c9b789
title: "تحدٍّ 166: مدقق الألوان الست عشرية"
challengeType: 28
dashedName: challenge-166
---

# --description--

معطاة سلسلة نصية، حدّد ما إذا كانت لون CSS سداسي عشري صالحًا. يجب أن يكون لون CSS السداسي العشري الصالح:

- يبدأ بـ `#`، و
- يتبعه إما 3 أو 6 أحرف سداسية عشرية.

الأحرف السداسية العشرية هي الأرقام من `0` إلى `9` والحروف من `a` إلى `f` (غير حساسة لحالة الحروف).

# --hints--

يجب أن تُرجع الدالة `isValidHex("#123")` القيمة `true`.

```js
assert.isTrue(isValidHex("#123"));
```

يجب أن تُرجع الدالة `isValidHex("#123abc")` القيمة `true`.

```js
assert.isTrue(isValidHex("#123abc"));
```

يجب أن تُرجع الدالة `isValidHex("#ABCDEF")` القيمة `true`.

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

يجب أن تُرجع الدالة `isValidHex("#0a1B2c")` القيمة `true`.

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

يجب أن تُرجع الدالة `isValidHex("#12G")` القيمة `false`.

```js
assert.isFalse(isValidHex("#12G"));
```

يجب أن تُرجع الدالة `isValidHex("#1234567")` القيمة `false`.

```js
assert.isFalse(isValidHex("#1234567"));
```

يجب أن تُرجع الدالة `isValidHex("#12 3")` القيمة `false`.

```js
assert.isFalse(isValidHex("#12 3"));
```

يجب أن تُرجع الدالة `isValidHex("fff")` القيمة `false`.

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
