---
id: 696655d24b614176d4c9b78b
title: "تحدٍّ 168: صورة بمقياس"
challengeType: 28
dashedName: challenge-168
---

# --description--

معطى سلسلة نصية تمثل عرض وارتفاع صورة، ورقم لتكبير الصورة، أرجع العرض والارتفاع بعد التكبير.

- السلسلة النصية المُدخلة تكون بصيغة `"WxH"`. على سبيل المثال، `"800x600"`.
- المقياس هو رقم يُستخدم في الضرب مع العرض والارتفاع.

أرجع الأبعاد بعد التكبير بنفس صيغة `"WxH"`.

# --hints--

يجب أن تُرجع الدالة `scaleImage("800x600", 2)` القيمة `"1600x1200"`.

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

يجب أن تُرجع الدالة `scaleImage("100x100", 10)` القيمة `"1000x1000"`.

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

يجب أن تُرجع الدالة `scaleImage("1024x768", 0.5)` القيمة `"512x384"`.

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

يجب أن تُرجع الدالة `scaleImage("300x200", 1.5)` القيمة `"450x300"`.

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
