---
id: 696655d24b614176d4c9b78c
title: "تحدٍّ 169: نسخة مصغرة من FizzBuzz"
challengeType: 28
dashedName: challenge-169
---

# --description--

بالنظر إلى عدد صحيح، أعد سلسلة نصية بناءً على القواعد التالية:

- إذا كان العدد يقبل القسمة على 3، أعد `"Fizz"`.
- إذا كان العدد يقبل القسمة على 5، أعد `"Buzz"`.
- إذا كان العدد يقبل القسمة على كل من 3 و5، أعد `"FizzBuzz"`.
- خلاف ذلك، أعد العدد المعطى كسلسلة نصية.

# --hints--

`fizzBuzzMini(3)` يجب أن تُعيد `"Fizz"`.

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` يجب أن تُعيد `"4"`.

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` يجب أن تُعيد `"Buzz"`.

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` يجب أن تُعيد `"FizzBuzz"`.

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` يجب أن تُعيد `"98"`.

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
