---
id: 69738771fb5a7b8b24cca2a5
title: "تحدٍّ 179: النقود في الجيب"
challengeType: 28
dashedName: challenge-179
---

# --description--

معطاة مصفوفة من الأعداد الصحيحة تمثل العملات في جيبك، حيث كل عدد صحيح هو قيمة عملة بالسنت، أرجع المبلغ الإجمالي بالتنسيق `"$D.CC"`.

- 100 سنت تساوي دولارًا واحدًا.
- في قيمة الإرجاع، أدرج صفرًا بادئًا للمبالغ الأقل من دولار واحد ودوماً رقمين بالضبط للسنتات.

# --hints--

`countChange([25, 10, 5, 1])` يجب أن تُرجع `"$0.41"`.

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` يجب أن تُرجع `"$1.43"`.

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` يجب أن تُرجع `"$37.55"`.

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` يجب أن تُرجع `"$0.70"`.

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` يجب أن تُرجع `"$0.01"`.

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` يجب أن تُرجع `"$1.00"`.

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
