---
id: 69738771fb5a7b8b24cca2a4
title: "تحدٍّ 178: اقتطاع النص"
challengeType: 28
dashedName: challenge-178
---

# --description--

أعطِ سلسلة نصية، أعدها كما هي إذا كانت 20 حرفًا أو أقل. إذا كانت أطول من 20 حرفًا، اقتطع أول 17 حرفًا وألحق بها `"..."` في النهاية (بحيث يصبح المجموع 20 حرفًا) ثم أعد النتيجة.

# --hints--

`truncateText("Hello, world!")` يجب أن تُعيد `"Hello, world!"`.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` يجب أن تُعيد `"This string shoul..."`.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` يجب أن تُعيد `"Exactly twenty chars"`.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` يجب أن تُعيد `"...................."`.

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
