---
id: 69738771fb5a7b8b24cca2a2
title: "تحدٍّ 176: يوم الغراب"
challengeType: 28
dashedName: challenge-176
---

# --description--

اليوم هو يوم الغراب، حيث يتنبأ الغراب بالطقس بناءً على ما إذا كان يرى ظله أم لا.

معطى قيمة تمثل ظهور الغراب، أرجع التنبؤ الصحيح:

- إذا كانت القيمة المعطاة هي القيمة المنطقية `true` (رأى القندس ظله)، فأرجع `"Looks like we'll have six more weeks of winter."`.
- إذا كانت القيمة هي القيمة المنطقية `false` (لم يرَ القندس ظله)، فأرجع `"It's going to be an early spring."`.
- إذا كانت القيمة أي شيء آخر (لم يظهر القندس)، فأرجع `"No prediction this year."`.

# --hints--

يجب أن تُعيد الدالة `groundhogDayPrediction(true)` القيمة `"Looks like we'll have six more weeks of winter."`.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

يجب أن تُعيد الدالة `groundhogDayPrediction(false)` القيمة `"It's going to be an early spring."`.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

يجب أن تُعيد الدالة `groundhogDayPrediction(null)` القيمة `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` يجب أن تُعيد `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` يجب أن تُعيد القيمة `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
