---
id: 69738771fb5a7b8b24cca29f
title: "تحدٍّ 173: تحركات البيدق الصالحة"
challengeType: 28
dashedName: challenge-173
---

# --description--

بالنظر إلى موقع أحد بيدقك على رقعة الشطرنج، أرجع مصفوفة بكل المربعات الصالحة التي يمكنه التحرك إليها بترتيب تصاعدي.

رقعة الشطرنج القياسية هي 8×8، مع أعمدة معنونة من `A` إلى `H` (من اليسار إلى اليمين) وصفوف معنونة من `1` إلى `8` (من الأسفل إلى الأعلى). تبدو هكذا:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

لهذا التحدٍّ:

- أنت اللاعب في أسفل الرقعة.
- يمكن للبيدق التحرك مربعًا واحدًا فقط "أعلى".
- إلا إذا كان البيدق في الصف الابتدائي (الصف 2)، حينها يمكنه التحرك مربعًا أو مربعين لأعلى.

على سبيل المثال، إذا كان الإدخال `"D4"`، أرجع `["D5"]`، وهو المربع الوحيد الذي يمكن لبيدقك التحرك إليه.
إذا كان الإدخال `"B2"`، أرجع `["B3", "B4"]`، لأنه في الصف الابتدائي ويجب أن يكون الترتيب تصاعديًا.

# --hints--

`findPawnMoves("D4")` يجب أن تُرجع `["D5"]`.

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")` يجب أن تُرجع `["B3", "B4"]`.

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")` يجب أن تُرجع `["A8"]`.

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")` يجب أن تُرجع `["G3", "G4"]`.

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")` يجب أن تُرجع `["E4"]`.

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
