---
id: 69738771fb5a7b8b24cca2a3
title: "تحدٍّ 177: مرآة السلسلة النصية"
challengeType: 28
dashedName: challenge-177
---

# --description--

بالنظر إلى سلسلة نصية، أعد سلسلة جديدة تتكون من السلسلة النصية المعطاة مع نسخة معكوسة منها ملحقة في نهايتها.

# --hints--

`mirror("freeCodeCamp")` يجب أن تُعيد `"freeCodeCamppmaCedoCeerf"`.

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` يجب أن تُعيد `"RaceCarraCecaR"`.

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` يجب أن تُعيد `"helloworlddlrowolleh"`.

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` يجب أن تُعيد `"The quick brown fox......xof nworb kciuq ehT"`.

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
