---
id: 69738771fb5a7b8b24cca2a1
title: "تحدٍّ 175: التخلص الرقمي"
challengeType: 28
dashedName: challenge-175
---

# --description--

معطاة مصفوفة من سجلات تسجيل الدخول الخاصة بك، حدّد ما إذا كنت قد حققت هدف التخلص الرقمي.

كل سجل هو سلسلة نصية بالصيغة `"YYYY-MM-DD HH:mm:ss"`.

لقد حققت هدف التخلص الرقمي إذا كانت العبارتان التاليتان صحيحتين:

- لم تسجل الدخول أكثر من مرة واحدة خلال أي فترة أربع ساعات.
- لم تسجل الدخول أكثر من مرتين في أي يوم واحد.

# --hints--

`digitalDetox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"])` يجب أن تُرجع `true`.

```js
assert.isTrue(digitalDetox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"]));
```

`digitalDetox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"])` يجب أن تُرجع `false`.

```js
assert.isFalse(digitalDetox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"]));
```

`digitalDetox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"])` يجب أن تُرجع `true`.

```js
assert.isTrue(digitalDetox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"]));
```

`digitalDetox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"])` يجب أن تُرجع `false`.

```js
assert.isFalse(digitalDetox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"]));
```

`digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])` يجب أن تُرجع `true`.

```js
assert.isTrue(digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 20:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]));
```

`digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])` يجب أن تُرجع `false`.

```js
assert.isFalse(digitalDetox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]));
```

# --seed--

## --seed-contents--

```js
function digitalDetox(logs) {

  return logs;
}
```

# --solutions--

```js
function digitalDetox(logs) {
  if (logs.length === 0) return true;

  const times = logs
    .map(log => new Date(log.replace(" ", "T") + "Z").getTime())
    .sort((a, b) => a - b);

  const ONE_DAY = 24 * 60 * 60 * 1000;
  const FOUR_HOURS = 4 * 60 * 60 * 1000;

  for (let i = 1; i < times.length; i++) {
    if (times[i] - times[i - 1] < FOUR_HOURS) {
      return false;
    }
  }

  const dailyCounts = {};
  for (const time of times) {
    const day = new Date(time).toISOString().slice(0, 10);
    dailyCounts[day] = (dailyCounts[day] || 0) + 1;
    if (dailyCounts[day] > 2) {
      return false;
    }
  }

  return true;
}
```
