---
id: 69738771fb5a7b8b24cca2a0
title: "تحدٍّ 174: مكتشف الأبراج"
challengeType: 28
dashedName: challenge-174
---

# --description--

بالنظر إلى سلسلة نصية للتاريخ بصيغة `"YYYY-MM-DD"`، أرجع برج الأبراج لذلك التاريخ باستخدام المخطط التالي:

| نطاق التاريخ | برج الأبراج |
| - | - |
| 21 مارس - 19 أبريل | `"Aries"` |
| 20 أبريل - 20 مايو | `"Taurus"` |
| 21 مايو - 20 يونيو | `"Gemini"` |
| 21 يونيو - 22 يوليو | `"Cancer"` |
| 23 يوليو - 22 أغسطس | `"Leo"` |
| 23 أغسطس - 22 سبتمبر | `"Virgo"` |
| 23 سبتمبر - 22 أكتوبر | `"Libra"` |
| 23 أكتوبر - 21 نوفمبر | `"Scorpio"` |
| 22 نوفمبر - 21 ديسمبر | `"Sagittarius"` |
| 22 ديسمبر - 19 يناير | `"Capricorn"` |
| 20 يناير - 18 فبراير | `"Aquarius"` |
| 19 فبراير - 20 مارس | `"Pisces"` |

- تعتمد الأبراج فقط على الشهر واليوم، يمكنك تجاهل السنة.

# --hints--

`getSign("2026-01-31")` يجب أن تُرجع `"Aquarius"`.

```js
assert.equal(getSign("2026-01-31"), "Aquarius");
```

`getSign("2001-06-10")` يجب أن تُرجع `"Gemini"`.

```js
assert.equal(getSign("2001-06-10"), "Gemini");
```

`getSign("1985-09-07")` يجب أن تُرجع `"Virgo"`.

```js
assert.equal(getSign("1985-09-07"), "Virgo");
```

`getSign("2023-03-19")` يجب أن تُرجع `"Pisces"`.

```js
assert.equal(getSign("2023-03-19"), "Pisces");
```

`getSign("2045-11-05")` يجب أن تُرجع `"Scorpio"`.

```js
assert.equal(getSign("2045-11-05"), "Scorpio");
```

`getSign("1985-12-06")` يجب أن تُرجع `"Sagittarius"`.

```js
assert.equal(getSign("1985-12-06"), "Sagittarius");
```

`getSign("2025-12-30")` يجب أن تُرجع `"Capricorn"`.

```js
assert.equal(getSign("2025-12-30"), "Capricorn");
```

`getSign("2018-10-08")` يجب أن تُرجع `"Libra"`.

```js
assert.equal(getSign("2018-10-08"), "Libra");
```

`getSign("1958-05-04")` يجب أن تُرجع `"Taurus"`.

```js
assert.equal(getSign("1958-05-04"), "Taurus");
```

# --seed--

## --seed-contents--

```js
function getSign(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getSign(dateStr) {
  const [, monthStr, dayStr] = dateStr.split("-");
  const month = Number(monthStr);
  const day = Number(dayStr);

  if ((month === 3 && day >= 21) || (month === 4 && day <= 19)) return "Aries";
  if ((month === 4 && day >= 20) || (month === 5 && day <= 20)) return "Taurus";
  if ((month === 5 && day >= 21) || (month === 6 && day <= 20)) return "Gemini";
  if ((month === 6 && day >= 21) || (month === 7 && day <= 22)) return "Cancer";
  if ((month === 7 && day >= 23) || (month === 8 && day <= 22)) return "Leo";
  if ((month === 8 && day >= 23) || (month === 9 && day <= 22)) return "Virgo";
  if ((month === 9 && day >= 23) || (month === 10 && day <= 22)) return "Libra";
  if ((month === 10 && day >= 23) || (month === 11 && day <= 21)) return "Scorpio";
  if ((month === 11 && day >= 22) || (month === 12 && day <= 21)) return "Sagittarius";
  if ((month === 12 && day >= 22) || (month === 1 && day <= 19)) return "Capricorn";
  if ((month === 1 && day >= 20) || (month === 2 && day <= 18)) return "Aquarius";
  if ((month === 2 && day >= 19) || (month === 3 && day <= 20)) return "Pisces";
}
```
