---
id: 697a49e9860d24853adef67f
title: "تحدٍّ 194: ألعاب الشتاء 2026 اليوم 15: التزلج الحر"
challengeType: 28
dashedName: challenge-194
---

# --description--

معطى اسم خدعة يتكون من كلمتين، حدّد إذا كان اسم خدعة التزلج الحر هذا صحيحًا.

الخدعة صحيحة إذا كانت الكلمة الأولى في قائمة الكلمات الأولى الصحيحة، والكلمة الثانية في قائمة الكلمات الثانية الصحيحة.

- ستفصل بين الكلمتين مسافة واحدة فقط.

الكلمات الأولى الصحيحة:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

الكلمات الثانية الصحيحة:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`isValidTrick("Polar Vortex")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidTrick("Polar Vortex"));
```

`isValidTrick("Solar Icequake")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidTrick("Solar Icequake"));
```

`isValidTrick("Thunder Blizzard")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidTrick("Thunder Blizzard"));
```

`isValidTrick("Phantom Frostbite")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidTrick("Phantom Frostbite"));
```

`isValidTrick("Ghost Avalanche")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidTrick("Ghost Avalanche"));
```

`isValidTrick("Snowstorm Shadow")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidTrick("Snowstorm Shadow"));
```

`isValidTrick("Solar Sky")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidTrick("Solar Sky"));
```

# --seed--

## --seed-contents--

```js
function isValidTrick(trickName) {

  return trickName;
}
```

# --solutions--

```js
function isValidTrick(trickName) {
  const validFirst = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"];
  const validSecond = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"];

  const words = trickName.split(" ");
  const [first, second] = words;
  return validFirst.includes(first) && validSecond.includes(second);
}
```
