---
id: 697a49e6ff50d756c9b6935d
title: "تحدٍّ 180: ألعاب الشتاء اليوم 1: يوم الافتتاح"
challengeType: 28
dashedName: challenge-180
---

# --description--

ستقدم لك الأيام السبعة عشر القادمة تحديات برمجية مستوحاة من ألعاب الشتاء.

في التحدي الأول، يُعطى لك رمز دولة مكون من حرفين وتحتاج إلى إرجاع رمز العلم الخاص بتلك الدولة.

استخدم هذه القائمة:

| الدولة | الرمز | العلم |
| - | -: | :--: |
| ألبانيا                    |            `"AL"` | `"🇦🇱"` |
| أندورا                    |            `"AD"` | `"🇦🇩"` |
| الأرجنتين                  |            `"AR"` | `"🇦🇷"` |
| أرمينيا                    |            `"AM"` | `"🇦🇲"` |
| أستراليا                  |            `"AU"` | `"🇦🇺"` |
| النمسا                    |            `"AT"` | `"🇦🇹"` |
| أذربيجان                 |            `"AZ"` | `"🇦🇿"` |
| بلجيكا                    |            `"BE"` | `"🇧🇪"` |
| بنين                      |            `"BJ"` | `"🇧🇯"` |
| بوليفيا                    |            `"BO"` | `"🇧🇴"` |
| البوسنة والهرسك     |            `"BA"` | `"🇧🇦"` |
| البرازيل                     |            `"BR"` | `"🇧🇷"` |
| بلغاريا                   |            `"BG"` | `"🇧🇬"` |
| كندا                     |            `"CA"` | `"🇨🇦"` |
| تشيلي                      |            `"CL"` | `"🇨🇱"` |
| الصين                      |            `"CN"` | `"🇨🇳"` |
| كولومبيا                   |            `"CO"` | `"🇨🇴"` |
| كرواتيا                    |            `"HR"` | `"🇭🇷"` |
| قبرص                     |            `"CY"` | `"🇨🇾"` |
| جمهورية التشيك             |            `"CZ"` | `"🇨🇿"` |
| الدنمارك                    |            `"DK"` | `"🇩🇰"` |
| الإكوادور                    |            `"EC"` | `"🇪🇨"` |
| إريتريا                    |            `"ER"` | `"🇪🇷"` |
| إستونيا                    |            `"EE"` | `"🇪🇪"` |
| فنلندا                    |            `"FI"` | `"🇫🇮"` |
| فرنسا                     |            `"FR"` | `"🇫🇷"` |
| جورجيا                    |            `"GE"` | `"🇬🇪"` |
| ألمانيا                    |            `"DE"` | `"🇩🇪"` |
| بريطانيا العظمى              |            `"GB"` | `"🇬🇧"` |
| اليونان                     |            `"GR"` | `"🇬🇷"` |
| غينيا بيساو              |            `"GW"` | `"🇬🇼"` |
| هايتي                      |            `"HT"` | `"🇭🇹"` |
| هونغ كونغ                  |            `"HK"` | `"🇭🇰"` |
| المجر                    |            `"HU"` | `"🇭🇺"` |
| أيسلندا                    |            `"IS"` | `"🇮🇸"` |
| الهند                      |            `"IN"` | `"🇮🇳"` |
| إيران                       |            `"IR"` | `"🇮🇷"` |
| أيرلندا                    |            `"IE"` | `"🇮🇪"` |
| إسرائيل                     |            `"IL"` | `"🇮🇱"` |
| إيطاليا                      |            `"IT"` | `"🇮🇹"` |
| جامايكا                    |            `"JM"` | `"🇯🇲"` |
| اليابان                      |            `"JP"` | `"🇯🇵"` |
| كازاخستان                 |            `"KZ"` | `"🇰🇿"` |
| كينيا                      |            `"KE"` | `"🇰🇪"` |
| كوسوفو                     |            `"XK"` | `"🇽🇰"` |
| قيرغيزستان                 |            `"KG"` | `"🇰🇬"` |
| لاتفيا                     |            `"LV"` | `"🇱🇻"` |
| لبنان                    |            `"LB"` | `"🇱🇧"` |
| ليختنشتاين              |            `"LI"` | `"🇱🇮"` |
| ليتوانيا                  |            `"LT"` | `"🇱🇹"` |
| لوكسمبورغ                 |            `"LU"` | `"🇱🇺"` |
| مدغشقر                 |            `"MG"` | `"🇲🇬"` |
| ماليزيا                   |            `"MY"` | `"🇲🇾"` |
| مالطا                      |            `"MT"` | `"🇲🇹"` |
| المكسيك                     |            `"MX"` | `"🇲🇽"` |
| مولدوفا                    |            `"MD"` | `"🇲🇩"` |
| موناكو                     |            `"MC"` | `"🇲🇨"` |
| منغوليا                   |            `"MN"` | `"🇲🇳"` |
| الجبل الأسود                 |            `"ME"` | `"🇲🇪"` |
| المغرب                    |            `"MA"` | `"🇲🇦"` |
| هولندا                |            `"NL"` | `"🇳🇱"` |
| نيوزيلندا                |            `"NZ"` | `"🇳🇿"` |
| نيجيريا                    |            `"NG"` | `"🇳🇬"` |
| مقدونيا الشمالية            |            `"MK"` | `"🇲🇰"` |
| النرويج                     |            `"NO"` | `"🇳🇴"` |
| باكستان                   |            `"PK"` | `"🇵🇰"` |
| الفلبين                |            `"PH"` | `"🇵🇭"` |
| بولندا                     |            `"PL"` | `"🇵🇱"` |
| البرتغال                   |            `"PT"` | `"🇵🇹"` |
| بورتو ريكو                |            `"PR"` | `"🇵🇷"` |
| رومانيا                    |            `"RO"` | `"🇷🇴"` |
| سان مارينو                 |            `"SM"` | `"🇸🇲"` |
| السعودية               |            `"SA"` | `"🇸🇦"` |
| صربيا                     |            `"RS"` | `"🇷🇸"` |
| سنغافورة                  |            `"SG"` | `"🇸🇬"` |
| سلوفاكيا                   |            `"SK"` | `"🇸🇰"` |
| سلوفينيا                   |            `"SI"` | `"🇸🇮"` |
| جنوب أفريقيا               |            `"ZA"` | `"🇿🇦"` |
| كوريا الجنوبية                |            `"KR"` | `"🇰🇷"` |
| إسبانيا                      |            `"ES"` | `"🇪🇸"` |
| السويد                     |            `"SE"` | `"🇸🇪"` |
| سويسرا                |            `"CH"` | `"🇨🇭"` |
| تايلاند                   |            `"TH"` | `"🇹🇭"` |
| ترينيداد وتوباغو          |            `"TT"` | `"🇹🇹"` |
| تركيا                     |            `"TR"` | `"🇹🇷"` |
| أوكرانيا                    |            `"UA"` | `"🇺🇦"` |
| الإمارات العربية المتحدة       |            `"AE"` | `"🇦🇪"` |
| الولايات المتحدة              |            `"US"` | `"🇺🇸"` |
| أوروغواي                    |            `"UY"` | `"🇺🇾"` |
| أوزبكستان                 |            `"UZ"` | `"🇺🇿"` |
| فنزويلا                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` يجب أن تُرجع `"🇦🇱"`.

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` يجب أن تُرجع `"🇦🇩"`.

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` يجب أن تُرجع `"🇦🇷"`.

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` يجب أن تُرجع `"🇦🇲"`.

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` يجب أن تُرجع `"🇦🇺"`.

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` يجب أن تُرجع `"🇦🇹"`.

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` يجب أن تُرجع `"🇦🇿"`.

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` يجب أن تُرجع `"🇧🇪"`.

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` يجب أن تُرجع `"🇧🇯"`.

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` يجب أن تُرجع `"🇧🇴"`.

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` يجب أن تُرجع `"🇧🇦"`.

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` يجب أن تُرجع `"🇧🇷"`.

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` يجب أن تُرجع `"🇧🇬"`.

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` يجب أن تُرجع `"🇨🇦"`.

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` يجب أن تُرجع `"🇨🇱"`.

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` يجب أن تُرجع `"🇨🇳"`.

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` يجب أن تُرجع `"🇨🇴"`.

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` يجب أن تُرجع `"🇭🇷"`.

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` يجب أن تُرجع `"🇨🇾"`.

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` يجب أن تُرجع `"🇨🇿"`.

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` يجب أن تُرجع `"🇩🇰"`.

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` يجب أن تُرجع `"🇪🇨"`.

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` يجب أن تُرجع `"🇪🇷"`.

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` يجب أن تُرجع `"🇪🇪"`.

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` يجب أن تُرجع `"🇫🇮"`.

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` يجب أن تُرجع `"🇫🇷"`.

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` يجب أن تُرجع `"🇬🇪"`.

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` يجب أن تُرجع `"🇩🇪"`.

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` يجب أن تُرجع `"🇬🇧"`.

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` يجب أن تُرجع `"🇬🇷"`.

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` يجب أن تُرجع `"🇬🇼"`.

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` يجب أن تُرجع `"🇭🇹"`.

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` يجب أن تُرجع `"🇭🇰"`.

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` يجب أن تُرجع `"🇭🇺"`.

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` يجب أن تُرجع `"🇮🇸"`.

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` يجب أن تُرجع `"🇮🇳"`.

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` يجب أن تُرجع `"🇮🇷"`.

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` يجب أن تُرجع `"🇮🇪"`.

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` يجب أن تُرجع `"🇮🇱"`.

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` يجب أن تُرجع `"🇮🇹"`.

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` يجب أن تُرجع `"🇯🇲"`.

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` يجب أن تُرجع `"🇯🇵"`.

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` يجب أن تُرجع `"🇰🇿"`.

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` يجب أن تُرجع `"🇰🇪"`.

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` يجب أن تُرجع `"🇽🇰"`.

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` يجب أن تُرجع `"🇰🇬"`.

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` يجب أن تُرجع `"🇱🇻"`.

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` يجب أن تُرجع `"🇱🇧"`.

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` يجب أن تُرجع `"🇱🇮"`.

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` يجب أن تُرجع `"🇱🇹"`.

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` يجب أن تُرجع `"🇱🇺"`.

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` يجب أن تُرجع `"🇲🇬"`.

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` يجب أن تُرجع `"🇲🇾"`.

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` يجب أن تُرجع `"🇲🇹"`.

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` يجب أن تُرجع `"🇲🇽"`.

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` يجب أن تُرجع `"🇲🇩"`.

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` يجب أن تُرجع `"🇲🇨"`.

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` يجب أن تُرجع `"🇲🇳"`.

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` يجب أن تُرجع `"🇲🇪"`.

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` يجب أن تُرجع `"🇲🇦"`.

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` يجب أن تُرجع `"🇳🇱"`.

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` يجب أن تُرجع `"🇳🇿"`.

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` يجب أن تُرجع `"🇳🇬"`.

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` يجب أن تُرجع `"🇲🇰"`.

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` يجب أن تُرجع `"🇳🇴"`.

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` يجب أن تُرجع `"🇵🇰"`.

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` يجب أن تُرجع `"🇵🇭"`.

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` يجب أن تُرجع `"🇵🇱"`.

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` يجب أن تُرجع `"🇵🇹"`.

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` يجب أن تُرجع `"🇵🇷"`.

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` يجب أن تُرجع `"🇷🇴"`.

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` يجب أن تُرجع `"🇸🇲"`.

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` يجب أن تُرجع `"🇸🇦"`.

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` يجب أن تُرجع `"🇷🇸"`.

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` يجب أن تُرجع `"🇸🇬"`.

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` يجب أن تُرجع `"🇸🇰"`.

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` يجب أن تُرجع `"🇸🇮"`.

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` يجب أن تُرجع `"🇿🇦"`.

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` يجب أن تُرجع `"🇰🇷"`.

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` يجب أن تُرجع `"🇪🇸"`.

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` يجب أن تُرجع `"🇸🇪"`.

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` يجب أن تُرجع `"🇨🇭"`.

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` يجب أن تُرجع `"🇹🇭"`.

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` يجب أن تُرجع `"🇹🇹"`.

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` يجب أن تُرجع `"🇹🇷"`.

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` يجب أن تُرجع `"🇺🇦"`.

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` يجب أن تُرجع `"🇦🇪"`.

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` يجب أن تُرجع `"🇺🇸"`.

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` يجب أن تُرجع `"🇺🇾"`.

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` يجب أن تُرجع `"🇺🇿"`.

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` يجب أن تُرجع `"🇻🇪"`.

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
