---
id: 697a49e9860d24853adef67e
title: "تحدٍّ 193: ألعاب الشتاء 2026 اليوم 14: تسلق الجبال بالتزلج"
challengeType: 28
dashedName: challenge-193
---

# --description--

بما أن لديك عمق الثلج وانحدار الجبل، حدّد إذا كان هناك خطر انزلاق جليدي.

- قيم عمق الثلج هي `"Shallow"`، `"Moderate"`، أو `"Deep"`.
- قيم الانحدار هي `"Gentle"`، `"Steep"`، أو `"Very Steep"`.

أرجع `"Safe"` أو `"Risky"` بناءً على هذا الجدول:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalancheRisk("Shallow", "Gentle")` يجب أن تُرجع `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Gentle"), "Safe");
```

`avalancheRisk("Shallow", "Steep")` يجب أن تُرجع `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Steep"), "Safe");
```

`avalancheRisk("Shallow", "Very Steep")` يجب أن تُرجع `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Very Steep"), "Safe");
```

`avalancheRisk("Moderate", "Gentle")` يجب أن تُرجع `"Safe"`.

```js
assert.equal(avalancheRisk("Moderate", "Gentle"), "Safe");
```

`avalancheRisk("Moderate", "Steep")` يجب أن تُرجع `"Risky"`.

```js
assert.equal(avalancheRisk("Moderate", "Steep"), "Risky");
```

`avalancheRisk("Moderate", "Very Steep")` يجب أن تُرجع `"Risky"`.

```js
assert.equal(avalancheRisk("Moderate", "Very Steep"), "Risky");
```

`avalancheRisk("Deep", "Gentle")` يجب أن تُرجع `"Safe"`.

```js
assert.equal(avalancheRisk("Deep", "Gentle"), "Safe");
```

`avalancheRisk("Deep", "Steep")` يجب أن تُرجع `"Risky"`.

```js
assert.equal(avalancheRisk("Deep", "Steep"), "Risky");
```

`avalancheRisk("Deep", "Very Steep")` يجب أن تُرجع `"Risky"`.

```js
assert.equal(avalancheRisk("Deep", "Very Steep"), "Risky");
```

# --seed--

## --seed-contents--

```js
function avalancheRisk(snowDepth, slope) {

  return snowDepth;
}
```

# --solutions--

```js
function avalancheRisk(snowDepth, slope) {
  const riskTable = {
    Gentle: {
      Shallow: "Safe",
      Moderate: "Safe",
      Deep: "Safe",
    },
    Steep: {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
    "Very Steep": {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
  };

  return riskTable[slope][snowDepth];
}
```
