---
id: 697a49e9860d24853adef67d
title: "تحدٍّ 192: ألعاب الشتاء 2026 اليوم 13: النورديك كومبايند"
challengeType: 28
dashedName: challenge-192
---

# --description--

معطاة مصفوفة من درجات القفز للرياضيين، احسب أوقات تأخير البداية لكل منهم لجزء التزلج الريفي في النورديك كومبايند.

الرياضي الذي يملك أعلى درجة قفز يبدأ أولًا (تأخير 0 ثانية). يبدأ جميع الرياضيين الآخرين لاحقًا بناءً على مدى تأخر درجة قفزهم مقارنة بأفضل قفزة.

لحساب التأخير لكل رياضي، اطرح درجة قفز الرياضي من أفضل درجة قفز إجمالية واضرب الناتج في 1.5. قرب التأخير لأقرب عدد صحيح للأعلى.

# --hints--

`calculateStartDelays([120, 110, 125])` يجب أن تُرجع `[8, 23, 0]`.

```js
assert.deepEqual(calculateStartDelays([120, 110, 125]), [8, 23, 0]);
```

`calculateStartDelays([118, 125, 122, 120])` يجب أن تُرجع `[11, 0, 5, 8]`.

```js
assert.deepEqual(calculateStartDelays([118, 125, 122, 120]), [11, 0, 5, 8]);
```

`calculateStartDelays([100, 105, 95, 110, 120, 115, 108])` يجب أن تُرجع `[30, 23, 38, 15, 0, 8, 18]`.

```js
assert.deepEqual(calculateStartDelays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18]);
```

`calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` يجب أن تُرجع `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
assert.deepEqual(calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12]);
```

# --seed--

## --seed-contents--

```js
function calculateStartDelays(jumpScores) {

  return jumpScores
}
```

# --solutions--

```js
function calculateStartDelays(jumpScores) {
  const bestJump = Math.max(...jumpScores);
  return jumpScores.map(score => Math.ceil((bestJump - score) * 1.5));
}
```
