---
id: 697a49e6ff50d756c9b69362
title: "تحدٍّ 185: ألعاب الشتاء 2026 اليوم 6: التزلج الفني"
challengeType: 28
dashedName: challenge-185
---

# --description--

معطى مصفوفة من درجات الحكام وعقوبات اختيارية، احسب النتيجة النهائية لروتين التزلج الفني.

المعلمة الأولى هي مصفوفة من 10 درجات حكام، كل منها رقم من 0 إلى 10. احذف أعلى وأدنى درجات الحكام واجمع الدرجات الثمانية المتبقية لتحصل على النتيجة الأساسية.

أي معلمات إضافية تُمرر إلى الدالة هي عقوبات. اطرح كل العقوبات من النتيجة الأساسية لتحصل على النتيجة النهائية.

# --hints--

`computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` يجب أن تُرجع `64`.

```js
assert.equal(computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64);
```

`computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` يجب أن تُرجع `80`.

```js
assert.equal(computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80);
```

`computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` يجب أن تُرجع `67`.

```js
assert.equal(computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67);
```

`computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` يجب أن تُرجع `67.5`.

```js
assert.equal(computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5);
```

`computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` يجب أن تُرجع `59`.

```js
assert.equal(computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59);
```

# --seed--

## --seed-contents--

```js
function computeScore(judgeScores, ...penalties) {

  return judgeScores;
}
```

# --solutions--

```js
function computeScore(judgeScores, ...penalties) {
  const sortedScores = [...judgeScores].sort((a, b) => a - b);
  const baseScore = sortedScores.slice(1, 9).reduce((sum, s) => sum + s, 0);
  const totalPenalty = penalties.reduce((sum, p) => sum + p, 0);

  return baseScore - totalPenalty;
}
```
