---
id: 697a49e6ff50d756c9b69363
title: "تحدٍّ 186: ألعاب الشتاء 2026 اليوم 7: التزلج السريع"
challengeType: 28
dashedName: challenge-186
---

# --description--

معطى مصفوفتان تمثلان أوقات اللفات (بالثواني) لمتزلجين اثنين، أرجع رقم اللفة التي يكون فيها الفرق في أوقات اللفات هو الأكبر.

العنصر الأول من كل مصفوفة يمثل اللفة 1، والعنصر الثاني يمثل اللفة 2، وهكذا.

# --hints--

`largestDifference([26.11, 25.80, 25.92, 26.23, 26.07], [25.93, 25.74, 26.53, 26.11, 26.30])` يجب أن تُرجع `3`.

```js
assert.equal(largestDifference([26.11, 25.80, 25.92, 26.23, 26.07], [25.93, 25.74, 26.53, 26.11, 26.30]), 3);
```

`largestDifference([27.04, 25.94, 26.22, 26.07, 26.18], [25.59, 25.80, 26.11, 26.01, 26.23])` يجب أن تُرجع `1`.

```js
assert.equal(largestDifference([27.04, 25.94, 26.22, 26.07, 26.18], [25.59, 25.80, 26.11, 26.01, 26.23]), 1);
```

`largestDifference([25.82, 25.90, 26.05, 26.00, 26.48], [25.85, 25.92, 26.07, 25.98, 25.95])` يجب أن تُرجع `5`.

```js
assert.equal(largestDifference([25.82, 25.90, 26.05, 26.00, 26.48], [25.85, 25.92, 26.07, 25.98, 25.95]), 5);
```

`largestDifference([25.88, 26.10, 25.95, 26.05, 26.00], [25.90, 26.55, 25.92, 26.03, 26.01])` يجب أن تُرجع `2`.

```js
assert.equal(largestDifference([25.88, 26.10, 25.95, 26.05, 26.00], [25.90, 26.55, 25.92, 26.03, 26.01]), 2);
```

`largestDifference([25.92, 26.01, 26.05, 25.88, 26.12], [25.95, 26.00, 26.03, 26.45, 26.10])` يجب أن تُرجع `4`.

```js
assert.equal(largestDifference([25.92, 26.01, 26.05, 25.88, 26.12], [25.95, 26.00, 26.03, 26.45, 26.10]), 4);
```

# --seed--

## --seed-contents--

```js
function largestDifference(skater1, skater2) {

  return skater1;
}
```

# --solutions--

```js
function largestDifference(skater1, skater2) {
  let maxDiff = -1;
  let lapNumber = 1;

  for (let i = 0; i < skater1.length; i++) {
    const diff = Math.abs(skater1[i] - skater2[i]);
    if (diff > maxDiff) {
      maxDiff = diff;
      lapNumber = i + 1;
    }
  }

  return lapNumber;
}
```
