---
id: 697a49e6ff50d756c9b69365
title: "تحدٍّ 188: ألعاب الشتاء 2026 اليوم 9: الهيكل العظمي"
challengeType: 28
dashedName: challenge-188
---

# --description--

معطاة سلسلة نصية تمثل المنحنيات على مسار الهيكل العظمي، حدّد صعوبة المسار.

- تتكون السلسلة النصية المعطاة فقط من الحروف:
  - `"L"` للدوران إلى اليسار
  - `"R"` للدوران إلى اليمين
  - `"S"` لجزء مستقيم

- كل تغيير في الاتجاه يضيف 15 نقطة (عندما يتبع `"L"` حرف `"R"` أو العكس).
- كل منحنى آخر يضيف 5 نقاط لكل منهما (كل حروف `"L"` أو `"R"` الأخرى).
- الأجزاء المستقيمة تضيف 0 نقاط.

تعتمد صعوبة المسار على المجموع الكلي للنقاط. أرجع:

- `"Easy"` إذا كان المجموع بين 0 و100
- `"Medium"` إذا كان المجموع بين 101 و200
- `"Hard"` إذا كان المجموع أكثر من 200

# --hints--

يجب أن تُرجع الدالة `getDifficulty("SLSLLSRRLSRLRL")` القيمة `"Easy"`.

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

يجب أن تُرجع الدالة `getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` القيمة `"Hard"`.

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

يجب أن تُرجع الدالة `getDifficulty("SRRRRLSLLRLRSSRLSRL")` القيمة `"Medium"`.

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

يجب أن تُرجع الدالة `getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` القيمة `"Hard"`.

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

يجب أن تُرجع الدالة `getDifficulty("SLLSSLRLSLSLRSLSSLRL")` القيمة `"Medium"`.

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

يجب أن تُرجع الدالة `getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` القيمة `"Easy"`.

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
