---
id: 697a49e9860d24853adef67c
title: "تحدٍّ 191: ألعاب الشتاء 2026 اليوم 12: التزحلق على الزلاجات"
challengeType: 28
dashedName: challenge-191
---

# --description--

معطى مصفوفة تمثل أوزان الرياضيين في فريق التزحلق على الزلاجات ورقم يمثل وزن الزلاجة، حدّد ما إذا كان الفريق مؤهلاً للمشاركة في السباق.

- طول المصفوفة يحدد حجم الفريق: فرق من 1 أو 2 أو 4 أشخاص.
- جميع قيم الأوزان المعطاة بالكيلوغرام (كغ).

يجب أن يكون للزلاجة (الزلاجة وحدها) وزن الحد الأدنى التالي:

- 162 كغ لفريق من شخص واحد
- 170 كغ لفريق من شخصين
- 210 كغ لفريق من أربعة أشخاص

يجب ألا يتجاوز الوزن الكلي للزلاجة (الرياضيون بالإضافة إلى الزلاجة) ما يلي:

- 247 كغ لفريق من شخص واحد
- 390 كغ لفريق من شخصين
- 630 كغ لفريق من أربعة أشخاص

أرجع `"Eligible"` إذا استوفى الفريق جميع المتطلبات، أو `"Not Eligible"` إذا لم يستوفِ الفريق واحدًا أو أكثر من المتطلبات.

# --hints--

`checkEligibility([78], 165)` يجب أن تُرجع `"Eligible"`.

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` يجب أن تُرجع `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` يجب أن تُرجع `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` يجب أن تُرجع `"Eligible"`.

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` يجب أن تُرجع `"Not Eligible"`.

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` يجب أن تُرجع `"Not Eligible"`.

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` يجب أن تُرجع `"Eligible"`.

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` يجب أن تُرجع `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` يجب أن تُرجع `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
