---
id: 697a49e6ff50d756c9b6935d
title: "鬲丨丿賺賾 180: 兀賱毓丕亘 丕賱卮鬲丕亍 2026 丕賱賷賵賲 1: 賷賵賲 丕賱丕賮鬲鬲丕丨"
challengeType: 28
dashedName: challenge-180
---

# --description--

丕賱賷賵賲 賷亘丿兀 兀賱毓丕亘 丕賱卮鬲丕亍 2026. 爻鬲鬲賱賯賶 禺賱丕賱 丕賱兀賷丕賲 丕賱爻亘毓丞 毓卮乇 丕賱賯丕丿賲丞 鬲丨丿賷丕鬲 卮賮乇丞 賲爻鬲賵丨丕丞 賲賳賴丕.

賮賷 丕賱鬲丨丿賺賾 丕賱兀賵賱貙 賷購毓胤賶 乇賲夭 丿賵賱丞 賲賰賵賳 賲賳 丨乇賮賷賳 賵賷噩亘 毓賱賷賰 廿乇噩丕毓 乇賲夭 丕賱毓賱賲 丕賱禺丕氐 亘鬲賱賰 丕賱丿賵賱丞.

丕爻鬲禺丿賲 賴匕賴 丕賱賯丕卅賲丞:

| 丕賱丿賵賱丞 | 丕賱乇賲夭 | 丕賱毓賱賲 |
| - | -: | :--: |
| 兀賱亘丕賳賷丕                    |            `"AL"` | `"馃嚘馃嚤"` |
| 兀賳丿賵乇丕                    |            `"AD"` | `"馃嚘馃嚛"` |
| 丕賱兀乇噩賳鬲賷賳                  |            `"AR"` | `"馃嚘馃嚪"` |
| 兀乇賲賷賳賷丕                    |            `"AM"` | `"馃嚘馃嚥"` |
| 兀爻鬲乇丕賱賷丕                  |            `"AU"` | `"馃嚘馃嚭"` |
| 丕賱賳賲爻丕                    |            `"AT"` | `"馃嚘馃嚬"` |
| 兀匕乇亘賷噩丕賳                 |            `"AZ"` | `"馃嚘馃嚳"` |
| 亘賱噩賷賰丕                    |            `"BE"` | `"馃嚙馃嚜"` |
| 亘賳賷賳                      |            `"BJ"` | `"馃嚙馃嚡"` |
| 亘賵賱賷賮賷丕                    |            `"BO"` | `"馃嚙馃嚧"` |
| 丕賱亘賵爻賳丞 賵丕賱賴乇爻賰     |            `"BA"` | `"馃嚙馃嚘"` |
| 丕賱亘乇丕夭賷賱                     |            `"BR"` | `"馃嚙馃嚪"` |
| 亘賱睾丕乇賷丕                   |            `"BG"` | `"馃嚙馃嚞"` |
| 賰賳丿丕                     |            `"CA"` | `"馃嚚馃嚘"` |
| 鬲卮賷賱賷                      |            `"CL"` | `"馃嚚馃嚤"` |
| 丕賱氐賷賳                      |            `"CN"` | `"馃嚚馃嚦"` |
| 賰賵賱賵賲亘賷丕                   |            `"CO"` | `"馃嚚馃嚧"` |
| 賰乇賵丕鬲賷丕                    |            `"HR"` | `"馃嚟馃嚪"` |
| 賯亘乇氐                     |            `"CY"` | `"馃嚚馃嚲"` |
| 噩賲賴賵乇賷丞 丕賱鬲卮賷賰             |            `"CZ"` | `"馃嚚馃嚳"` |
| 丕賱丿賳賲丕乇賰                    |            `"DK"` | `"馃嚛馃嚢"` |
| 丕賱廿賰賵丕丿賵乇                    |            `"EC"` | `"馃嚜馃嚚"` |
| 廿乇賷鬲乇賷丕                    |            `"ER"` | `"馃嚜馃嚪"` |
| 廿爻鬲賵賳賷丕                    |            `"EE"` | `"馃嚜馃嚜"` |
| 賮賳賱賳丿丕                    |            `"FI"` | `"馃嚝馃嚠"` |
| 賮乇賳爻丕                     |            `"FR"` | `"馃嚝馃嚪"` |
| 噩賵乇噩賷丕                    |            `"GE"` | `"馃嚞馃嚜"` |
| 兀賱賲丕賳賷丕                    |            `"DE"` | `"馃嚛馃嚜"` |
| 亘乇賷胤丕賳賷丕 丕賱毓馗賲賶              |            `"GB"` | `"馃嚞馃嚙"` |
| 丕賱賷賵賳丕賳                     |            `"GR"` | `"馃嚞馃嚪"` |
| 睾賷賳賷丕 亘賷爻丕賵              |            `"GW"` | `"馃嚞馃嚰"` |
| 賴丕賷鬲賷                      |            `"HT"` | `"馃嚟馃嚬"` |
| 賴賵賳睾 賰賵賳睾                  |            `"HK"` | `"馃嚟馃嚢"` |
| 丕賱賲噩乇                    |            `"HU"` | `"馃嚟馃嚭"` |
| 丌賷爻賱賳丿丕                    |            `"IS"` | `"馃嚠馃嚫"` |
| 丕賱賴賳丿                      |            `"IN"` | `"馃嚠馃嚦"` |
| 廿賷乇丕賳                       |            `"IR"` | `"馃嚠馃嚪"` |
| 兀賷乇賱賳丿丕                    |            `"IE"` | `"馃嚠馃嚜"` |
| 廿爻乇丕卅賷賱                     |            `"IL"` | `"馃嚠馃嚤"` |
| 廿賷胤丕賱賷丕                      |            `"IT"` | `"馃嚠馃嚬"` |
| 噩丕賲丕賷賰丕                    |            `"JM"` | `"馃嚡馃嚥"` |
| 丕賱賷丕亘丕賳                      |            `"JP"` | `"馃嚡馃嚨"` |
| 賰丕夭丕禺爻鬲丕賳                 |            `"KZ"` | `"馃嚢馃嚳"` |
| 賰賷賳賷丕                      |            `"KE"` | `"馃嚢馃嚜"` |
| 賰賵爻賵賮賵                     |            `"XK"` | `"馃嚱馃嚢"` |
| 賯賷乇睾賷夭爻鬲丕賳                 |            `"KG"` | `"馃嚢馃嚞"` |
| 賱丕鬲賮賷丕                     |            `"LV"` | `"馃嚤馃嚮"` |
| 賱亘賳丕賳                    |            `"LB"` | `"馃嚤馃嚙"` |
| 賱賷禺鬲賳卮鬲丕賷賳              |            `"LI"` | `"馃嚤馃嚠"` |
| 賱賷鬲賵丕賳賷丕                  |            `"LT"` | `"馃嚤馃嚬"` |
| 賱賵賰爻賲亘賵乇睾                 |            `"LU"` | `"馃嚤馃嚭"` |
| 賲丿睾卮賯乇                 |            `"MG"` | `"馃嚥馃嚞"` |
| 賲丕賱賷夭賷丕                   |            `"MY"` | `"馃嚥馃嚲"` |
| 賲丕賱胤丕                      |            `"MT"` | `"馃嚥馃嚬"` |
| 丕賱賲賰爻賷賰                     |            `"MX"` | `"馃嚥馃嚱"` |
| 賲賵賱丿賵賮丕                    |            `"MD"` | `"馃嚥馃嚛"` |
| 賲賵賳丕賰賵                     |            `"MC"` | `"馃嚥馃嚚"` |
| 賲賳睾賵賱賷丕                   |            `"MN"` | `"馃嚥馃嚦"` |
| 丕賱噩亘賱 丕賱兀爻賵丿                 |            `"ME"` | `"馃嚥馃嚜"` |
| 丕賱賲睾乇亘                    |            `"MA"` | `"馃嚥馃嚘"` |
| 賴賵賱賳丿丕                |            `"NL"` | `"馃嚦馃嚤"` |
| 賳賷賵夭賷賱賳丿丕                |            `"NZ"` | `"馃嚦馃嚳"` |
| 賳賷噩賷乇賷丕                    |            `"NG"` | `"馃嚦馃嚞"` |
| 賲賯丿賵賳賷丕 丕賱卮賲丕賱賷丞            |            `"MK"` | `"馃嚥馃嚢"` |
| 丕賱賳乇賵賷噩                     |            `"NO"` | `"馃嚦馃嚧"` |
| 亘丕賰爻鬲丕賳                   |            `"PK"` | `"馃嚨馃嚢"` |
| 丕賱賮賱亘賷賳                |            `"PH"` | `"馃嚨馃嚟"` |
| 亘賵賱賳丿丕                     |            `"PL"` | `"馃嚨馃嚤"` |
| 丕賱亘乇鬲睾丕賱                   |            `"PT"` | `"馃嚨馃嚬"` |
| 亘賵乇鬲賵 乇賷賰賵                |            `"PR"` | `"馃嚨馃嚪"` |
| 乇賵賲丕賳賷丕                    |            `"RO"` | `"馃嚪馃嚧"` |
| 爻丕賳 賲丕乇賷賳賵                 |            `"SM"` | `"馃嚫馃嚥"` |
| 丕賱爻毓賵丿賷丞               |            `"SA"` | `"馃嚫馃嚘"` |
| 氐乇亘賷丕                     |            `"RS"` | `"馃嚪馃嚫"` |
| 爻賳睾丕賮賵乇丞                  |            `"SG"` | `"馃嚫馃嚞"` |
| 爻賱賵賮丕賰賷丕                   |            `"SK"` | `"馃嚫馃嚢"` |
| 爻賱賵賮賷賳賷丕                   |            `"SI"` | `"馃嚫馃嚠"` |
| 噩賳賵亘 兀賮乇賷賯賷丕               |            `"ZA"` | `"馃嚳馃嚘"` |
| 賰賵乇賷丕 丕賱噩賳賵亘賷丞                |            `"KR"` | `"馃嚢馃嚪"` |
| 廿爻亘丕賳賷丕                      |            `"ES"` | `"馃嚜馃嚫"` |
| 丕賱爻賵賷丿                     |            `"SE"` | `"馃嚫馃嚜"` |
| 爻賵賷爻乇丕                |            `"CH"` | `"馃嚚馃嚟"` |
| 鬲丕賷賱丕賳丿                   |            `"TH"` | `"馃嚬馃嚟"` |
| 鬲乇賷賳賷丿丕丿 賵鬲賵亘丕睾賵          |            `"TT"` | `"馃嚬馃嚬"` |
| 鬲乇賰賷丕                     |            `"TR"` | `"馃嚬馃嚪"` |
| 兀賵賰乇丕賳賷丕                    |            `"UA"` | `"馃嚭馃嚘"` |
| 丕賱廿賲丕乇丕鬲 丕賱毓乇亘賷丞 丕賱賲鬲丨丿丞       |            `"AE"` | `"馃嚘馃嚜"` |
| 丕賱賵賱丕賷丕鬲 丕賱賲鬲丨丿丞              |            `"US"` | `"馃嚭馃嚫"` |
| 兀賵乇賵睾賵丕賷                    |            `"UY"` | `"馃嚭馃嚲"` |
| 兀賵夭亘賰爻鬲丕賳                 |            `"UZ"` | `"馃嚭馃嚳"` |
| 賮賳夭賵賷賱丕                  |            `"VE"` | `"馃嚮馃嚜"` |

# --hints--

`getFlag("AL")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚘馃嚤"`.

```js
assert.equal(getFlag("AL"), "馃嚘馃嚤");
```

`getFlag("AD")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚘馃嚛"`.

```js
assert.equal(getFlag("AD"), "馃嚘馃嚛");
```

`getFlag("AR")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚘馃嚪"`.

```js
assert.equal(getFlag("AR"), "馃嚘馃嚪")
```

`getFlag("AM")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚘馃嚥"`.

```js
assert.equal(getFlag("AM"), "馃嚘馃嚥")
```

`getFlag("AU")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚘馃嚭"`.

```js
assert.equal(getFlag("AU"), "馃嚘馃嚭")
```

`getFlag("AT")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚘馃嚬"`.

```js
assert.equal(getFlag("AT"), "馃嚘馃嚬")
```

`getFlag("AZ")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚘馃嚳"`.

```js
assert.equal(getFlag("AZ"), "馃嚘馃嚳")
```

`getFlag("BE")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚙馃嚜"`.

```js
assert.equal(getFlag("BE"), "馃嚙馃嚜")
```

`getFlag("BJ")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚙馃嚡"`.

```js
assert.equal(getFlag("BJ"), "馃嚙馃嚡")
```

`getFlag("BO")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚙馃嚧"`.

```js
assert.equal(getFlag("BO"), "馃嚙馃嚧")
```

`getFlag("BA")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚙馃嚘"`.

```js
assert.equal(getFlag("BA"), "馃嚙馃嚘")
```

`getFlag("BR")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚙馃嚪"`.

```js
assert.equal(getFlag("BR"), "馃嚙馃嚪")
```

`getFlag("BG")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚙馃嚞"`.

```js
assert.equal(getFlag("BG"), "馃嚙馃嚞")
```

`getFlag("CA")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚚馃嚘"`.

```js
assert.equal(getFlag("CA"), "馃嚚馃嚘")
```

`getFlag("CL")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚚馃嚤"`.

```js
assert.equal(getFlag("CL"), "馃嚚馃嚤")
```

`getFlag("CN")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚚馃嚦"`.

```js
assert.equal(getFlag("CN"), "馃嚚馃嚦")
```

`getFlag("CO")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚚馃嚧"`.

```js
assert.equal(getFlag("CO"), "馃嚚馃嚧")
```

`getFlag("HR")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚟馃嚪"`.

```js
assert.equal(getFlag("HR"), "馃嚟馃嚪")
```

`getFlag("CY")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚚馃嚲"`.

```js
assert.equal(getFlag("CY"), "馃嚚馃嚲")
```

`getFlag("CZ")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚚馃嚳"`.

```js
assert.equal(getFlag("CZ"), "馃嚚馃嚳")
```

`getFlag("DK")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚛馃嚢"`.

```js
assert.equal(getFlag("DK"), "馃嚛馃嚢")
```

`getFlag("EC")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚜馃嚚"`.

```js
assert.equal(getFlag("EC"), "馃嚜馃嚚")
```

`getFlag("ER")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚜馃嚪"`.

```js
assert.equal(getFlag("ER"), "馃嚜馃嚪")
```

`getFlag("EE")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚜馃嚜"`.

```js
assert.equal(getFlag("EE"), "馃嚜馃嚜")
```

`getFlag("FI")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚝馃嚠"`.

```js
assert.equal(getFlag("FI"), "馃嚝馃嚠")
```

`getFlag("FR")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚝馃嚪"`.

```js
assert.equal(getFlag("FR"), "馃嚝馃嚪")
```

`getFlag("GE")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚞馃嚜"`.

```js
assert.equal(getFlag("GE"), "馃嚞馃嚜")
```

`getFlag("DE")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚛馃嚜"`.

```js
assert.equal(getFlag("DE"), "馃嚛馃嚜")
```

`getFlag("GB")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚞馃嚙"`.

```js
assert.equal(getFlag("GB"), "馃嚞馃嚙")
```

`getFlag("GR")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚞馃嚪"`.

```js
assert.equal(getFlag("GR"), "馃嚞馃嚪")
```

`getFlag("GW")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚞馃嚰"`.

```js
assert.equal(getFlag("GW"), "馃嚞馃嚰")
```

`getFlag("HT")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚟馃嚬"`.

```js
assert.equal(getFlag("HT"), "馃嚟馃嚬")
```

`getFlag("HK")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚟馃嚢"`.

```js
assert.equal(getFlag("HK"), "馃嚟馃嚢")
```

`getFlag("HU")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚟馃嚭"`.

```js
assert.equal(getFlag("HU"), "馃嚟馃嚭")
```

`getFlag("IS")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚠馃嚫"`.

```js
assert.equal(getFlag("IS"), "馃嚠馃嚫")
```

`getFlag("IN")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚠馃嚦"`.

```js
assert.equal(getFlag("IN"), "馃嚠馃嚦")
```

`getFlag("IR")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚠馃嚪"`.

```js
assert.equal(getFlag("IR"), "馃嚠馃嚪")
```

`getFlag("IE")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚠馃嚜"`.

```js
assert.equal(getFlag("IE"), "馃嚠馃嚜")
```

`getFlag("IL")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚠馃嚤"`.

```js
assert.equal(getFlag("IL"), "馃嚠馃嚤")
```

`getFlag("IT")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚠馃嚬"`.

```js
assert.equal(getFlag("IT"), "馃嚠馃嚬")
```

`getFlag("JM")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚡馃嚥"`.

```js
assert.equal(getFlag("JM"), "馃嚡馃嚥")
```

`getFlag("JP")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚡馃嚨"`.

```js
assert.equal(getFlag("JP"), "馃嚡馃嚨")
```

`getFlag("KZ")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚢馃嚳"`.

```js
assert.equal(getFlag("KZ"), "馃嚢馃嚳")
```

`getFlag("KE")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚢馃嚜"`.

```js
assert.equal(getFlag("KE"), "馃嚢馃嚜")
```

`getFlag("XK")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚱馃嚢"`.

```js
assert.equal(getFlag("XK"), "馃嚱馃嚢")
```

`getFlag("KG")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚢馃嚞"`.

```js
assert.equal(getFlag("KG"), "馃嚢馃嚞")
```

`getFlag("LV")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚤馃嚮"`.

```js
assert.equal(getFlag("LV"), "馃嚤馃嚮")
```

`getFlag("LB")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚤馃嚙"`.

```js
assert.equal(getFlag("LB"), "馃嚤馃嚙")
```

`getFlag("LI")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚤馃嚠"`.

```js
assert.equal(getFlag("LI"), "馃嚤馃嚠")
```

`getFlag("LT")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚤馃嚬"`.

```js
assert.equal(getFlag("LT"), "馃嚤馃嚬")
```

`getFlag("LU")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚤馃嚭"`.

```js
assert.equal(getFlag("LU"), "馃嚤馃嚭")
```

`getFlag("MG")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚞"`.

```js
assert.equal(getFlag("MG"), "馃嚥馃嚞")
```

`getFlag("MY")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚲"`.

```js
assert.equal(getFlag("MY"), "馃嚥馃嚲")
```

`getFlag("MT")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚬"`.

```js
assert.equal(getFlag("MT"), "馃嚥馃嚬")
```

`getFlag("MX")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚱"`.

```js
assert.equal(getFlag("MX"), "馃嚥馃嚱")
```

`getFlag("MD")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚛"`.

```js
assert.equal(getFlag("MD"), "馃嚥馃嚛")
```

`getFlag("MC")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚚"`.

```js
assert.equal(getFlag("MC"), "馃嚥馃嚚")
```

`getFlag("MN")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚦"`.

```js
assert.equal(getFlag("MN"), "馃嚥馃嚦")
```

`getFlag("ME")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚜"`.

```js
assert.equal(getFlag("ME"), "馃嚥馃嚜")
```

`getFlag("MA")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚘"`.

```js
assert.equal(getFlag("MA"), "馃嚥馃嚘")
```

`getFlag("NL")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚦馃嚤"`.

```js
assert.equal(getFlag("NL"), "馃嚦馃嚤")
```

`getFlag("NZ")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚦馃嚳"`.

```js
assert.equal(getFlag("NZ"), "馃嚦馃嚳")
```

`getFlag("NG")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚦馃嚞"`.

```js
assert.equal(getFlag("NG"), "馃嚦馃嚞")
```

`getFlag("MK")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚥馃嚢"`.

```js
assert.equal(getFlag("MK"), "馃嚥馃嚢")
```

`getFlag("NO")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚦馃嚧"`.

```js
assert.equal(getFlag("NO"), "馃嚦馃嚧")
```

`getFlag("PK")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚨馃嚢"`.

```js
assert.equal(getFlag("PK"), "馃嚨馃嚢")
```

`getFlag("PH")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚨馃嚟"`.

```js
assert.equal(getFlag("PH"), "馃嚨馃嚟")
```

`getFlag("PL")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚨馃嚤"`.

```js
assert.equal(getFlag("PL"), "馃嚨馃嚤")
```

`getFlag("PT")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚨馃嚬"`.

```js
assert.equal(getFlag("PT"), "馃嚨馃嚬")
```

`getFlag("PR")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚨馃嚪"`.

```js
assert.equal(getFlag("PR"), "馃嚨馃嚪")
```

`getFlag("RO")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚪馃嚧"`.

```js
assert.equal(getFlag("RO"), "馃嚪馃嚧")
```

`getFlag("SM")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚫馃嚥"`.

```js
assert.equal(getFlag("SM"), "馃嚫馃嚥")
```

`getFlag("SA")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚫馃嚘"`.

```js
assert.equal(getFlag("SA"), "馃嚫馃嚘")
```

`getFlag("RS")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚪馃嚫"`.

```js
assert.equal(getFlag("RS"), "馃嚪馃嚫")
```

`getFlag("SG")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚫馃嚞"`.

```js
assert.equal(getFlag("SG"), "馃嚫馃嚞")
```

`getFlag("SK")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚫馃嚢"`.

```js
assert.equal(getFlag("SK"), "馃嚫馃嚢")
```

`getFlag("SI")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚫馃嚠"`.

```js
assert.equal(getFlag("SI"), "馃嚫馃嚠")
```

`getFlag("ZA")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚳馃嚘"`.

```js
assert.equal(getFlag("ZA"), "馃嚳馃嚘")
```

`getFlag("KR")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚢馃嚪"`.

```js
assert.equal(getFlag("KR"), "馃嚢馃嚪")
```

`getFlag("ES")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚜馃嚫"`.

```js
assert.equal(getFlag("ES"), "馃嚜馃嚫")
```

`getFlag("SE")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚫馃嚜"`.

```js
assert.equal(getFlag("SE"), "馃嚫馃嚜")
```

`getFlag("CH")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚚馃嚟"`.

```js
assert.equal(getFlag("CH"), "馃嚚馃嚟")
```

`getFlag("TH")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚬馃嚟"`.

```js
assert.equal(getFlag("TH"), "馃嚬馃嚟")
```

`getFlag("TT")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚬馃嚬"`.

```js
assert.equal(getFlag("TT"), "馃嚬馃嚬")
```

`getFlag("TR")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚬馃嚪"`.

```js
assert.equal(getFlag("TR"), "馃嚬馃嚪")
```

`getFlag("UA")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚭馃嚘"`.

```js
assert.equal(getFlag("UA"), "馃嚭馃嚘")
```

`getFlag("AE")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚘馃嚜"`.

```js
assert.equal(getFlag("AE"), "馃嚘馃嚜")
```

`getFlag("US")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚭馃嚫"`.

```js
assert.equal(getFlag("US"), "馃嚭馃嚫")
```

`getFlag("UY")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚭馃嚲"`.

```js
assert.equal(getFlag("UY"), "馃嚭馃嚲")
```

`getFlag("UZ")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚭馃嚳"`.

```js
assert.equal(getFlag("UZ"), "馃嚭馃嚳")
```

`getFlag("VE")` 賷噩亘 兀賳 鬲購乇噩毓 `"馃嚮馃嚜"`.

```js
assert.equal(getFlag("VE"), "馃嚮馃嚜")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
