---
id: 697a49e6ff50d756c9b6935f
title: "تحدٍّ 182: ألعاب الشتاء 2026 اليوم 3: البياثلون"
challengeType: 28
dashedName: challenge-182
---

# --description--

معطاة مصفوفة من الأعداد الصحيحة، حيث تمثل كل قيمة عدد الأهداف التي تم إصابتها في جولة واحدة من البياثلون، أرجع المسافة الإجمالية للعقوبة التي يجب على الرياضي التزلج عليها.

- تتكون كل جولة من 5 أهداف.
- كل هدف فائت يؤدي إلى دورة عقوبة بطول 150 مترًا.

# --hints--

`calculatePenaltyDistance([4, 4])` يجب أن تُرجع `300`.

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` يجب أن تُرجع `0`.

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` يجب أن تُرجع `450`.

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` يجب أن تُرجع `150`.

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` يجب أن تُرجع `1500`.

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
