---
id: 697a49e6ff50d756c9b69366
title: "تحدٍّ 189: ألعاب الشتاء 2026 اليوم 10: التزلج على المنحدرات"
challengeType: 28
dashedName: challenge-189
---

# --description--

معطى الانخفاض الرأسي، والمسافة الأفقية، ونوع تل التزلج، حدّد تصنيف صعوبة التل.

لتحديد التصنيف:

- احسب انحدار التل بقسمة الانخفاض على المسافة.
- ثم احسب الانحدار المعدل بناءً على نوع التل:
  - `"Downhill"`: اضرب الانحدار في 1.2
  - `"Slalom"`: اضرب الانحدار في 0.9
  - `"Giant Slalom"`: اضرب الانحدار في 1.0

أرجع:

- `"Green"` إذا كان الانحدار المعدل أقل من أو يساوي 0.1
- `"Blue"` إذا كان الانحدار المعدل أكبر من 0.1 وأقل من أو يساوي 0.25
- `"Black"` إذا كان الانحدار المعدل أكبر من 0.25

# --hints--

`getHillRating(95, 900, "Slalom")` يجب أن تُرجع `"Green"`.

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")` يجب أن تُرجع `"Black"`.

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")` يجب أن تُرجع `"Blue"`.

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")` يجب أن تُرجع `"Green"`.

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")` يجب أن تُرجع `"Blue"`.

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")` يجب أن تُرجع `"Black"`.

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```
