---
id: 697a49e6ff50d756c9b6935e
title: "تحدٍّ 181: ألعاب الشتاء 2026 اليوم 2: التزلج على الثلج"
challengeType: 28
dashedName: challenge-181
---

# --description--

معطى وضعية البداية للمتزلج على الثلج وزاوية الدوران بالدرجات، حدّد وضعية الهبوط.

- وضعية المتزلج على الثلج تكون إما `"Regular"` أو `"Goofy"`.
- دوران الحركات هو مضاعفات 90 درجة. القيمة الموجبة تعني دورانًا مع عقارب الساعة، والقيمة السالبة تعني دورانًا عكس عقارب الساعة.
- تتغير وضعية الهبوط كل 180 درجة من الدوران.

على سبيل المثال، إذا كانت القيمة المعطاة هي `"Regular"` و `90`، فأرجع `"Regular"`. وإذا كانت القيمة المعطاة هي `"Regular"` و `180` درجة، فأرجع `"Goofy"`.

# --hints--

`getLandingStance("Regular", 90)` يجب أن تُرجع `"Regular"`.

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` يجب أن تُرجع `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` يجب أن تُرجع `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` يجب أن تُرجع `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` يجب أن تُرجع `"Goofy"`.

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` يجب أن تُرجع `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

`getLandingStance("Goofy", 90)` يجب أن تُرجع `"Goofy"`.

```js
assert.equal(getLandingStance("Goofy", 90), "Goofy");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
