---
id: 697a49e6ff50d756c9b69364
title: "تحدٍّ 187: ألعاب الشتاء 2026 اليوم 8: التزحلق على الزلاجات"
challengeType: 28
dashedName: challenge-187
---

# --description--

معطاة مصفوفة من خمسة أرقام، كل رقم يمثل الوقت (بالثواني) الذي استغرقه المتزحلق لإكمال جزء من المسار، حدّد أي جزء كان أسرع سرعة وما هي تلك السرعة.

المسار مقسّم إلى الأجزاء التالية:

- الجزء 1: 320 مترًا
- الجزء 2: 280 مترًا
- الجزء 3: 350 مترًا
- الجزء 4: 300 مترًا
- الجزء 5: 250 مترًا

القيمة الأولى في المصفوفة المعطاة تمثل الوقت للجزء 1، والقيمة الثانية للجزء 2، وهكذا.

لحساب السرعة (بالمتر في الثانية) لجزء معين، اقسم المسافة على الوقت.

أرجع `"The luger's fastest speed was X m/s on segment Y."`. حيث `X` هي أسرع سرعة، مقربة إلى منزلتين عشريتين، و`Y` هو رقم الجزء الذي حدثت فيه أسرع سرعة.

# --hints--

`getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128])` يجب أن تُرجع `"The luger's fastest speed was 35.07 m/s on segment 5."`

```js
assert.equal(getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.");
```

`getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284])` يجب أن تُرجع `"The luger's fastest speed was 37.75 m/s on segment 2."`

```js
assert.equal(getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.");
```

`getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912])` يجب أن تُرجع `"The luger's fastest speed was 38.49 m/s on segment 3."`

```js
assert.equal(getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.");
```

`getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840])` يجب أن تُرجع `"The luger's fastest speed was 37.69 m/s on segment 1."`

```js
assert.equal(getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.");
```

`getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508])` يجب أن تُرجع `"The luger's fastest speed was 39.24 m/s on segment 4."`

```js
assert.equal(getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.");
```

# --seed--

## --seed-contents--

```js
function getFastestSpeed(times) {

  return times;
}
```

# --solutions--

```js
function getFastestSpeed(times) {
  const distances = [320, 280, 350, 300, 250];
  let fastestSpeed = 0;
  let fastestSegment = 0;

  for (let i = 0; i < times.length; i++) {
    const speed = distances[i] / times[i];
    if (speed > fastestSpeed) {
      fastestSpeed = speed;
      fastestSegment = i + 1;
    }
  }

  return `The luger's fastest speed was ${fastestSpeed.toFixed(2)} m/s on segment ${fastestSegment}.`;
}
```
