---
id: 698a1a73ade5ac0e19180fa1
title: "تحدٍّ 198: حساب أيام العمل"
challengeType: 28
dashedName: challenge-198
---

# --description--

عند إعطائك تاريخ بداية وتاريخ نهاية، أرجع عدد أيام العمل بينهما.

- التواريخ المعطاة بصيغة `"YYYY-MM-DD"`.
- أيام الأسبوع هي أيام العمل (من الإثنين إلى الجمعة).
- عطلات نهاية الأسبوع ليست أيام عمل (السبت والأحد).
- ضمن العد كلا التاريخين البداية والنهاية.

# --hints--

يجب أن تُرجع الدالة `countBusinessDays("2026-02-24", "2026-02-26")` القيمة `3`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

يجب أن تُرجع الدالة `countBusinessDays("2026-02-24", "2026-02-28")` القيمة `4`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

يجب أن تُرجع الدالة `countBusinessDays("2026-02-21", "2026-03-01")` القيمة `5`.

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

يجب أن تُرجع الدالة `countBusinessDays("2026-03-08", "2026-03-17")` القيمة `7`.

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

يجب أن تُرجع الدالة `countBusinessDays("2026-02-24", "2027-02-24")` القيمة `262`.

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
