---
id: 698a1a863194f1f4e63f6460
title: "تحدٍّ 209: حجم العنصر"
challengeType: 28
dashedName: challenge-209
---

# --description--

معطى حجم النافذة، وعرض عنصر بوحدات عرض الإطار `"vw"`، وارتفاع عنصر بوحدات ارتفاع الإطار `"vh"`، حدّد حجم العنصر بالبكسل.

- حجم النافذة المعطى وحجم العنصر المرجع عبارة عن سلاسل نصية بصيغة `"width x height"`، مثلًا `"1200 x 800"`.

- وحدات `"vw"` تمثل نسبة مئوية من عرض النافذة. على سبيل المثال، `"50vw"` تعني 50% من عرض النافذة.
- وحدات `"vh"` تمثل نسبة مئوية من ارتفاع النافذة. على سبيل المثال، `"50vh"` تعني 50% من ارتفاع النافذة.

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` يجب أن تُرجع `"600 x 400"`.

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` يجب أن تُرجع `"80 x 240"`.

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` يجب أن تُرجع `"70 x 15"`.

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` يجب أن تُرجع `"1824 x 1080"`.

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` يجب أن تُرجع `"0 x 0"`.

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` يجب أن تُرجع `"1440 x 1026"`.

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
