---
id: 698a1a73ade5ac0e19180fa3
title: "تحدٍّ 200: عد الحروف والأرقام"
challengeType: 28
dashedName: challenge-200
---

# --description--

أعطِ سلسلة نصية، أرجع رسالة بعدد الحروف والأرقام التي تحتويها.

- الحروف هي `A-Z` و `a-z`.
- الأرقام هي `0-9`.
- تجاهل كل الأحرف الأخرى.

أرجع `"The string has X letters and Y numbers."`، حيث `"X"` هو عدد الحروف و`"Y"` هو عدد الأرقام. إذا كان العدد 1، استخدم الصيغة المفردة لذلك العنصر. مثلاً: `"1 letter"` بدلًا من `"1 letters"` و`"1 number"` بدلًا من `"1 numbers"`.

# --hints--

`countLettersAndNumbers("helloworld123")` يجب أن تُرجع `"The string has 10 letters and 3 numbers."`.

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")` يجب أن تُرجع `"The string has 5 letters and 2 numbers."`.

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")` يجب أن تُرجع `"The string has 1 letter and 1 number."`.

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")` يجب أن تُرجع `"The string has 0 letters and 5 numbers."`.

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")` يجب أن تُرجع `"The string has 8 letters and 0 numbers."`.

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
