---
id: 698a1a863194f1f4e63f645e
title: "تحدٍّ 207: أصغر فجوة"
challengeType: 28
dashedName: challenge-207
---

# --description--

معطاة سلسلة نصية، أرجع الجزء الفرعي بين الحرفين المتطابقين اللذين يفصل بينهما أقل عدد من الأحرف (أصغر فجوة).

- سيكون هناك دائمًا زوج واحد على الأقل من الأحرف المتطابقة.
- يجب أن يستبعد الجزء الفرعي المرجع الأحرف المتطابقة.
- إذا تساوت فجوتان أو أكثر في الطول، أرجع الأحرف من الفجوة الأولى.

على سبيل المثال، معطاة `"ABCDAC"`، أرجع `"DA"` لأن:

- فقط `"A"` و `"C"` يتكرران في السلسلة النصية.
- عدد الأحرف بين الحرفين `"A"` هو 3، وبين الحرفين `"C"` هو 2.
- لذا أرجع السلسلة النصية بين الحرفين `"C"`.

# --hints--

`smallestGap("ABCDAC")` يجب أن تُرجع `"DA"`.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` يجب أن تُرجع `"e"`.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` يجب أن تُرجع `"#q6e&rkf(p"`.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` يجب أن تُرجع `""`.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` يجب أن تُرجع `"fox"`.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
