---
id: 698a1a73ade5ac0e19180fa2
title: "تحدٍّ 199: الفرق التتابعي"
challengeType: 28
dashedName: challenge-199
---

# --description--

معطاة مصفوفة من الأرقام، أرجع مصفوفة جديدة تحتوي القيمة اللازمة للانتقال من كل رقم إلى الرقم التالي.

- للرقم الأخير، استخدم `0` لأنه لا يوجد رقم تالي.

على سبيل المثال، إذا كانت المصفوفة `[1, 2, 4, 7]`، أرجع `[1, 2, 3, 0]`.

# --hints--

يجب أن تُرجع الدالة `findDifferences([1, 2, 4, 7])` القيمة `[1, 2, 3, 0]`.

```js
assert.deepEqual(findDifferences([1, 2, 4, 7]), [1, 2, 3, 0]);
```

يجب أن تُرجع الدالة `findDifferences([10, 15, 19, 22, 24, 25])` القيمة `[5, 4, 3, 2, 1, 0]`.

```js
assert.deepEqual(findDifferences([10, 15, 19, 22, 24, 25]), [5, 4, 3, 2, 1, 0]);
```

يجب أن تُرجع الدالة `findDifferences([25, 20, 16, 13, 11, 10])` القيمة `[-5, -4, -3, -2, -1, 0]`.

```js
assert.deepEqual(findDifferences([25, 20, 16, 13, 11, 10]), [-5, -4, -3, -2, -1, 0]);
```

يجب أن تُرجع الدالة `findDifferences([0, 1, 2, 2, 3, 3, 4, 5])` القيمة `[ 1, 1, 0, 1, 0, 1, 1, 0 ]`.

```js
assert.deepEqual(findDifferences([0, 1, 2, 2, 3, 3, 4, 5]), [ 1, 1, 0, 1, 0, 1, 1, 0 ]);
```

يجب أن تُرجع الدالة `findDifferences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1])` القيمة `[1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]`.

```js
assert.deepEqual(findDifferences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1]), [1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]);
```

# --seed--

## --seed-contents--

```js
function findDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function findDifferences(arr) {
  const result = [];

  for (let i = 0; i < arr.length; i++) {
    result.push(arr[i + 1] - arr[i] || 0);
  }

  return result;
}
```
