---
id: 698a1a73ade5ac0e19180fa6
title: "تحدٍّ 203: مسطحة"
challengeType: 28
dashedName: challenge-203
---

# --description--

معطاة مصفوفة، حدّد إذا كانت مسطحة.

- المصفوفة مسطحة إذا لم يكن أي من عناصرها مصفوفات.

# --hints--

`isFlat([1, 2, 3, 4])` يجب أن تُرجع `true`.

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` يجب أن تُرجع `false`.

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` يجب أن تُرجع `true`.

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` يجب أن تُرجع `false`.

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` يجب أن تُرجع `false`.

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
