---
id: 698a1a73ade5ac0e19180fa8
title: "تحدٍّ 205: عد المكعبات المثالية"
challengeType: 28
dashedName: challenge-205
---

# --description--

معطى عددان صحيحان، حدّد كم عدد المكعبات المثالية الموجودة في النطاق بين العددين وشاملهما.

- لا يضمن أن يكون العدد الأصغر هو المعلمة الأولى.
- العدد مكعب مثالي إذا وُجد عدد صحيح (`n`) بحيث `n * n * n = number`. على سبيل المثال، 27 هو مكعب مثالي لأن `3 * 3 * 3 = 27`.

# --hints--

`countPerfectCubes(3, 30)` يجب أن تُرجع `2`.

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` يجب أن تُرجع `3`.

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` يجب أن تُرجع `4`.

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` يجب أن تُرجع `9`.

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` يجب أن تُرجع `41`.

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
