---
id: 698a1a863194f1f4e63f645f
title: "تحدٍّ 208: تتبع المسار"
challengeType: 28
dashedName: challenge-208
---

# --description--

معطاة مصفوفة من السلاسل النصية تمثل خريطة مسارك، أعد سلسلة نصية بالحركات اللازمة للوصول إلى هدفك.

السلاسل النصية المعطاة ستحتوي على القيم التالية:

- `"C"`: موقعك الحالي
- `"G"`: هدفك
- `"T"`: أجزاء المسار القابلة للعبور
- `"-"`: أجزاء الخريطة غير القابلة للعبور

أعد سلسلة نصية بالحركات اللازمة لاتباع المسار من موقعك إلى هدفك حيث:

- `"R"` تعني حركة إلى اليمين
- `"D"` تعني حركة إلى الأسفل
- `"L"` تعني حركة إلى اليسار
- `"U"` تعني حركة إلى الأعلى

- سيكون هناك دائمًا مسار واحد مستمر، بدون أي تفرعات، من موقعك الحالي إلى هدفك.
- كل موقع على المسار سيتصل بحد أقصى بموقعين قابلين للعبور.

على سبيل المثال، معطى:

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

أعد `"RDDRDD"`.

# --hints--

`navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` يجب أن تعيد `"RDDRDD"`.

```js
assert.equal(navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD");
```

`navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` يجب أن تعيد `"RRUUURR"`.

```js
assert.equal(navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR");
```

`navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` يجب أن تعيد `"DLDDRRRRD"`.

```js
assert.equal(navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD");
```

`navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` يجب أن تعيد `"RRRDDL"`.

```js
assert.equal(navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL");
```

`navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` يجب أن تعيد `"ULLLUUURRRRRRDDDR"`.

```js
assert.equal(navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR");
```

# --seed--

## --seed-contents--

```js
function navigateTrail(map) {

  return map;
}
```

# --solutions--

```js
function navigateTrail(map) {
  const rows = map.length;
  const cols = map[0].length;

  const directions = [
    [0, 1, "R"],
    [1, 0, "D"],
    [0, -1, "L"],
    [-1, 0, "U"],
  ];

  let row, col;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (map[r][c] === "C") {
        row = r;
        col = c;
      }
    }
  }

  let result = "";
  let prev = null;

  while (map[row][col] !== "G") {
    for (let [dr, dc, move] of directions) {
      const newRow = row + dr;
      const newCol = col + dc;

      if (
        newRow >= 0 &&
        newRow < rows &&
        newCol >= 0 &&
        newCol < cols &&
        (map[newRow][newCol] === "T" ||
          map[newRow][newCol] === "G") &&
        (!prev || newRow !== prev[0] || newCol !== prev[1])
      ) {
        result += move;
        prev = [row, col];
        row = newRow;
        col = newCol;
        break;
      }
    }
  }

  return result;
}
```
