---
id: 698a1a73ade5ac0e19180fa5
title: "تحدٍّ 202: إضافة علامات الترقيم"
challengeType: 28
dashedName: challenge-202
---

# --description--

أعطِ سلسلة نصية من الجمل تفتقد الفواصل، أضف نقطة (`"."`) في الأماكن التالية:

- قبل كل مسافة تسبق مباشرة حرفًا كبيرًا
- وفي نهاية السلسلة النصية

أرجع السلسلة النصية الناتجة.

# --hints--

`addPunctuation("Hello world")` يجب أن تُرجع `"Hello world."`

```js
assert.equal(addPunctuation("Hello world"), "Hello world.");
```

`addPunctuation("Hello world It's nice today")` يجب أن تُرجع `"Hello world. It's nice today."`

```js
assert.equal(addPunctuation("Hello world It's nice today"), "Hello world. It's nice today.");
```

`addPunctuation("JavaScript is great Sometimes")` يجب أن تُرجع `"JavaScript is great. Sometimes."`

```js
assert.equal(addPunctuation("JavaScript is great Sometimes"), "JavaScript is great. Sometimes.");
```

`addPunctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z")` يجب أن تُرجع `"A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z."`

```js
assert.equal(addPunctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z"), "A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z.");
```

`addPunctuation("Wait.. For it")` يجب أن تُرجع `"Wait... For it."`

```js
assert.equal(addPunctuation("Wait.. For it"), "Wait... For it.");
```

# --seed--

## --seed-contents--

```js
function addPunctuation(sentences) {

  return sentences;
}
```

# --solutions--

```js
function addPunctuation(sentences) {
  let result = "";

  for (let i = 0; i < sentences.length; i++) {
    if (sentences[i] === " " && i + 1 < sentences.length && /[A-Z]/.test(sentences[i + 1])) {
      result += ".";
    }
    result += sentences[i];
  }
  result += ".";

  return result;
}
```
