---
id: 698a1a73ade5ac0e19180fa0
title: "تحدٍّ 197: توافق فصائل الدم"
challengeType: 28
dashedName: challenge-197
---

# --description--

معطى فصيلة دم المتبرع وفصيلة دم المتلقي، حدّد ما إذا كان المتبرع يستطيع التبرع بالدم للمتلقي.

تتكون كل فصيلة دم من:

- حرف: `"A"`، `"B"`، `"AB"`، أو `"O"`
- وعامل Rh: `"+"` أو `"-"`

ستكون فصائل الدم واحدة من الحروف الصحيحة تليها عامل Rh. على سبيل المثال، `"AB+"` و`"O-"` هما فصائل دم صحيحة.

قواعد الحروف:

- `"O"` يمكنه التبرع لأي نوع حرف آخر.
- `"A"` يمكنه التبرع لـ `"A"` و`"AB"`.
- `"B"` يمكنه التبرع لـ `"B"` و`"AB"`.
- `"AB"` يمكنه التبرع فقط لـ `"AB"`.

قواعد عامل Rh:

- السالب (`"-"`) يمكنه التبرع لكل من `"-"` و`"+"`.
- الموجب (`"+"`) يمكنه التبرع فقط لـ `"+"`.

يجب أن تنطبق كل من قواعد الحروف وقواعد Rh ليتمكن المتبرع من التبرع للمتلقي.

# --hints--

`canDonate("B+", "B+")` يجب أن تُرجع `true`.

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` يجب أن تُرجع `true`.

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` يجب أن تُرجع `false`.

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` يجب أن تُرجع `true`.

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` يجب أن تُرجع `false`.

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` يجب أن تُرجع `true`.

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` يجب أن تُرجع `false`.

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` يجب أن تُرجع `true`.

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` يجب أن تُرجع `false`.

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` يجب أن تُرجع `false`.

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
