---
id: 698a1a73ade5ac0e19180fa7
title: "تحدٍّ 204: مجموع الحروف"
challengeType: 28
dashedName: challenge-204
---

# --description--

أعطِ سلسلة نصية، أرجع مجموع حروفها.

- الحروف هي A-Z بأحرف كبيرة أو أحرف صغيرة
- قيم الحروف هي: `"A"` = 1، `"B"` = 2، ...، `"Z"` = 26
- الحروف الكبيرة والصغيرة لها نفس القيمة.
- تجاهل كل الأحرف غير الحروف.

# --hints--

`sumLetters("Hello")` يجب أن تُرجع `52`.

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")` يجب أن تُرجع `94`.

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")` يجب أن تُرجع `473`.

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` يجب أن تُرجع `1681`.

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")` يجب أن تُرجع `0`.

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```
