---
id: 698a1a73ade5ac0e19180fa9
title: "تحدٍّ 206: قيم أوراق اللعب"
challengeType: 28
dashedName: challenge-206
---

# --description--

معطاة مصفوفة من أوراق اللعب، أرجع مصفوفة جديدة تحتوي على القيمة الرقمية لكل ورقة.

قيم الأوراق:

- الآس (`"A"`) قيمته 1.
- الأوراق المرقمة (`"2"` - `"10"`) قيمتها تساوي الرقم المكتوب عليها: من 2 إلى 10 على التوالي.
- أوراق الوجه: الولد (`"J"`)، الملكة (`"Q"`)، والملك (`"K"`) كل منها قيمته 10.

الرموز:

- كل ورقة لها رمز: البستوني (`"S"`)، الكبة (`"C"`)، الماس (`"D"`)، أو القلوب (`"H"`).

صيغة الورقة:

- تمثّل كل ورقة كسلسلة نصية: `"valueSuit"`. على سبيل المثال: `"AS"` هو آس البستوني، `"10H"` هو عشرة القلوب، و`"QC"` هي ملكة الكبة.

# --hints--

`cardValues(["3H", "4D", "5S"])` يجب أن تُرجع `[3, 4, 5]`.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` يجب أن تُرجع `[1, 10, 10, 6, 7]`.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` يجب أن تُرجع `[8, 10, 2, 10, 9]`.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` يجب أن تُرجع `[1, 10]`.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` يجب أن تُرجع `[10, 10, 10, 10, 1]`.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
