---
id: 6994cff2290543b3aec9f510
title: "تحدٍّ 211: مجموع المصفوفة"
challengeType: 28
dashedName: challenge-211
---

# --description--

معطاة مصفوفة من الأرقام، أرجع مجموع كل الأرقام.

# --hints--

يجب أن تُرجع الدالة `sumArray([1, 2, 3, 4, 5])` القيمة `15`.

```js
assert.equal(sumArray([1, 2, 3, 4, 5]), 15);
```

يجب أن تُرجع الدالة `sumArray([42])` القيمة `42`.

```js
assert.equal(sumArray([42]), 42);
```

يجب أن تُرجع الدالة `sumArray([5, -2, 7, -3])` القيمة `7`.

```js
assert.equal(sumArray([5, -2, 7, -3]), 7);
```

يجب أن تُرجع الدالة `sumArray([203, 145, -129, 6293, 523, -919, 845, 2434])` القيمة `9395`.

```js
assert.equal(sumArray([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395);
```

يجب أن تُرجع الدالة `sumArray([0, 0])` القيمة `0`.

```js
assert.equal(sumArray([0, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function sumArray(numbers) {

  return numbers;
}
```

# --solutions--

```js
function sumArray(numbers) {

  return numbers.reduce((a, c) => c + a, 0);
}
```
