---
id: 6994cff2290543b3aec9f50f
title: "تحدٍّ 210: مدقق HSL"
challengeType: 28
dashedName: challenge-210
---

# --description--

بناءً على سلسلة نصية، حدّد ما إذا كانت قيمة لون `hsl()` في CSS صحيحة.

- يجب أن تكون قيمة HSL صحيحة بالتنسيق `"hsl(h, s%, l%)"`، حيث:
  - `h` (درجة اللون) يجب أن تكون رقمًا بين 0 و360 (شاملة).
  - `s` (التشبع) يجب أن تكون نسبة مئوية بين 0% و100%.
  - `l` (فاتح اللون) يجب أن تكون نسبة مئوية بين 0% و100%.

- يُسمح بالمسافات فقط:
  - بعد القوس المفتوح
  - قبل وبعد الفواصل أو بعدهم
  - قبل وبعد القوس المغلق أو بعده

- يمكن أن تنتهي القيمة اختياريًا بفاصلة منقوطة (`";"`).

على سبيل المثال، `"hsl(240, 50%, 50%)"` هي قيمة HSL صحيحة.

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
