---
id: 699c8e045ee7cb94ed2322d5
title: "تحدٍّ 214: مدقق سلسلة الدومينو"
challengeType: 28
dashedName: challenge-214
---

# --description--

معطى مصفوفة ثنائية الأبعاد تمثل متتالية من قطع الدومينو، حدّد ما إذا كانت تشكل سلسلة صحيحة.

- كل عنصر في المصفوفة يمثل قطعة دومينو ويكون مصفوفة من رقمين من 1 إلى 6، (شاملة).
- لكي تكون السلسلة صحيحة، يجب أن يطابق الرقم الثاني من كل قطعة الرقم الأول من القطعة التالية.
- لا يلزم أن يطابق الرقم الأول من أول قطعة أو الرقم الأخير من آخر قطعة أي شيء.

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
