---
id: 699c8e045ee7cb94ed2322d6
title: "تحدٍّ 215: حاسبة رسوم الموقف"
challengeType: 28
dashedName: challenge-215
---

# --description--

معطيتان تمثلان وقت ركن سيارتك ووقت استلامها، احسب رسوم الموقف.

- السلسلتان المعطيتان ستكونان بصيغة `"HH:MM"` باستخدام ساعة 24 ساعة. فمثلاً `"14:00"` تعني الساعة 2 ظهرًا.
- السلسلة الأولى تمثل وقت ركن السيارة، والثانية تمثل وقت استلامها.
- إذا كان وقت الاستلام أبكر من وقت الدخول، فهذا يعني أن الموقف امتد بعد منتصف الليل إلى اليوم التالي.

قواعد الرسوم:

- كل ساعة ركن تكلف 3 دولارات.
- الساعات الجزئية تُقرب إلى الساعة الكاملة التالية.
- إذا امتد الموقف طوال الليل (بعد منتصف الليل)، تُضاف رسوم إضافية قدرها 10 دولارات.
- هناك حد أدنى للرسوم بقيمة 5 دولارات (يُستخدم فقط إذا كان المجموع أقل من 5 دولارات).

أرجع التكلفة الإجمالية بصيغة `"$cost"`، مثلًا `"$5"`.

# --hints--

`calculateParkingFee("09:00", "11:00")` يجب أن تُرجع `"$6"`.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` يجب أن تُرجع `"$5"`.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` يجب أن تُرجع `"$9"`.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` يجب أن تُرجع `"$27"`.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` يجب أن تُرجع `"$34"`.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` يجب أن تُرجع `"$82"`.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
