---
id: 699c8e045ee7cb94ed2322d8
title: "تحدٍّ 217: قطع الشطرنج المأسورة"
challengeType: 28
dashedName: challenge-217
---

# --description--

معطاة مصفوفة من السلاسل النصية تمثل قطع الشطرنج التي لا تزال لديك على الرقعة، احسب قيمة القطع التي أسرها خصمك.

في الشطرنج، تبدأ بـ 16 قطعة:

| القطعة  | الاختصار   | الكمية | القيمة |
| ------ | -------------- | -------- | ----- |
| جندي   | `"P"`          | 8        | 1     |
| رخ     | `"R"`          | 2        | 5     |
| حصان   | `"N"`          | 2        | 3     |
| فيل    | `"B"`          | 2        | 3     |
| وزير   | `"Q"`          | 1        | 9     |
| ملك    | `"K"`          | 1        | 0     |

- المصفوفة المعطاة ستحتوي فقط على الاختصارات أعلاه.
- أي من القطع الـ 16 غير المدرجة في المصفوفة المعطاة قد أُسرت.
- أرجع القيمة الإجمالية لكل القطع المأسورة، إلا إذا...
- إذا أُسر الملك، أرجع `"Checkmate"`.

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` يجب أن تُرجع `8`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])` يجب أن تُرجع `26`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` يجب أن تُرجع `16`.

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` يجب أن تُرجع `4`.

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])` يجب أن تُرجع `38`.

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` يجب أن تُرجع `0`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` يجب أن تُرجع `"Checkmate"`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```
