---
id: 699c8e045ee7cb94ed2322d9
title: "تحدٍّ 218: قابل للقسمة بالتساوي"
challengeType: 28
dashedName: challenge-218
---

# --description--

بما أن لديك عددين صحيحين، حدّد إذا كان يمكنك قسمة الأول بالتساوي على الثاني.

# --hints--

يجب أن تُرجع الدالة `isEvenlyDivisible(4, 2)` القيمة `true`.

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

يجب أن تُرجع الدالة `isEvenlyDivisible(7, 3)` القيمة `false`.

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

يجب أن تُرجع الدالة `isEvenlyDivisible(5, 10)` القيمة `false`.

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

يجب أن تُرجع الدالة `isEvenlyDivisible(48, 6)` القيمة `true`.

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

يجب أن تُرجع الدالة `isEvenlyDivisible(3186, 9)` القيمة `true`.

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

يجب أن تُرجع الدالة `isEvenlyDivisible(4192, 11)` القيمة `false`.

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
