---
id: 699c8e045ee7cb94ed2322dd
title: "تحدٍّ 222: ظلال الاعتدال"
challengeType: 28
dashedName: challenge-222
---

# --description--

اليوم هو يوم الاعتدال، حيث تكون الشمس مباشرة فوق خط الاستواء ومباشرة في الأعلى عند الظهر. بناءً على وقت معين، حدّد الظل الذي يلقيه عمود عمودي طوله 4 أقدام.

- سيكون الوقت سلسلة نصية بصيغة `"HH:MM"` بنظام 24 ساعة (على سبيل المثال، `"15:00"` تعني الساعة 3 مساءً).
- سيُعطى الوقت فقط بفواصل 30 دقيقة.

القواعد:

- تشرق الشمس عند الساعة 6 صباحًا مباشرةً من جهة `"east"`، وتغرب عند الساعة 6 مساءً مباشرةً من جهة `"west"`.
- يشير الظل دائمًا في الاتجاه المعاكس للشمس.
- طول الظل (بالأقدام) هو عدد الساعات بعيدًا عن الظهر مرفوعًا للقوة الثالثة.
- لا يوجد ظل قبل شروق الشمس (قبل الساعة 6 صباحًا)، أو بعد غروب الشمس (الساعة 6 مساءً أو بعدها)، أو عند الظهر.

الإرجاع:

- إذا وُجد ظل، أرجع `"(length)ft (direction)"`. على سبيل المثال، `"8ft west"`.
- وإلا، أرجع `"No shadow"`.

على سبيل المثال، عند إعطاء `"10:00"`، أرجع `"8ft west"` لأن الساعة 10 صباحًا تبعد ساعتين عن الظهر، إذن 2<sup>3</sup> = 8 أقدام، والظل يشير إلى الغرب لأن الشمس في الشرق عند الساعة 10 صباحًا.

# --hints--

`getShadow("10:00")` يجب أن تُرجع `"8ft west"`.

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")` يجب أن تُرجع `"27ft east"`.

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")` يجب أن تُرجع `"No shadow"`.

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")` يجب أن تُرجع `"166.375ft east"`.

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")` يجب أن تُرجع `"No shadow"`.

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")` يجب أن تُرجع `"216ft west"`.

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")` يجب أن تُرجع `"No shadow"`.

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")` يجب أن تُرجع `"91.125ft west"`.

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")` يجب أن تُرجع `"No shadow"`.

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
