---
id: 699c8e045ee7cb94ed2322db
title: "تحدٍّ 220: أكبر رقم"
challengeType: 28
dashedName: challenge-220
---

# --description--

أعطيت سلسلة نصية من الأرقام مفصولة بعلامات ترقيم مختلفة، أعد أكبر رقم.

- ستحتوي السلسلة النصية المعطاة على أرقام وفواصل فقط.
- يمكن أن تكون الفواصل عبارة عن `","` أو `"!"` أو `"?"` أو `":"` أو `";"`.

# --hints--

`largestNumber("1,2")` يجب أن تُعيد `2`.

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` يجب أن تُعيد `60`.

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` يجب أن تُعيد `-402`.

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` يجب أن تُعيد `99.9`.

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
