---
id: 699c8e045ee7cb94ed2322da
title: "تحدٍّ 219: معالم الذكرى السنوية"
challengeType: 28
dashedName: challenge-219
---

# --description--

معطى عدد صحيح يمثل عدد سنوات زواج الزوجين، أعد أحدث معلم ذكرى سنوية لهما وفقًا لهذا المخطط:

| سنوات الزواج | المعلم      |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- إذا لم يصلوا إلى أول معلم، أعد `"Newlyweds"`.

# --hints--

`getMilestone(0)` يجب أن يعيد `"Newlyweds"`.

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` يجب أن يعيد `"Paper"`.

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` يجب أن يعيد `"Wood"`.

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` يجب أن يعيد `"Tin"`.

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` يجب أن يعيد `"Silver"`.

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` يجب أن يعيد `"Ruby"`.

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` يجب أن يعيد `"Gold"`.

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` يجب أن يعيد `"Diamond"`.

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` يجب أن يعيد `"Platinum"`.

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
