---
id: 699c8e045ee7cb94ed2322d4
title: "تحدٍّ 213: محول طول الكلمات"
challengeType: 28
dashedName: challenge-213
---

# --description--

بالنظر إلى سلسلة نصية من الكلمات، أعد سلسلة جديدة حيث تُستبدَل كل كلمة بطولها.

- الكلمات في السلسلة النصية المعطاة ستكون مفصولة بمسافة واحدة
- احتفظ بالمسافات في السلسلة المعادة.

على سبيل المثال، إذا كانت القيمة المعطاة `"hello world"`، فأعد `"5 5"`.

# --hints--

`convertWords("hello world")` يجب أن تُعيد `"5 5"`.

```js
assert.equal(convertWords("hello world"), "5 5");
```

`convertWords("Thanks and happy coding")` يجب أن تُعيد `"6 3 5 6"`.

```js
assert.equal(convertWords("Thanks and happy coding"), "6 3 5 6");
```

`convertWords("The quick brown fox jumps over the lazy dog")` يجب أن تُعيد `"3 5 5 3 5 4 3 4 3"`.

```js
assert.equal(convertWords("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3");
```

`convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` يجب أن تُعيد `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`.

```js
assert.equal(convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4");
```

# --seed--

## --seed-contents--

```js
function convertWords(str) {

  return str;
}
```

# --solutions--

```js
function convertWords(str) {
  return str
    .split(" ")
    .map(word => word.length)
    .join(" ");
}
```
