---
id: 69a890af247de743333bd4cd
title: "تحدٍّ 224: كاشف تحميص القهوة"
challengeType: 28
dashedName: challenge-224
---

# --description--

معطى سلسلة نصية تمثل الحبوب المستخدمة لصنع فنجان قهوة، حدّد درجة تحميص الفنجان.

- ستحتوي السلسلة النصية المعطاة على الأحرف التالية، كل منها يمثل نوعًا من الحبوب:

- الفاصلة العليا (`'`) تمثل حبة تحميص خفيف بقيمة نقطة واحدة لكل منها.
  - الشرطة (`-`) تمثل حبة تحميص متوسط بقيمة نقطتين لكل منها.
  - النقطة (`.`) تمثل حبة تحميص داكن بقيمة ثلاث نقاط لكل منها.

- تُحدد درجة التحميص بناءً على متوسط كل الحبوب.

أرجع:

- `"Light"` إذا كان المتوسط أقل من 1.75.
- `"Medium"` إذا كان المتوسط بين 1.75 و2.5.
- `"Dark"` إذا كان المتوسط أكبر من 2.5.

# --hints--

`detectRoast("''-''''''-'-''--''''")` يجب أن تُرجع `"Light"`.

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` يجب أن تُرجع `"Medium"`.

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` يجب أن تُرجع `"Medium"`.

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` يجب أن تُرجع `"Dark"`.

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` يجب أن تُرجع `"Medium"`.

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` يجب أن تُرجع `"Dark"`.

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` يجب أن تُرجع `"Light"`.

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
