---
id: 69a890af247de743333bd4d2
title: "تحدٍّ 229: اقتطاع النص 2"
challengeType: 28
dashedName: challenge-229
---

# --description--

بالنظر إلى سلسلة نصية، أعد سلسلة جديدة مقتطعة بحيث لا يتجاوز العرض الكلي للأحرف 50 وحدة.

لكل حرف عرض محدد:

| الأحرف | العرض |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

الجدول أعلاه يشمل جميع الأحرف الكبيرة والصغيرة. بالإضافة إلى ذلك:

- الفراغات (`" "`) عرضها 2
- النقاط (`"."`) عرضها 1

- إذا كانت السلسلة النصية المعطاة 50 وحدة أو أقل، أعد السلسلة كما هي، وإلا
- اقتطع السلسلة وأضف ثلاث نقاط في النهاية (`"..."`) بحيث يكون عرضها الكلي، بما في ذلك النقاط الثلاث، قريبًا قدر الإمكان من 50 وحدة دون تجاوزها.

# --hints--

`truncateText("The quick brown fox")` يجب أن تعيد `"The quick brown f..."`.

```js
assert.equal(truncateText("The quick brown fox"), "The quick brown f...");
```

`truncateText("The silky smooth sloth")` يجب أن تعيد `"The silky smooth s..."`.

```js
assert.equal(truncateText("The silky smooth sloth"), "The silky smooth s...");
```

`truncateText("THE LOUD BRIGHT BIRD")` يجب أن تعيد `"THE LOUD BRIG..."`.

```js
assert.equal(truncateText("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...");
```

`truncateText("The fast striped zebra")` يجب أن تعيد `"The fast striped z..."`.

```js
assert.equal(truncateText("The fast striped zebra"), "The fast striped z...");
```

`truncateText("The big black bear")` يجب أن تعيد `"The big black bear"`.

```js
assert.equal(truncateText("The big black bear"), "The big black bear");
```

# --seed--

## --seed-contents--

```js
function truncateText(str) {

  return str;
}
```

# --solutions--

```js
function truncateText(str) {
  const MAX_WIDTH = 50;
  const ELLIPSIS = "...";
  const ELLIPSIS_WIDTH = 3;
  const TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH;

  const charGroups = {
    'ilI.': 1,
    'fjrt ': 2,
    'abcdeghkmnopqrstuvwxyzJL': 3,
    'ABCDEFGHKMNOPQRSTUVWXYZ': 4
  };

  function getCharWidth(char) {
    for (const key in charGroups) {
      if (key.includes(char)) return charGroups[key];
    }
    return 3;
  }

  function stringWidth(str) {
    let totalWidth = 0;
    for (const char of str) {
      totalWidth += getCharWidth(char);
    }
    return totalWidth;
  }

  if (stringWidth(str) <= MAX_WIDTH) return str;

  let result = "";
  let totalWidth = 0;

  for (const char of str) {
    const charWidth = getCharWidth(char);

    if (totalWidth + charWidth > TRUNCATE_LIMIT) break;

    result += char;
    totalWidth += charWidth;
  }

  return result + ELLIPSIS;
}
```
