---
id: 69a890af247de743333bd4cc
title: "تحدٍّ 223: فك شفرة QR"
challengeType: 28
dashedName: challenge-223
---

# --description--

معطى مصفوفة 6x6 (مصفوفة من السلاسل النصية) تمثل رمز QR، أرجع سلسلة البيانات الثنائية في الرمز.

- قد يُعطى رمز QR بأي دوران بزيادات 90 درجة.
- الرمز الموجه بشكل صحيح يحتوي على مجموعة 2x2 من `1` (علامات التوجيه) في الزوايا السفلى اليسرى، العلوية اليسرى، والعلوية اليمنى.
- علامات التوجيه الثلاث 2x2 ليست جزءًا من البيانات الثنائية.
- تُقرأ البيانات الثنائية من اليسار إلى اليمين، ومن الأعلى إلى الأسفل (مثل الكتاب) عندما يكون رمز QR موجهًا بشكل صحيح.
- سيكون للرمز دائمًا اتجاه صحيح واحد فقط.

على سبيل المثال، معطى:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

أو معطى نفس الرمز مع اتجاه مختلف:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

أرجع `"000000000000000000000001"`، كل البيانات الثنائية باستثناء علامات التوجيه الثلاث 2x2.

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])` يجب أن تُرجع `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])` يجب أن تُرجع `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])` يجب أن تُرجع `"001101000011000000110100"`.

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])` يجب أن تُرجع `"010001000100010101010110"`.

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])` يجب أن تُرجع `"010000100100100101001110"`.

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
