---
id: 69a890af247de743333bd4d1
title: "تحدٍّ 228: ليلة الفيلم"
challengeType: 28
dashedName: challenge-228
---

# --description--

معطى سلسلة نصية لليوم من الأسبوع، وسلسلة نصية أخرى لوقت العرض، وعدد صحيح من التذاكر، أرجع التكلفة الإجمالية لتذاكر الفيلم لذلك العرض.

اليوم المعطى سيكون واحدًا من:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

سيُعطى وقت العرض بالتنسيق `"H:MMam"` أو `"H:MMpm"`. على سبيل المثال `"10:00am"` أو `"10:00pm"`.

أرجع التكلفة الإجمالية بالتنسيق `"$D.CC"` باستخدام القواعد التالية:

- عطلة نهاية الأسبوع (الجمعة - الأحد): 12.00$ لكل تذكرة.
- أيام الأسبوع (الاثنين - الخميس): 10.00$ لكل تذكرة.
- عروض النهار (قبل الساعة 5:00 مساءً): اخصم 2.00$ لكل تذكرة (باستثناء أيام الثلاثاء).
- أيام الثلاثاء: كل التذاكر بسعر 5.00$ لكل واحدة.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` يجب أن تُرجع `"$12.00"`.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` يجب أن تُرجع `"$10.00"`.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` يجب أن تُرجع `"$10.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` يجب أن تُرجع `"$30.00"`.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` يجب أن تُرجع `"$32.00"`.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` يجب أن تُرجع `"$50.00"`.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` يجب أن تُرجع `"$5.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
