---
id: 69a890af247de743333bd4ce
title: "تحدٍّ 225: لا تكرار متتالي"
challengeType: 28
dashedName: challenge-225
---

# --description--

أعطِ سلسلة نصية، حدّد إذا لم تحتوي على حروف مكررة.

- لا تحتوي السلسلة النصية على تكرار إذا لم يظهر نفس الحرف مرتين أو أكثر على التوالي.

# --hints--

`hasNoRepeats("hi world")` يجب أن تُرجع `true`.

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")` يجب أن تُرجع `false`.

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")` يجب أن تُرجع `true`.

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")` يجب أن تُرجع `false`.

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")` يجب أن تُرجع `true`.

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")` يجب أن تُرجع `false`.

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```
