---
id: 69a890af247de743333bd4d0
title: "تحدٍّ 227: وقت التهدئة"
challengeType: 28
dashedName: challenge-227
---

# --description--

معطى طابعان زمنيّان، الأول يمثل وقت انتهاء المستخدم من الاختبار، والثاني يمثل الوقت الحالي، حدّد ما إذا كان يمكن للمستخدم أن يخوض الاختبار مرة أخرى.

- سيُعطى كلا الطابعين الزمنيّين بالتنسيق: `"YYYY-MM-DDTHH:MM:SS"`، مثلًا `"2026-03-25T14:00:00"`. لاحظ أن الوقت بنظام 24 ساعة.
- يجب على المستخدم الانتظار 48 ساعة على الأقل قبل إعادة الاختبار.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` يجب أن تُرجع `true`.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` يجب أن تُرجع `false`.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` يجب أن تُرجع `true`.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` يجب أن تُرجع `false`.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
