---
id: 69b1028d6e265413d0198a2d
title: "تحدٍّ 234: رقم المقلب"
challengeType: 28
dashedName: challenge-234
---

# --description--

أعطِ مصفوفة من الأرقام حيث يتبع كل رقم ما عدا رقم واحد نمطًا معينًا، أرجع مصفوفة جديدة مع تصحيح الرقم الذي لا يتبع النمط.

سيكون النمط واحدًا من التالي:

- تزداد الأرقام من رقم إلى آخر بمقدار ثابت (جمع).
- تنقص الأرقام من رقم إلى آخر بمقدار ثابت (طرح).

على سبيل المثال، إذا كانت المصفوفة `[2, 4, 7, 8, 10]` أرجع `[2, 4, 6, 8, 10]`.

# --hints--

يجب أن تُرجع الدالة `fixPrankNumber([2, 4, 7, 8, 10])` القيمة `[2, 4, 6, 8, 10]`.

```js
assert.deepEqual(fixPrankNumber([2, 4, 7, 8, 10]), [2, 4, 6, 8, 10]);
```

يجب أن تُرجع الدالة `fixPrankNumber([10, 10, 8, 7, 6])` القيمة `[10, 9, 8, 7, 6]`.

```js
assert.deepEqual(fixPrankNumber([10, 10, 8, 7, 6]), [10, 9, 8, 7, 6]);
```

يجب أن تُرجع الدالة `fixPrankNumber([12, 24, 36, 48, 61, 72, 84, 96])` القيمة `[12, 24, 36, 48, 60, 72, 84, 96]`.

```js
assert.deepEqual(fixPrankNumber([12, 24, 36, 48, 61, 72, 84, 96]), [12, 24, 36, 48, 60, 72, 84, 96]);
```

يجب أن تُرجع الدالة `fixPrankNumber([4, 1, -2, -5, -8, -5])` القيمة `[4, 1, -2, -5, -8, -11]`.

```js
assert.deepEqual(fixPrankNumber([4, 1, -2, -5, -8, -5]), [4, 1, -2, -5, -8, -11]);
```

يجب أن تُرجع الدالة `fixPrankNumber([0, 100, 200, 300, 150, 500])` القيمة `[0, 100, 200, 300, 400, 500]`.

```js
assert.deepEqual(fixPrankNumber([0, 100, 200, 300, 150, 500]), [0, 100, 200, 300, 400, 500]);
```

يجب أن تُرجع الدالة `fixPrankNumber([400, 425, 400, 375, 350, 325, 300])` القيمة `[450, 425, 400, 375, 350, 325, 300]`.

```js
assert.deepEqual(fixPrankNumber([400, 425, 400, 375, 350, 325, 300]), [450, 425, 400, 375, 350, 325, 300]);
```

يجب أن تُرجع الدالة `fixPrankNumber([-5, 5, 10, 15, 20])` القيمة `[0, 5, 10, 15, 20]`.

```js
assert.deepEqual(fixPrankNumber([-5, 5, 10, 15, 20]), [0, 5, 10, 15, 20]);
```

# --seed--

## --seed-contents--

```js
function fixPrankNumber(arr) {

  return arr;
}
```

# --solutions--

```js
function fixPrankNumber(arr) {
  const diffs = arr.slice(1).map((v, i) => v - arr[i]);

  let step = diffs[0];
  for (let i = 0; i < diffs.length - 1; i++) {
    if (diffs[i] === diffs[i + 1]) {
      step = diffs[i];
      break;
    }
  }

  const result = [...arr];

  for (let i = 0; i < result.length - 1; i++) {
    if (result[i + 1] - result[i] !== step) {
      if (i === 0 && result[2] - result[1] === step) {
        result[0] = result[1] - step;
      } else {
        result[i + 1] = result[i] + step;
      }
      break;
    }
  }

  return result;
}
```
