---
id: 69b1028d6e265413d0198a2c
title: "تحدٍّ 233: منبه الاستيقاظ"
challengeType: 28
dashedName: challenge-233
---

# --description--

معطى سلسلة نصية تمثّل الوقت الذي ضبطت فيه المنبه وسلسلة نصية تمثّل الوقت الذي استيقظت فيه فعليًا، حدّد إذا استيقظت مبكرًا، في الوقت المحدد، أو متأخرًا.

- كلا الوقتين سيكونان بصيغة 24 ساعة `"HH:MM"`.

أرجع:

- `"early"` إذا استيقظت قبل وقت المنبه.
- `"on time"` إذا استيقظت في وقت المنبه، أو خلال نافذة الغفوة التي تستمر 10 دقائق بعد وقت المنبه.
- `"late"` إذا استيقظت بعد أكثر من 10 دقائق من وقت المنبه.

كلا الوقتين في نفس اليوم.

# --hints--

`alarmCheck("07:00", "06:45")` يجب أن تُرجع `"early"`.

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")` يجب أن تُرجع `"on time"`.

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")` يجب أن تُرجع `"on time"`.

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")` يجب أن تُرجع `"late"`.

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")` يجب أن تُرجع `"on time"`.

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")` يجب أن تُرجع `"late"`.

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")` يجب أن تُرجع `"early"`.

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
