---
id: 69b1028d6e265413d0198a2a
title: "تحدٍّ 231: مدقق ISBN-10"
challengeType: 28
dashedName: challenge-231
---

# --description--

معطاة سلسلة نصية، حدّد إذا كانت تمثّل رقم ISBN-10 صالحًا.

يتكوّن رقم ISBN-10 من شرطات (`"-"`) و10 أحرف أخرى. بعد إزالة الشرطات (`"-"`):

- يجب أن تكون الأحرف التسعة الأولى أرقامًا، و
- قد يكون الحرف الأخير رقمًا أو الحرف `"X"`، الذي يمثل الرقم 10.

للتحقق من صحته:

- اضرب كل رقم (أو قيمة) في موقعه (اضرب الرقم الأول في 1، والثاني في 2، وهكذا).
- اجمع كل النتائج معًا.
- إذا كان المجموع قابلًا للقسمة على 11، فهو صالح.

# --hints--

`isValidIsbn10("0-306-40615-2")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
