---
id: 69b58ce40693f140c84c8559
title: "تحدٍّ 240: حروف الكلمة المتماثلة"
challengeType: 28
dashedName: challenge-240
---

# --description--

معطاة سلسلة نصية، حدّد إذا كانت كلمة متماثلة وأرجع الحرف الأوسط (إذا كان طولها فرديًا) أو الحرفين الأوسطين (إذا كان طولها زوجيًا).

- الكلمة المتماثلة هي سلسلة نصية تتطابق عند قراءتها من الأمام إلى الخلف.
- إذا لم تكن كلمة متماثلة، أرجع `"none"`.

# --hints--

`palindromeLocator("racecar")` يجب أن تُرجع `"e"`.

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` يجب أن تُرجع `"v"`.

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` يجب أن تُرجع `"none"`.

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` يجب أن تُرجع `"oo"`.

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` يجب أن تُرجع `"00"`.

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
