---
id: 69b559d2903b9e4afe9075f7
title: "تحدٍّ 237: التحقق من المعادلة"
challengeType: 28
dashedName: challenge-237
---

# --description--

أعطيت سلسلة نصية تمثل معادلة رياضية، حدّد ما إذا كانت صحيحة.

- قد يحتوي الجانب الأيسر على ما يصل إلى ثلاثة أعداد صحيحة موجبة والعمليات `+`، `-`، `*`، و`/`.
- ستُعطى المعادلة بالتنسيق: `"number operator number = number"` (مع وجود رقمين أو ثلاثة أرقام على اليسار). على سبيل المثال: `"2 + 2 = 4"` أو `"2 + 3 - 1 = 4"`.
- الجانب الأيمن سيكون دائمًا عددًا صحيحًا واحدًا.

اتبع ترتيب العمليات القياسي: تُقيّم الضرب والقسمة قبل الجمع والطرح، من اليسار إلى اليمين.

# --hints--

`isValidEquation("2 + 2 = 4")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidEquation("2 + 2 = 4"));
```

`isValidEquation("2 + 3 - 1 = 4")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidEquation("2 + 3 - 1 = 4"));
```

`isValidEquation("8 / 2 = 4")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidEquation("8 / 2 = 4"));
```

`isValidEquation("10 * 5 = 50")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidEquation("10 * 5 = 50"));
```

`isValidEquation("2 - 2 = 0")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidEquation("2 - 2 = 0"));
```

`isValidEquation("2 + 9 / 3 = 5")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidEquation("2 + 9 / 3 = 5"));
```

`isValidEquation("20 - 2 * 3 = 14")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidEquation("20 - 2 * 3 = 14"));
```

`isValidEquation("2 + 5 = 6")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidEquation("2 + 5 = 6"));
```

`isValidEquation("10 - 2 * 3 = 24")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidEquation("10 - 2 * 3 = 24"));
```

`isValidEquation("3 + 9 / 3 = 4")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidEquation("3 + 9 / 3 = 4"));
```

# --seed--

## --seed-contents--

```js
function isValidEquation(equation) {

  return equation;
}
```

# --solutions--

```js
function isValidEquation(equation) {
  const [left, right] = equation.split(" = ");
  const tokens = left.split(" ");

  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    if (operator === "*" || operator === "/") {
      const leftOperand = Number(tokens[i - 1]);
      const rightOperand = Number(tokens[i + 1]);
      const computed = operator === "*" ? leftOperand * rightOperand : leftOperand / rightOperand;
      tokens.splice(i - 1, 3, String(computed));
      i -= 2;
    }
  }

  let result = Number(tokens[0]);
  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    const operand = Number(tokens[i + 1]);
    result = operator === "+" ? result + operand : result - operand;
  }

  return result === Number(right);
}
```
