---
id: 69b58ce40693f140c84c855a
title: "تحدٍّ 241: مدقق FizzBuzz"
challengeType: 28
dashedName: challenge-241
---

# --description--

معطاة مصفوفة من أعداد صحيحة متتالية، مع استبدال مضاعفات 3 و5، حدّد ما إذا كانت متتالية FizzBuzz صحيحة.

في متتالية FizzBuzz صحيحة:

- تُستبدَل مضاعفات 3 بـ "Fizz".
- تُستبدَل مضاعفات 5 بـ "Buzz".
- تُستبدَل مضاعفات كل من 3 و5 بـ "FizzBuzz".
- تبقى جميع الأعداد الأخرى كأعداد صحيحة.

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz"])` يجب أن تُعيد `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz"]));
```

`isFizzBuzz([13, 14, "FizzBuzz", 16, 17])` يجب أن تُعيد `true`.

```js
assert.isTrue(isFizzBuzz([13, 14, "FizzBuzz", 16, 17]));
```

`isFizzBuzz([1, 2, "Fizz", 4, 5])` يجب أن تُعيد `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, 5]));
```

`isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` يجب أن تُعيد `true`.

```js
assert.isTrue(isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", "Buzz", 5])` يجب أن تُعيد `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", "Buzz", 5]));
```

`isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` يجب أن تُعيد `false`.

```js
assert.isFalse(isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]));
```

`isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` يجب أن تُعيد `true`.

```js
assert.isTrue(isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(arr) {

  return arr;
}
```

# --solutions--

```js
function isFizzBuzz(arr) {
  let start = null;
  for (let i = 0; i < arr.length; i++) {
    if (typeof arr[i] === "number") {
      start = arr[i] - i;
      break;
    }
  }

  if (start === null) return false;

  for (let i = 0; i < arr.length; i++) {
    const n = start + i;
    const expected = n % 15 === 0 ? "FizzBuzz" : n % 3 === 0 ? "Fizz" : n % 5 === 0 ? "Buzz" : n;
    if (arr[i] !== expected) return false;
  }

  return true;
}
```
