---
id: 69b559d2903b9e4afe9075f8
title: "تحدٍّ 238: الهروب بدوران الأرقام"
challengeType: 28
dashedName: challenge-238
---

# --description--

أعطِ عددًا صحيحًا موجبًا، حدّد إذا كان هو أو أي من دورانه يقبل القسمة بالتساوي على عدد أرقامه.

الدوران يعني نقل الرقم الأول إلى النهاية. على سبيل المثال، بعد دوران واحد، يصبح 123 هو 231.

- افحص الدوران `0` (العدد المعطى) أولًا.
- الأعداد المعطاة لا تحتوي على أصفار.
- أرجع رقم أول دوران إذا وُجد، أو `"none"` إذا لم يوجد.

# --hints--

`getRotation(123)` يجب أن تُرجع `0`.

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` يجب أن تُرجع `3`.

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` يجب أن تُرجع `"none"`.

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` يجب أن تُرجع `6`.

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
