---
id: 69b559d2903b9e4afe9075f9
title: "تحدٍّ 239: ما هو اليوم؟"
challengeType: 28
dashedName: challenge-239
---

# --description--

بالنظر إلى طابع زمني من نوع Unix بوحدة المللي ثانية، أعد اليوم من الأسبوع.

الأيام الصحيحة للإرجاع هي:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

تأكد من تجاهل المناطق الزمنية.

# --hints--

`getDayOfWeek(1775492249000)` يجب أن تُعيد `"Monday"`.

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` يجب أن تُعيد `"Saturday"`.

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` يجب أن تُعيد `"Tuesday"`.

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` يجب أن تُعيد `"Sunday"`.

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` يجب أن تُعيد `"Thursday"`.

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
