---
id: 69b58ce40693f140c84c855b
title: "تحدٍّ 242: الرقم التالي في لعبة البنغو"
challengeType: 28
dashedName: challenge-242
---

# --description--

عند إعطائك رقم بنغو، أعد الرقم التالي في التسلسل.

رقم البنغو هو حرف واحد يتبعه رقم ضمن نطاقه حسب هذا المخطط:

| الحرف | نطاق الرقم |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

على سبيل المثال، عند إعطائك `"B10"`، أعد `"B11"`، الرقم التالي في البنغو. إذا كان الرقم المعطى هو آخر رقم بنغو، أعد `"B1"`.

# --hints--

`getNextBingoNumber("B10")` يجب أن يعيد `"B11"`.

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` يجب أن يعيد `"N34"`.

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` يجب أن يعيد `"N31"`.

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` يجب أن يعيد `"O61"`.

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` يجب أن يعيد `"B1"`.

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
