---
id: 69b5b2be76ec8135a7fbe973
title: "تحدٍّ 244: هجوم الرخ والأسقف"
challengeType: 28
dashedName: challenge-244
---

# --description--

معطاة سلسلة نصية لموقع رخ على رقعة الشطرنج، وأخرى لموقع أسقف، حدّد إذا كان بإمكان قطعة واحدة مهاجمة الأخرى.

رقعة الشطرنج القياسية 8×8، مع أعمدة معنونة من `A` إلى `H` (من اليسار إلى اليمين) وصفوف معنونة من `1` إلى `8` (من الأسفل إلى الأعلى). تبدو هكذا:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- يمكن للرُّخ التحرك بعدد مربعات غير محدود في الاتجاه الأفقي أو العمودي.
- يمكن للأسقف التحرك بعدد مربعات غير محدود في أي اتجاه قطري.
- يمكن لقطعة مهاجمة أخرى إذا استطاعت التحرك إلى موقع تلك القطعة.

أرجع:

- `"rook"` إذا كان الرخ يستطيع مهاجمة الأسقف.
- `"bishop"` إذا كان الأسقف يستطيع مهاجمة الرخ.
- `"neither"` إذا لم تستطع أي من القطعتين مهاجمة الأخرى.

# --hints--

`rookBishopAttack("A1", "A5")` يجب أن تُرجع `"rook"`.

```js
assert.equal(rookBishopAttack("A1", "A5"), "rook");
```

`rookBishopAttack("C3", "F6")` يجب أن تُرجع `"bishop"`.

```js
assert.equal(rookBishopAttack("C3", "F6"), "bishop");
```

`rookBishopAttack("D4", "D7")` يجب أن تُرجع `"rook"`.

```js
assert.equal(rookBishopAttack("D4", "D7"), "rook");
```

`rookBishopAttack("B7", "H1")` يجب أن تُرجع `"bishop"`.

```js
assert.equal(rookBishopAttack("B7", "H1"), "bishop");
```

`rookBishopAttack("B3", "C5")` يجب أن تُرجع `"neither"`.

```js
assert.equal(rookBishopAttack("B3", "C5"), "neither");
```

`rookBishopAttack("G3", "E8")` يجب أن تُرجع `"neither"`.

```js
assert.equal(rookBishopAttack("G3", "E8"), "neither");
```

# --seed--

## --seed-contents--

```js
function rookBishopAttack(rook, bishop) {

  return rook;
}
```

# --solutions--

```js
function rookBishopAttack(rook, bishop) {
  const colDiff = Math.abs(rook.charCodeAt(0) - bishop.charCodeAt(0));
  const rowDiff = Math.abs(rook[1] - bishop[1]);

  if (rook[0] === bishop[0] || rook[1] === bishop[1]) return "rook";
  if (colDiff === rowDiff) return "bishop";
  return "neither";
}
```
