---
id: 69bc6cb30c1d112a2e110a08
title: "تحدٍّ 251: العثور على مجموع مصفوفة"
challengeType: 28
dashedName: challenge-251
---

# --description--

معطاة مصفوفة من الأرقام ورقم هدف، أعد أول مجموعة جزئية من رقمين أو أكثر مجموعها يساوي الهدف.

- المجموعة الجزئية "الأولى" هي التي تحتوي عناصرها على أدنى فهارس ممكنة، مع إعطاء الأولوية لأدنى فهرس أولاً.
- يمكن استخدام كل رقم في المصفوفة مرة واحدة فقط.
- إذا لم توجد مجموعة جزئية صالحة، أعد `"Sum not found"`.

أعد الأرقام المطابقة كمصفوفة بالترتيب الذي تظهر به في المصفوفة الأصلية.

# --hints--

`findSum([1, 3, 5, 7], 6)` يجب أن تُعيد `[1, 5]`.

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)` يجب أن تُعيد `[1, 4]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)` يجب أن تُعيد `[1, 2, 3]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)` يجب أن تُعيد `[-1, -2, 4]`.

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)` يجب أن تُعيد `[3, 1, 4, 2]`.

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` يجب أن تُعيد `[1, 2, 3, 5, 9]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)` يجب أن تُعيد `"Sum not found"`.

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```
