---
id: 69bc6cb30c1d112a2e110a06
title: "تحدٍّ 249: حساب السلاسل النصية"
challengeType: 28
dashedName: challenge-249
---

# --description--

معطى سلسلة نصية تحتوي أرقامًا وحروفًا أخرى، نفذ عمليات حسابية على الأرقام بناءً على عدد الأحرف غير الرقمية بين الأرقام.

- إذا كان عدد الأحرف التي تفصل بين رقمين زوجيًا، استخدم الجمع.
- إذا كان فرديًا، استخدم الطرح.
- الأرقام المتتالية تشكل رقمًا واحدًا.
- تُطبق العمليات من اليسار إلى اليمين.
- تجاهل الأحرف غير الرقمية في البداية والنهاية.

على سبيل المثال، معطى `"3ab10c8"`، أرجع `5`. اجمع 3 و10 لتحصل على 13 لأن عدد الأحرف بينهما زوجي. ثم اطرح 8 من 13 لأن عدد الأحرف بين النتيجة و8 فردي.

# --hints--

`doMath("3ab10c8")` يجب أن تُرجع `5`.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` يجب أن تُرجع `2`.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` يجب أن تُرجع `12`.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` يجب أن تُرجع `15`.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` يجب أن تُرجع `67`.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
