---
id: 69bc6cb30c1d112a2e110a09
title: "تحدٍّ 252: متسلّق الدرج الفريد"
challengeType: 28
dashedName: challenge-252
---

# --description--

عند إعطائك عددًا من الدرجات، أرجع عدد الطرق المختلفة التي يمكن لشخص أن يصعدها بخطوة واحدة أو خطوتين في كل مرة.

# --hints--

`getUniqueClimbs(4)` يجب أن تُرجع `5`.

```js
assert.equal(getUniqueClimbs(4), 5);
```

`getUniqueClimbs(5)` يجب أن تُرجع `8`.

```js
assert.equal(getUniqueClimbs(5), 8);
```

`getUniqueClimbs(10)` يجب أن تُرجع `89`.

```js
assert.equal(getUniqueClimbs(10), 89);
```

`getUniqueClimbs(18)` يجب أن تُرجع `4181`.

```js
assert.equal(getUniqueClimbs(18), 4181);
```

`getUniqueClimbs(29)` يجب أن تُرجع `832040`.

```js
assert.equal(getUniqueClimbs(29), 832040);
```

`getUniqueClimbs(50)` يجب أن تُرجع `20365011074`.

```js
assert.equal(getUniqueClimbs(50), 20365011074);
```

# --seed--

## --seed-contents--

```js
function getUniqueClimbs(steps) {

  return steps;
}
```

# --solutions--

```js
function getUniqueClimbs(steps) {
  if (steps <= 0) return 0;
  if (steps === 1) return 1;
  if (steps === 2) return 2;

  let prev2 = 1, prev1 = 2;
  for (let i = 3; i <= steps; i++) {
    [prev2, prev1] = [prev1, prev2 + prev1];
  }
  return prev1;
}
```
