---
id: 69bc6cb30c1d112a2e110a07
title: "تحدٍّ 250: المفتاح المخفي"
challengeType: 28
dashedName: challenge-250
---

# --description--

مرحبًا بك في تحدٍّ اليومي رقم 250!

معطاة سلسلة نصية مشفرة، فك شفرتها باستخدام مفتاح تشفير وأعدها.

للعثور على المفتاح:

- اطلع على جميع تحديات اليوميات حتى اليوم التي رقم تحديها من مضاعفات 25 (بما في ذلك هذا التحدي).
- خذ الحرف الأول من كل عنوان من تلك العناوين وادمجها في سلسلة نصية. إذا بدأ العنوان بحرف غير حرف أبجدي، ابحث عن أول حرف أبجدي فيه.

لفك شفرة الرسالة، استعرض كل حرف في الرسالة المشفرة و:

1. انظر إلى الحرف المقابل في المفتاح (كرر المفتاح إذا كانت الرسالة أطول من المفتاح).
  2. حوّل حرف المفتاح إلى رقمه المقابل: `"A"` = 1، `"B"` = 2، ...، `"Z"` = 26.
  3. حرك الحرف المشفر إلى الوراء في الأبجدية بعدد تلك القيمة.
  4. إذا تجاوزت الحركة الحرف `"A"`، عد إلى `"Z"`.

على سبيل المثال، إذا بدأت الرسالة المشفرة بـ `"Y"` وكان أول حرف في المفتاح هو `"V"` (22)، حرك `"Y"` إلى الوراء 22 خانة لتحصل على `"C"`. كرر هذه العملية لكل حرف لفك شفرة الرسالة كاملة.

- تُحرك الحروف فقط، وتُعاد الفراغات كما هي.
- كل الحروف المعطاة والمعادة تكون بحروف كبيرة.

# --hints--

`decode("YAVJYNXE")` يجب أن تُعيد `"CONGRATS"`.

```js
assert.equal(decode("YAVJYNXE"), "CONGRATS");
```

`decode("YALLUT PQUMJP")` يجب أن تُعيد `"CODING LEGEND"`.

```js
assert.equal(decode("YALLUT PQUMJP"), "CODING LEGEND");
```

`decode("UAC DYR EISAKYM")` يجب أن تُعيد `"YOU ARE AWESOME"`.

```js
assert.equal(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME");
```

`decode("GQMS NBMZU")` يجب أن تُعيد `"KEEP GOING"`.

```js
assert.equal(decode("GQMS NBMZU"), "KEEP GOING");
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` يجب أن تُعيد `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
assert.equal(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP");
```

# --seed--

## --seed-contents--

```js
function decode(message) {

  return message;
}
```

# --solutions--

```js
function decode(message) {
  const key = "VLHCGMDLNH";
  let keyIndex = 0;
  return message.split('').map(char => {
    if (char === ' ') return char;
    const shift = key[keyIndex % key.length].charCodeAt(0) - 64;
    keyIndex++;
    return String.fromCharCode(((char.charCodeAt(0) - 65 - shift + 26) % 26) + 65);
  }).join('');
}
```
