---
id: 69bc6cb30c1d112a2e110a04
title: "تحدٍّ 247: الحرف الأخير"
challengeType: 28
dashedName: challenge-247
---

# --description--

بالنظر إلى سلسلة نصية، أعد الحرف من السلسلة الذي يظهر آخرًا في الأبجدية.

- إذا تعادلت حرفان أو أكثر في الترتيب الأخير في الأبجدية، أعد الأول منها.
- تجاهل كل الأحرف غير الحروف.

# --hints--

`getLastLetter("world")` يجب أن تُعيد `"w"`.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` يجب أن تُعيد `"W"`.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` يجب أن تُعيد `"z"`.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` يجب أن تُعيد `"L"`.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` يجب أن تُعيد `"T"`.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
