---
id: 69bc6cb30c1d112a2e110a03
title: "تحدٍّ 246: الأحرف الأولى من الاسم"
challengeType: 28
dashedName: challenge-246
---

# --description--

بالنظر إلى اسم كامل كسلسلة نصية، أعد الأحرف الأولى منه.

- الأسماء التي تُؤخذ منها الأحرف الأولى مفصولة بمسافة.
- يجب أن تكون الأحرف الأولى بحروف كبيرة.
- يجب أن تُفصل الأحرف الأولى بنقاط.

على سبيل المثال، `"Tommy Millwood"` تُعيد `"T.M."`.

# --hints--

`getInitials("Tommy Millwood")` يجب أن تُعيد `"T.M."`.

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` يجب أن تُعيد `"S.P."`.

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` يجب أن تُعيد `"F.C.C."`.

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` يجب أن تُعيد `"D."`.

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` يجب أن تُعيد `"D.V.C.H.N."`.

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
