---
id: 69c5f3d787b1725d5f00c8ba
title: "تحدٍّ 256: أقرب اتجاه زمني"
challengeType: 28
dashedName: challenge-256
---

# --description--

معطى وقتان، حدّد ما إذا كان بإمكانك الانتقال من الأول إلى الثاني أسرع بالحركة للأمام أو للخلف.

- الأوقات معطاة بصيغة 24 ساعة (`"HH:MM"`)
- الساعة تدور بشكل دائري (الساعة 23:59 تليها 00:00 عند الحركة للأمام، و00:00 تسبقها 23:59 عند الحركة للخلف)

أرجع:

- `"forward"` إذا كانت الحركة للأمام أقصر
- `"backward"` إذا كانت الحركة للخلف أقصر
- `"equal"` إذا استغرق الاتجاهان نفس الوقت

# --hints--

`getDirection("10:00", "12:00")` يجب أن تُرجع `"forward"`.

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` يجب أن تُرجع `"backward"`.

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` يجب أن تُرجع `"equal"`.

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` يجب أن تُرجع `"forward"`.

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` يجب أن تُرجع `"backward"`.

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` يجب أن تُرجع `"equal"`.

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
