---
id: 69c5f3d787b1725d5f00c8bd
title: "تحدٍّ 259: انفجار FizzBuzz"
challengeType: 28
dashedName: challenge-259
---

# --description--

بالنظر إلى عدد صحيح، أعد عدد الخطوات اللازمة لتحويل الكلمة `"fizzbuzz"` إلى سلسلة نصية تحتوي على الأقل على العدد المعطى من الحروف `"z"` باستخدام القواعد التالية:

- ابدأ بالسلسلة النصية `"fizzbuzz"`.
- في كل خطوة، طبق قواعد FizzBuzz القياسية باستخدام موقع الحرف في السلسلة النصية (الحرف الأول `"f"` في الموقع 1).
  - إذا كان موقع الحرف قابلًا للقسمة على 3، استبدل الحرف بـ `"fizz"`
  - إذا كان قابلًا للقسمة على 5، استبدل الحرف بـ `"buzz"`
  - إذا كان قابلًا للقسمة على 3 و 5، استبدل الحرف بـ `"fizzbuzz"`

لذا بعد خطوة واحدة، تتحول `"fizzbuzz"` إلى `"fifizzzbuzzfizzzz"`، والتي تحتوي على 9 حروف `"z"`.

# --hints--

`explodeFizzbuzz(9)` يجب أن تُعيد `1`.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` يجب أن تُعيد `2`.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` يجب أن تُعيد `3`.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` يجب أن تُعيد `4`.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` يجب أن تُعيد `5`.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` يجب أن تُعيد `6`.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` يجب أن تُعيد `11`.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` يجب أن تُعيد `14`.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
