---
id: 69c5f3d787b1725d5f00c8b9
title: "تحدٍّ 255: فريق تنظيف يوم الأرض"
challengeType: 28
dashedName: challenge-255
---

# --description--

اليوم هو يوم الأرض. معطى مصفوفة من العناصر التي نظفتها، أرجع مجموع نقاط التنظيف بناءً على القواعد التالية.

العناصر المعطاة ستكون واحدة من:

| العنصر | القيمة الأساسية |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

العنصر النادر يُمثّل كـ `["rare", value]`. على سبيل المثال، `["rare", 80]`. العناصر النادرة لا تحصل على مكافأة التتابع.

- مكافأة التتابع: إذا ظهر نفس العنصر بشكل متتالي، يحصل على نقاط مكافأة متزايدة.
  - الظهور المتتالي الأول: القيمة الأساسية
  - الثاني: القيمة الأساسية + 1
  - الثالث: القيمة الأساسية + 2
  - وهكذا.
- مضاعف العنصر الخامس: كل عنصر خامس يتم جمعه يحصل على مضاعف.
  - العنصر الخامس: *2
  - العنصر العاشر: *3
  - وهكذا.

- طبق المضاعف بعد حساب أي مكافآت.

# --hints--

`getCleanupScore(["bottle", "straw", "shoe", "battery"])` يجب أن تُرجع `44`.

```js
assert.equal(getCleanupScore(["bottle", "straw", "shoe", "battery"]), 44);
```

`getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"])` يجب أن تُرجع `58`.

```js
assert.equal(getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"]), 58);
```

`getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` يجب أن تُرجع `79`.

```js
assert.equal(getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79);
```

`getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` يجب أن تُرجع `358`.

```js
assert.equal(getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358);
```

`getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` يجب أن تُرجع `383`.

```js
assert.equal(getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383);
```

# --seed--

## --seed-contents--

```js
function getCleanupScore(items) {

  return items;
}
```

# --solutions--

```js
function getCleanupScore(items) {
  const values = {
    bottle: 10,
    can: 6,
    bag: 8,
    tire: 35,
    straw: 4,
    cardboard: 3,
    newspaper: 3,
    shoe: 12,
    electronics: 25,
    battery: 18,
    mattress: 38
  };

  let total = 0;
  let prevItem = null;
  let streak = 0;

  for (let i = 0; i < items.length; i++) {
    let item = items[i];
    let base;
    let isRare = Array.isArray(item) && item[0] === "rare";

    if (isRare) {
      base = item[1];
      prevItem = null;
      streak = 0;
    } else {
      base = values[item];
      streak = item === prevItem ? streak + 1 : 0;
      prevItem = item;
    }

    let points = base;
    if (!isRare) points += streak;

    if ((i + 1) % 5 === 0) {
      const multiplier = (i + 1) / 5 + 1;
      points *= multiplier;
    }

    total += points;
  }

  return total;
}
```
