---
id: 69c5f3d787b1725d5f00c8b8
title: "تحدٍّ 254: كلمات ذات عدد فردي"
challengeType: 28
dashedName: challenge-254
---

# --description--

معطاة سلسلة نصية من الكلمات، أعد فقط الكلمات التي تحتوي على عدد فردي من الحروف.

- الكلمات في السلسلة النصية المعطاة ستكون مفصولة بمسافة واحدة.
- أعد الكلمات مفصولة بمسافة واحدة.

# --hints--

`getOddWords("This is a super good test")` يجب أن تُعيد `"a super"`.

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` يجب أن تُعيد `"one two three"`.

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` يجب أن تُعيد `"split rainbow sprinkles top"`.

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` يجب أن تُعيد `"The quick brown fox the river"`.

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
