---
id: 69c6ff713a52713463aa7928
title: "تحدٍّ 262: محلل استعلام URL"
challengeType: 28
dashedName: challenge-262
---

# --description--

عند وجود URL يحتوي على سلسلة استعلام، حلل سلسلة الاستعلام إلى كائن (أو قاموس) من أزواج المفتاح-القيمة.

- تبدأ سلسلة الاستعلام بعد `"?"`،
- تُفصل كل معلمة بـ `"&"`،
- يُفصل كل زوج مفتاح/قيمة بـ `"="`

على سبيل المثال، عند إعطاء `"https://example.com/search?name=Alice&age=30"`، أرجع:

```json
{
  "name": "Alice",
  "age": "30"
}
```

يجب إرجاع كل القيم كسلاسل نصية.

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` يجب أن تُرجع `{"name": "Alice", "age": "30"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` يجب أن تُرجع `{"skill": "programming", "language": "python"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` يجب أن تُرجع `{"category": "books", "sort": "asc", "page": "2"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` يجب أن تُرجع `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
