---
id: 69c6ff713a52713463aa7926
title: "تحدٍّ 260: نقاط الكلمة"
challengeType: 28
dashedName: challenge-260
---

# --description--

أعطِ كلمة، أرجع نقاطها باستخدام جدول قياسي لقيمة الحروف:

| الحرف | القيمة |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- الحروف الكبيرة والأحرف الصغيرة لها نفس القيمة.

# --hints--

`getWordScore("hi")` يجب أن تُرجع `17`.

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` يجب أن تُرجع `52`.

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` يجب أن تُرجع `169`.

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` يجب أن تُرجع `94`.

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
