---
id: 69c6ff713a52713463aa7927
title: "تحدٍّ 261: كلمات الأعداد"
challengeType: 28
dashedName: challenge-261
---

# --description--

أعطِ عددًا صحيحًا من 0 إلى 99، وأرجع تمثيله بالكلمات الإنجليزية.

- `0` تُرجع `"zero"`.
- الأعداد من 1 إلى 19 لها أسماء فريدة (`"one"`، `"two"`، ...، `"ten"`، `"eleven"`، ...، `"eighteen"`، `"nineteen"`).
- مضاعفات 10 من 20 إلى 90 لها أسماء خاصة بها (`"twenty"`، `"thirty"`، ...، `"eighty"`، `"ninety"`).
- الأعداد من 21 إلى 99 التي ليست مضاعفات 10 تُكتب ككلمتين مرتبطتين بشرطة. على سبيل المثال `"forty-two"` و `"fifty-three"`.

# --hints--

`getNumberWords(0)` يجب أن تُرجع `"zero"`.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` يجب أن تُرجع `"ten"`.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` يجب أن تُرجع `"nineteen"`.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` يجب أن تُرجع `"thirty"`.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` يجب أن تُرجع `"fifty-three"`.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` يجب أن تُرجع `"seven"`.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` يجب أن تُرجع `"twelve"`.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` يجب أن تُرجع `"sixty"`.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` يجب أن تُرجع `"sixty-seven"`.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` يجب أن تُرجع `"ninety-eight"`.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
