---
id: 69c6ff713a52713463aa7929
title: "تحدٍّ 263: الكلمات المتقاطعة الثنائية"
challengeType: 28
dashedName: challenge-263
---

# --description--

بالنظر إلى حرف، حدّد إذا كان تمثيله الثنائي المكوّن من 8 بت يمكن إيجاده في الشبكة التالية، أفقيًا أو عموديًا في أي اتجاه:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

على سبيل المثال، `"A"` له التمثيل الثنائي `01000001`، الذي يظهر في الصف الأول من اليسار إلى اليمين.

# --hints--

يجب أن تُرجع الدالة `isInCrossword("I")` القيمة `true`.

```js
assert.isTrue(isInCrossword("I"));
```

يجب أن تُرجع الدالة `isInCrossword("D")` القيمة `true`.

```js
assert.isTrue(isInCrossword("D"));
```

يجب أن تُرجع الدالة `isInCrossword("0")` القيمة `true`.

```js
assert.isTrue(isInCrossword("0"));
```

يجب أن تُرجع الدالة `isInCrossword("u")` القيمة `true`.

```js
assert.isTrue(isInCrossword("u"));
```

يجب أن تُرجع الدالة `isInCrossword("Y")` القيمة `false`.

```js
assert.isFalse(isInCrossword("Y"));
```

يجب أن تُرجع الدالة `isInCrossword("p")` القيمة `false`.

```js
assert.isFalse(isInCrossword("p"));
```

يجب أن تُرجع الدالة `isInCrossword("1")` القيمة `false`.

```js
assert.isFalse(isInCrossword("1"));
```

يجب أن تُرجع الدالة `isInCrossword("Q")` القيمة `false`.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
