---
id: 69cfca90e8a0a6d4d6871c54
title: "تحدٍّ 270: أطول سلسلة نصية مشتركة"
challengeType: 28
dashedName: challenge-270
---

# --description--

أعطِ سلسلة نصية، أعد أطول سلسلة نصية فرعية تظهر أكثر من مرة.

- يمكن أن تتداخل السلاسل النصية الفرعية.

# --hints--

`getLongestSubstring("abracadabra")` يجب أن تُعيد `"abra"`.

```js
assert.equal(getLongestSubstring("abracadabra"), "abra");
```

`getLongestSubstring("hello world hello")` يجب أن تُعيد `"hello"`.

```js
assert.equal(getLongestSubstring("hello world hello"), "hello");
```

`getLongestSubstring("mississippi")` يجب أن تُعيد `"issi"`.

```js
assert.equal(getLongestSubstring("mississippi"), "issi");
```

`getLongestSubstring("ha ha ha ha ha ha ha")` يجب أن تُعيد `"ha ha ha ha ha ha"`.

```js
assert.equal(getLongestSubstring("ha ha ha ha ha ha ha"), "ha ha ha ha ha ha");
```

`getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over")` يجب أن تُعيد `"the quick brown fox jumped over"`.

```js
assert.equal(getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over"), "the quick brown fox jumped over");
```

# --seed--

## --seed-contents--

```js
function getLongestSubstring(str) {

  return str;
}
```

# --solutions--

```js
function getLongestSubstring(str) {
  let longest = '';

  for (let len = str.length - 1; len >= 1; len--) {
    for (let i = 0; i <= str.length - len; i++) {
      const sub = str.slice(i, i + len);
      if (str.indexOf(sub) !== str.lastIndexOf(sub)) {
        return sub;
      }
    }
  }

  return longest;
}
```
