---
id: 69cfca90e8a0a6d4d6871c52
title: "تحدٍّ 268: الرقم النرجسي"
challengeType: 28
dashedName: challenge-268
---

# --description--

أعطِ عددًا صحيحًا موجبًا، حدّد ما إذا كان رقمًا نرجسيًا.

- الرقم نرجسي إذا كان مجموع كل رقم من أرقامه مرفوعًا لقوة عدد الأرقام يساوي الرقم نفسه.

على سبيل المثال، الرقم 153 يحتوي على 3 أرقام، و1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153، إذًا هو نرجسي.

# --hints--

`isNarcissistic(153)` يجب أن تُرجع `true`.

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` يجب أن تُرجع `false`.

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` يجب أن تُرجع `true`.

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` يجب أن تُرجع `false`.

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` يجب أن تُرجع `true`.

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` يجب أن تُرجع `false`.

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` يجب أن تُرجع `true`.

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` يجب أن تُرجع `false`.

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
