---
id: 69cfca90e8a0a6d4d6871c50
title: "تحدٍّ 266: يوم سعيد"
challengeType: 28
dashedName: challenge-266
---

# --description--

بالنظر إلى سلسلة نصية للوقت بصيغة `"HH:MM"` (ساعة 24)، أعد:

- `"Good morning"` للأوقات من `05:00` إلى `11:59`
- `"Good afternoon"` للأوقات من `12:00` إلى `17:59`
- `"Good evening"` للأوقات من `18:00` إلى `21:59`
- `"Good night"` للأوقات من `22:00` إلى `04:59`

# --hints--

`getGreeting("06:30")` يجب أن تعيد `"Good morning"`.

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")` يجب أن تعيد `"Good afternoon"`.

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")` يجب أن تعيد `"Good evening"`.

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")` يجب أن تعيد `"Good night"`.

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")` يجب أن تعيد `"Good morning"`.

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
