---
id: 69cfca90e8a0a6d4d6871c4f
title: "تحدٍّ 265: أعمق أقواس"
challengeType: 28
dashedName: challenge-265
---

# --description--

معطاة سلسلة نصية تحتوي على أقواس متوازنة، أعد محتوى الأقواس المسكنة الأعمق.

- يمكن أن تكون الأقواس من أي نوع من الأنواع الثلاثة: `()`, `[]`, و `{}`.
- سيكون المُدخَل دائمًا يحتوي على مجموعة واحدة أعمق.

على سبيل المثال، معطى `"(hello (world))"`، أعد `"world"`.

# --hints--

`getDeepestBrackets("(hello (world))")` يجب أن تُعيد `"world"`.

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` يجب أن تُعيد `"inner"`.

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` يجب أن تُعيد `"e"`.

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` يجب أن تُعيد `"fox"`.

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` يجب أن تُعيد `"C"`.

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
