---
id: 69cfca90e8a0a6d4d6871c4e
title: "تحدٍّ 264: مجموعات الكلمات المتشابهة"
challengeType: 28
dashedName: challenge-264
---

# --description--

بالنظر إلى مصفوفة من الكلمات، أعد مصفوفة ثنائية الأبعاد تحتوي الكلمات مجمعة في مجموعات كلمات متشابهة.

- الكلمات متشابهة إذا احتوت على نفس الحروف بأي ترتيب.
- تنتمي كل كلمة إلى مجموعة واحدة فقط.
- ترتيب الإرجاع غير مهم.

على سبيل المثال، إذا كانت المصفوفة `["listen", "silent", "hello", "enlist", "world"]`، فأعد `[["listen", "silent", "enlist"], ["hello"], ["world"]]`.

# --hints--

`groupAnagrams(["listen", "silent", "hello", "enlist", "world"])` يجب أن تُعيد `[["listen", "silent", "enlist"], ["hello"], ["world"]]`.

```js
const groups = groupAnagrams(["listen", "silent", "hello", "enlist", "world"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["enlist", "listen", "silent"], ["hello"], ["world"]]);
```

`groupAnagrams(["eat", "tea", "tan", "ate", "nat", "bat"])` يجب أن تُعيد `[["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]`.

```js
const groups = groupAnagrams(["eat", "tea", "tan", "ate", "nat", "bat"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]);
```

`groupAnagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])` يجب أن تُعيد `[["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]`.

```js
const groups = groupAnagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]);

```

`groupAnagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])` يجب أن تُعيد `[["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]`.

```js
const groups = groupAnagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]);
```

# --seed--

## --seed-contents--

```js
function groupAnagrams(words) {

  return words;
}
```

# --solutions--

```js
function groupAnagrams(words) {
  const map = {};
  for (const word of words) {
    const key = word.split('').sort().join('');
    if (!map[key]) map[key] = [];
    map[key].push(word);
  }
  return Object.values(map);
}
```
