---
id: 69cfca90e8a0a6d4d6871c51
title: "تحدٍّ 267: محول Parsec"
challengeType: 28
dashedName: challenge-267
---

# --description--

في مجرة بعيدة، تُستخدم الـ parsecs لقياس كل من الزمن والمسافة. معطى عدد صحيح من الـ parsecs، أعد ما يعادله من الزمن أو المسافة.

- إذا كان العدد الصحيح المعطى فرديًا، فهو يمثل الزمن. وإذا كان زوجيًا، فهو يمثل المسافة.

استخدم معدلات التحويل هذه:

| Parsecs | الزمن/المسافة |
| - | - |
| 1 | ساعتان |
| 2 | 6 سنوات ضوئية |

أعد القيمة المحولة كعدد صحيح.

# --hints--

يجب أن تُعيد `convertParsecs(1)` القيمة `2`.

```js
assert.equal(convertParsecs(1), 2);
```

يجب أن تُعيد `convertParsecs(2)` القيمة `6`.

```js
assert.equal(convertParsecs(2), 6);
```

يجب أن تُعيد `convertParsecs(31)` القيمة `62`.

```js
assert.equal(convertParsecs(31), 62);
```

يجب أن تُعيد `convertParsecs(88)` القيمة `264`.

```js
assert.equal(convertParsecs(88), 264);
```

يجب أن تُعيد `convertParsecs(17)` القيمة `34`.

```js
assert.equal(convertParsecs(17), 34);
```

يجب أن تُعيد `convertParsecs(14)` القيمة `42`.

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
