---
id: 69d03e549613fbdbe21d11e0
title: "تحدٍّ 271: تذكير الدواء"
challengeType: 28
dashedName: challenge-271
---

# --description--

معطاة مصفوفة من الأدوية وسلسلة نصية تمثل الوقت الحالي، أرجع الدواء التالي الذي تحتاج إلى تناوله وكم من الوقت تبقى حتى تحتاج لتناوله.

- كل دواء يكون بالشكل `[name, lastTaken]`، حيث `name` هو اسم الدواء و`lastTaken` هو الوقت الذي تم تناوله فيه آخر مرة.
- كل الأوقات المعطاة ستكون بصيغة `"HH:MM"` (24 ساعة).

استخدم جدول مواعيد الأدوية التالي:

| الاسم | الجدول |
| - | - |
| Deployxitrin | 08:00، 16:00 |
| Debuggamanizole | 07:00، 13:00، 21:00 |
| Mergeflictamine | كل 4 ساعات |

أرجع سلسلة نصية بالشكل `"{name} in Hh Mm"`. على سبيل المثال، `"Debuggamanizole in 2h 0m"` أو `"Deployxitrin in 1h 5m"`.

# --hints--

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00")` يجب أن تُرجع `"Debuggamanizole in 2h 0m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00"), "Debuggamanizole in 2h 0m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55")` يجب أن تُرجع `"Deployxitrin in 1h 5m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55"), "Deployxitrin in 1h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15")` يجب أن تُرجع `"Mergeflictamine in 0h 45m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15"), "Mergeflictamine in 0h 45m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59")` يجب أن تُرجع `"Debuggamanizole in 0h 1m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59"), "Debuggamanizole in 0h 1m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55")` يجب أن تُرجع `"Debuggamanizole in 0h 5m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55"), "Debuggamanizole in 0h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00")` يجب أن تُرجع `"Mergeflictamine in 3h 30m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00"), "Mergeflictamine in 3h 30m");
```

# --seed--

## --seed-contents--

```js
function medicationReminder(medications, currentTime) {

  return medications;
}
```

# --solutions--

```js
function medicationReminder(medications, currentTime) {
  const schedule = {
    Deployxitrin: { type: 'fixed', times: ['08:00', '16:00'] },
    Debuggamanizole: { type: 'fixed', times: ['07:00', '13:00', '21:00'] },
    Mergeflictamine: { type: 'interval', hours: 4 }
  };

  const toMinutes = t => {
    const [h, m] = t.split(':').map(Number);
    return h * 60 + m;
  };

  const current = toMinutes(currentTime);
  let earliest = Infinity;
  let nextMed = '';

  for (const [name, lastTaken] of medications) {
    const med = schedule[name];
    let nextDue;

    if (med.type === 'fixed') {
      const futureTimes = med.times
        .map(toMinutes)
        .filter(t => t > current);
      if (futureTimes.length === 0) continue;
      nextDue = Math.min(...futureTimes);
    } else {
      nextDue = toMinutes(lastTaken) + med.hours * 60;
    }

    const diff = nextDue - current;
    if (diff > 0 && diff < earliest) {
      earliest = diff;
      nextMed = name;
    }
  }

  const h = Math.floor(earliest / 60);
  const m = earliest % 60;
  return `${nextMed} in ${h}h ${m}m`;
}
```
