---
id: 69e2383af7832c8032603b90
title: "تحدٍّ 274: أقدم شخص"
challengeType: 28
dashedName: challenge-274
---

# --description--

بالنظر إلى مصفوفة من الكائنات، كل منها يحتوي على خاصيتي `"name"` و`"age"`، أرجع مصفوفة تحتوي على أسماء أقدم الأشخاص.

إذا شارك عدة أشخاص في أكبر عمر، أرجع جميع أسمائهم بالترتيب الذي يظهرون به في المُدخَل.

# --hints--

`getOldest([{ name: "Brenda", age: 40 }])` يجب أن تُرجع `["Brenda"]`.

```js
assert.deepEqual(getOldest([{ name: "Brenda", age: 40 }]), ["Brenda"]);
```

`getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }])` يجب أن تُرجع `["Alice"]`.

```js
assert.deepEqual(getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }]), ["Alice"]);
```

`getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }])` يجب أن تُرجع `["Bill", "Carol"]`.

```js
assert.deepEqual(getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }]), ["Bill", "Carol"]);
```

`getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }])` يجب أن تُرجع `["George", "Holly", "Zach"]`.

```js
assert.deepEqual(getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }]), ["George", "Holly", "Zach"]);
```

# --seed--

## --seed-contents--

```js
function getOldest(people) {

  return people;
}
```

# --solutions--

```js
function getOldest(people) {
  const maxAge = Math.max(...people.map(p => p.age));
  return people.filter(p => p.age === maxAge).map(p => p.name);
}
```
