---
id: 69e2383af7832c8032603b93
title: "تحدٍّ 277: الصورة المعكوسة"
challengeType: 28
dashedName: challenge-277
---

# --description--

معطيتان سلسلتا نص، حدّد إذا كانت السلسلة الثانية صورة معكوسة للسلسلة الأولى.

تتكوّن الصورة المعكوسة بعكس السلسلة واستبدال كل حرف بما يعادله في المرآة.

- الأحرف المتماثلة تبدو كما هي في المرآة:

`W`، `T`، `Y`، `U`، `I`، `O`، `H`، `A`، `X`، `V`، `M`، `w`، `o`، `x`، `v`، `0`، `8`، `=`، `+`، `:`، `|`، `-`، `_`، `*`، `^`، `!`، `.`، والمسافة (` `).

- أزواج الأحرف المعكوسة تتبادل أماكنها في المرآة:

| الحرف | يتبادل مع |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

إذا احتوت أي من السلسلتين على حرف غير موجود في القوائم أعلاه، فلا يمكن تكوين صورة معكوسة من تلك الأحرف.

على سبيل المثال، الصورة المعكوسة لـ `"[HOW]"` هي `"[WOH]"`.

# --hints--

`isMirrorImage("[HOW]", "[WOH]")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMirrorImage("[HOW]", "[WOH]"));
```

`isMirrorImage("MOM", "MOM")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMirrorImage("MOM", "MOM"));
```

`isMirrorImage("vow", "wov")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMirrorImage("vow", "wov"));
```

`isMirrorImage("TIM", "TIM")` يجب أن تُرجع `false`.

```js
assert.isFalse(isMirrorImage("TIM", "TIM"));
```

`isMirrorImage("{WOW}", "}WOW{")` يجب أن تُرجع `false`.

```js
assert.isFalse(isMirrorImage("{WOW}", "}WOW{"));
```

`isMirrorImage("XXVII", "IIV%X")` يجب أن تُرجع `false`.

```js
assert.isFalse(isMirrorImage("XXVII", "IIV%X"));
```

`isMirrorImage("><(((*>", "<*)))><")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMirrorImage("><(((*>", "<*)))><"));
```

`isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")` يجب أن تُرجع `true`.

```js
assert.isTrue(isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"));
```

# --seed--

## --seed-contents--

```js
function isMirrorImage(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirrorImage(str1, str2) {
  const symmetric = new Set('WTYUIOHA XVMwoxv08=+:|-_*^!.');
  const pairs = { '[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<', 'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '(' };
  
  const mirrored = str1.split('').reverse().map(c => {
    if (symmetric.has(c)) return c;
    if (pairs[c]) return pairs[c];
    return null;
  });
  
  if (mirrored.includes(null)) return false;
  return mirrored.join('') === str2;
}
```
