---
id: 69e2383af7832c8032603b8f
title: "تحدٍّ 273: مدقق ISBN-13"
challengeType: 28
dashedName: challenge-273
---

# --description--

معطاة سلسلة نصية، حدّد إذا كانت رقم ISBN-13 صالحًا.

رقم ISBN-13 صالح:

- يحتوي فقط على أرقام وواصلات
- يحتوي على 13 رقمًا بالضبط بعد إزالة الواصلات
- يجتاز الفحص التالي:
  1. اضرب كل رقم في 1 أو 3 بالتناوب (اضرب الرقم الأول في 1، والثاني في 3، والثالث في 1، وهكذا).
  2. يجب أن يكون مجموع النتائج قابلاً للقسمة على 10.

# --hints--

`isValidIsbn13("9780306406157")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
