---
id: 69f35a5bb823ed620fcb7cbb
title: "تحدٍّ 282: دين النوم"
challengeType: 28
dashedName: challenge-282
---

# --description--

معطاة مصفوفة من ساعات النوم كل ليلة حتى اليوم، وعدد ساعات مستهدف لكل ليلة، أرجع عدد ساعات النوم التي تحتاجها الليلة لتسديد دين النوم.

- أدرج ساعات الليلة ضمن إجمالي الوقت اللازم للتعويض.
- إذا نمت بما يكفي لتغطية الهدف الليلي أو أكثر، أرجع `0`.

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)` يجب أن تُرجع `20`.

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)` يجب أن تُرجع `10`.

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)` يجب أن تُرجع `4`.

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)` يجب أن تُرجع `0`.

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)` يجب أن تُرجع `0`.

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
