---
id: 69f35a5bb823ed620fcb7cbc
title: "تحدٍّ 283: دمج سلسلتي نص"
challengeType: 28
dashedName: challenge-283
---

# --description--

بالنظر إلى سلسلتين نصيتين، أعد سلسلة نصية جديدة تدمج حروفهما بالتناوب حرفًا حرفًا. إذا كانت إحدى السلسلتين أطول، أضف الحروف المتبقية في النهاية.

ابدأ بالحرف الأول من السلسلة النصية الأولى.

# --hints--

`zipStrings("abc", "123")` يجب أن تُعيد `"a1b2c3"`.

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` يجب أن تُعيد `"abcdefghijklmnopqrstuvwxyz"`.

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` يجب أن تُعيد `"dnaiyght"`.

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` يجب أن تُعيد `"pjyatvhaosncript"`.

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` يجب أن تُعيد `"freeCodeCamp"`.

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
