---
id: 69f35a5bb823ed620fcb7cbd
title: "تحدٍّ 284: I قبل E"
challengeType: 28
dashedName: challenge-284
---

# --description--

عند إعطائك كلمة أو جملة، أرجع نسخة مصححة حيث تتبع كل كلمة قاعدة "I قبل E إلا بعد C".

- إذا احتوت الكلمة على `"ei"` ولم تسبقها `"c"`، استبدلها بـ `"ie"`.
- إذا احتوت الكلمة على `"ie"` وسبقها `"c"`، استبدلها بـ `"ei"`.
- تُترك جميع الكلمات الأخرى دون تغيير.

# --hints--

`iBeforeE("beleive")` يجب أن تُرجع `"believe"`.

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` يجب أن تُرجع `"receive"`.

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` يجب أن تُرجع `"we received a brief"`.

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` يجب أن تُرجع `"she believed the friendly niece could perceive the grief"`.

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` يجب أن تُرجع `"we received relief after the thief gave us a brief piece of fierce deceit"`.

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
