---
id: 69f35a5bb823ed620fcb7cbe
title: "تحدٍّ 285: وقت الاجتماع"
challengeType: 28
dashedName: challenge-285
---

# --description--

معطى مصفوفة ثلاثية الأبعاد تمثل نوافذ التوفر لعدة أشخاص، أعد الوقت المبكر حيث يكون لدى الجميع ساعة واحدة متاحة. إذا لم يوجد مثل هذا الوقت، أعد `"None"`.

- توفر كل شخص هو مصفوفة من أزواج صحيحة `[start, end]` بتوقيت 24 ساعة. على سبيل المثال، `[10, 12]` تعني أن الشخص متوفر من الساعة 10 إلى 12. تبدأ أوقات البداية من 0 إلى 23، وأوقات النهاية من 1 إلى 24.

على سبيل المثال، معطى:

```json
[
  [[10, 12], [15, 16]], // person 1
  [[11, 14], [15, 16]]  // person 2
]
```

أعد `11`، بداية أول ساعة مجانية مشتركة بينهم.

# --hints--

`getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]])` يجب أن تعيد `11`.

```js
assert.equal(getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]]), 11);
```

`getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]])` يجب أن تعيد `13`.

```js
assert.equal(getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]]), 13);
```

`getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]])` يجب أن تعيد `9`.

```js
assert.equal(getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]]), 9);
```

`getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]])` يجب أن تعيد `None`.

```js
assert.equal(getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]]), "None");
```

`getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]])` يجب أن تعيد `21`.

```js
assert.equal(getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]]), 21);
```

# --seed--

## --seed-contents--

```js
function getMeetingTime(availability) {

  return availability;
}
```

# --solutions--

```js
function getMeetingTime(availability) {
  for (let hour = 0; hour < 24; hour++) {
    const allFree = availability.every(person =>
      person.some(([start, end]) => start <= hour && hour + 1 <= end)
    );
    if (allFree) return hour;
  }
  return "None";
}
```
