---
id: 69f8c998d78ad3171a0713bd
title: "تحدٍّ 290: حفلة بيتزا"
challengeType: 28
dashedName: challenge-290
---

# --description--

معطاة مصفوفة بساعات العمل اليوم لكل شخص، أرجع عدد البيتزا التي يجب طلبها لحفلة بيتزا.

- اقسم ساعات العمل لكل شخص على 3 لتحصل على عدد شرائح البيتزا الخاصة به.
- لا يمكنك أكل شريحة جزئية، لذا قرب عدد الشرائح لكل شخص إلى أقرب عدد صحيح أكبر.
- يحصل كل شخص على الحد الأدنى من شريحتين.
- تحتوي كل بيتزا على 8 شرائح. قرب العدد الإجمالي للبيتزا إلى أقرب بيتزا كاملة أكبر.

# --hints--

`getPizzasToOrder([8, 8, 8])` يجب أن تُرجع `2`.

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` يجب أن تُرجع `3`.

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` يجب أن تُرجع `2`.

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` يجب أن تُرجع `3`.

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` يجب أن تُرجع `1`.

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` يجب أن تُرجع `5`.

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
