---
id: 69f8c998d78ad3171a0713b9
title: "تحدٍّ 286: القضايا المفتوحة"
challengeType: 28
dashedName: challenge-286
---

# --description--

بالنظر إلى مصفوفة من أرقام القضايا ومصفوفة أخرى من أرقام طلبات السحب (PR)، أعد مصفوفة بالقضايا التي تبقى مفتوحة بعد دمج جميع طلبات السحب.

- يغلق طلب السحب قضية إذا كانت أرقامهما تدويرًا لبعضها البعض. على سبيل المثال، تُغلق القضية 123 بواسطة طلب السحب 231 أو 312.
- لا يغلق طلب السحب قضية تحمل نفس الرقم بالضبط. على سبيل المثال، طلب السحب 123 لا يغلق القضية 123. لذلك لا يمكن إغلاق قضية أرقامها كلها متطابقة.
- قد يُحذف الصفر الرائد من أي رقم. على سبيل المثال، يغلق طلب السحب 201 القضية 12 (012، تدوير للرقم 201). وبالمثل، تُغلق القضية 201 بواسطة طلب السحب 12.

أعد القضايا المفتوحة المتبقية بنفس الترتيب الذي أُعطيت به.

# --hints--

`getOpenIssues([123, 234], [231])` يجب أن تُعيد `[234]`.

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` يجب أن تُعيد `[345, 16]`.

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` يجب أن تُعيد `[456, 332, 15]`.

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` يجب أن تُعيد `[115, 296, 24]`.

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` يجب أن تُعيد `[95, 395, 754, 296, 709, 237, 1802]`.

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
