---
id: 69f8c998d78ad3171a0713bb
title: "تحدٍّ 288: الرقم السري"
challengeType: 28
dashedName: challenge-288
---

# --description--

معطى رقم سري وتخمين، حدّد إذا كان التخمين صحيحًا.

أرجع:

- `"higher"` إذا كان الرقم السري أعلى من التخمين.
- `"lower"` إذا كان الرقم السري أقل من التخمين.
- `"you got it!"` إذا كان التخمين صحيحًا.

# --hints--

`guessNumber(50, 30)` يجب أن تُرجع `"higher"`.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` يجب أن تُرجع `"lower"`.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` يجب أن تُرجع `"you got it!"`.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` يجب أن تُرجع `"higher"`.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` يجب أن تُرجع `"lower"`.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` يجب أن تُرجع `"you got it!"`.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
