---
id: 69f8c998d78ad3171a0713ba
title: "تحدٍّ 287: مُصلح الأعداد الرومانية"
challengeType: 28
dashedName: challenge-287
---

# --description--

بالنظر إلى سلسلة نصية من الأعداد الرومانية المشوّهة، أعد القيمة إلى التدوين القياسي للأعداد الرومانية.

سيستخدم المُدخَل فقط التدوين الجمعي، بحيث يضيف كل رمز قيمته إلى المجموع. للتذكير، هذه هي الرموز والقيم:

| الرمز | القيمة |
| ------ | ----- |
| `"I"`  | 1     |
| `"V"`  | 5     |
| `"X"`  | 10    |
| `"L"`  | 50    |
| `"C"`  | 100   |
| `"D"`  | 500   |
| `"M"`  | 1000  |

عند إعادة الترميز، استخدم أكبر رمز ممكن في كل خطوة، مع استخدام الأزواج الطارحة (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`) عند الحاجة.

# --hints--

`fixNumerals("XIIIII")` يجب أن تُعيد `"XV"`.

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")` يجب أن تُعيد `"LXIV"`.

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")` يجب أن تُعيد `"XL"`.

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")` يجب أن تُعيد `"MDCCXCIX"`.

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")` يجب أن تُعيد `"MCDLXIV"`.

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")` يجب أن تُعيد `"MMCMLXXXIV"`.

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```
