---
id: 69f8c998d78ad3171a0713be
title: "تحدٍّ 291: عد FizzBuzz"
challengeType: 28
dashedName: challenge-291
---

# --description--

بالنظر إلى رقم بداية ورقم نهاية، احسب عدد مرات ظهور fizz و buzz في النطاق (شامل).

- الأرقام القابلة للقسمة على 3 تُحتسب كـ fizz.
- الأرقام القابلة للقسمة على 5 تُحتسب كـ buzz.
- الأرقام القابلة للقسمة على كل من 3 و 5 تُحتسب كـ fizz و buzz معًا.

أرجع كائنًا أو قاموسًا يحتوي على العدادات بالشكل: `{ fizz, buzz }`.

# --hints--

`fizzBuzzCount(1, 11)` يجب أن تُرجع `{fizz: 3, buzz: 2}`.

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` يجب أن تُرجع `{fizz: 9, buzz: 6}`.

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` يجب أن تُرجع `{fizz: 26, buzz: 16}`.

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` يجب أن تُرجع `{fizz: 207, buzz: 125}`.

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` يجب أن تُرجع `{fizz: 4074, buzz: 2444}`.

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
