---
id: 69f8c998d78ad3171a0713bf
title: "تحدٍّ 292: نسبة العرض إلى الارتفاع الأوسع"
challengeType: 28
dashedName: challenge-292
---

# --description--

بالنظر إلى سلسلتي نصية لأبعاد صورة مختلفة، أعد نسبة العرض إلى الارتفاع للصورة التي لها نسبة عرض إلى ارتفاع أكبر.

- ستكون السلاسل النصية المعطاة بالتنسيق `"WxH"`، على سبيل المثال، `"1920x1080"`.
- نسبة العرض إلى الارتفاع هي نسبة العرض إلى الارتفاع، مخفضة إلى أصغر أعداد صحيحة. على سبيل المثال، `"1920x1080"` تُختصر إلى `"16:9"`.
- أعد سلسلة نصية بالتنسيق `"W:H"`، على سبيل المثال، `"16:9"`.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` يجب أن تُعيد `"16:9"`.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` يجب أن تُعيد `"4:3"`.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` يجب أن تُعيد `"2:1"`.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` يجب أن تُعيد `"9:16"`.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` يجب أن تُعيد `"43:18"`.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` يجب أن تُعيد `"2089:8539"`.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
