---
id: 69f90c1329a94b37e2a2086d
title: "تحدٍّ 294: تراكيب الأقواس العادية"
challengeType: 28
dashedName: challenge-294
---

# --description--

معطى عدد صحيح `n`، أعد عدد التراكيب الصالحة لـ `n` أزواج من الأقواس العادية.

- التركيبة الصالحة هي سلسلة نصية حيث كل قوس عادي مفتوح له قوس عادي مغلق مطابق، ولا يظهر أي قوس عادي مغلق قبل قوسه العادي المفتوح المقابل.

على سبيل المثال، عند إعطاء `2`، هناك تراكيب صالحة عددها 2:

```json
(())
()()
```

# --hints--

`getCombinations(2)` يجب أن تُعيد `2`.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` يجب أن تُعيد `5`.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` يجب أن تُعيد `42`.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` يجب أن تُعيد `1430`.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` يجب أن تُعيد `742900`.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
