---
id: 6a0dcc730cb92a616f86f0c2
title: "تحدٍّ 298: مدقق المخطط الجزء 4"
challengeType: 28
dashedName: challenge-298
---

# --description--

معطى كائن (JavaScript) أو قاموس (Python)، حدّد ما إذا كان يطابق المخطط التالي:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean
}
```

- رمز الأنبوب (`|`) يعني "أو". يجب أن يكون `role` واحدًا من قيم `Roles` المدرجة.
- علامة الاستفهام (`?`) بعد `supporter` تعني أن الحقل اختياري، لكنه من النوع المحدد إذا وُجد.
- يُسمح بالمفاتيح الإضافية

# --hints--

`isValidSchema({ username: "vivian", posts: 1, verified: false, role: "user", supporter: true })` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidSchema({ username: "vivian", posts: 1, verified: false, role: "user", supporter: true }));
```

`isValidSchema({ username: "rudolph", posts: 15, verified: true, role: "creator" })` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidSchema({ username: "rudolph", posts: 15, verified: true, role: "creator" }));
```

`isValidSchema({ username: "hernandez", posts: 35, verified: true, role: "moderator", supporter: false, followers: 55 })` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidSchema({ username: "hernandez", posts: 35, verified: true, role: "moderator", supporter: false, followers: 55 }));
```

`isValidSchema({ username: "julia", posts: 50, verified: true, role: "admin", supporter: "true" })` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidSchema({ username: "julia", posts: 50, verified: true, role: "admin", supporter: "true" }));
```

`isValidSchema({ username: "bernard", posts: 0, verified: true, role: "friend", supporter: true })` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidSchema({ username: "bernard", posts: 0, verified: true, role: "friend", supporter: true }));
```

`isValidSchema({ username: "felix", posts: 40, verified: "yes", role: "staff", supporter: false })` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidSchema({ username: "felix", posts: 40, verified: "yes", role: "staff", supporter: false }));
```

`isValidSchema({ username: "jimmy", posts: true, verified: false, role: "creator", supporter: true })` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidSchema({ username: "jimmy", posts: true, verified: false, role: "creator", supporter: true }));
```

`isValidSchema({ username: true, posts: 30, verified: true, role: "moderator", supporter: false })` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidSchema({ username: true, posts: 30, verified: true, role: "moderator", supporter: false }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role) &&
    (obj.supporter === undefined || typeof obj.supporter === 'boolean')
  );
}
```
