---
id: 6a0dcc730cb92a616f86f0c5
title: "تحدٍّ 301: آخر تحميل"
challengeType: 28
dashedName: challenge-301
---

# --description--

بما أن لديك عددًا معينًا من مغارف مسحوق الغسيل المتبقية ومصفوفة تبين عدد المغارف التي استخدمتها في الأيام السابقة، أرجع عدد الأيام الكاملة التي يمكنك استخدام المسحوق خلالها.

احسب متوسط استهلاكك اليومي من تاريخ الاستخدام وافترض نفس كمية الاستهلاك يوميًا في المستقبل.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` يجب أن تُرجع `5`.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` يجب أن تُرجع `7`.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` يجب أن تُرجع `11`.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` يجب أن تُرجع `12`.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` يجب أن تُرجع `1`.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
