---
id: 6a0dcd03ee4e68698080ef6b
title: "تحدٍّ 307: تنظيمات التقسيم"
challengeType: 28
dashedName: challenge-307
---

# --description--

معطى شبكة ثنائية الأبعاد (مصفوفة من مصفوفات) تمثل تخطيط مباني المدينة، أرجع إحداثيات كل المباني التي تنتهك قواعد التقسيم.

كل خلية في الشبكة تحتوي على أحد المسميات من الجدول أدناه. المبنى يكون مخالفًا إذا كان أي من جيرانه الأربعة (أفقيًا أو عموديًا) من نوع لا يمكن أن يكون مجاورًا له.

| مسمى | النوع | لا يمكن أن يكون مجاورًا لـ |
| - | - | - |
| `"i"` | صناعي | `"R"`، `"I"` |
| `"A"` | زراعي | `"C"` |
| `"R"` | سكني | `"i"`، `"C"` |
| `"I"` | مؤسسي | `"i"` |
| `"C"` | تجاري | `"R"`، `"A"` |
| `""` (سلسلة نصية فارغة) | غير مطور | لا قيود |

أرجع إحداثيات كل الخلايا المخالفة كمصفوفة من أزواج `[row, col]` بأي ترتيب. إذا لم توجد مخالفات، أرجع مصفوفة فارغة.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` يجب أن تُرجع `[[0, 0], [0, 1]]`.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` يجب أن تُرجع `[[0, 1], [1, 1]]`.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` يجب أن تُرجع `[]`.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` يجب أن تُرجع `[[0, 1], [0, 2], [0, 3]]`.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` يجب أن تُرجع `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
