---
id: 6a0dcd03ee4e68698080ef68
title: "تحدٍّ 304: ترتيبات مسار الرحلة"
challengeType: 28
dashedName: challenge-304
---

# --description--

معطاة مصفوفة تحتوي على توقفين اختياريين على الأقل لرحلة يومية، أعد عدد ترتيبات المسار الصالحة.

يتضمن المسار دائمًا `"breakfast"` و`"lunch"` و`"dinner"`، ولن تُمرر هذه كمعلمات. يمكن وضع التوقفات الاختيارية في أي مكان في المسار، مع مراعاة القواعد التالية:

- `"breakfast"` دائمًا في البداية، مع وجود توقف واحد على الأقل قبل `"lunch"`.
- يجب أن يظهر `"lunch"` قبل `"dinner"`، مع وجود توقف واحد على الأقل بينهما.
- لا يجوز أن يظهر أكثر من توقف اختياري واحد بعد `"dinner"`.

أعد عدد الترتيبات الصالحة.

# --hints--

`getItineraryCount(["library", "park"])` يجب أن تُعيد `2`.

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` يجب أن تُعيد `18`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` يجب أن تُعيد `120`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` يجب أن تُعيد `840`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` يجب أن تُعيد `55440`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
