---
id: 6a0dcd03ee4e68698080ef66
title: "تحدٍّ 302: تعب السفر بسبب اختلاف التوقيت"
challengeType: 28
dashedName: challenge-302
---

# --description--

معطى مدينة المغادرة، ومدينة الوصول، ومدة الرحلة بالساعات، واتجاه السفر، أعد عدد ساعات تعب السفر التي يعاني منها المسافر.

المدن المعطاة ستكون من القائمة التالية التي تتضمن إزاحتها عن توقيت UTC:

| المدينة | الإزاحة |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

لحساب ساعات تعب السفر:

1. احسب فرق التوقيت بالساعات بين المدينتين.
2. حدد معامل الاتجاه. إذا كان السفر نحو `"east"`، فالمعامل 1.5، وإلا فهو 1.0.
3. احصل على ساعات تعب السفر بالمعادلة: فرق التوقيت + (مدة الرحلة * 0.1) * معامل الاتجاه

أعد ساعات تعب السفر مقربة إلى منزلة عشرية واحدة.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` يجب أن تُعيد `6.5`.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` يجب أن تُعيد `5.8`.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` يجب أن تُعيد `1.6`.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` يجب أن تُعيد `4.7`.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` يجب أن تُعيد `17.5`.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` يجب أن تُعيد `5.9`.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` يجب أن تُعيد `9.5`.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
