---
id: 6a0dcd03ee4e68698080ef67
title: "تحدٍّ 303: رفقاء الغرفة"
challengeType: 28
dashedName: challenge-303
---

# --description--

معطاة مصفوفة من الأشخاص ومجموعات رفقاء الغرفة الخاصة بهم، أرجع تعيينات الغرف لإقامة في فندق باستخدام القواعد التالية:

- كل شخص لديه خاصية `name` و`group`:

```json
[
  { "name": "Alice", "group": "A" },
  { "name": "Bob", "group": "B" },
  { "name": "Carol", "group": "A" }
]
```

- يمكن للأشخاص مشاركة الغرفة فقط مع شخص من نفس المجموعة ويتم إقرانهم حسب الترتيب الذي يُعطون به.
- أرجع مصفوفة من السلاسل النصية مع الأسماء مفصولة بـ `" and "` للغرفة المشتركة، واسم واحد فقط للغرفة الفردية. يجب أن تظهر الأسماء بالترتيب الذي تم إقرانهم به. للمثال أعلاه، أرجع `["Alice and Carol", "Bob"]`.

# --hints--

`getRoommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }])` يجب أن تُرجع `["Alice and Carol", "Bob"]`.

```js
assert.deepEqual(getRoommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }]).sort(), ["Alice and Carol", "Bob"].sort());
```

`getRoommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }])` يجب أن تُرجع `["John and Julia", "Jim"]`.

```js
assert.deepEqual(getRoommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }]).sort(), ["John and Julia", "Jim"].sort());
```

`getRoommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }])` يجب أن تُرجع `["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"]`.

```js
assert.deepEqual(getRoommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }]).sort(), ["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"].sort());
```

`getRoommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }])` يجب أن تُرجع `["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"]`.

```js
assert.deepEqual(getRoommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }]).sort(), ["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"].sort());
```

`getRoommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }])` يجب أن تُرجع `["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"]`.

```js
assert.deepEqual(getRoommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }]).sort(), ["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"].sort());
```

# --seed--

## --seed-contents--

```js
function getRoommates(people) {

  return people;
}
```

# --solutions--

```js
function getRoommates(people) {
  const groups = {};

  for (const { name, group } of people) {
    (groups[group] = groups[group] || []).push(name);
  }

  return Object.values(groups).flatMap(group => {
    const rooms = [];
    for (let i = 0; i < group.length; i += 2) {
      rooms.push(group.slice(i, i + 2).join(" and "));
    }
    return rooms;
  });
}
```
