---
id: 6a0dcc730cb92a616f86f0bf
title: "تحدٍّ 295: مدقق المخطط الجزء 1"
challengeType: 28
dashedName: challenge-295
---

# --description--

معطى كائن (JavaScript) أو قاموس (Python)، حدّد إذا كان يتطابق مع المخطط التالي:

```json
{
  username: string
}
```

- يُسمح بالمفاتيح الإضافية

# --hints--

`isValidSchema({ username: "bob" })` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
