---
id: 6a15cadf5f240d05a264955a
title: "تحدٍّ 316: لغة 1337"
challengeType: 28
dashedName: challenge-316
---

# --description--

بالنظر إلى سلسلة نصية من أحرف صغيرة، أعدها مترجمة إلى لغة 1337 باستبدال الحروف أدناه بنظائرها في لغة 1337:

| الحرف | 1337 |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- تُترك الأحرف التي لا يوجد لها بديل كما هي دون تغيير.

# --hints--

`makeLeet("cool")` يجب أن تُعيد `"c001"`.

```js
assert.equal(makeLeet("cool"), "c001");
```

`makeLeet("leet")` يجب أن تُعيد `"1337"`.

```js
assert.equal(makeLeet("leet"), "1337");
```

`makeLeet("hacker")` يجب أن تُعيد `"h4ck3r"`.

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")` يجب أن تُعيد `"547311173"`.

```js
assert.equal(makeLeet("satellite"), "547311173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")` يجب أن تُعيد `"4bcd3f9h1jk1mn0pqr57uvwxyz"`.

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return str;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", g: "9", i: "1", l: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
