---
id: 6a15cadf5f240d05a2649557
title: "تحدٍّ 313: تكلفة الإيجار"
challengeType: 28
dashedName: challenge-313
---

# --description--

معطى طابع زمني للإيجار، وطابع زمني للإرجاع، ودرجة الإيجار، أرجع التكلفة الإجمالية للإيجار بما في ذلك أي رسوم تأخير.

- الطوابع الزمنية المعطاة هي سلاسل نصية بصيغة UTC ISO، على سبيل المثال: `"2026-06-18T18:30:00Z"`.
- درجة الإيجار هي عدد الأيام قبل موعد إرجاع الإيجار: `1`، `3`، أو `7`.
- يجب إرجاع الإيجارات بحلول الساعة 12:00 ظهرًا بتوقيت UTC أو قبلها في آخر يوم من فترة الإيجار. على سبيل المثال، إيجار ليوم واحد تم استعارته في أي وقت من يوم 15 مارس يجب إرجاعه بحلول الساعة 12:00 ظهرًا بتوقيت UTC في 16 مارس.
- كل يوم يتجاوز موعد الإرجاع يترتب عليه رسوم تأخير.

التسعير كما يلي:

| الدرجة | التكلفة الأساسية | رسوم التأخير لكل يوم |
| - | - | - |
| يوم واحد | 4.99$ | 3.99$ |
| 3 أيام | 3.99$ | 2.99$ |
| 7 أيام | 2.99$ | 0.99$ |

أرجع التكلفة الإجمالية مقربة إلى منزلتين عشريتين بالشكل `"$D.CC"`.

# --hints--

`getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` يجب أن تُرجع `"$4.99"`.

```js
assert.equal(getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99");
```

`getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` يجب أن تُرجع `"$12.97"`.

```js
assert.equal(getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97");
```

`getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` يجب أن تُرجع `"$3.99"`.

```js
assert.equal(getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99");
```

`getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` يجب أن تُرجع `"$9.97"`.

```js
assert.equal(getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97");
```

`getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` يجب أن تُرجع `"$2.99"`.

```js
assert.equal(getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99");
```

`getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` يجب أن تُرجع `"$358.40"`.

```js
assert.equal(getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40");
```

# --seed--

## --seed-contents--

```js
function getRentalCost(rented, returned, tier) {

  return rented;
}
```

# --solutions--

```js
function getRentalCost(rented, returned, tier) {
  const pricing = {
    1: { base: 4.99, late: 3.99 },
    3: { base: 3.99, late: 2.99 },
    7: { base: 2.99, late: 0.99 },
  };

  const rentedDate = new Date(rented);
  const returnDate = new Date(returned);
  const due = new Date(rentedDate);
  due.setUTCDate(due.getUTCDate() + tier);
  due.setUTCHours(12, 0, 0, 0);
  const { base, late } = pricing[tier];

  if (returnDate <= due) {
    return `$${base.toFixed(2)}`;
  }

  const msPerDay = 1000 * 60 * 60 * 24;
  const daysLate = Math.ceil((returnDate - due) / msPerDay);
  const total = base + daysLate * late;

  return `$${total.toFixed(2)}`;
}
```
