---
id: 6a151d32d772271dd2448c2d
title: "تحدٍّ 311: ساحر التعاويذ"
challengeType: 28
dashedName: challenge-311
---

# --description--

معطاة سلسلة من رموز التعاويذ التي تلقيها، احسب المجموع الكلي للنقاط.

كل حرف في السلسلة يُمثّل تعويذة:

| الرمز | التعويذة | الفئة | النقاط الأساسية |
| - | - | - | - |
| `"f"` | نار | التدمير | 3 |
| `"l"` | برق | التدمير | 3 |
| `"i"` | جليد | السيطرة | 2 |
| `"w"` | ريح | السيطرة | 2 |
| `"h"` | شفاء | الاستعادة | 1 |
| `"s"` | درع | الاستعادة | 1 |

يُطبّق مضاعف التتابع بناءً على عدد التعاويذ المتتالية التي أُلقيت من فئات مختلفة:

- التعويذة الأولى دائمًا تُحسب بقيمتها الأساسية.
- كل تعويذة متتالية من فئة مختلفة عن السابقة تزيد المضاعف بمقدار 1.
- إلقاء تعويذة من نفس الفئة كالسابقة يعيد المضاعف إلى 1.
- النقاط لكل تعويذة تساوي نقاطها الأساسية مضروبة في المضاعف الحالي.

أرجع المجموع الكلي للنقاط من متتالية التعاويذ.

# --hints--

`cast("fihwl")` يجب أن تُرجع `33`.

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` يجب أن تُرجع `45`.

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` يجب أن تُرجع `37`.

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` يجب أن تُرجع `50`.

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` يجب أن تُرجع `101`.

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
