---
id: 6a15cadf5f240d05a2649556
title: "تحدٍّ 312: تكلفة البث"
challengeType: 28
dashedName: challenge-312
---

# --description--

معطاة مصفوفة تمثل الأفلام في عربة التسوق لخدمة البث الخاصة بك، وسلسلة نصية لمستوى الاشتراك، أرجع التكلفة الإجمالية للأفلام.

كل عنصر في العربة هو كائن يحتوي على `"format"` (`"HD"` أو `"4K"`) و`"type"` (`"rent"` أو `"buy"`). تكاليفها هي:

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

طبق الخصومات التالية لمستوى الاشتراك:

- `"none"`: السعر الكامل
- `"basic"`: خصم 10%
- `"premium"`: خصم 25%

أرجع التكلفة الإجمالية مقربة إلى منزلتين عشريتين بالشكل `"$D.CC"`.

# --hints--

`getStreamingBill([{ format: "HD", type: "rent" }], "none")` يجب أن تُرجع `"$3.99"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }], "none"), "$3.99");
```

يجب أن تُعيد الدالة `getStreamingBill([{ format: "HD", type: "buy" }], "premium")` القيمة `"$9.74"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "buy" }], "premium"), "$9.74");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic")` يجب أن تُرجع `"$18.87"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic"), "$18.87");
```

يجب أن تُعيد الدالة `getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium")` القيمة `"$44.98"`.

```js
assert.equal(getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium"), "$44.98");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none")` يجب أن تُرجع `"$42.96"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none"), "$42.96");
```

يجب أن تُعيد الدالة `getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic")` القيمة `"$38.66"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic"), "$38.66");
```

# --seed--

## --seed-contents--

```js
function getStreamingBill(cart, subscription) {

  return cart;
}
```

# --solutions--

```js
function getStreamingBill(cart, subscription) {
  const pricing = {
    HD: { rent: 3.99, buy: 12.99 },
    "4K": { rent: 5.99, buy: 19.99 },
  };

  const discounts = {
    none: 1,
    basic: 0.9,
    premium: 0.75,
  };

  const multiplier = discounts[subscription];

  const total = cart.reduce((sum, item) => {
    const price = pricing[item.format][item.type];
    return sum + price * multiplier;
  }, 0);

  return `$${total.toFixed(2)}`;
}
```
