---
id: 6a15cadf5f240d05a264955b
title: "تحدٍّ 317: حاسبة مؤشر كتلة الجسم"
challengeType: 28
dashedName: challenge-317
---

# --description--

بالنظر إلى وزن بوحدة الباوند وطول بوحدة الإنش، أرجع مؤشر كتلة الجسم (BMI) مقربًا إلى منزلة عشرية واحدة.

لحساب مؤشر كتلة الجسم: اقسم الوزن على مربع الطول، ثم اضرب الناتج في 703.

# --hints--

`calculateBmi(180, 70)` يجب أن تُرجع `25.8`.

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)` يجب أن تُرجع `24.0`.

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)` يجب أن تُرجع `19.5`.

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)` يجب أن تُرجع `39.1`.

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)` يجب أن تُرجع `22.8`.

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
