---
id: 6a15cadf5f240d05a2649558
title: "تحدٍّ 314: التحليل إلى عوامل أولية"
challengeType: 28
dashedName: challenge-314
---

# --description--

أعطِ عددًا صحيحًا أكبر من 1، أرجع تحليله إلى عوامل أولية كمصفوفة أعداد مرتبة تصاعديًا.

التحليل إلى عوامل أولية هو مجموعة الأعداد الأولية التي تضرب معًا لإنتاج العدد الصحيح المعطى. لكل عدد مجموعة واحدة فقط. على سبيل المثال، التحليل إلى عوامل أولية للعدد 20 هو `[2, 2, 5]` لأن 2 * 2 * 5 = 20.

إذا كان العدد الصحيح المعطى عددًا أوليًا بحد ذاته، أرجعه في مصفوفة ذات عنصر واحد.

# --hints--

`primeFactorization(20)` يجب أن تُرجع `[2, 2, 5]`.

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` يجب أن تُرجع `[17]`.

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` يجب أن تُرجع `[3, 5]`.

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` يجب أن تُرجع `[5, 7]`.

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` يجب أن تُرجع `[3, 3, 3, 37]`.

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` يجب أن تُرجع `[2, 2, 2, 3, 3, 5]`.

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` يجب أن تُرجع `[2, 3, 5, 7, 11, 13, 17]`.

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
