---
id: 6a19b1062d1b153d8ac76d70
title: "تحدٍّ 322: ربط 3"
challengeType: 28
dashedName: challenge-322
---

# --description--

معطاة مصفوفة من السلاسل النصية تمثل القطع على شبكة لعبة، حدّد إذا كان أي لاعب لديه ثلاثة على التوالي.

- كل خلية تحتوي على `"R"` أو `"Y"` أو `""` (سلسلة نصية فارغة).
- ثلاثة على التوالي تعني ثلاث خلايا متتالية غير فارغة من نفس النوع أفقيًا أو عموديًا أو قطريًا.

أرجع:

- مصفوفة مسطحة تحتوي على الفائز وإحداثيات خلايا الفوز الثلاث بصيغة: `["R", [0,2], [1,3], [2,4]]`. تُرجع الإحداثيات من الأعلى إلى الأسفل، ثم من اليسار إلى اليمين.
- مصفوفة فارغة إذا لم يكن هناك فائز.

# --hints--

`connectThree([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]])` يجب أن تُرجع `["R", [3, 1], [3, 2], [3, 3]]`.

```js
assert.deepEqual(connectThree([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]]), ["R", [3, 1], [3, 2], [3, 3]]);
```

`connectThree([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]])` يجب أن تُرجع `["Y", [1, 1], [2, 1], [3, 1]]`.

```js
assert.deepEqual(connectThree([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]]), ["Y", [1, 1], [2, 1], [3, 1]]);
```

`connectThree([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]])` يجب أن تُرجع `["R", [0, 3], [1, 2], [2, 1]]`.

```js
assert.deepEqual(connectThree([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]]), ["R", [0, 3], [1, 2], [2, 1]]);
```

`connectThree([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]])` يجب أن تُرجع `["Y", [0, 1], [1, 2], [2, 3]]`.

```js
assert.deepEqual(connectThree([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]]), ["Y", [0, 1], [1, 2], [2, 3]]);
```

`connectThree([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]])` يجب أن تُرجع `[]`.

```js
assert.deepEqual(connectThree([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]]), []);
```

# --seed--

## --seed-contents--

```js
function connectThree(matrix) {

  return matrix;
}
```

# --solutions--

```js
function connectThree(matrix) {
  const rows = matrix.length;
  const cols = matrix[0].length;

  const directions = [[0, 1], [1, 0], [1, 1], [1, -1]];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const piece = matrix[r][c];
      if (!piece) continue;

      for (const [dr, dc] of directions) {
        const cells = [[r, c]];

        for (let i = 1; i < 3; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols) break;
          if (matrix[nr][nc] !== piece) break;
          cells.push([nr, nc]);
        }

        if (cells.length === 3) {
          cells.sort(([r1, c1], [r2, c2]) => r1 !== r2 ? r1 - r2 : c1 - c2);
          return [piece, ...cells];
        }
      }
    }
  }

  return [];
}
```
