---
id: 6a19b1062d1b153d8ac76d73
title: "تحدٍّ 325: الرقم المحظوظ"
challengeType: 28
dashedName: challenge-325
---

# --description--

معطاة سلسلة نصية تحتوي على الاسم الأول واسم العائلة لشخص، احسب رقمه المحظوظ باستخدام القواعد التالية:

- الاسم الأول واسم العائلة مفصولان بمسافة
- احسب عدد الحروف المتحركة والحروف الساكنة لكل اسم
- اضرب العدد الأصغر من الحروف المتحركة والساكنة ببعضهما ثم في طول الاسم الأصغر
- افعل الشيء نفسه للعددين الأكبر وللاسم الأكبر
- اطرح القيمة الأصغر من الأكبر لتحصل على الرقم المحظوظ

إذا كانت القيمة النهائية صفرًا (`0`)، أرجع `13`.

# --hints--

`getLuckyNumber("John Doe")` يجب أن تُرجع `21`.

```js
assert.equal(getLuckyNumber("John Doe"), 21);
```

`getLuckyNumber("Olivia Lewis")` يجب أن تُرجع `52`.

```js
assert.equal(getLuckyNumber("Olivia Lewis"), 52);
```

`getLuckyNumber("James Wilson")` يجب أن تُرجع `18`.

```js
assert.equal(getLuckyNumber("James Wilson"), 18);
```

`getLuckyNumber("Elizabeth Hernandez")` يجب أن تُرجع `81`.

```js
assert.equal(getLuckyNumber("Elizabeth Hernandez"), 81);
```

`getLuckyNumber("Mike Walker")` يجب أن تُرجع `32`.

```js
assert.equal(getLuckyNumber("Mike Walker"), 32);
```

`getLuckyNumber("Chloe Perez")` يجب أن تُرجع `13`.

```js
assert.equal(getLuckyNumber("Chloe Perez"), 13);
```

# --seed--

## --seed-contents--

```js
function getLuckyNumber(name) {

  return name;
}
```

# --solutions--

```js
function getLuckyNumber(name) {
  const vowels = "aeiouAEIOU";

  function nameStats(str) {
    const vowelCount = str.split("").filter(c => vowels.includes(c)).length;
    const consonantCount = str.split("").filter(c => c.match(/[a-zA-Z]/) && !vowels.includes(c)).length;
    return { vowelCount, consonantCount, length: str.length };
  }

  const [first, last] = name.split(" ");
  const firstStats = nameStats(first);
  const lastStats = nameStats(last);

  const largerVowels = Math.max(firstStats.vowelCount, lastStats.vowelCount);
  const smallerVowels = Math.min(firstStats.vowelCount, lastStats.vowelCount);
  const largerConsonants = Math.max(firstStats.consonantCount, lastStats.consonantCount);
  const smallerConsonants = Math.min(firstStats.consonantCount, lastStats.consonantCount);
  const largerLength = Math.max(firstStats.length, lastStats.length);
  const smallerLength = Math.min(firstStats.length, lastStats.length);

  const largerValue = largerVowels * largerConsonants * largerLength;
  const smallerValue = smallerVowels * smallerConsonants * smallerLength;
  const result = largerValue - smallerValue;

  return result === 0 ? 13 : result;
}
```
