---
id: 6a19b1062d1b153d8ac76d72
title: "تحدٍّ 324: عدّ الحروف المكررة"
challengeType: 28
dashedName: challenge-324
---

# --description--

بالنظر إلى سلسلتين نصيتين، أعد عدد الحروف من السلسلة النصية الثانية التي يمكن إيجادها في الأولى.

- تُحتسب الحروف المكررة في السلسلة النصية الثانية بشكل منفصل.

# --hints--

`duplicateCharacterCount("aloha", "hei")` يجب أن تُعيد `1`.

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` يجب أن تُعيد `4`.

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` يجب أن تُعيد `3`.

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` يجب أن تُعيد `0`.

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` يجب أن تُعيد `5`.

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` يجب أن تُعيد `2`.

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` يجب أن تُعيد `26`.

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
