---
id: 6a1d9f98e819ed70a0e994da
title: "تحدٍّ 328: روتين كابريكار"
challengeType: 28
dashedName: challenge-328
---

# --description--

أعطِ رقمًا مكونًا من 4 أرقام، أرجع عدد المرات التي تحتاج فيها لتطبيق روتين كابريكار حتى تصل إلى 6174.

يعمل روتين كابريكار كما يلي:

- رتب الأرقام ترتيبًا تنازليًا لتكوين أكبر رقم
- رتب الأرقام ترتيبًا تصاعديًا لتكوين أصغر رقم (أضف أصفارًا بادئة إذا لزم الأمر)
- اطرح الرقم الأصغر من الأكبر
- كرر مع الرقم الجديد

# --hints--

يجب أن تُرجع الدالة `kaprekar(1234)` القيمة `3`.

```js
assert.equal(kaprekar(1234), 3);
```

يجب أن تُرجع الدالة `kaprekar(2025)` القيمة `6`.

```js
assert.equal(kaprekar(2025), 6);
```

يجب أن تُرجع الدالة `kaprekar(7173)` القيمة `4`.

```js
assert.equal(kaprekar(7173), 4);
```

يجب أن تُرجع الدالة `kaprekar(3164)` القيمة `7`.

```js
assert.equal(kaprekar(3164), 7);
```

يجب أن تُرجع الدالة `kaprekar(8082)` القيمة `2`.

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
