---
id: 6a1d9f98e819ed70a0e994dd
title: "تحدٍّ 331: أقرب مضاعف"
challengeType: 28
dashedName: challenge-331
---

# --description--

بالنظر إلى عددين صحيحين، قرب الأول إلى أقرب مضاعف للثاني.

# --hints--

`roundToNearestMultiple(5, 3)` يجب أن تُرجع `6`.

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` يجب أن تُرجع `16`.

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` يجب أن تُرجع `45`.

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` يجب أن تُرجع `33`.

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` يجب أن تُرجع `96`.

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
