---
id: 6a1d9f98e819ed70a0e994e0
title: "تحدٍّ 334: التغيير الدقيق"
challengeType: 28
dashedName: challenge-334
---

# --description--

معطى مبلغ صحيح بالسنتات، أرجع عدد الطرق المختلفة لصنع التغيير الدقيق باستخدام البنسات (1 سنت)، النيكل (5 سنتات)، الدايم (10 سنتات)، والكوارتير (25 سنتًا).

# --hints--

`exactChange(3)` يجب أن تُرجع `1`.

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` يجب أن تُرجع `2`.

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` يجب أن تُرجع `6`.

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` يجب أن تُرجع `24`.

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` يجب أن تُرجع `73`.

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` يجب أن تُرجع `213`.

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
