---
id: 6a1d9f98e819ed70a0e994e1
title: "تحدٍّ 335: خمسة نردات"
challengeType: 28
dashedName: challenge-335
---

# --description--

معطاة مصفوفة من خمسة نردات بقيم من 1 إلى 6، أرجع أفضل يد ممكنة.

هذه هي الأيدي مرتبة من الأدنى إلى الأعلى:

| اليد | الوصف |
|------|-------------|
| `"no pair"` | لا يوجد زوج أو أفضل |
| `"pair"` | نردان لهما نفس القيمة |
| `"two pair"` | زوجان مختلفان |
| `"three of a kind"` | ثلاثة نردات لها نفس القيمة |
| `"small straight"` | أربع قيم متتالية |
| `"large straight"` | خمس قيم متتالية |
| `"full house"` | ثلاثة من نوع واحد وزوج |
| `"four of a kind"` | أربعة نردات لها نفس القيمة |
| `"five of a kind"` | جميع النردات الخمسة لها نفس القيمة |

# --hints--

`fiveDice([1, 1, 1, 1, 1])` يجب أن تُرجع `"five of a kind"`.

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])` يجب أن تُرجع `"four of a kind"`.

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])` يجب أن تُرجع `"large straight"`.

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])` يجب أن تُرجع `"three of a kind"`.

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])` يجب أن تُرجع `"pair"`.

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])` يجب أن تُرجع `"no pair"`.

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])` يجب أن تُرجع `"small straight"`.

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])` يجب أن تُرجع `"full house"`.

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])` يجب أن تُرجع `"two pair"`.

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
