---
id: 6a1d9f98e819ed70a0e994df
title: "تحدٍّ 333: تصنيف القضايا 2"
challengeType: 28
dashedName: challenge-333
---

# --description--

بالنظر إلى عنوان قضية ومصفوفة من المسميات الحالية، أرجع مصفوفة محدثة من المسميات بناءً على القواعد التالية:

إذا لم يكن لدى القضية أي مسميات، أضف:

- `"bug"` و `"needs triage"` إذا كان العنوان يحتوي على `"error"` أو `"bug"`
- `"enhancement"` و `"discussing"` إذا كان العنوان يحتوي على `"feature"` أو `"add"`

أما إذا كانت المسميات المعطاة تحتوي على:

- `"needs triage"` وكان العنوان يحتوي على `"simple"` أو `"easy"`، أزل `"needs triage"` وأضف `"good first issue"`
- `"discussing"` وكان العنوان يحتوي على `"planned"` أو `"next"`، أزل `"discussing"` وأضف `"on the roadmap"`
- وإلا، إذا كان `"needs triage"` أو `"discussing"` موجودًا، أزلها وأضف `"help wanted"`

إذا كان العنوان يحتوي على:

- `"security"`، أضف مسمى `"critical"`

# --hints--

`triageIssue("app crashes with error", [])` يجب أن تُرجع `["bug", "needs triage"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", []).sort(), ["bug", "needs triage"].sort());
```

`triageIssue("app crashes with error", ["bug", "needs triage"])` يجب أن تُرجع `["bug", "help wanted"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", ["bug", "needs triage"]).sort(), ["bug", "help wanted"].sort());
```

`triageIssue("add dark mode", [])` يجب أن تُرجع `["enhancement", "discussing"]`.

```js
assert.deepEqual(triageIssue("add dark mode", []).sort(), ["enhancement", "discussing"].sort());
```

`triageIssue("add dark mode", ["enhancement", "discussing"])` يجب أن تُرجع `["enhancement", "help wanted"]`.

```js
assert.deepEqual(triageIssue("add dark mode", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted"].sort());
```

`triageIssue("xss security bug", [])` يجب أن تُرجع `["bug", "needs triage", "critical"]`.

```js
assert.deepEqual(triageIssue("xss security bug", []).sort(), ["bug", "needs triage", "critical"].sort());
```

`triageIssue("security vulnerability in auth", [])` يجب أن تُرجع `["critical"]`.

```js
assert.deepEqual(triageIssue("security vulnerability in auth", []).sort(), ["critical"].sort());
```

`triageIssue("easy a11y fix", ["bug", "needs triage"])` يجب أن تُرجع `["bug", "good first issue"]`.

```js
assert.deepEqual(triageIssue("easy a11y fix", ["bug", "needs triage"]).sort(), ["bug", "good first issue"].sort());
```

`triageIssue("planned api migration", ["enhancement", "discussing"])` يجب أن تُرجع `["enhancement", "on the roadmap"]`.

```js
assert.deepEqual(triageIssue("planned api migration", ["enhancement", "discussing"]).sort(), ["enhancement", "on the roadmap"].sort());
```

`triageIssue("improve security", ["enhancement", "discussing"])` يجب أن تُرجع `["enhancement", "help wanted", "critical"]`.

```js
assert.deepEqual(triageIssue("improve security", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted", "critical"].sort());
```

# --seed--

## --seed-contents--

```js
function triageIssue(title, labels) {

  return title;
}
```

# --solutions--

```js
function triageIssue(title, labels) {
  const result = [...labels];
  const t = title.toLowerCase();

  const has = label => result.includes(label);
  const add = label => { if (!has(label)) result.push(label); };
  const remove = label => { const i = result.indexOf(label); if (i > -1) result.splice(i, 1); };

  if (result.length === 0) {
    if (t.includes("error") || t.includes("bug")) {
      add("bug"); add("needs triage");
    } else if (t.includes("feature") || t.includes("add")) {
      add("enhancement"); add("discussing");
    }
  } else {
    if (has("needs triage") && (t.includes("simple") || t.includes("easy"))) {
      remove("needs triage"); add("good first issue");
    } else if (has("discussing") && (t.includes("planned") || t.includes("next"))) {
      remove("discussing"); add("on the roadmap");
    } else if (has("needs triage") || has("discussing")) {
      remove("needs triage"); remove("discussing"); add("help wanted");
    }
  }

  if (t.includes("security")) {
    add("critical");
  }

  return result;
}
```
