---
id: 6a1d9f98e819ed70a0e994e2
title: "تحدٍّ 336: توافق الأبراج"
challengeType: 28
dashedName: challenge-336
---

# --description--

بالنظر إلى سلسلتي نصوص لعلامتي برج، أرجع نسبة توافقهما.

ترتب العلامات في عجلة مكونة من 12 موقعًا بهذا الترتيب: `"Aries"`، `"Taurus"`، `"Gemini"`، `"Cancer"`، `"Leo"`، `"Virgo"`، `"Libra"`، `"Scorpio"`، `"Sagittarius"`، `"Capricorn"`، `"Aquarius"`، `"Pisces"`، وتعود إلى `"Aries"` بعد `"Pisces"`. جد أقصر مسافة بين العلامتين وأرجع نسبة التوافق:

| المسافة | التوافق |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

يجب أن تُرجع الدالة `horoscopeMatch("Libra", "Sagittarius")` القيمة `"80%"`.

```js
assert.equal(horoscopeMatch("Libra", "Sagittarius"), "80%");
```

يجب أن تُرجع الدالة `horoscopeMatch("Gemini", "Scorpio")` القيمة `"20%"`.

```js
assert.equal(horoscopeMatch("Gemini", "Scorpio"), "20%");
```

يجب أن تُرجع الدالة `horoscopeMatch("Pisces", "Aries")` القيمة `"40%"`.

```js
assert.equal(horoscopeMatch("Pisces", "Aries"), "40%");
```

يجب أن تُرجع الدالة `horoscopeMatch("Capricorn", "Cancer")` القيمة `"50%"`.

```js
assert.equal(horoscopeMatch("Capricorn", "Cancer"), "50%");
```

يجب أن تُرجع الدالة `horoscopeMatch("Aquarius", "Aquarius")` القيمة `"100%"`.

```js
assert.equal(horoscopeMatch("Aquarius", "Aquarius"), "100%");
```

يجب أن تُرجع الدالة `horoscopeMatch("Virgo", "Taurus")` القيمة `"90%"`.

```js
assert.equal(horoscopeMatch("Virgo", "Taurus"), "90%");
```

يجب أن تُرجع الدالة `horoscopeMatch("Leo", "Scorpio")` القيمة `"30%"`.

```js
assert.equal(horoscopeMatch("Leo", "Scorpio"), "30%");
```

# --seed--

## --seed-contents--

```js
function horoscopeMatch(sign1, sign2) {

  return sign1;
}
```

# --solutions--

```js
function horoscopeMatch(sign1, sign2) {
  const signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"];
  
  const a = signs.indexOf(sign1) + 1;
  const b = signs.indexOf(sign2) + 1;
  const diff = Math.abs(a - b);
  const distance = Math.min(diff, 12 - diff);

  const compatibility = { 0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%" };
  return compatibility[distance];
}
```
