---
id: 6a1d9f98e819ed70a0e994de
title: "تحدٍّ 332: فرز القضايا"
challengeType: 28
dashedName: challenge-332
---

# --description--

معطى عدد من المللي ثانية منذ آخر منشور على قضية، وآخر رسالة نُشرت على القضية، حدّد ما يجب فعله مع القضية وفقًا لهذه القواعد:

- إذا كانت آخر رسالة قبل أقل من 7 أيام، أعد `"leave it"`
- إذا كانت آخر رسالة قبل 7 أيام أو أكثر ومحتواها يحتوي على `"bump"` (بغض النظر عن حالة الأحرف)، أعد `"close it"`
- خلاف ذلك، أعد `"bump it"`

# --hints--

`triageIssue(86400000, "Lets fix it")` يجب أن تُعيد `"leave it"`.

```js
assert.equal(triageIssue(86400000, "Lets fix it"), "leave it");
```

`triageIssue(1209600000, "still waiting")` يجب أن تُعيد `"bump it"`.

```js
assert.equal(triageIssue(1209600000, "still waiting"), "bump it");
```

`triageIssue(864000000, "bump")` يجب أن تُعيد `"close it"`.

```js
assert.equal(triageIssue(864000000,  "bump"), "close it");
```

`triageIssue(604800000, "Do we still want this?")` يجب أن تُعيد `"bump it"`.

```js
assert.equal(triageIssue(604800000, "Do we still want this?"), "bump it");
```

`triageIssue(604800000, "Bumping this")` يجب أن تُعيد `"close it"`.

```js
assert.equal(triageIssue(604800000, "Bumping this"), "close it");
```

`triageIssue(345600000, "I'll make a PR")` يجب أن تُعيد `"leave it"`.

```js
assert.equal(triageIssue(345600000, "I'll make a PR"), "leave it");
```

# --seed--

## --seed-contents--

```js
function triageIssue(ms, message) {

  return ms;
}
```

# --solutions--

```js
function triageIssue(ms, message) {
  const sevenDays = 7 * 24 * 60 * 60 * 1000;

  if (ms < sevenDays) return "leave it";
  if (message.toLowerCase().includes("bump")) return "close it";
  return "bump it";
}
```
