---
id: 6a1d9f98e819ed70a0e994dc
title: "تحدٍّ 330: كلمات بأحرف صغيرة"
challengeType: 28
dashedName: challenge-330
---

# --description--

بالنظر إلى سلسلة نصية، أعد فقط الكلمات التي تحتوي على أحرف صغيرة بالكامل، بترتيبها الأصلي وبفاصل مسافة بين كل كلمة.

# --hints--

`getLowercaseWords("hello GOOD world")` يجب أن تُعيد `"hello world"`.

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` يجب أن تُعيد `"these are all lowercase"`.

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` يجب أن تُعيد `"less is more"`.

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` يجب أن تُعيد `"eat pizza every day"`.

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` يجب أن تُعيد `"the quick brown fox jumped over the lazy dog"`.

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
