---
id: 6a2037a68a0bc2aef0006000
title: "تحدٍّ 338: حاسبة عمر الحيوانات الأليفة"
challengeType: 28
dashedName: challenge-338
---

# --description--

عند إعطائك نوع حيوان أليف وعمره بالسنوات البشرية، أرجع العمر المكافئ له في سنوات الحيوان الأليف باستخدام جدول التحويل التالي:

| الحيوان الأليف | المضاعف |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

يجب أن تُرجع الدالة `petYears("dog", 5)` القيمة `35`.

```js
assert.equal(petYears("dog", 5), 35);
```

يجب أن تُرجع الدالة `petYears("cat", 9)` القيمة `54`.

```js
assert.equal(petYears("cat", 9), 54);
```

يجب أن تُرجع الدالة `petYears("rabbit", 3)` القيمة `24`.

```js
assert.equal(petYears("rabbit", 3), 24);
```

يجب أن تُرجع الدالة `petYears("hamster", 4)` القيمة `120`.

```js
assert.equal(petYears("hamster", 4), 120);
```

يجب أن تُرجع الدالة `petYears("guinea pig", 5)` القيمة `60`.

```js
assert.equal(petYears("guinea pig", 5), 60);
```

يجب أن تُرجع الدالة `petYears("goldfish", 2)` القيمة `12`.

```js
assert.equal(petYears("goldfish", 2), 12);
```

يجب أن تُرجع الدالة `petYears("bird", 1)` القيمة `5`.

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
