---
id: 6a2037a68a0bc2aef0006002
title: "تحدٍّ 340: محول Pig Latin"
challengeType: 28
dashedName: challenge-340
---

# --description--

معطاة سلسلة نصية، حوّلها إلى Pig Latin باستخدام القواعد التالية:

- إذا بدأت الكلمة بحرف متحرك (`"a"`، `"e"`، `"i"`، `"o"`، أو `"u"`)، أضف `"way"` في النهاية. على سبيل المثال، تتحول `"universe"` إلى `"universeway"`.
- إذا بدأت الكلمة بحرف أو أكثر من الحروف الساكنة، انقلها إلى النهاية وأضف `"ay"`. على سبيل المثال، تتحول `"hello"` إلى `"ellohay"`.
- حافظ على حالة الحرف الأول. على سبيل المثال، تتحول `"Hello"` إلى `"Ellohay"`.

# --hints--

`pigLatin("universe")` يجب أن تُرجع `"universeway"`.

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` يجب أن تُرجع `"ellohay"`.

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` يجب أن تُرجع `"ellohay universeway"`.

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` يجب أن تُرجع `"Ellohay universeway"`.

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` يجب أن تُرجع `"Igpay Atinlay isway unfay"`.

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` يجب أن تُرجع `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
