---
id: 6a2037a68a0bc2aef0006004
title: "تحدٍّ 342: احتمالات النرد"
challengeType: 28
dashedName: challenge-342
---

# --description--

معطى عدد من نرد سداسي الأوجه للرمي ومجموع مستهدف، أرجع احتمالات رمي ذلك المجموع كسلسلة نصية بالتنسيق `"1 in X"`.

- عدد النرد سيكون بين 1 و6.
- المجموع المستهدف دائمًا قابل للتحقيق بعدد النرد المعطى.
- قرب `"X"` لأقرب عدد صحيح.

# --hints--

`getOdds(1, 5)` يجب أن تُرجع `"1 in 6"`.

```js
assert.equal(getOdds(1, 5), "1 in 6");
```

`getOdds(2, 4)` يجب أن تُرجع `"1 in 12"`.

```js
assert.equal(getOdds(2, 4), "1 in 12");
```

`getOdds(3, 10)` يجب أن تُرجع `"1 in 8"`.

```js
assert.equal(getOdds(3, 10), "1 in 8");
```

`getOdds(4, 7)` يجب أن تُرجع `"1 in 65"`.

```js
assert.equal(getOdds(4, 7), "1 in 65");
```

`getOdds(5, 26)` يجب أن تُرجع `"1 in 111"`.

```js
assert.equal(getOdds(5, 26), "1 in 111");
```

`getOdds(6, 35)` يجب أن تُرجع `"1 in 7776"`.

```js
assert.equal(getOdds(6, 35), "1 in 7776");
```

# --seed--

## --seed-contents--

```js
function getOdds(dice, target) {

  return dice;
}
```

# --solutions--

```js
function getOdds(dice, target) {
  let validCombos = 0;
  const totalCombos = Math.pow(6, dice);

  function roll(diceLeft, sum) {
    if (diceLeft === 0) {
      if (sum === target) validCombos++;
      return;
    }
    for (let face = 1; face <= 6; face++) {
      roll(diceLeft - 1, sum + face);
    }
  }

  roll(dice, 0);
  return `1 in ${Math.round(totalCombos / validCombos)}`;
}
```
