---
id: 6a2037a68a0bc2aef0006003
title: "تحدٍّ 341: العد التنازلي لعيد الميلاد"
challengeType: 28
dashedName: challenge-341
---

# --description--

بما أن لديك تاريخ اليوم وتاريخ عيد ميلاد، أرجع عدد الأيام حتى عيد ميلاد الشخص التالي.

- يُعطى تاريخ اليوم كسلسلة نصية بصيغة `"YYYY-MM-DD"` مع أصفار بادئة، مثل: `"2026-07-16"`.
- يُعطى عيد الميلاد كسلسلة نصية بصيغة `"M/D"` بدون أصفار بادئة، مثل: `"9/7"`.
- إذا كان اليوم هو عيد ميلادهم، أرجع عدد الأيام حتى عيد ميلادهم التالي (ليس `0`).
- يجب احتساب سنوات الكبيسة.

# --hints--

`daysUntilBirthday("2026-07-16", "9/7")` يجب أن تُرجع `53`.

```js
assert.equal(daysUntilBirthday("2026-07-16", "9/7"), 53);
```

`daysUntilBirthday("2026-07-16", "3/22")` يجب أن تُرجع `249`.

```js
assert.equal(daysUntilBirthday("2026-07-16", "3/22"), 249);
```

`daysUntilBirthday("2026-07-16", "7/16")` يجب أن تُرجع `365`.

```js
assert.equal(daysUntilBirthday("2026-07-16", "7/16"), 365);
```

`daysUntilBirthday("2024-02-28", "3/1")` يجب أن تُرجع `2`.

```js
assert.equal(daysUntilBirthday("2024-02-28", "3/1"), 2);
```

`daysUntilBirthday("2023-04-24", "12/30")` يجب أن تُرجع `250`.

```js
assert.equal(daysUntilBirthday("2023-04-24", "12/30"), 250);
```

`daysUntilBirthday("2024-03-01", "2/29")` يجب أن تُرجع `1460`.

```js
assert.equal(daysUntilBirthday("2024-03-01", "2/29"), 1460);
```

`daysUntilBirthday("2096-03-01", "2/29")` يجب أن تُرجع `2920`.

```js
assert.equal(daysUntilBirthday("2096-03-01", "2/29"), 2920);
```

# --seed--

## --seed-contents--

```js
function daysUntilBirthday(today, birthday) {

  return today;
}
```

# --solutions--

```js
function daysUntilBirthday(today, birthday) {
  const [year, month, day] = today.split("-").map(Number);
  const [bMonth, bDay] = birthday.split("/").map(Number);

  const todayDate = new Date(year, month - 1, day);

  function isLeapYear(y) {
    return (y % 4 === 0 && y % 100 !== 0) || y % 400 === 0;
  }

  function nextValidBirthday(fromYear) {
    for (let y = fromYear; y <= fromYear + 8; y++) {
      if (bMonth === 2 && bDay === 29 && !isLeapYear(y)) continue;
      const date = new Date(y, bMonth - 1, bDay);
      if (date > todayDate) return date;
    }
  }

  const nextBirthday = nextValidBirthday(year);
  const diff = nextBirthday - todayDate;
  return Math.round(diff / (1000 * 60 * 60 * 24));
}
```
