---
id: 6a2037a68a0bc2aef0006006
title: "تحدٍّ 344: النسبة الذهبية"
challengeType: 28
dashedName: challenge-344
---

# --description--

معطى رقمان، حدّد إذا كانت نسبتُهما تقارب النسبة الذهبية.

- استخدم النسبة الذهبية `1.618`
- اقبل هامش خطأ `0.01`

# --hints--

يجب أن تُرجع الدالة `isGoldenRatio(21, 34)` القيمة `true`.

```js
assert.isTrue(isGoldenRatio(21, 34));
```

يجب أن تُرجع الدالة `isGoldenRatio(15, 20)` القيمة `false`.

```js
assert.isFalse(isGoldenRatio(15, 20));
```

يجب أن تُرجع الدالة `isGoldenRatio(8, 13)` القيمة `true`.

```js
assert.isTrue(isGoldenRatio(8, 13));
```

يجب أن تُرجع الدالة `isGoldenRatio(10, 16)` القيمة `false`.

```js
assert.isFalse(isGoldenRatio(10, 16));
```

يجب أن تُرجع الدالة `isGoldenRatio(1618, 1000)` القيمة `true`.

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

يجب أن تُرجع الدالة `isGoldenRatio(88, 55)` القيمة `false`.

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
