---
id: 6a2037a68a0bc2aef0006007
title: "تحدٍّ 345: مزج الكلمات"
challengeType: 28
dashedName: challenge-345
---

# --description--

بالنظر إلى كلمتين، أرجع كلمة جديدة بدمج النصف الأول من الكلمة الأولى مع النصف الثاني من الكلمة الثانية.

- بالنسبة للكلمات ذات الطول الفردي، يكون النصف الأول هو النصف الأقصر.

# --hints--

`blendWords("turtle", "toucan")` يجب أن تُرجع `"turcan"`.

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")` يجب أن تُرجع `"chipingo"`.

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")` يجب أن تُرجع `"falican"`.

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")` يجب أن تُرجع `"hyana"`.

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")` يجب أن تُرجع `"scorilla"`.

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")` يجب أن تُرجع `"platerine"`.

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
