---
id: 6a2037a68a0bc2aef0005fff
title: "تحدٍّ 337: عداد العلامات"
challengeType: 28
dashedName: challenge-337
---

# --description--

بالنظر إلى سلسلة نصية من علامات العد، أعد العدد الإجمالي المُمثل.

- كل أنبوب `"|"` يُمثل عدداً واحداً.
- كل علامة خامسة تُمثل بشرطة مائلة للأمام `"/"`، لتُكمل مجموعة من خمسة (`"||||/"`).
- تُفصل المجموعات بمسافة.

# --hints--

`getTallyCount("||||")` يجب أن تُعيد `4`.

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` يجب أن تُعيد `5`.

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` يجب أن تُعيد `8`.

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` يجب أن تُعيد `17`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` يجب أن تُعيد `41`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
