---
id: 6a22d77ddf034bc4e35b1d58
title: "تحدٍّ 351: العدد الزوجي المتتابع"
challengeType: 28
dashedName: challenge-351
---

# --description--

أعطِ رقمًا، حدّد ما إذا كان عددًا زوجيًا متتابعًا.

العدد الزوجي المتتابع هو حاصل ضرب عددين صحيحين متتاليين. على سبيل المثال، 6 عدد زوجي متتابع لأن 2 * 3 = 6.

# --hints--

`isPronic(6)` يجب أن تُرجع `true`.

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` يجب أن تُرجع `false`.

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` يجب أن تُرجع `true`.

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` يجب أن تُرجع `true`.

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` يجب أن تُرجع `false`.

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` يجب أن تُرجع `true`.

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
