---
id: 6a22d77ddf034bc4e35b1d54
title: "تحدٍّ 347: نظرية الألعاب"
challengeType: 28
dashedName: challenge-347
---

# --description--

معطيتان سلسلتا نص متساويتان في الطول تمثلان استراتيجيات لاعبين في لعبة، أرجع النقاط كمصفوفة `[player1, player2]`.

- السلسلتان المعطيتان تحتويان فقط على أحد الحرفين: `"C"` (تعاون) أو `"D"` (خيانة).
- كل حرف يمثل جولة واحدة، تُحتسب النقاط كما يلي:
  - إذا تعاون اللاعبان، يحصل كل منهما على 3 نقاط.
  - إذا خانهما اللاعبان، يحصل كل منهما على نقطة واحدة.
  - إذا خانه لاعب وتعاون الآخر، يحصل الخائن على 5 نقاط والمتعاون على 0.

# --hints--

`playGame("CCCC", "CCCC")` يجب أن تُرجع `[12, 12]`.

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` يجب أن تُرجع `[4, 4]`.

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` يجب أن تُرجع `[5, 10]`.

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` يجب أن تُرجع `[24, 34]`.

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` يجب أن تُرجع `[66, 21]`.

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
