---
id: 6a22d77ddf034bc4e35b1d56
title: "تحدٍّ 349: إشارة الخلية"
challengeType: 28
dashedName: challenge-349
---

# --description--

معطى شبكة تحتوي على قراءات من ثلاثة أبراج خلوية، حدّد موقع الهاتف.

- كل خلية في الشبكة إما `0` (لا يوجد برج) أو عدد صحيح موجب يمثل عدد الخلايا إلى الهاتف، مقاسًا في خط مستقيم: أفقي، عمودي، أو قطري.
- أرجع `[row, col]` للخلية التي تبعد العدد الصحيح من الخلايا عن جميع الأبراج الثلاثة.
- يوجد دائمًا حل واحد فقط.

# --hints--

يجب أن تُرجع `findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` القيمة `[1, 2]`.

```js
assert.deepEqual(findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2]);
```

يجب أن تُرجع `findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` القيمة `[2, 1]`.

```js
assert.deepEqual(findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1]);
```

يجب أن تُرجع `findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` القيمة `[2, 2]`.

```js
assert.deepEqual(findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2]);
```

يجب أن تُرجع `findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` القيمة `[3, 4]`.

```js
assert.deepEqual(findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4]);
```

يجب أن تُرجع `findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` القيمة `[3, 3]`.

```js
assert.deepEqual(findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3]);
```

# --seed--

## --seed-contents--

```js
function findSignal(grid) {

  return grid;
}
```

# --solutions--

```js
function findSignal(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const towers = [];
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) towers.push([r, c, grid[r][c]]);
    }
  }

  function isValid(r, c, tr, tc, dist) {
    const dr = Math.abs(r - tr);
    const dc = Math.abs(c - tc);
    return (dr === dist || dc === dist) && (dr === 0 || dc === 0 || dr === dc);
  }

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) continue;
      if (towers.every(([tr, tc, dist]) => isValid(r, c, tr, tc, dist))) {
        return [r, c];
      }
    }
  }
}
```
