---
id: 6a22d77ddf034bc4e35b1d5a
title: "تحدٍّ 353: تقييم التباين 2"
challengeType: 28
dashedName: challenge-353
---

# --description--

بالنظر إلى قيمتين للسطوع النسبي وقيمة منطقية تشير إلى ما إذا كان النص كبيرًا، أرجع تقييم التباين وفقًا لمعايير WCAG باستخدام الطريقة التالية:

احسب نسبة التباين بإضافة 0.05 إلى كل قيمة سطوع، ثم اقسم القيمة الأعلى على القيمة الأدنى. ستكون القيمة الأعلى دائمًا هي المعلمة الأولى.

أرجع التقييم بناءً على نسبة التباين باستخدام الجدول التالي:

| التقييم | النص العادي | النص الكبير |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | أقل من 4.5 | أقل من 3.0 |

# --hints--

`getContrastRating(1.0, 0.0, false)` يجب أن تُرجع `"AAA"`.

```js
assert.equal(getContrastRating(1.0, 0.0, false), "AAA");
```

`getContrastRating(0.9015, 0.1364, false)` يجب أن تُرجع `"AA"`.

```js
assert.equal(getContrastRating(0.9015, 0.1364, false), "AA");
```

`getContrastRating(0.8965, 0.1628, false)` يجب أن تُرجع `"Fail"`.

```js
assert.equal(getContrastRating(0.8965, 0.1628, false), "Fail");
```

`getContrastRating(0.7469, 0.0957, true)` يجب أن تُرجع `"AAA"`.

```js
assert.equal(getContrastRating(0.7469, 0.0957, true), "AAA");
```

`getContrastRating(0.7489, 0.2018, true)` يجب أن تُرجع `"AA"`.

```js
assert.equal(getContrastRating(0.7489, 0.2018, true), "AA");
```

`getContrastRating(0.6571, 0.1974, true)` يجب أن تُرجع `"Fail"`.

```js
assert.equal(getContrastRating(0.6571, 0.1974, true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(l1, l2, isLargeText) {

  return l1;
}
```

# --solutions--

```js
function getContrastRating(l1, l2, isLargeText) {
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
