---
id: 6a22d77ddf034bc4e35b1d5b
title: "تحدٍّ 354: تقييم نسبة التباين 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

معطى مصفوفتان تمثلان قيم RGB وقيمة منطقية تشير إلى ما إذا كان النص كبيرًا، أرجع تقييم نسبة التباين وفقًا لمعايير WCAG باستخدام الطريقة التالية:

أولًا، حوّل كل قيمة RGB إلى الإضاءة النسبية:

- اقسم كل قناة `[R, G, B]` على 255 لتحصل على قيمة بين `0` و `1`
- طبق صيغة تصحيح جاما على كل قناة:
  - إذا كانت قيمة القناة أقل من أو تساوي `0.04045`: `channel / 12.92`
  - وإلا: `((channel + 0.055) / 1.055) ^ 2.4`
- احسب الإضاءة: `0.2126 * R + 0.7152 * G + 0.0722 * B`

بعدها، احسب نسبة التباين بإضافة 0.05 إلى كل قيمة إضاءة، ثم اقسم القيمة الأعلى على الأقل. القيمة الأعلى ستكون دائمًا المعلمة الأولى.

أرجع التقييم بناءً على نسبة التباين باستخدام الجدول التالي:

| التقييم | النص العادي | النص الكبير |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | أقل من 4.5 | أقل من 3.0 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)` يجب أن تُرجع `"AAA"`.

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)` يجب أن تُرجع `"AA"`.

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)` يجب أن تُرجع `"Fail"`.

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)` يجب أن تُرجع `"AAA"`.

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)` يجب أن تُرجع `"AA"`.

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)` يجب أن تُرجع `"Fail"`.

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
