---
id: 6a22d77ddf034bc4e35b1d57
title: "تحدٍّ 350: مسافة الحروف"
challengeType: 28
dashedName: challenge-350
---

# --description--

بالنظر إلى سلسلتين نصيتين متساويتين في الطول، أرجع مجموع أقصر المسافات بين كل زوج من الحروف.

- المُدخَل سيحتوي فقط على أحرف صغيرة
- يُعامل الأبجدية كدائرة، لذا المسافة بين `a` و `z` هي 1.

# --hints--

`letterDistance("abc", "bcd")` يجب أن تُرجع `3`.

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` يجب أن تُرجع `9`.

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` يجب أن تُرجع `10`.

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` يجب أن تُرجع `43`.

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` يجب أن تُرجع `47`.

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` يجب أن تُرجع `55`.

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
