---
id: 6a26df3e2988bcdded204893
title: "تحدٍّ 357: السلسلة الغذائية"
challengeType: 28
dashedName: challenge-357
---

# --description--

معطى مصفوفة من أزواج `[predator, prey]`، أرجع السلسلة الغذائية من المفترس الأعلى حتى القاع.

- المفترس الأعلى هو الحيوان الذي لا يكون فريسة لحيوان آخر.
- أرجع السلسلة كمصفوفة من سلاسل نصية.

# --hints--

`getFoodChain([["cat", "mouse"]])` يجب أن تُرجع `["cat", "mouse"]`.

```js
assert.deepEqual(getFoodChain([["cat", "mouse"]]), ["cat", "mouse"]);
```

`getFoodChain([["wolf", "deer"], ["deer", "grass"]])` يجب أن تُرجع `["wolf", "deer", "grass"]`.

```js
assert.deepEqual(getFoodChain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"]);
```

`getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` يجب أن تُرجع `["hawk", "snake", "frog", "fly"]`.

```js
assert.deepEqual(getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"]);
```

`getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` يجب أن تُرجع `["eagle", "fox", "rabbit", "grass"]`.

```js
assert.deepEqual(getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"]);
```

`getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` يجب أن تُرجع `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`.

```js
assert.deepEqual(getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"]);
```

# --seed--

## --seed-contents--

```js
function getFoodChain(pairs) {

  return pairs;
}
```

# --solutions--

```js
function getFoodChain(pairs) {
  const prey = new Set(pairs.map(([, p]) => p));
  const map = Object.fromEntries(pairs);
  const apex = pairs.map(([p]) => p).find(p => !prey.has(p));
  const chain = [apex];
  while (map[chain.at(-1)]) chain.push(map[chain.at(-1)]);
  return chain;
}
```
