---
id: 6a26df95efa55a2524399743
title: "تحدٍّ 362: مدقق النونوغرام"
challengeType: 28
dashedName: challenge-362
---

# --description--

معطى مصفوفة أرقام التلميحات ومصفوفة خلايا، حدّد ما إذا كانت الخلايا تحقق تلميح النونوغرام.

- التلميح هو مصفوفة أرقام تمثل أطوال الخلايا المتتالية المملوءة، بالترتيب. على سبيل المثال، التلميح `[3, 2]` يعني أن هناك 3 خلايا متتالية مملوءة تليها 2 خلايا متتالية مملوءة، مفصولة بخلية فارغة واحدة على الأقل.
- الصف هو مصفوفة من 1 (مملوءة) و0 (فارغة).

# --hints--

`isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1])` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1]));
```

`isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1])` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1]));
```

`isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]));
```

`isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])` يجب أن تُرجع `true`.

```js
assert.isTrue(isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])` يجب أن تُرجع `false`.

```js
assert.isFalse(isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]));
```

# --seed--

## --seed-contents--

```js
function isValidNonogram(clue, cells) {

  return clue;
}
```

# --solutions--

```js
function isValidNonogram(clue, cells) {
  const runs = [];
  let count = 0;

  for (let i = 0; i <= cells.length; i++) {
    if (cells[i] === 1) {
      count++;
    } else if (count > 0) {
      runs.push(count);
      count = 0;
    }
  }

  if (runs.length !== clue.length) return false;
  return runs.every((run, i) => run === clue[i]);
}
```
