---
id: 6a26df95efa55a2524399745
title: "تحدٍّ 364: بين دلوين"
challengeType: 28
dashedName: challenge-364
---

# --description--

مع وجود دلوين من الطلاء، كل منهما يحتوي على لون RGB ومستوى امتلاء، أرجع اللون RGB المختلط كمصفوفة من ثلاثة أعداد صحيحة.

- كل دلو هو كائن (JavaScript) أو قاموس (Python) يحتوي على خاصية `color` (مصفوفة من ثلاثة أعداد صحيحة `[r, g, b]`) وخاصية `fullness` (من 0 إلى 100).
- اللون المختلط هو متوسط مرجح لكل قناة في اللونين بناءً على مستوى الامتلاء، مع تقريب كل قناة إلى أقرب عدد صحيح.

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })` يجب أن تُرجع `[125, 125, 125]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })` يجب أن تُرجع `[200, 200, 200]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })` يجب أن تُرجع `[100, 150, 200]`.

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })` يجب أن تُرجع `[114, 184, 170]`.

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })` يجب أن تُرجع `[68, 163, 123]`.

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
