---
id: 6a26df3e2988bcdded204894
title: "تحدٍّ 358: مترجم الرموز التعبيرية"
challengeType: 28
dashedName: challenge-358
---

# --description--

بالنظر إلى سلسلة نصية من الرموز التعبيرية، أعد العبارة باستخدام الجدول التالي:

| رمز تعبيري | كلمة |
|-------|------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

أعد الكلمات مفصولة بمسافات.

# --hints--

`getEmojiPhrase("🪨⭐")` يجب أن تُعيد `"rock star"`.

```js
assert.equal(getEmojiPhrase("🪨⭐"), "rock star");
```

`getEmojiPhrase("🥵🐕")` يجب أن تُعيد `"hot dog"`.

```js
assert.equal(getEmojiPhrase("🥵🐕"), "hot dog");
```

`getEmojiPhrase("👶🦈")` يجب أن تُعيد `"baby shark"`.

```js
assert.equal(getEmojiPhrase("👶🦈"), "baby shark");
```

`getEmojiPhrase("⭐🐟")` يجب أن تُعيد `"star fish"`.

```js
assert.equal(getEmojiPhrase("⭐🐟"), "star fish");
```

`getEmojiPhrase("🧊🧊👶")` يجب أن تُعيد `"ice ice baby"`.

```js
assert.equal(getEmojiPhrase("🧊🧊👶"), "ice ice baby");
```

`getEmojiPhrase("🐱🐟🍲")` يجب أن تُعيد `"cat fish soup"`.

```js
assert.equal(getEmojiPhrase("🐱🐟🍲"), "cat fish soup");
```

# --seed--

## --seed-contents--

```js
function getEmojiPhrase(str) {

  return str;
}
```

# --solutions--

```js
function getEmojiPhrase(str) {
  const table = {
    '👶': 'baby',
    '🐱': 'cat',
    '🐕': 'dog',
    '🐟': 'fish',
    '🥵': 'hot',
    '🧊': 'ice',
    '🪨': 'rock',
    '🦈': 'shark',
    '🍲': 'soup',
    '⭐': 'star',
  };

  return [...str].map(emoji => table[emoji]).join(' ');
}
```
