---
id: 6a26df3e2988bcdded204897
title: "تحدٍّ 361: الوقت المنطوق"
challengeType: 28
dashedName: challenge-361
---

# --description--

بالنظر إلى زوايا عقارب الساعة والدقائق في ساعة تناظرية بالدرجات (باتجاه عقارب الساعة من الرقم 12)، أرجع الوقت باللغة الإنجليزية المنطوقة.

حوّل زاوية عقرب الدقائق إلى دقائق (360° = 60 دقيقة)، ثم استخدم القواعد التالية:

| الدقائق | المنطوق |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (باستثناء 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (باستثناء 45) | `"X minutes to Z"` (حيث `X` هو 60 - الدقائق) |

حيث `Y` هو الساعة الحالية و`Z` هي الساعة التالية، كلاهما مستخلص من زاوية عقرب الساعة (360° = 12 ساعة).

ملاحظة: قد لا تقع زوايا العقارب بالضبط على رقم، فكر في تقريبها بطريقة ما.

# --hints--

`getSpokenTime(90, 0)` يجب أن تُرجع `"3 o'clock"`.

```js
assert.equal(getSpokenTime(90, 0), "3 o'clock");
```

`getSpokenTime(160, 120)` يجب أن تُرجع `"20 minutes past 5"`.

```js
assert.equal(getSpokenTime(160, 120), "20 minutes past 5");
```

`getSpokenTime(255, 180)` يجب أن تُرجع `"half past 8"`.

```js
assert.equal(getSpokenTime(255, 180), "half past 8");
```

`getSpokenTime(67.5, 92)` يجب أن تُرجع `"quarter past 2"`.

```js
assert.equal(getSpokenTime(67.5, 92), "quarter past 2");
```

`getSpokenTime(200, 240)` يجب أن تُرجع `"20 minutes to 7"`.

```js
assert.equal(getSpokenTime(200, 240), "20 minutes to 7");
```

`getSpokenTime(322.5, 273)` يجب أن تُرجع `"quarter to 11"`.

```js
assert.equal(getSpokenTime(322.5, 273), "quarter to 11");
```

`getSpokenTime(117.5, 335)` يجب أن تُرجع `"5 minutes to 4"`.

```js
assert.equal(getSpokenTime(117.5, 335), "5 minutes to 4");
```

# --seed--

## --seed-contents--

```js
function getSpokenTime(hourAngle, minuteAngle) {

  return hourAngle;
}
```

# --solutions--

```js
function getSpokenTime(hourAngle, minuteAngle) {
  const minutes = Math.floor(minuteAngle / 6);
  const hour = Math.floor(hourAngle / 30) || 12;
  const nextHour = (hour % 12) + 1;

  if (minutes === 0) return `${hour} o'clock`;
  if (minutes === 15) return `quarter past ${hour}`;
  if (minutes === 30) return `half past ${hour}`;
  if (minutes === 45) return `quarter to ${nextHour}`;
  if (minutes < 30) return `${minutes} minutes past ${hour}`;
  return `${60 - minutes} minutes to ${nextHour}`;
}
```
