---
id: 69b83e35f19ba26ba1fa517a
title: 构建交通信号灯序列器
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

在本实验中，你将模拟可配置的交通信号灯周期并记录异常。

你将使用描述交通信号灯阶段的配置对象。每个配置对象具有以下属性：

- `fault`：一个布尔标记，当为 `true` 时触发提前终止。
- `phases`：一个阶段对象的数组。

`phases` 中的每个阶段对象具有以下属性：

- `color`：表示灯光颜色的字串（`"green"`、`"yellow"` 或 `"red"`）。
- `duration`：表示阶段持续时间（秒）的正整数。

你可以参考提供的 `config1`、`config2`、`config3` 和 `config4` 对象，作为可能的配置对象示例。

**注意**：不要添加额外的 `console.log()` 语句，因为它们可能导致测试失败。

**目标**：完成以下用户故事并通过所有测试以完成实验。

# 用户故事

1. 你应该有一个名为 `runSequence` 的函数，带有两个参数：`config` 和 `cycles`，其中 `config` 表示配置对象，`cycles` 表示序列运行的最大次数。

1. 你应该使用 `for` 或 `while` 循环实现 `runSequence(config, cycles)`，以迭代给定周期数内的每个阶段。

1. `runSequence` 函数应当：
    - 如果 `config.phases` 为空，记录 `No phases found` 并立即返回。
    - 如果 `config.fault` 为 `true`，记录 `Faulted phase!` 并提前停止模拟。
    - 如果 `duration <= 0`，记录 `Invalid phase detected`。
    - 对于每个有效阶段，记录 `Switching to [color] for [duration] s`。将 `[color]` 和 `[duration]` 替换为阶段对象的对应属性。
    - 例如，`runSequence(config1, 1)` 应记录：

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. 你应该有一个名为 `generateTimeline` 的函数，带有两个参数：`config` 和 `cycles`。

1. `generateTimeline` 函数应当：
    - 在各周期中记录每个阶段结束后的累计经过时间到数组中，迭代时将每个阶段的 `duration` 累加到总时间。
    - 即使 `config.fault` 为 `true` 或 `duration <= 0`，也应处理所有故障和无效阶段而不做验证。
    - 返回累计时间戳数组。
    - 例如，`generateTimeline(config1, 1)` 应返回数组 `[5, 7, 11]`。

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

你不应修改 `config1`、`config2`、`config3` 和 `config4` 对象。

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

你应该有一个名为 `runSequence` 的函数。

```js
assert.isFunction(runSequence);
```

你的 `runSequence` 函数应接受两个参数，`config` 和 `cycles`。

```js
const regex = __helpers.functionRegex('runSequence', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regex);
```

`runSequence(config1, 1)` 应按顺序记录 `Switching to green for 5 s`、`Switching to yellow for 2 s`、`Switching to red for 4 s`。

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` 应按顺序记录 `Switching to green for 5 s`、`Switching to yellow for 2 s`、`Switching to red for 4 s`、`Switching to green for 5 s`、`Switching to yellow for 2 s`、`Switching to red for 4 s`。

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` 应按顺序记录 `Switching to red for 3 s`、`Invalid phase detected`、`Switching to green for 6 s`。

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` 应记录 `Faulted phase!`。

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` 应记录 `No phases found`。

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

你应该有一个名为 `generateTimeline` 的函数。

```js
assert.isFunction(generateTimeline);
```

你的 `generateTimeline` 函数应接受两个参数，`config` 和 `cycles`。

```js
const regEx = __helpers.functionRegex('generateTimeline', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regEx);
```

`generateTimeline(config1, 1)` 应返回 `[5, 7, 11]`。

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` 应返回 `[5, 7, 11, 16, 18, 22]`。

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` 应返回 `[3, 1, 7, 10, 8, 14]`。

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` 应返回 `[5, 7, 13]`。

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` 应返回 `[]`。

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
