---
id: 68e6aff5325410118e76e7be
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 접근 가능한 오디오 컨트롤러를 만들어 ARIA를 다루는 연습을 합니다.

먼저, 텍스트가 `h1`인 `Audio Controls` 요소를 추가하세요.

# --hints--

`h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 요소는 `body` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector('body > h1'));
```

`h1` 요소는 텍스트가 `Audio Controls`이어야 합니다.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
