---
id: 68eac57fadcc662ce917e75e
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

다음으로, `button` 요소를 추가하고 `type` 속성을 `button` 값으로 설정하세요.

`button` 요소의 텍스트는 `Play`여야 합니다.

# --hints--

`button` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('button'));
```

`button` 요소에는 `type` 속성이 있어야 합니다.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

`button` 요소에는 `type` 속성이 `button` 값으로 설정되어야 합니다.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

`button` 요소의 텍스트는 `Play`여야 합니다.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
