---
id: 68eb7f2397da251a40e04aa2
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

다음으로, 볼륨 조절을 위한 모든 요소를 포함할 `div` 요소를 추가하세요.

`div` 요소 안에 `span` 요소를 중첩하고, `id` 속성을 `volume-label`로 설정하세요.
또한, `span` 요소의 텍스트는 `Volume`이어야 합니다.

`id`은 나중에 ARIA 속성에서 참조되어 보조 기술이 이 텍스트를 볼륨 조절 레이블로 인식할 수 있게 합니다.

# --hints--

`div` 요소 아래에 `button` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('button + div'));
```

`span` 요소 안에 `div` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('button + div > span'));
```

`span` 요소에 `id` 속성이 있어야 합니다.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

`span` 요소에 `id` 속성이 있고, 그 값은 `volume-label`이어야 합니다.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

`span` 요소의 텍스트는 `Volume`이어야 합니다.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
