---
id: 68ec0a48c9d6db3594d768c3
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

이제 또 다른 `span` 요소를 추가할 차례입니다.

이 `span` 요소는 `id` 속성이 `volume-description`로 설정되어야 합니다.
또한, 이 `span` 요소의 텍스트는 `Adjust the sound level`여야 합니다.

이전 `span` 요소와 마찬가지로, 이 `id`의 `span`는 나중에 ARIA 속성으로 참조되어 보조 기술이 이 텍스트를 볼륨 조절기의 설명으로 인식할 수 있게 합니다.

# --hints--

`span` 요소 안에 `div` 요소가 두 개 있어야 합니다.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

두 번째 `span` 요소에는 `id` 속성이 있어야 합니다.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

두 번째 `span` 요소에는 `id` 속성이 `volume-description` 값으로 있어야 합니다.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

두 번째 `span` 요소에는 `Adjust the sound level` 텍스트가 있어야 합니다.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
