---
id: 68ed4204ab7edc66b6cb2383
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

이전 수업에서 기억하셨겠지만, `aria-labelledby` 속성은 페이지에 이미 있는 텍스트를 레이블로 사용할 때 쓰입니다. 이번에는 `id` 속성을 `span` 요소에 추가했으며, 이 속성이 여기서 사용될 것입니다.

`aria-labelledby` 요소에 `input` 속성을 추가하고 값을 `volume-label volume-description`로 설정하세요.

# --hints--

`input` 요소에 `aria-labelledby` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

`input` 요소에 `aria-labelledby` 속성이 있어야 하며, 값은 `volume-label volume-description`입니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
