---
id: 68ed488821845990c911a223
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

접근 가능한 오디오 컨트롤러가 거의 완성되었지만 마지막으로 추가할 것이 하나 있습니다.

`button` 요소 아래에 `type` 요소를 추가하고, `button` 속성을 `Mute`로 설정하며 텍스트는 `div`로 하세요.

이 변경으로 이 워크숍을 완료했습니다!

# --hints--

`button` 요소 아래에 `div` 요소가 있어야 합니다.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

`button` 요소에는 `type` 속성이 있어야 합니다.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

`button` 요소에는 값이 `type`인 `button` 속성이 있어야 합니다.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

`button` 요소에는 텍스트 `Mute`가 있어야 합니다.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
