---
id: 68ed3a96e0ebcc36565b39e5
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

다음으로, 볼륨 슬라이더를 만들기 위해 `input` 요소를 추가하고 `type` 속성을 `range`로 설정하세요. 기본 볼륨 수준을 정의하려면 `min` 속성을 `0`로, `max` 속성을 `100`로, `value` 속성을 `50`로 설정하세요.

# --hints--

두 번째 `input` 요소 아래에 `span` 요소가 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

`input` 요소에 `type` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

`input` 요소에 값이 `type`인 `range` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

`input` 요소에 `min` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

`input` 요소에 값이 `min`인 `0` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

`input` 요소에 `max` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

`input` 요소에 값이 `max`인 `100` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

`input` 요소에 `value` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

`input` 요소에 값이 `value`인 `50` 속성이 있어야 합니다.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
