---
id: a5229172f011153519423690
title: Побудувати калькулятор суми непарних чисел Фібоначчі
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

У цій лабораторній роботі ви створите калькулятор суми непарних чисел Фібоначчі, який обчислює суму всіх непарних чисел Фібоначчі, що менші або дорівнюють заданому додатному числу.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача**

1. Ви повинні мати `sumFibs` функцію, яка приймає число як аргумент.
1. `sumFibs` функція повинна повертати суму всіх непарних чисел Фібоначчі, що менші або дорівнюють заданому числу.
1. Послідовність Фібоначчі починається з `0` та `1`, а кожне наступне число є сумою двох попередніх.
1. До суми слід додавати лише непарні числа Фібоначчі.

# --hints--

Ви повинні мати `sumFibs` функцію.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` має повертати число.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` має повертати `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` має повертати `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` має повертати `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` має повертати `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` має повертати `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
