---
id: 6720b7f3155a9278d6b6a8b8
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Ahora es momento de estilizar la tarjeta de felicitación.

Añade un selector para el elemento `body`, luego:

- cambia la `fuente` a `Arial` seguida de la genérica `sans-serif`,
- dale un relleno en todos los lados de `40px`,
- ajusta la propiedad `alineación de texto` a `centro`.

# --hints--

Debe tener un selector `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Debe asignar `font-family` a `"Arial", sans-serif` en el selector `body`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

Debe asignar `padding` a `40px` en el selector `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

Debe asignar `text-align` a `centro` en el selector `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
    </div>
  </body>
</html>

```

```css

--fcc-editable-region--


--fcc-editable-region--

```
