---
id: 6728fda085006772a247e779
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Asigne a `.card` un `max-width` de `400px`, un `padding` de `40px` en todos los lados, y un `margin` de `0` para arriba y abajo y `auto` para izquierda y derecha (utilice la propiedad abreviada).

# --hints--

`.card` debe tener un `max-width` de `400px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("max-width"), "400px");
```

`.card` debe tener un `padding` de `40px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("padding"), "40px");
```

`.card` debe tener un `margin` de `0 auto`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("margin"), ["0 auto", "0 auto 0 auto", "0 auto 0", "0px auto", "0px auto 0px auto", "0px auto 0px"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
    </div>
  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

--fcc-editable-region--

.card {
  background-color: white;

}

--fcc-editable-region--

```
