---
id: 687ad7be1ff45032ccf1d92f
title: Construir um Alternador de Tema
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

Neste laboratório, você construirá um aplicativo que alterna entre diferentes temas.Quando você alterna para um tema diferente, a cor de fundo mudará e uma mensagem diferente será exibida.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um elemento `button` com o texto `Alternar Tema`.
2. Seu elemento `button` deve ter os seguintes atributos:
   - Um atributo `id` definido como `"theme-switcher-button"`.
   - Um atributo `aria-haspopup` definido como `"true"`.
   - Um atributo `aria-expanded` definido como `"false"`.
   - Um atributo `aria-controls` definido como `"theme-dropdown"`.
3. Você deve ter um elemento `ul` com os seguintes atributos:
   - Um atributo `id` definido como `"theme-dropdown"`.
   - Um atributo `role` definido como `"menu"`.
   - Um atributo `aria-labelledby` definido como `"theme-switcher-button"`.
   - Um atributo `hidden`.
4. Seu elemento `ul` deve ter pelo menos dois elementos `li` com um atributo `role` definido como `"menuitem"` e um texto de sua escolha representando um tema diferente.
5. Cada um dos seus elementos `li` deve ter um atributo `id` que comece com `theme-` e termine com o tema que você definiu para o texto do elemento `li`. Por exemplo, se um dos seus temas for `Light`, então seu `id` deve ser `theme-light`.
6. Você deve ter um elemento com um atributo `aria-live` definido como `"polite"`.
7. Você deve ter um array `themes` com pelo menos dois objetos que contenham cada um uma propriedade `name` e `message`. O `name` representará um tema diferente e a `message` será exibida quando o tema for alterado. Você pode criar valores para `name` e `message` de sua escolha, mas os valores de `name` devem corresponder a uma das opções que você definiu em seus itens `li`. 
8. Quando um usuário clica no botão `#theme-switcher-button`, o elemento `#theme-dropdown` deve exibir os temas disponíveis.
9. Quando o elemento `#theme-dropdown` for exibido, o atributo `aria-expanded` deve ser definido como `"true"` para o elemento `button`.
10. Quando o usuário clicar no `#theme-switcher-button` enquanto o elemento `#theme-dropdown` estiver exibido, o menu deve ser fechado, o atributo `hidden` deve ser adicionado e o atributo `aria-expanded` deve ser definido como `"false"`.
11. Quando um usuário clica no `#theme-switcher-button` e seleciona um tema, a classe correspondente `theme-<name>` deve ser adicionada ao elemento `body` e a mensagem de tema relacionada que você definiu no array `themes` deve ser exibida no elemento com `aria-live="polite"`. 

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

Você deve ter um elemento `button`.

```js
assert.exists(document.querySelector("button"));
```

Seu elemento `button` deve conter o texto `Alternar Tema`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

Seu elemento `button` deve ter um atributo `id` definido como `"theme-switcher-button"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

Seu elemento `button` deve ter um atributo `aria-haspopup` definido como `"true"`. 

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

Seu elemento `button` deve ter um atributo `aria-expanded` definido como `"false"`. 

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

Seu elemento `button` deve ter um atributo `aria-controls` definido como `"theme-dropdown"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

Você deve ter um elemento `ul`.

```js
assert.exists(document.querySelector("ul"));
```

Seu elemento `ul` deve ter um atributo `id` definido como `"theme-dropdown"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

Seu elemento `ul` deve ter um atributo `role` definido como `"menu"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

Seu elemento `ul` deve ter um atributo `aria-labelledby` definido como `"theme-switcher-button"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

Seu elemento `ul` deve ter um atributo `hidden`.

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

Seu elemento `ul` deve conter pelo menos dois elementos `li`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

Cada um dos seus elementos `li` deve ter um atributo `role` definido como `"menuitem"`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

Cada um dos seus elementos `li` deve conter um texto que represente um tema.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

Cada um dos seus elementos `li` deve ter um atributo `id` que comece com `theme-` e termine com o tema correspondente para esse elemento.Certifique-se de que o valor do seu `id` esteja todo em letras minúsculas.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

Você deve ter um elemento com um atributo `aria-live` definido como `"polite"`.

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

Você deve ter um array `themes`.

```js
assert.isArray(themes);
```

Seu array `themes` deve conter pelo menos dois objetos.

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

Seu array `themes` deve conter pelo menos dois objetos, cada um com uma propriedade `name` e `message`. 

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

Seus valores das propriedades `name` e `message` devem ser strings que não estejam vazias.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

Cada um dos valores da sua propriedade `name` deve corresponder a um dos valores que você forneceu para os elementos `li`.Certifique-se de que esses valores estejam todos em letras minúsculas.

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

Quando um usuário clica no `#theme-switcher-button` para exibir as opções de tema, o atributo `hidden` deve ser removido do elemento `#theme-dropdown`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

Quando o elemento `#theme-dropdown` for exibido, o atributo `aria-expanded` deve ser definido como `"true"` para o elemento `button`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

Quando o usuário clicar no `#theme-switcher-button` enquanto o elemento `#theme-dropdown` estiver exibido, o menu dropdown deve ser definido como `hidden`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

Quando o usuário clicar no `#theme-switcher-button` enquanto o elemento `#theme-dropdown` estiver exibido, o atributo `aria-expanded` deve ser definido como `"false"`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

Quando um usuário clica no `#theme-switcher-button` e seleciona um tema, a classe correspondente `theme-<name>` deve ser adicionada ao elemento `body`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1]; 
clickedItem?.dispatchEvent(new Event("click"), {bubbles: true});

const expectedClass = clickedItem?.id; 
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

Quando um usuário clica no `#theme-switcher-button` e seleciona um tema, uma `message` relacionada ao `theme` selecionado do array `themes` deve ser exibida no elemento com `aria-live="polite"`.


```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click"), {bubbles: true});

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
