---
id: 699b6d91cde15459d3619c4a
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

`profile` オブジェクトと同様に、`userRoles` オブジェクトもコンソールに表示できるといいですね。

`console.log()` オブジェクトをログに出力するために `userRoles` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`userRoles` をコンソールにログ出力する必要があります。

```js
assert.equal(getLogs()[1], userRoles);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

--fcc-editable-region--

--fcc-editable-region--
```
