---
id: 699b4e284291adcfcb90df47
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、プリミティブ型とオブジェクト型および配列型を使って、型安全なユーザープロフィールを作成する練習をします。

まず、`profile` という変数を作成し、オブジェクトを代入してください。そのオブジェクトには次の3つのプロパティと値が含まれている必要があります。

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

`profile` という変数があるはずです。

```js
assert.exists(profile);
```

`profile` 変数はオブジェクトである必要があります。

```js
assert.isObject(profile);
```

`username` オブジェクトに `profile` プロパティがあるはずです。

```js
assert.property(profile, "username");
```

`username` プロパティの値は `"codeLearner"` である必要があります。

```js
assert.propertyVal(profile, "username", "codeLearner");
```

`profile` オブジェクトには `age` プロパティがあるはずです。

```js
assert.property(profile, "age");
```

`age` プロパティの値は `25` である必要があります。

```js
assert.propertyVal(profile, "age", 25);
```

`profile` オブジェクトには `isLoggedIn` プロパティがあるはずです。

```js
assert.property(profile, "isLoggedIn");
```

`isLoggedIn` プロパティの値は `false` である必要があります。

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
