---
id: 699b61ca8c9cb9bc1968364e
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

コンソールを開くと、次のエラーメッセージが表示されているはずです。

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

バニラのJavaScriptでは、`profile`オブジェクトに好きなプロパティや値を自由に追加できました。しかしTypeScriptでは、コンパイラがオブジェクトに定義したプロパティ（この場合は`username`、`age`、`isLoggedIn`）だけがあるかどうかをチェックします。`mood`が型に含まれていなかったため、TypeScriptはそれを無効なプロパティとみなし、エラーをスローします。

コードから`mood: null`を削除してください。

# --hints--

コード内に`mood: null`がもう存在しないことを確認してください。

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
