---
id: 699b6dd1a2961bd38491ef2e
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

`profile` オブジェクトは厳密なキーと型を持っていました。定義されていないプロパティを追加しようとしたり、互換性のない型を代入しようとすると、TypeScript はエラーをスローしました。

`userRoles` オブジェクトがそれほど厳密でなく、より柔軟なキーを持っていたらいいですね。

前のレッスンで覚えているように、この結果は `Record` 型を使うことで実現できます。

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

この `studentGrades` の例では、TypeScript はすべての科目を最初から知っている必要はありません。これにより、後から `Art` や `Music` のような科目を追加できます。TypeScript が気にするのは、各キーの値が数値であることだけです。

インラインの `userRoles` 型を追加して、`Record` オブジェクトを更新してください。各キーが任意の文字列で、各値が文字列であるように `<string, string>` に設定してください。

# --hints--

`userRoles` オブジェクトには、インラインの `Record` 型が `Record<string, string>` に設定されている必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
