---
id: 699b70b10ae344c802d106b0
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

`userRoles` オブジェクトにさらにいくつかのプロパティを追加する時間です。

`moderator` プロパティを値 `"medium-access"` で追加してください。

その下に、`guest` プロパティを値 `3` で追加してください。

# --hints--

`userRoles` オブジェクトには `moderator` プロパティがあるはずです。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

`moderator` プロパティの値は `"medium-access"` であるべきです。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

`userRoles` オブジェクトには `guest` プロパティがあるはずです。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

`guest` プロパティの値は数値の `3` であるべきです。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
