---
id: 699b71d98c820ac7ab9eea2a
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

もう一度コンソールを開くと、新しいエラーメッセージが表示されます。

```md
Type 'number' is not assignable to type 'string'.
```

これは、`guest` プロパティに数値の値を追加しようとしているために起こっています。

TypeScript のエラーを解消するには、`guest` プロパティの値を `"read-only"` に変更してください。

この最後の変更で、ユーザープロフィールのワークショップは完了です！

# --hints--

`guest` オブジェクトの `userRoles` プロパティは `"read-only"` の値にする必要があります。

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
