---
id: 699b5370a43e299b4673bcaa
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

`profile` オブジェクトのプロパティを作成しながらコンソールで確認できると便利です。

`console.log` をコンソールにログ出力する `profile` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`profile` をコンソールにログ出力してください。

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
