---
id: 699b62c99ec8c63470313365
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

今度は、`profile` オブジェクトに別のプロパティを追加する時間です。

まず、インラインオブジェクト型を更新して、型が `bio` に設定された `string` プロパティを含めてください。

# --hints--

インラインオブジェクト型を更新して、`bio` プロパティを `string` に設定する必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
