---
id: 699b65caf17a8e8d340af27d
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

いよいよ `profile` オブジェクトに最後のプロパティを追加する時です。

これまでのレッスンでは、配列型の扱い方を学びました。`string[]` や `number[]` のようなさまざまな配列型についても学びました。

```ts
const scores: number[] = [95, 87, 100];
```

この例では、`scores` 配列は数字のみを受け入れます。`string` や `null` のような別の型を追加しようとすると、TypeScript はエラーをスローします。

`profile` オブジェクトを更新して、`programmingLanguages` プロパティを `string[]` 型で追加してください。オブジェクトリテラルの中で、このプロパティを `["JavaScript", "Python", "C++"]` に設定してください。

# --hints--

`programmingLanguages` オブジェクトの中に `profile` プロパティがあるはずです。

```js
assert.property(profile, "programmingLanguages");
```

`programmingLanguages` プロパティの値は `["JavaScript", "Python", "C++"]` であるべきです。

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

`programmingLanguages` プロパティは `string[]` 型であるべきです。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};
--fcc-editable-region--

console.log(profile);
```
