---
id: 699b68f94a0d154c2a7a5ce0
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

`profile` オブジェクトが完成したので、次はユーザーの役割を表す新しいオブジェクトを作成しましょう。

`userRoles` という名前の新しい変数を作成し、オブジェクトを代入してください。オブジェクトには以下のプロパティと値を含める必要があります。

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

`userRoles` という名前の変数があるはずです。

```js
assert.isDefined(userRoles);
```

`userRoles` 変数はオブジェクトである必要があります。

```js
assert.isObject(userRoles);
```

`userRoles` オブジェクトには `admin` プロパティがあるはずです。

```js
assert.property(userRoles, "admin");
```

`admin` プロパティの値は `"full-access"` である必要があります。

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

`userRoles` オブジェクトには `editor` プロパティがあるはずです。

```js
assert.property(userRoles, "editor");
```

`editor` プロパティの値は `"limited-access"` である必要があります。

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

`userRoles` オブジェクトには `viewer` プロパティがあるはずです。

```js
assert.property(userRoles, "viewer");
```

`viewer` プロパティの値は `"read-only"` である必要があります。

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
