---
id: 66ed41f912d0bb1dc62da5dd
title: Construye un Bot de Trivia en JavaScript
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debería registrar `"Hello! I'm your coding fun fact guide!"` en la consola como un mensaje de saludo para el usuario.
1. Deberías crear tres variables: `botName`, `botLocation` y `favoriteLanguage`, que almacenan el nombre del bot, de dónde es y su lenguaje de programación favorito, respectivamente.
1. Debes registrar `"My name is (botName) and I live on (botLocation)."` en la consola.
1. Debes registrar `"My favorite programming language is (favoriteLanguage)."` en la consola.
1. Debes usar `let` para crear una variable `codingFact` y asignarle una cadena que sea un dato curioso sobre el lenguaje de programación favorito de tu bot e incluir el uso de la variable `favoriteLanguage`.
1. Deberías registrar `codingFact` en la consola.
1. Deberías reasignar la variable `codingFact` a un nuevo dato sobre el lenguaje favorito del bot usando nuevamente la variable `favoriteLanguage`.
1. Deberías registrar `codingFact` en la consola nuevamente.
1. Deberías reasignar la variable `codingFact` nuevamente a otro nuevo dato sobre el lenguaje favorito del bot usando la variable `favoriteLanguage`.
1. Deberías registrar `codingFact` en la consola por tercera vez.
1. Debes registrar `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` en la consola como una declaración de despedida del bot.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debería registrar <code>"Hello! ¡Soy tu guía de datos curiosos de programación!"</code> en la consola.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

Deberías declarar una variable `botName`. Verifique nuevamente si hay errores ortográficos o de mayúsculas.

```js
assert.exists(botName);
```

Su variable `botName` debe ser una cadena.

```js
assert.isString(botName);
```

Deberías declarar una variable `botLocation`. Verifique nuevamente si hay errores ortográficos o de mayúsculas.

```js
assert.exists(botLocation);
```

Tu variable `botLocation` debería ser una cadena.

```js
assert.isString(botLocation);
```

Deberías declarar una variable `favoriteLanguage`. Verifique nuevamente si hay errores ortográficos o de mayúsculas.

```js
assert.exists(favoriteLanguage);
```

Tu variable `favoriteLanguage` debería ser una cadena.

```js
assert.isString(favoriteLanguage);
```

Debes registrar en la consola `"My name is (botName) and I live on (botLocation)."` y agregar las variables a la cadena.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

Debes registrar en la consola `"My favorite programming language is (favoriteLanguage)."` y agregar la variable a la cadena.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

Deberías usar `let` para declarar una nueva variable `codingFact`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

Debes darle a `codingFact` un valor que incluya `favoriteLanguage`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

Deberías registrar `codingFact` en la consola.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

Deberías asignar un nuevo valor a `codingFact` que también contenga `favoriteLanguage`, y registrarlo en la consola.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

Deberías asignar un valor a `codingFact` por tercera vez que también contenga `favoriteLanguage`, y registrarlo en la consola.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

Debes registrar en la consola `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` y agregar los valores de las variables.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
