---
id: 65c4f07da204f4d2f325fbd3
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

在 `__init__` 方法中，將節點的 `left` 和 `right` 屬性初始化為 `None`。這是因為當節點第一次被創建時，它沒有任何左子節點或右子節點。記得使用 `self` 關鍵字。 

# --hints--

你應該使用 `self.left` 將節點的 `left` 屬性指定為 `None`。

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

你應該使用 `self.right` 將節點的 `right` 屬性指定為 `None`。

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
