---
id: 65c4eb814cc977c95cd7df0e
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個專案中，你將要創建一個 Binary Search Tree（BST）。BST 是一種資料結構，其中每個節點最多有兩個子節點，左子節點包含小於父節點的值，右子節點包含大於父節點的值，從而允許高效的搜尋和排序操作。

這就是二元搜尋樹的樣子：

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

首先定義一個空的 `TreeNode` 類別。`TreeNode` 類別表述二元搜尋樹中的節點。使用 `pass` 關鍵字填充類別本體以避免錯誤。

# --hints--

你應該使用 `class` 關鍵字來宣告一個名為 `TreeNode` 的空類別。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
