---
id: 65c4ef8463f869d0eaf87c67
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

在 `TreeNode` 類別中，將 `pass` 取代為 `__init__` 方法，以便你可以初始化物件的屬性。暫時不要新增任何參數。

# --hints--

你應該從 `TreeNode` 類別中移除 `pass` 關鍵字，並將它移到 `__init__` 方法內。

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

你應該在你的 `TreeNode` 類別中定義一個 `__init__` 方法。記得在裡面使用 `pass` 關鍵字。

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
