---
id: 65c4f300da28d8da361bfa93
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

在你的 `elif` 語句中，將 `pass` 取代為呼叫 `_insert` 方法，並以目前節點的右子節點作為第一個引數、`key` 作為第二個引數。將結果指定給目前節點的 `right` 屬性。

# --hints--

你應該從 `elif` 區塊中移除 `pass` 關鍵字。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

你應該呼叫 `self._insert()` 方法，並傳入 `node.right` 和 `key` 作為引數。

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

你應該將 `self._insert()` 呼叫的結果指定到目前節點的右邊屬性（`node.right`）。

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
