---
id: 65c4f09e074dd8d37830ea00
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

創建另一個名為 `BinarySearchTree` 的空類別，用來表述二分搜尋樹。

# --hints--

你應該有一個名為 `BinarySearchTree` 的類別。記得使用 `pass` 關鍵字。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
