---
id: 65c4f3258d2e4cdacc919dfd
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

在你的 `_insert` 行為結束時，在插入處理程序完成後，傳回目前節點以更新遞迴呼叫堆疊較高層級的樹狀結構。


# --hints--

你應該在條件區塊外傳回目前的節點。

```js
({ test: () => assert.match(code, /^\s{8}return\s+node/m) })
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)

--fcc-editable-region--

```

