---
id: 65c4f37ba6cbcfdc77e24165
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

`insert` 方法將被用戶調用。 除了 `self` 參數之外，它還需要 `key` 參數。 此參數將是插入二叉搜索樹的關鍵值。

將 `key` 作爲第二個參數添加到函數定義中。

# --hints--

`insert` 方法應該包含 `self` 參數。

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

`insert` 方法應該包含 `key` 參數。

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
