---
id: 65c4f2b7178afed8e88f782f
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

如果 `key` < `node.key` 計算結果為 `True`，則新節點應該放置在左子樹中。

刪除 `pass` 並遞迴呼叫 `_insert` 方法，將左子節點作為第一個引數，`key` 作為第二個引數。將結果指定到目前節點的 `left` 屬性。

# --hints--

你應該從 `if` 區塊中移除 `pass` 關鍵字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

你應該呼叫 `self._insert` 方法，並傳入 `node.left` 和 `key` 作為引數。

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

你應該將 `self._insert()` 呼叫的結果指定到目前節點的 `left` 屬性。

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
