---
id: 65c4f33bc3c3d8db7f732787
title: 步驟 16
challengeType: 20
dashedName: step-16
---

# --description--

現在，爲了執行實際的插入，在 `BinarySearchTree` 類中定義一個空的 `insert` 方法併爲其提供 `self` 參數。

# --hints--

你應該在 `BinarySearchTree` 類中定義一個帶有 `self` 參數的 `insert` 方法。 記住 `pass` 關鍵字。

```js
({ test: () => assert.match(code, /^\s{4}def\s+insert\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node
--fcc-editable-region--    

--fcc-editable-region--

```
