---
id: 65c4f22498d22ed775ef8efb
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

現在你需要檢查 `node` 參數是否為 `None`。如果是，這表示該方法已經到達樹中的葉節點或應該插入新節點的空位。

在 `_insert` 方法主體內，將 `pass` 取代為檢查 `node is None` 的 `if` 述語。

在新的 `if` 區塊內，傳回 `TreeNode(key)` 以使用提供的 key 創建新的 `TreeNode` 實體。這將成為新的葉節點，有效地將 key 插入樹中。

# --hints--

你不應該在你的 `_insert` 方法中使用 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

你應該撰寫一個 `if` 述語來檢查 `node is None`。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

你應該從 `if` 區塊傳回 `TreeNode(key)`。

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
