---
id: 65c4f300da28d8da361bfa93
title: 步驟14
challengeType: 20
dashedName: step-14
---

# --description--

在 `elif` 子句中，將 `pass` 替換爲對 `_insert` 方法的調用，其中當前節點的右子節點作爲第一個參數，`key` 作爲第二個參數。 將結果賦值給當前節點的 `right` 屬性。

# --hints--

你應該從 `elif` 塊中刪除 `pass` 關鍵字。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

你應該調用 `self._insert()` 方法並傳遞 `node.right` 和 `key` 作爲參數。

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

你應該將 `self._insert()` 調用的結果分配給當前節點的正確屬性 (`node.right`)。

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
