---
id: 65c4f02dbd3557d23f12db62
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

在 `__init__` 方法內，刪除 `pass` 並使用 `self.key` 將 `key` 參數的值指定給節點的 `key` 屬性。

這表示 `TreeNode` 實體的 `key` 屬性將會被設定為在物件產生時傳入的值。

# --hints--

你應該從 `__init__` 方法中移除 `pass` 述語。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

你應該使用 `self.key` 將 `key` 參數的值指定到節點的 `key` 屬性。

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
