---
id: 65c4f2851645e8d84d14f5d1
title: 步驟11
challengeType: 20
dashedName: step-11
---

# --description--

現在你需要遞歸遍歷樹並使用二叉樹的原理插入值：

- 小於鍵的值放在左子樹中
- 大於鍵的值放在右子樹中

在現有的條件語句之後，編寫另一個 `if` 語句來檢查 `key` 是否小於 `node.key`。

# --hints--

你應該編寫另一個 if 語句來檢查 `key` 是否小於 `node.key`。 記得使用 `pass` 關鍵字。

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```

