---
id: 65c4f2d9fd872fd99ac659dd
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

為你的程式碼添加一個 `elif` 條件述語，用來檢查 `key` 是否大於 `node.key`。

# --hints--

你應該為你的程式碼添加一個 `elif` 條件述語來檢查是否 `key > node.key`。記得使用 `pass` 關鍵字。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[1].is_equivalent("key > node.key")`))})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)

--fcc-editable-region--    

```

