---
id: 65c4f0d32ca17ad4b1636b0e
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

在 `__init__` 方法中，刪除 `pass` 並將 `root` 初始化為值 `None`。 

`root` 屬性表述二元搜尋樹的根節點。由於這是當新的 `BinarySearchTree` 物件被建立時的建構式，它從空樹開始，因此 `root` 屬性被設為 `None`。

# --hints--

你應該從 `__init__` 方法中移除 `pass` 述語。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

你應該使用 `self.root` 將 `root` 屬性初始化為 `None`。

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
