---
id: 65c4f2851645e8d84d14f5d1
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

現在你需要遞迴地巡訪二元樹並使用二元樹的原則插入值：

- 比鍵值小的值會被放置在左子樹。
- 大於鍵的值會被放置在右子樹

在你現有的條件述語之後，撰寫另一個 `if` 述語以檢查 `key` 是否小於 `node.key`。

# --hints--

你應該撰寫另一個 if 述語來檢查 `key` 是否小於 `node.key`。記得使用 `pass` 關鍵字。

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```

