---
id: 65c4f195de7e2ad5932be717
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

接下來，你需要定義一個機制來插入節點到樹中。為此，你需要定義一個 `_insert` 方法，這是一個輔助函式，稍後會被實際的 `insert` 方法使用。 

此方法是遞迴的，表示它會呼叫自身以巡訪樹狀結構，直到找到新節點的適當位置。

定義一個帶有參數 `self`、`node` 和 `key` 的 `_insert` 方法。

# --hints--

你應該在 `BinarySearchTree` 類別中定義一個 `_insert` 方法。記得使用 `pass`。

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

你的 `_insert` 方法應該接受三個參數：`self`、`node` 和 `key`。

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
