---
id: 65c4f37ba6cbcfdc77e24165
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

`insert` 方法將由使用者呼叫。除了 `self` 參數外，它還需要一個 `key` 參數。此參數將是要插入到二元搜尋樹中的鍵值。

將 `key` 添加到函式定義的第二個參數。

# --hints--

`insert` 方法應該包含 `self` 參數。

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

`insert` 方法應該包含 `key` 參數。

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
