---
id: 65c4f0b82db314d3fcc0b8e1
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

在 `BinarySearchTree` 類別中，將 `pass` 取代為一個 `__init__` 方法，並為此方法添加一個 `self` 參數。

# --hints--

你應該從 `BinarySearchTree` 類別中移除 `pass` 關鍵字。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


你應該定義一個 `__init__` 方法，並為這個方法添加一個 `self` 參數。記得在方法主體內使用 `pass` 關鍵字。

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
