---
id: 65c4f013851cefd1a4fe4c96
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

`__init__` 方法接受兩個參數：`self`（表述正在被創建的類別實體）和 `key`（要儲存在節點中的值）。為你的 `__init__()` 方法添加這兩個參數。

# --hints--

你應該為你的方法添加 `self` 參數。

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

你應該為這個方法添加 `key` 參數。

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
