---
id: 65c63da2ed6769a10e141341
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

現在你將為遞迴搜尋定義一個基底案例。移除目前的 `pass`，並撰寫一個 `if` 述語來檢查兩個條件：

- 如果 `node` 是 `None`：這顯示搜尋已到達分支的末端，且未找到鍵。
- 如果 `node.key == key`：這表示在目前的節點中已找到該鍵。

將兩個條件用 `or` 運算子結合，並在 `if` 區塊內傳回目前節點。

# --hints--

你應該從 `_search` 方法中移除 `pass` 關鍵字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

你應該撰寫一個 `if` 述語來檢查 `node` 是否為 `None` 或 `node.key` 是否與 `key` 相等。

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

你應該從 `if` 區塊傳回 `node`。

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
