---
id: 65c63fd9b387f0a6c5bf6a72
title: 第 24 步
challengeType: 20
dashedName: step-24
---

# --description--

在 `search` 方法中，刪除 `pass` 並以以下引數呼叫輔助方法 `_search`。

- `self.root`：這是二元搜尋樹的根節點。搜尋從根節點開始。
- `key`：這是使用者想要在二元搜尋樹中尋找的值。

在內部，`search` 將實際的搜尋邏輯委派給 `_search` 輔助方法，該方法在二元搜尋樹中執行實際的遞迴搜尋。

# --hints--

你應該從 `search` 方法中移除 `pass` 關鍵字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

你應該在 `search` 方法中呼叫 `_search` 方法。

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

你應該呼叫 `_search` 方法，並傳入 `self.root` 和 `key` 作為引數。

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
