---
id: 65c63a4c4da62e9ae18e321a
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

是時候開始處理搜尋功能了。就像你為 `insert` 函式創建了輔助方法 `_insert` 一樣，你需要為 `search` 函式創建輔助方法 `_search`。

定義帶有三個參數的 `_search` 方法，分別是 `self`、`node` 和 `key`。

# --hints--

你應該宣告一個 `_search` 方法。請注意前導底線並記得使用 `pass`。

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

你的 `_search` 方法應該接受三個參數：`self`、`node` 和 `key`。

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
