---
id: 65c645b838e7deb080fc25e0
title: 第 29 步
challengeType: 20
dashedName: step-29
---

# --description--

現在，節點已經被插入。若要確認它們是否已正確插入，你可以在樹中搜尋節點。

在 for 迴圈外，在 `bst` 實體中搜尋節點 `80` 並將它添加到 `print` 呼叫中。`print` 函式的第一個引數應該是 `'Search for 80:'`。

# --hints--

你應該列印呼叫 `bst.search(80)` 的結果，且你的 `print` 述語的第一個引數應為 `'Search for 80:'`。

```js
assert.match(code, /^print\(('|")Search for 80:('|"),\s*bst\.search\(80\)/gm);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

--fcc-editable-region--
```
