---
id: 65c63da2ed6769a10e141341
title: 步驟 20
challengeType: 20
dashedName: step-20
---

# --description--

現在你將爲遞歸搜索定義一個基本情況。 刪除當前的 `pass` 並編寫一個檢查兩個條件的 `if` 語句：

- 如果 `node` 是 `None`：這表示搜索已到達分支的末尾，但尚未找到密鑰。
- 如果是 `node.key == key`: 這意味着在當前節點中已經找到了該鍵。

將兩個條件與 `or` 運算符結合起來，並返回 `if` 塊內的當前節點。

# --hints--

你應該從 `_search` 方法中刪除 `pass` 關鍵字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

你應該編寫一個 `if` 語句來檢查 `node` 是否爲 `None` 或者 `node.key` 是否等於 `key`。

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

你應該從 `if` 塊返回 `node`。

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
