---
id: 65c646d4148ae3b2d1cbcac4
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

請注意，你的搜尋會傳回類似 `80: <__main__.TreeNode object at 0x108b3e0>` 的結果。這是在列印類別實體時的預設字串表現。

要將其改為列印有用的值，請在 `TreeNode` 類別中定義另一個名為 `__str__` 的方法。它接受一個引數 `self`。

在定義 `__str__` 之後，你會在主控台看到例外，因為 `__str__` 方法尚未傳回任何東西。你會在下一步處理方法本體。

# --hints--

你應該定義一個方法 `__str__`，該方法帶有單一引數 `self`。記得使用 `pass`。

```js
assert.match(code, /^(\s+)def\s+__init__.+?^\1def\s+__str__\(\s*self\s*\)\s*:\s*\n^\1\1pass/ms)
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
