---
id: 65c63e6962159fa372ecdbec
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

如果第二個 `if` 述語不是 `True`，表示標的鍵大於或相等於目前節點鍵。 

在二元搜尋樹中，如果標的鍵大於目前節點的鍵，搜尋會繼續在右子樹進行。

在 `if` 區段之後，傳回以目前節點的右子節點和鍵作為引數呼叫 `_search` 方法的結果。

# --hints--

你應該遞迴地呼叫 `_search` 方法，使用 `self._search()`。

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

你應該在 `_search` 方法的呼叫中傳入兩個引數：`node.right` 和 `key`。

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

你應該傳回 `_search` 方法呼叫的結果。

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
