---
id: 65c63df529bd15a24c187c62
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

撰寫另一個 `if` 述語，檢查標的鍵是否小於目前節點的鍵。

在 `if` 區段內，傳回以目前節點的左子節點和 `key` 作為引數呼叫 `_search` 方法的結果。

# --hints--

你應該撰寫另一個 `if` 述語來檢查 `key` 是否小於 `node.key`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

你應該從你的新 `if` 區塊傳回 `self._search(node.left, key)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_bodies()[0].is_equivalent("return self._search(node.left, key)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
```
