---
id: 65c6409418806da8d0636ffc
title: 步驟 26
challengeType: 20
dashedName: step-26
---

# --description--

`insert` 和 `search` 功能已完成，現在是時候測試它們了。

請注意，此時節點尚未排序，只是被插入。你稍後會使用中序遍歷來進行排序。

你可以像這樣在 Python 中創建一個類別的實體：

```python
object_name = ClassName()
```

在類別定義之外，創建 `BinarySearchTree` 類別的實體並將它指定給變數 `bst`。


# --hints--

你應該創建 `BinarySearchTree` 類別的實體。

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

你應該將新的 `BinarySearchTree` 實體指定給變數 `bst`。

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
