---
id: 65c9ddd336596e30a4266a50
title: 步驟 31
challengeType: 20
dashedName: step-31
---

# --description--

在 `__str__` 方法的主體中，刪除 `pass` 並傳回以 `self.key` 作為引數呼叫 `str()` 函式的結果。這是目前節點物件中儲存與節點相關的值的屬性。

在傳回結果後，你應該會看到例外從主控台消失，且輸出現在應該顯示與節點相關的 `key` 的值。

# --hints--

你應該從 `__str__` 方法中移除 `pass` 關鍵字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__str__").has_pass()`
      )
    );
  },
});

```

你應該從你的 `__str__` 方法傳回 `self.key` 的字串值。

```js
({ test: () => assert.match(code, /^\s{8}return\s+str\(\s*self\.key\s*\)/m) })
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--
    def __str__(self):
        pass

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
