---
id: 65ca03bd8eb5faf24b250c56
title: 步驟 33
challengeType: 20
dashedName: step-33
---

# --description--

下一步是處理節點的刪除。為此，你會先定義一個輔助方法，然後再定義實際的方法。

在 `BinarySearchTree` 類別中，定義一個新的輔助方法 `_delete`，該方法帶有三個參數：`self`、`node` 和 `key`。

# --hints--

你應該在 `BinarySearchTree` 類別內定義 `_delete` 方法，並帶有 `self`、`node` 和 `key` 參數。記得使用 `pass` 關鍵字。

```js
({ test: () => assert.match(code, /def\s+_delete\(\s*self\s*,\s*node\s*,\s*key\s*\)/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))

```
