---
id: 65ca06864129380054b55dc0
title: 第 36 步
challengeType: 20
dashedName: step-36
---

# --description--

在 `if` 區段中，將 `pass` 取代為呼叫 `_delete` 方法，並將目前節點的左子節點和 `key` 作為引數傳入。將函式呼叫指定給左節點。

# --hints--

你應該從 `if` 述語中移除 `pass` 關鍵字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_delete").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

你應該以 `node.left` 和 `key` 作為引數來呼叫 `_delete` 方法。

```js
assert.match(code, /self\._delete\(\s*node\.left\s*,\s*key\s*\)/);
```

你應該將 `_delete()` 呼叫的結果指定給目前節點的左子節點（`node.left`）。

```js
assert.match(code, /node\.left\s*=\s*self\._delete\(\s*node\.left\s*,\s*key\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            pass

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
