---
id: 65ca0d5adf39c410cd1177cc
title: 步驟 50
challengeType: 20
dashedName: step-50
---

# --description--

現在，你將根據中序遍歷方法來遍歷樹。中序遍歷是一種深度優先二元樹遍歷演算法，會先造訪左子樹、當前節點，然後是右子樹。

在 `BinarySearchTree` 類別中定義 `_inorder_traversal` 方法，並給它三個參數：`self`、`node` 和 `result`。其中 `node` 是遍歷過程中目前正在考慮的節點，`result` 是用來以排序順序附加鍵的列表。

# --hints--

你應該定義一個方法 `_inorder_traversal`，該方法接受三個參數：`self`、`node` 和 `result`。記得使用 `pass`。

```js
assert.match(code, /def\s+_inorder_traversal\(\s*self\s*,\s*node\s*,\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
            
            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)   
        
        return node

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

--fcc-editable-region--

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
