---
id: 65ca089e848eca0672b9cd77
title: 步驟 45
challengeType: 20
dashedName: step-45
---

# --description--

現在，回到 `_delete` 方法，你必須使用 `_min_value` 輔助函式來選擇繼任者。

在你的 `elif` 區段後面添加一個 `_min_value` 呼叫，並將 `node.right` 作為引數傳入。將其指定給目前的節點鍵。

# --hints--

你應該在你的 `elif` 區段之後，將 `self._min_value(node.right)` 指定給 `node.key`。

```js
({ test: () => (runPython(`
import ast

assign_after_elif = _Node(_code).find_class("BinarySearchTree").find_function("_delete").find_ifs()[1].tree.orelse[0].orelse[1]
assign_code = ast.get_source_segment(_code, assign_after_elif)
assert _Node(assign_code).is_equivalent("node.key = self._min_value(node.right)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
            
--fcc-editable-region--
    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))



```
