---
id: 65cb45d130c97cb459439fac
title: 步驟 42
challengeType: 20
dashedName: step-42
---

# --description--

爲了找到右子樹中的最小值，你需要遍歷給定節點的左子節點，直到到達子樹中最左邊（最小）的節點。

爲此，編寫一個 `while` 循環，該循環在 `node.left is not None` 時運行，並將 `pass` 移動到 `while` 塊內。 此條件檢查是否存在左孩子。 只要有左孩子，循環就繼續，並會找到一個較小的值。

# --hints--

你應該從 `_min_value` 方法中刪除 `pass` 關鍵字。

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_min_value").has_pass()`
      )
    );
  },
});
```

你應該在 `while` 循環中使用條件 `node.left is not None`。

```js
({ test: () => assert.match(code, /while\s+node\.left\s+is\s+not\s+None\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left  

--fcc-editable-region--
    def _min_value(self, node):
        pass

--fcc-editable-region--   

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
