---
id: 6a57c331ae94fa4631b60870
title: Hatua 17
challengeType: 1
dashedName: step-17
---

# --description--

Mwishowe, tengeneza muhtasari unaoweza kusomwa kwa kipande chochote cha sanaa. Anza kwa kushughulikia hali ambapo kipande cha sanaa hakipo.

Tengeneza kitendakazi cha `buildSummary` kinachopokea kigezo cha `id`. Tumia `collection[id]` kutafuta kipande cha sanaa, na kihifadhi katika kigezo kinachoitwa `artifact`. Ikiwa kipande cha sanaa hakipatikani, rudisha mfuatano wa herufi `"Artifact not found"`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `buildSummary` chenye kigezo kinachoitwa `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary);
const params = explorer.allFunctions.buildSummary.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

Unapaswa kuweka thamani ya `collection[id]` katika kigezo kinachoitwa `artifact`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary?.variables.artifact);
assert.isTrue(explorer.allFunctions.buildSummary?.variables.artifact?.value.matches("collection[id]"));
```

`buildSummary` inapaswa kurudisha `"Artifact not found"` wakati hakuna kipande cha sanaa kinacholingana na `id`.

```js
assert.strictEqual(buildSummary(999), "Artifact not found");
```

`buildSummary` haipaswi kurudisha `"Artifact not found"` wakati kipande cha sanaa kipo.

```js
assert.notStrictEqual(buildSummary(101), "Artifact not found");
assert.notStrictEqual(buildSummary(102), "Artifact not found");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
