---
id: 6a57c331ae94fa4631b60863
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Makumbusho ina maingizo zaidi ya moja ya sanamu zilizoandikwa. Ongeza sanamu ya pili katika `collection` chini ya ufunguo `102`, ukifuata muundo uleule kama wa kwanza.

Tumia maadili haya:

```js
{
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
}
```

# --hints--

`collection` inapaswa kuwa na maingizo mawili.

```js
assert.strictEqual(Object.keys(collection).length, 2);
```

Sanamu iliyohifadhiwa chini ya `102` inapaswa kuendana na kitu kilichoelezwa.

```js
assert.deepEqual(collection[102], {
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
});
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
--fcc-editable-region--
  
--fcc-editable-region--
};
```
