---
id: 6a57c331ae94fa4631b60865
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Wafanyakazi huchunguza vitu vya kale mara kadhaa kila siku, kwa hivyo mkaguzi wako anahitaji njia ya kutafuta inayoweza kutumika tena. Mabano ya mraba pia yanaweza kuchukua kigezo, kwa hivyo `collection[id]` hurudisha kitu chochote cha kale ambacho `id` inarejelea.

Tengeneza kitendakazi kinachoitwa `getArtifactTitle` kinachopokea kigezo `id`. Tumia `collection[id]` kupata kitu cha kale, kihifadhi katika kigezo kinachoitwa `artifact`, na rudisha `artifact.title`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `getArtifactTitle` chenye kigezo kinachoitwa `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

Unapaswa kuweka thamani ya `collection[id]` katika kigezo kinachoitwa `artifact`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle?.variables.artifact);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.variables.artifact?.value.matches("collection[id]"));
```

`getArtifactTitle` inapaswa kurudisha `artifact.title`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.hasReturn("artifact.title"));
```

`getArtifactTitle(101)` inapaswa kurudisha `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle(102)` inapaswa kurudisha `"Bronze Tablet"`.

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
