---
id: 6a57c331ae94fa4631b60864
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Uhakiki huanza kwa kusoma kile kilicho tayari kwenye rekodi. Kwa kuwa kila kifaa kimehifadhiwa chini ya kitambulisho chake, unaweza kupata kifaa kwa kuweka kitambulisho hicho ndani ya mabano ya mraba, kama `collection[101]`.

Fanya wito wa `console.log()` ukiambatanisha na `collection[101].title` ili kurekodi kichwa cha kifaa cha kwanza. Kisha, fanya wito wa `console.log()` ukiambatanisha na `collection[101].curator.name` ili kurekodi jina la msimamizi wake.

# --hints--

Unapaswa kurekodi `title` ya kifaa cha kwanza kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)/);
```

Unapaswa kuandika `name` ya mkusanyaji wa kifaa cha kwanza kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

Unapaswa kuandika kwanza `title`, na kisha `name` ya mhifadhi.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
