---
id: 6a57c331ae94fa4631b60861
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Kila kifaa kina mtu anayehusika nacho.

Ongeza sifa ya `curator` kwa kifaa kilichohifadhiwa chini ya ufunguo `101`, na weka thamani yake kuwa kitu chenye `id` ya `201` na `name` ya `"Earl Sinclair"`.

# --hints--

Kifaa kilichohifadhiwa chini ya `101` kinapaswa kuwa na sifa ya `curator`.

```js
assert.property(collection[101], "curator");
```

`curator` inapaswa kuwa kitu.

```js
assert.isObject(collection[101].curator);
```

`curator` inapaswa kuwa na `id` ya `201`.

```js
assert.strictEqual(collection[101].curator.id, 201);
```

`curator` inapaswa kuwa na `name` ya `"Earl Sinclair"`.

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
