---
id: 6a57c265abbc95dfebc05098
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda mkaguzi wa asili ya vitu vya kale. Utahifadhi vitu vya kale vya makumbusho katika kitu na kuandika vitendakazi vya kuvitafuta, kuviweka lebo, kuvihamisha kati ya makumbusho, na kufupisha historia yao.

Kila kitu cha kale huhifadhiwa chini ya kitambulisho chake (ID). Unda kitu kinachoitwa `collection` na ongeza kipengee ndani yake. Weka ufunguo kuwa `101` na thamani iwe kitu chenye `title` ya `"Golden Mask"` na `category` ya `"Ceremonial"`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `collection`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` inapaswa kuwa kitu.

```js
assert.isObject(collection);
```

`collection` inapaswa kuwa na kipengee kimoja.

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

Kitu cha kale kilichohifadhiwa chini ya `101` kinapaswa kuwa kitu.

```js
assert.isObject(collection[101]);
```

Kitu cha kale kilichohifadhiwa chini ya `101` kinapaswa kuwa na `title` ya `"Golden Mask"`.

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

Kitu cha kale kilichohifadhiwa chini ya `101` kinapaswa kuwa na `category` ya `"Ceremonial"`.

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
