---
id: 6a57c331ae94fa4631b6086b
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Hamisha kifaa chenye `id` `102` kwenda `"Hall B"` katika mwaka wa `2026`. Kisha, fanya wito wa `console.log()` ukiwa na `collection[102].locations` ili kurekodi safu ya data ya maeneo ya kifaa hicho.

# --hints--

Unapaswa kufanya wito wa `moveArtifact` ukiwa na `102`, `"Hall B"`, na `2026`.

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)/);
```

Unapaswa kurekodi `locations` za kifaa cha pili kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

Unapaswa kurekodi `locations` baada ya kuhamisha kifaa.

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

Mahali pa hivi karibuni pa kifaa cha pili kinapaswa kuwa `{ gallery: "Hall B", year: 2026 }`.

```js
assert.deepEqual(collection[102].locations[collection[102].locations.length - 1], { gallery: "Hall B", year: 2026 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
