---
id: 6a57c331ae94fa4631b60866
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Wakati mwingine wafanyakazi wataomba kipande cha sanaa ambacho makumbusho hayana. Ikiwa hakuna kipande cha sanaa kilichohifadhiwa chini ya `id` hiyo, `collection[id]` itakuwa `undefined`, na kusoma `.title` kutoka kwake kutatupa kosa.

Sahihisha taarifa ya return ili itumie hali ya masharti. Rudisha `artifact.title` ikiwa kipande cha sanaa kimepatikana, na mfuatano wa herufi `"Artifact not found"` ikiwa hakipo.

# --hints--

`getArtifactTitle(101)` inapaswa bado kurudisha `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle` inapaswa kurudisha `"Artifact not found"` ikiwa hakuna kipande cha sanaa kinacholingana na `id`.

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle` inapaswa kurudisha kichwa cha kipande chochote cha sanaa kinachokipata, siyo tu `101`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
