---
id: 6a57c331ae94fa4631b60862
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Kila kifaa pia hufuatilia mahali kilipokuwepo, jinsi kinavyoelezwa, na kama kiko kwenye maonyesho.

Ongeza vigezo vitatu zaidi kwa kifaa kilichohifadhiwa chini ya `101`:

- `locations`: safu ya data inayoshikilia vitu viwili, `{ gallery: "Hall A", year: 2020 }` na `{ gallery: "Hall C", year: 2024 }`
- `tags`: safu ya data inayoshikilia mfuatano wa herufi `"gold"` na `"egypt"`
- `onDisplay`: thamani ya kweli au siyo kweli `true`

# --hints--

`locations` za kifaa zinapaswa kuwa safu ya data ya vitu viwili vya mahali vilivyotolewa.

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

`tags` za kifaa zinapaswa kuwa `["gold", "egypt"]`.

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

`onDisplay` ya kifaa inapaswa kuwa `true`.

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
