---
id: 66bba6fff611169359d9d36a
title: Створення конфіденційної сторінки електронної пошти
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---


# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати `main` елемент з `id` `email`.
1. Ваш `#email` елемент повинен мати `padding` `50px`, верхнє поле `50px`, `width` `500px` і `border`, який є `2px`.
1. Загальна ширина вашого `#email` елемента, включно з відступами та кордонами, має бути `500px`.
1. Ви повинні мати два `div` елементи, один з ідентифікатором `confidential`, а інший з ідентифікатором `top-secret`, всередині вашого `#email` елемента.
1. Ваші `#confidential` та `#top-secret` елементи повинні мати `display` `inline-block`.
1. Ваші `#confidential` та `#top-secret` елементи повинні мати `padding`, ліве поле та `border`.
1. Елемент `#confidential` повинен містити текст `CONFIDENTIAL`.
1. Елемент `#top-secret` повинен містити текст `TOP SECRET`.
1. Ваші `#confidential` та `#top-secret` елементи повинні бути повернені за допомогою CSS трансформації.
1. Ви повинні мати щонайменше три абзаци всередині вашого `#email` елемента.
1. Ви повинні мати щонайменше три `span` елементи з класом `blurred` всередині ваших абзаців.
1. Ви повинні мати `blurred` селектор, який розмиває елемент `3px` за допомогою CSS фільтра.

# --hints--

Ви повинні мати `main` елемент з `id` `email`.

```js
assert.exists(document.querySelector('main#email'));
```

Ви повинні мати `#email` селектор, який встановлює `padding` своїх елементів з усіх боків на `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

Ви повинні мати `#email` селектор, який встановлює `margin-top` своїх елементів на `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

Ви повинні мати `#email` селектор, який встановлює `width` своїх елементів на `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

Ваш `#email` елемент повинен мати `2px` кордон.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

Ваш `#email` елемент повинен мати `box-sizing` `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

Ви повинні мати `div` елемент з `id` `confidential` всередині вашого `#email` елемента.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

Ви повинні мати `div` елемент з `id` `top-secret` всередині вашого `#email` елемента.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

Ваш `#confidential` елемент повинен мати `display`, встановлений на `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

Ваш `#top-secret` елемент повинен мати `display`, встановлений на `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

Ваш `#confidential` елемент повинен мати `padding` з усіх боків.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Ваш `#top-secret` елемент повинен мати `padding` з усіх боків.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Ваш `#confidential` елемент повинен мати ліве поле.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

Ваш `#top-secret` елемент повинен мати ліве поле.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

Ваш `#confidential` елемент повинен мати `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

Ваш `#top-secret` елемент повинен мати `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

Текст вашого `#confidential` елемента повинен бути `CONFIDENTIAL`. Ви або пропустили текст, або зробили помилку.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

Текст вашого `#top-secret` елемента повинен бути `TOP SECRET`. Ви або пропустили текст, або зробили помилку.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

Ваш `#confidential` елемент повинен використовувати CSS `transform` для повороту елемента.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

Ваш `#top-secret` елемент повинен використовувати CSS `transform` для повороту елемента.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

Ви повинні мати щонайменше три абзаци всередині вашого `#email` елемента.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

Ви повинні мати щонайменше три `span` елементи з класом `blurred` всередині ваших абзаців.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

Ваші `.blurred` елементи не повинні бути порожніми.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

Ви повинні мати `.blurred` селектор, який встановлює `filter` своїх елементів на `blur(3px)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
