---
id: 685239f0cad43c7a909cbd98
title: Depurar uma Página de Adoção de Animais
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Sally, proprietária de uma loja de adoção de animais de estimação, criou sua primeira página web, mas há alguns problemas.

Seu trabalho é corrigir todos os erros para que a Sally possa continuar construindo sua página. 

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Sally quer usar uma imagem de alguns gatos, mas ela não está sendo exibida corretamente. Você precisará corrigir o seguinte no elemento `img`:
   - Substitua o atributo `href` pelo atributo correto para a fonte da imagem.
   - Substitua o atributo `att` pelo atributo correto que representa um texto curto e descritivo para imagens.
   - Remova a tag de fechamento `</img>` porque os elementos `img` são elementos void e não possuem tags de fechamento.
2. Sally quer usar alguns links para direcionar os usuários para as páginas de cachorro e gato. Mas os links não estão funcionando corretamente. Você precisará corrigir o seguinte nos elementos `a`:
   - Substitua ambos os atributos `src` pelos atributos corretos usados para especificar URLs.

# --hints--

Seu elemento `img` deve ter um atributo `src` em vez do atributo `href`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

Seu elemento `img` deve ter um atributo `alt` em vez do atributo inexistente `att`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

Seu elemento `img` não deve ter uma tag de fechamento `</img>`.

```js
assert.notMatch(code, /<\/img>/);
```

Seu elemento `a` com o texto `Visit cats page` precisa ter um atributo `href` em vez de um atributo `src`.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

Seu elemento `a` com o texto `Visit dogs page` precisa ter um atributo `href` em vez de um atributo `src`.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
