---
id: 68a23004e96a105df3a21fc4
title: Побудова помічника для підрахунку карт
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

У казино в грі Blackjack гравець може визначити, чи має він перевагу в наступній роздачі над казино, відстежуючи відносну кількість високих і низьких карт, що залишилися в колоді. Це називається підрахунком карт.

Наявність більшої кількості високих карт у колоді надає перевагу гравцю. Коли підрахунок позитивний, гравець повинен ставити більше. Коли підрахунок дорівнює нулю або є від’ємним, гравець повинен ставити менше.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні використати `let` для оголошення глобальної змінної з ім’ям `count` і встановити її в `0`.
1. Ви повинні мати функцію з назвою `cardCounter`.
1. Функція `cardCounter` повинна отримувати параметр `card`, який може бути числом або рядком.
    - Для значень від `2` до `10` параметр `card` буде числом.
    - Для всіх інших значень параметр `card` буде рядком.
1. Функція `cardCounter` повинна змінювати глобальну змінну `count` відповідно до певних критеріїв.
1. Глобальна змінна `count` повинна збільшуватися на `1` для карт `2`, `3`, `4`, `5` або `6`.
1. Глобальна змінна `count` повинна залишатися незмінною для карт `7`, `8`, `9`.
1. Глобальна змінна `count` повинна зменшуватися на `1` для карт `10`, `"J"`, `"Q"`, `"K"`, `"A"`.
1. Функція `cardCounter` повинна повертати рядок із поточним підрахунком і рядком `Bet`, якщо підрахунок позитивний.
1. Функція `cardCounter` повинна повертати рядок із поточним підрахунком і рядком `Hold`, якщо підрахунок менший або дорівнює `0`.
1. У вихідних даних функції поточний підрахунок і рішення гравця (`Bet` або `Hold`) повинні бути розділені пробілом. Наприклад, `-3 Hold`.

# --hints--

Ви повинні використати `let` для оголошення глобальної змінної з ім’ям `count` і встановити її в `0`. 

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

Ви повинні мати функцію з назвою `cardCounter`. 

```js
assert.isFunction(cardCounter);
```

Ваша функція повинна повертати значення `count` і текст (`Bet` або `Hold`) з одним пробілом між ними.

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

Після карт `2`, `3`, `4`, `5` виклик `cardCounter(6)` повинен повертати рядок `5 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

Після карт `7`, `8` виклик `cardCounter(9)` повинен повертати рядок `0 Hold`.

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

Після карт `10`, `"J"`, `"Q"`, `"K"` виклик `cardCounter("A")` повинен повертати рядок `-5 Hold`.

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

Після карт `3`, `7`, `"Q"`, `8` виклик `cardCounter("A")` повинен повертати рядок `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

Після карт `2`, `"J"`, `9`, `2` виклик `cardCounter(7)` повинен повертати рядок `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

Після карт `2`, `2` виклик `cardCounter(10)` повинен повертати рядок `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

Після карт `3`, `2`, `"A"`, `10` виклик `cardCounter("K")` повинен повертати рядок `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
