---
id: 66ba762af611169359d9d369
title: Crie um artigo de jornal
challengeType: 14
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

Fulfill the user stories below and get all the tests to pass to complete the lab.

**Histórias de usuário:**

1. Você deve definir o `font-size` base do documento HTML como `24px`.
1. Você deve ter um elemento com a classe `newspaper` que contenha todos os outros elementos.
1. O elemento `.newspaper` deve ter um `font-size` de `16px` e uma fonte `Open Sans` com uma fonte de alternativa `sans-serif`.
1. Dentro do elemento `.newspaper`, você deve ter pelo menos sete elementos a mais: um para o nome do jornal, com uma classe de `name`, um para a data do artigo, com uma classe de `date`, um para a manchete, com uma classe de `headline`, um para a submanchete, com uma classe de `sub-headline`, um para o autor, com uma classe de `author` e dois parágrafos, cada um com uma classe de `text`. Todos esses elementos devem estar preenchidos com as informações do artigo.
1. O elemento `.name` deve ter um `font-size` que seja duas vezes o `font-size` do elemento raiz e deve usar a fonte `Times New Roman`, com uma fonte alternativa de `serif`.
1. Os elementos `.name` e `.author` devem usar CSS para colocar todos os seus caracteres em maiúsculas.
1. O elemento `.headline` deve ter um `font-size` que seja duas vezes o `font-size` do elemento pai e deve estar em negrito.
1. O elemento `.sub-headline` deve ter um `font-weight` de `100`, um `font-size` que seja `1.5` vezes o `font-size` do elemento pai e deve ser em itálico.
1. O elemento `.author` deve usar CSS para deixá-lo em negrito.
1. Os elementos `.text` devem ter um `text-indent` de `20px`.
1. Os elementos `.text` devem ter um `line-height` que seja duas vezes o `font-size` do elemento pai.
1. A primeira letra dos elementos `.text` deve estar em negrito e ser duas vezes o tamanho do `font-size` do elemento pai. Use o seletor `::first-letter` para isso.

# --hints--

Você deve ter um seletor `html` que defina o `font-size` como `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

Você deve ter um elemento com a classe `newspaper`.

```js
assert.exists(document.querySelector('.newspaper'));
```

Você deve ter um seletor `.newspaper` que defina o `font-size` dos elementos como `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

Você deve ter um seletor `.newspaper` que defina a `font-family` dos elementos como `'Open Sans', sans-serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

Você deve ter um elemento com a classe `name` dentro do elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

O elemento `.name` deve ter o nome do seu jornal.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

Você deve ter um elemento com a classe `date` dentro do elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

O elemento `.date` deve ter a data.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

Você deve ter um elemento com a classe `headline` dentro do elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

O elemento `.headline` deve ter a manchete do artigo.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

Você deve ter um elemento com a classe `sub-headline` dentro do elemento `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

O elemento `.sub-headline` deve ter a submanchete do artigo.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

Você deve ter um elemento com a classe `author` dentro do elemento `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

O elemento `.author` deve ter o autor do artigo.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

Você deve ter pelo menos dois elementos de parágrafo, cada um com a classe `text`, dentro do elemento `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

Os elementos `.text` devem ter o texto do artigo.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

Você deve ter um seletor `.name` que defina o `font-size` dos elementos como `2rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

Você deve ter um seletor `.name` que defina a `font-family` dos elementos como `'Times New Roman', serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

Você deve ter um seletor `.name` que defina o `text-transform` dos elementos como `uppercase`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.textTransform, 'uppercase');
```

Você deve ter um seletor `.author` que defina o `text-transform` dos elementos como `uppercase`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.author')?.textTransform, 'uppercase');
```

Você deve ter um seletor `.headline` que defina o `font-size` dos elementos como `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

Você deve ter um seletor `.headline` que defina o `font-weight` dos elementos como `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

Você deve ter um seletor `.sub-headline` que defina o `font-weight` dos elementos como `100`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

Você deve ter um seletor `.sub-headline` que defina o `font-size` dos elementos como `1.5em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

Você deve ter um seletor `.sub-headline` que defina o `font-style` dos elementos como `italic`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

Você deve ter um seletor `.author` que defina o `font-weight` dos elementos como `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

Você deve ter um seletor `.text` que defina o `text-indent` dos elementos como `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

Você deve ter um seletor `.text` que defina a `line-height` dos elementos como `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

Você deve ter um seletor `.text::first-letter` que defina o `font-weight` dos elementos como `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

Você deve ter um seletor `.text::first-letter` que defina o `font-size` dos elementos como `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
