---
id: 65b7cd2b7bd9a684ccf1dc16
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

要遍历字典的键，你只需将字典放入 `for` 循环中。 下面的代码将打印字典 `dict` 中的每个键：

```py
for i in dict:
   print(i)
```

将 `print()` 调用替换为 `for` 循环，该循环遍历 `copper` 并打印每个键。

# --hints--

你的代码中不应该有 `print(copper)`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

你应该创建一个 `for` 循环来遍历 `copper` 字典。

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

你应该在 `for` 循环内打印 `copper` 的每个键。

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
