---
id: 56533eb9ac21ba0edf2244bc
title: 购物清单
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

创建一个名叫 `myList` 的购物清单。 清单的数据格式就是多维数组。

每个子数组中的第一个元素应该是购买的物品名称。 第二个元素应该是物品的数量，类似于：

```js
["Chocolate Bar", 15]
```

列表中应至少有 5 个子数组。

# --hints--

`myList` 应该是一个数组。

```js
assert(Array.isArray(myList));
```

每个子数组的第一个元素都应该是字符串。

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

每个子数组的第二个元素都应该是数字。

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

列表中至少要包含 5 个元素。

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
