---
id: 5cfa3679138e7d9595b9d9d4
title: 使用递归代替循环
challengeType: 1
videoUrl: >-
  https://www.freecodecamp.org/news/how-recursion-works-explained-with-flowcharts-and-a-video-de61f40cb7f9/
forumTopicId: 301175
dashedName: replace-loops-using-recursion
---

# --description--

递归是函数调用自身的操作。 为了便于理解，有如下任务：计算数组内元素前 `n` 的元素乘积。 使用 `for` 循环， 可以这样做：

```js
  function multiply(arr, n) {
    let product = 1;
    for (let i = 0; i < n; i++) {
      product *= arr[i];
    }
    return product;
  }
```

下面是递归写法，注意代码里的 `multiply(arr, n) == multiply(arr, n - 1) * arr[n - 1]`。 这意味着可以重写 `multiply` 以调用自身而无需依赖循环。

```js
  function multiply(arr, n) {
    if (n <= 0) {
      return 1;
    } else {
      return multiply(arr, n - 1) * arr[n - 1];
    }
  }
```

递归版本的 `multiply` 分解如下。 在 <dfn>base case</dfn>（基本情形）里，也就是 `n <= 0` 时，它返回 1。 对于较大的 `n` 值，它会调用自身，但参数为 `n - 1`。 该函数调用以相同的方式求值，不断调用 `multiply`，直到 `n <= 0`。 此时，所有函数都可以返回，原始的 `multiply` 返回最终结果。

**注意：**递归的函数必须有一个基本情形，在该情形下它们会返回而不再调用函数（在此示例中，当 `n <= 0` 时），否则它们永远无法完成执行。

# --instructions--

写一个递归函数，`sum(arr, n)`，返回递归调用数组 `arr` 从前 `n` 个元素和。

# --hints--

`sum([1], 0)` 应该返回 0 。

```js
assert.equal(sum([1], 0), 0);
```

`sum([2, 3, 4], 1)` 应该返回 2 。

```js
assert.equal(sum([2, 3, 4], 1), 2);
```

`sum([2, 3, 4, 5], 3)` 应该等于 9。

```js
assert.equal(sum([2, 3, 4, 5], 3), 9);
```

代码不能包含任意形式的循环（`for`、`while` 或者高阶函数如：`forEach`、`map`、`filter` 以及 `reduce`）。

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

应该用递归解决这个问题。

```js
assert(
  sum.toString().match(/sum\(.*\)/g).length > 1
);
```

# --seed--

## --seed-contents--

```js
function sum(arr, n) {
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function sum(arr, n) {
  // Only change code below this line
  if(n <= 0) {
    return 0;
  } else {
    return sum(arr, n - 1) + arr[n - 1];
  }
  // Only change code above this line
}
```
