---
id: a97fd23d9b809dac9921074f
title: Eine Funktion zum Summieren mit optionalen Argumenten erstellen
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die bis zu zwei Argumente akzeptiert und diese summiert. Wenn nur ein Argument übergeben wird, soll die Funktion eine Funktion zurückgeben, die auf die zweite Zahl wartet, um die Summe zu bilden.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Erstellen Sie eine Funktion `addTogether`.
2. Wenn die Funktion zwei Argumente erhält, soll sie die Summe der beiden Argumente zurückgeben.
3. Wenn nur ein Argument übergeben wird, soll `addTogether` eine Funktion zurückgeben.
   - Wenn die zurückgegebene Funktion mit einem einzelnen Argument aufgerufen wird, soll sie die Summe zurückgeben.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. Wenn eines der Argumente keine gültige Zahl ist, geben Sie `undefined` zurück.

# --hints--

Sie sollten die `addTogether`-Funktion haben.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` sollte 5 zurückgeben.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` sollte 53,4 zurückgeben.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` sollte `undefined` zurückgeben.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` sollte `undefined` zurückgeben.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` sollte `undefined` zurückgeben.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` sollte eine Funktion zurückgeben.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` sollte 12 zurückgeben.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` sollte `undefined` zurückgeben.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` sollte `undefined` zurückgeben.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
