---
id: 686383abac6c625a84458ea9
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви дізнаєтеся більше про Tailwind CSS, використовуючи його утиліти сітки для створення компонента ціноутворення.

Для вас уже підготовлено базовий шаблонний код HTML. Почніть зі створення елемента `h1` з текстом `Choose your listening plan`.

# --hints--

Ви повинні створити елемент `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Ваш елемент `h1` має містити текст `Choose your listening plan`.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
