---
id: 686386e76197458b7dfd52eb
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Створіть елемент `div` під `h1`. Всередині цього `div` створіть ще один елемент `div`, а всередині нього вкладіть ще один елемент `div`.

І всередині найглибшого `div` створіть елемент `h2` з текстом `Listener`. Це представить перший тарифний план.

# --hints--

Ви повинні створити елемент `div`.

```js
assert.exists(document.querySelector("div"));
```

Ви повинні створити ще один елемент `div` всередині зовнішнього `div`.

```js
assert.exists(document.querySelector("div > div"));
```

Ви повинні створити ще один елемент `div` всередині внутрішнього `div`.

```js
assert.exists(document.querySelector("div > div > div"));
```

Ви повинні створити `h2` з текстом `Listener` всередині найглибшого елемента `div`.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
