---
id: 686386e76197458b7dfd52ee
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Всередині `ul` створіть два елементи списку. Всередині першого елемента списку створіть `span` елемент і додайте HTML-сутність `&#10003;` як його текст. Одразу після цього `span`, але всередині елемента списку, додайте текст `Ad-supported streaming`.

Всередині другого елемента списку створіть `span` елемент і додайте сутність `&#10003;` як його текст. Одразу після цього `span`, але всередині елемента списку, додайте текст `Curated playlists`.

# --hints--

У вашому `ul` елементі має бути два елементи списку.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

Ваш перший `li` елемент має містити `span` елемент з HTML-символом `&#10003;` всередині.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

У вашому першому `li` елементі текст `Ad-supported streaming` має йти одразу після `span` елемента.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

Ваш другий `li` елемент має містити `span` елемент з HTML-символом `&#10003;` всередині.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

У вашому другому `li` елементі текст `Curated playlists` має йти одразу після `span` елемента.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
