---
id: 686386e76197458b7dfd52ef
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Тепер ви створите кнопку CTA. Call to Action (CTA) — це кнопка або посилання, призначене для того, щоб спонукати користувача виконати певну дію, у цьому випадку — почати слухати.

Поза межами найвнутрішнішого `div` створіть елемент посилання з `href`, рівним `#`, і текстом `Start listening`.

На цьому завершується HTML для першої картки з цінами.

# --hints--

Ви повинні створити елемент посилання.

```js
assert.exists(document.querySelector("a"))
```

Ваш елемент посилання має мати множину `href`, рівну `#`.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

Ваш елемент посилання має містити текст `Start listening`.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
