---
id: 6582507654b3ed5712341382
title: 第 38 步
challengeType: 20
dashedName: step-38
---

# --description--

你將使用條件述語來檢查使用者的選擇。如果選擇是 `'1'`，表示使用者想要添加一筆支出。

仍在 `while` 迴圈中，在 `choice` 變數下，撰寫一個 `if` 述語來檢查 `choice` 是否相等於字串 `'1'`。如果為真，將會是添加新支出的起點。

在 `if` 述語主體內，宣告一個變數 `amount` 並指定其為空的 `input()` 呼叫。

# --hints--

你應該在你的 `while` 迴圈中創建一個 `if` 述語來檢查 `choice` 是否相等於 `'1'`。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
cond.is_equivalent("choice == '1'") or cond.is_equivalent("'1' == choice")
`)) })
```

你應該在你的 `if` 述語主體中宣告一個名為 `amount` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("amount")
`)) })
```

你應該將一個空的 `input()` 呼叫指定給你的 `amount` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input()")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

--fcc-editable-region--
```
