---
id: 6582575b8089f85b8b92d7c8
title: 步驟 42
challengeType: 20
dashedName: step-42
---

# --description--

一旦你獲得了費用詳情，你需要調用 `add_expense` 函數將新費用添加到 `expenses` 列表中。

使用 `input()` 獲取 `amount` 和 `category` 後，調用 `add_expense` 函數，傳遞三個參數：`expenses`、`amount` 和 `category`。

- `expenses` 是本項目前面在主函數中創建的空列表。
- `amount` 是費用金額。
- `category` 是費用的類別。

# --hints--

你應該在 `if` 語句中調用 `add_expense()` 函數。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

你應該向 `add_expense()` 調用傳遞三個參數：`expenses`、`amount` 和 `category`。 記住參數傳遞的順序是很重要的。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

你的 `add_expense()` 調用應該在 `amount` 和 `category` 分配之後進行。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)


def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')

--fcc-editable-region--
```
