---
id: 65824dfdb6815d563b2d3256
title: 步驟 37
challengeType: 20
dashedName: step-37
---

# --description--

`input()` 函式（程式）會取得使用者輸入，並以字串形式傳回該使用者輸入。

在你的 `while` 迴圈中，呼叫 `input()` 函式，並將字串 `'Enter your choice: '` 作為引數傳入，然後將結果指定給名為 `choice` 的變數。

# --hints--

你應該在你的 `while` 迴圈中宣告一個名為 `choice` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].has_variable("choice")
`)) })
```

你應該將 `input()` 函式（程式）的結果，並以字串 `'Enter your choice: '` 作為其引數，指定到你的 `choice` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_variable("choice").is_equivalent("choice = input('Enter your choice: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
--fcc-editable-region--
```
