---
id: 65821c1c40ccac44adf38e6b
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

創建一個名為 `expenses` 的空列表。你將用它來儲存你的每一筆支出。

# --hints--

你應該宣告一個名為 `expenses` 的新變數。

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

你應該為你的 `expenses` 變數指定一個空的列表。

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
