---
id: 6582324b5d37fe4b622bc3c0
title: 第 19 步
challengeType: 20
dashedName: step-19
---

# --description--

Lambda 函式（程式）是在 Python 中簡短且匿名的函式（程式），非常適合用於簡單且一次性的任務。它們由 `lambda` 關鍵字定義，並使用以下語法：

```py
lambda x: expr
```

在上述範例中，`x` 表述要用在運算式 `expr` 的參數，且它的作用就像傳統函式中的任何參數一樣。`expr` 是當 lambda 函式被呼叫時會被評估並傳回的運算式。

創建一個名為 `test` 的變數，並指定一個帶有 `x` 參數且傳回 `x * 2` 的 lambda 函式（程式）。

# --hints--

你應該宣告一個名為 `test` 的變數。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

你應該將 `lambda x: x * 2` 指定到你的 `test` 變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
