---
id: 65822773f4c60a479ec89f66
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

在 `print_expenses` 函式（程式）內，創建一個 `for` 迴圈，對 `expenses` 列表中的每一個項目進行迭代。使用 `expense` 作為迴圈變數，並將 `pass` 移動到迴圈主體內。

# --hints--

你應該在你的 `print_expenses` 函式（程式）中創建一個 `for` 迴圈來迭代 `expenses`。記得在迴圈主體中使用 `pass`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

你應該使用 `expense` 作為迴圈變數。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

你應該只在你的迴圈主體內使用 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
