---
id: 65822308ded1ba4632f66aa3
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

字典是 Python 中另一種內建的資料型別。字典是一種以 *key*-*value* 對組形式存在的資料群集。字典使用大括號（`{}`）定義，並包含以逗號分隔的鍵值對。每個鍵後面都跟著冒號（`:`）和對應的值：

```py
{'amount': 50.0, 'category': 'Food'}
```

在上述範例中，`'amount'` 和 `'category'` 是鍵，而 `50.0` 和 `'Food'` 是它們對應的值。

創建一個字典，其鍵為 `'amount'`，值為 `amount` 參數，並將你的新字典傳遞給 `.append()` 呼叫。

# --hints--

你應該將字典作為引數傳遞給 `.append()` 方法。

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

你應該將 `{'amount': amount}` 作為引數傳遞給 `.append()` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
