---
id: 658212ba98182f3e855e85f9
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

在這個專案中，你將使用 Python 建置一個簡單但功能完整的支出追蹤器。

首先定義一個名為 `add_expense` 的函式（程式），該函式接受三個參數：`expenses`、`amount` 和 `category`。使用 `pass` 關鍵字填充函式主體。

# --hints--

你應該使用 `def` 關鍵字來宣告一個名為 `add_expense` 的函式（程式）。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

你的 `add_expense` 函式（程式）應該接受三個參數：`expenses`、`amount` 和 `category`。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
