---
id: 65823cfc74aa564ffc460489
title: 步驟 26
challengeType: 20
dashedName: step-26
---

# --description--

現在，呼叫 `map()`，將你的 `lambda` 函式作為第一個引數，並將 `expenses` 列表作為第二個引數。

# --hints--

你應該在 `total_expenses` 函式（程式）內呼叫 `map()` 函式（程式）。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

你應該將你的 `lambda` 函式作為第一個引數傳遞給 `map()` 呼叫。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

你應該將 `expenses` 列表作為第二個引數傳遞給 `map()` 呼叫。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
