---
id: 65822308ded1ba4632f66aa3
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

字典是 Python 中的另一種內置數據類型。 字典是以*鍵*-*值*對形式存在的數據集合。 字典用花括號（`{}`）定義，並且包含用逗號分隔的鍵值對。 每個鍵後面跟着一個冒號（`:`）和值：

```py
{'amount': 50.0, 'category': 'Food'}
```

在上面的例子中，`'amount'` 和 `'category'` 是鍵，`50.0` 和 `'Food'` 是它們對應的值。

創建一個帶有鍵 `'amount'` 和值 `amount` 參數的字典，並將新字典傳遞給 `.append()` 調用。

# --hints--

你應該將字典作爲參數傳遞給 `.append()` 方法。

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

你應該將 `{'amount': amount}` 作爲參數傳遞給 `.append()` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
