---
id: 65824dfdb6815d563b2d3256
title: 步驟 37
challengeType: 20
dashedName: step-37
---

# --description--

`input()` 函數接受用戶輸入並以字符串的形式返回用戶輸入。

在 `while` 循環內，調用 `input()` 函數，傳遞字符串 `'Enter your choice: '` 作爲參數，並將結果分配給名爲 `choice` 的變量。

# --hints--

你應該在 `while` 循環中聲明一個名爲 `choice` 的變量。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].has_variable("choice")
`)) })
```

你應該將 `input()` 函數的結果（以字符串 `'Enter your choice: '` 作爲其參數）分配給 `choice` 變量。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_variable("choice").is_equivalent("choice = input('Enter your choice: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

--fcc-editable-region--
```
