---
id: 65822773f4c60a479ec89f66
title: 步驟15
challengeType: 20
dashedName: step-15
---

# --description--

在 `print_expenses` 函數內部，創建一個 `for` 循環，迭代 `expenses` 列表中的每個項目。 使用 `expense` 作爲循環變量，並將 `pass` 移動到循環體內。

# --hints--

你應該在 `print_expenses` 函數中創建一個 `for` 循環，該循環遍歷 `expenses`。 記住在循環體內使用 `pass`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

你應該使用 `expense` 作爲循環變量。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

你應該只在循環體內有 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
