---
id: 65823634a308cb4d6bcc5fa6
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

Lambda 函式（程式）可以與 `map()` 函式（程式）有價值地結合使用，`map()` 會對群集中的每個物件元素（例如列表）執行指定的函式（程式）：

```py
map(lambda x: x * 2, [1, 2, 3])
```

要執行的函式（程式）作為第一個引數傳入，而可迭代物件作為第二個引數傳入。

上述範例的結果將會是 `[2, 4, 6]`，其中傳遞給 `map()` 的列表中的每個項目都被 lambda 函式的動作加倍了。

修改你的 `print()` 呼叫，以列印呼叫 `map()` 的結果，並將 `test` 作為第一個引數，`[2, 3, 5, 8]` 作為第二個引數。你還無法看到可讀的輸出。 

# --hints--

你應該在你的程式碼中有 `print(map(test, [2, 3, 5, 8]))`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
