---
id: 65822bd82d708c4895080c35
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在 Python 中，一個重要的事情是，定義字串時使用的引號型別不能在字串內部使用。 例如，字串 `'I'm a string!'` 是無效的。 若要在該字串內使用單引號，你應該：

- 透過在引號前加上反斜線來跳脫它：`'I\'m a string!'`
- 或者使用雙引號來定義字串：`"I'm a string!"`（建議使用）。

你可以透過字典的鍵來存取值。你需要使用括號表示法，並將鍵放在方括號中：

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

你目前正在你的 f-string 中插入 `expense` 字典。修改 f-string 運算式以存取 `expense` 字典中 `'amount'` 鍵和 `'category'` 鍵的值。

# --hints--

你應該將 `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` 傳遞給你的 `print()` 呼叫。請記得在你的單引號 f-string 中使用雙引號，反之亦然。

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
