---
id: 65825a9520a0795c8afbef50
title: 步驟 44
challengeType: 20
dashedName: step-44
---

# --description--

在 `print()` 呼叫之後，呼叫 `print_expenses` 函式（程式）以顯示目前為止已添加的所有支出。將 `expenses` 列表作為引數傳遞。

# --hints--

你應該在你的 `elif` 語句中呼叫 `print_expenses()` 函式（程式）。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].find_calls("print_expenses")) == 1`)) })
```

你應該將 `expenses` 列表作為引數傳遞給你的 `print_expenses()` 呼叫。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("print_expenses(expenses)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)
--fcc-editable-region--
        elif choice == '2':
            print('\nAll Expenses:')
            
--fcc-editable-region--
```
