---
id: 65823b0167fd894f4a7ea60a
title: 第 24 步
challengeType: 20
dashedName: step-24
---

# --description--

接下來，你將在 `total_expenses` 函式中實作相同的邏輯。

目前，刪除 `test` 函式和 `print()` 呼叫。

# --hints--

你應該刪除你的 `test` 變數和相關的 `print()` 呼叫。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_variable("test")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(sum(map(test, [2, 3, 5, 8])))
--fcc-editable-region--

expenses = []
```
