---
id: 658246c28575d653d1f89b59
title: 步驟 33
challengeType: 20
dashedName: step-33
---

# --description--

`while` 迴圈是另一種迴圈，只要指定的條件為 `True`，就會執行一段程式碼。當條件變為 `False` 時，迴圈終止：

```py
while condition:
    <code>
```

在 `expenses` 列表下方，創建一個 `while` 迴圈。使用 `True` 作為條件，並在迴圈主體內列印字串 `'\nExpense Tracker'` 以顯示程式標題。

# --hints--

你應該在你的 `main()` 函式中使用 `True` 作為條件來創建一個 `while` 迴圈。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

你應該在你的 `while` 迴圈中列印 `'\nExpense Tracker'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
