---
id: 6582267ef39f5047411c1743
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

接著，定義一個名為 `print_expenses` 的函式（程式），它帶有一個參數 `expenses`。這個函式（程式）稍後將用來顯示你列表中的每一筆支出。

用 `pass` 述語填充你新函式的主體。

# --hints--

你應該定義一個名為 `print_expenses` 的新函式（程式）。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

你的 `print_expenses` 函式（程式）應該接受一個單一的參數 `expenses`。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
