---
id: 65824872894f59546e3084e2
title: 步驟 34
challengeType: 20
dashedName: step-34
---

# --description--

你在前一步建立的 `while` 迴圈是一個無窮迴圈，將允許程式持續顯示選單選項，直到使用者決定離開。 

在 `print()` 呼叫之後，為你的程式碼添加另一個 `print()` 呼叫以列印字串 `'1. Add an expense'`。

# --hints--

你應該在現有的 `print()` 呼叫之後，在你的 `while` 迴圈中列印 `'1. Add an expense'`。

```js
({ test: () => assert(runPython(`
call_lst = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")
test_lst = ["print('\\\\nExpense Tracker')", "print('1. Add an expense')"]
all(call_lst[n].is_equivalent(i) for n, i in enumerate(test_lst))
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
--fcc-editable-region--
    while True:
        print('\nExpense Tracker')
        
--fcc-editable-region--
```
