---
id: 65821fcc010c3245718f2a06
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

你的 `add_expense` 函式的 `expenses` 參數將會是一個費用列表。你想要能夠在你的列表末端添加項目。為此你將使用 `.append()` 列表方法。

為你的 `expenses` 列表添加對 `.append()` 方法的呼叫。暫時不要傳遞任何引數給 `.append()`。

# --hints--

你應該在你的函式中使用 `expenses.append()`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

你的函式中不應該有 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
