---
id: 658238f7604f154ea9a23e1e
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

`sum()` 函式傳回作為其引數傳入的可疊代物件中項目的總和。你將會將 `sum()` 與 `map()` 和 `lambda` 函式一起使用，以取得費用的總金額。

現在，做一個小測試，並修改你目前的 `print()` 呼叫，將 `list()` 呼叫替換為呼叫 `sum()` 函式，並將目前的 `map()` 呼叫作為引數傳入。

# --hints--

你應該列印 `sum(map(test, [2, 3, 5, 8]))`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
