---
id: 6582352cbdbcde4d0c4f7b0b
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

要呼叫 lambda 函式（程式），你可以在變數名稱後使用一對小括弧的常用函式語法。

呼叫你的 `test` lambda 函式，並傳入 `3` 作為引數。然後，列印結果。

# --hints--

你應該在你的程式碼中有 `print(test(3))`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
