---
id: 65823ff0d4b991510fade1a8
title: 第 29 步
challengeType: 20
dashedName: step-29
---

# --description--

在 `filter_expenses_by_category` 函式（程式）中，將 `pass` 替換為 `lambda` 函式（程式）。使用 `expense` 作為參數，並使用等號運算子計算 `expense` 字典中 `'category'` 鍵的值與 `category` 之間的比較。

# --hints--

你應該創建一個使用名為 `expense` 的參數並在 `filter_expenses_by_category` 函式內計算表達式 `expense['category'] == category` 的 `lambda` 函式。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

你不應該在 `filter_expenses_by_category` 函式中使用 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
