---
id: 65824561f3478e5371a33ae5
title: 步驟 32
challengeType: 20
dashedName: step-32
---

# --description--

下一步是定義主要函式（程式），它將成為互動式支出追蹤程式的進入點。

定義一個名為 `main` 的無參數函式（程式）。在函式（程式）主體中填入你在此專案開始時創建的 `expenses` 列表。你將使用此列表來儲存支出記錄。

# --hints--

你應該定義一個名為 `main()` 且沒有參數的函式。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

你應該將 `expenses` 列表移動到 `main()` 函式（程式）主體內。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
