---
id: 65823dde36392f505a39f7c7
title: 步驟 27
challengeType: 20
dashedName: step-27
---

# --description--

最後，將你的 `map()` 呼叫傳遞給 `sum()` 函式，以取得總支出並傳回結果。

# --hints--

你應該呼叫 `sum()` 函式，將你目前的 `map()` 呼叫作為引數傳入，並從你的 `total_expenses` 函式傳回結果。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
