---
id: 658225d560369446ece5518b
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

為你正在附加到 `expense` 列表的字典添加另一個鍵值對。使用字串 `'category'` 作為鍵，並使用 `category` 參數作為值。

# --hints--

你應該為你作為引數傳遞給 `.append()` 方法的字典添加第二個鍵值對組。

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

你應該將 `{'amount': amount, 'category': category}` 作為引數傳遞給 `.append()` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
