---
id: 6582324b5d37fe4b622bc3c0
title: 步驟 19
challengeType: 20
dashedName: step-19
---

# --description--

Lambda 函數是 Python 中的簡短、匿名函數，非常適合簡單的一次性任務。 它們由 `lambda` 關鍵字定義，並使用以下語法：

```py
lambda x: expr
```

在上面的例子中，`x` 表示在表達式 `expr` 中使用的參數，它的作用就像傳統函數中的任何參數一樣。 `expr` 是在調用 lambda 函數時進行評估並返回的表達式。

創建一個名爲 `test` 的變量，併爲其分配一個 lambda 函數，該函數接受 `x` 參數並返回 `x * 2`。

# --hints--

你應該聲明一個名爲 `test` 的變量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

你應該將 `lambda x: x * 2` 分配給你的 `test` 變量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
