---
id: 65823e87c9741750a22085a7
title: 步驟 28
challengeType: 20
dashedName: step-28
---

# --description--

接著，定義一個名為 `filter_expenses_by_category` 的函式（程式），它接受兩個參數：`expenses` 和 `category`。使用 `pass` 來填充函式（程式）主體。

# --hints--

你應該定義一個名為 `filter_expenses_by_category` 的函式（程式）。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

你的 `filter_expenses_by_category` 函式（程式）應該以 `expenses` 和 `category` 作為參數。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
