---
id: 65823040bdb32949d4c1a96f
title: 步驟18
challengeType: 20
dashedName: step-18
---

# --description--

你將需要一個函數來計算總支出金額。

定義一個名爲 `total_expenses` 的函數，它接受一個參數 `expenses`。 現在用 `pass` 語句填充函數主體。

# --hints--

你應該定義一個名爲 `total_expenses` 的函數。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

你的 `total_expenses` 函數應該採用單個 `expenses` 參數。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
