---
id: 65822934c8a0454803e3ef12
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

接下來，你將顯示每筆費用的詳細資料。

在 `for` 迴圈中，將 `pass` 替換為 `print()` 呼叫，並傳入以下 f-字串：`f'Amount: {expense}, Category: {expense}'`。

# --hints--

你應該在你的 `for` 迴圈中列印 `f'Amount: {expense}, Category: {expense}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

你的迴圈主體中不應該有 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
