---
id: 65824111a09164518320088d
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

`filter()` 函式允許你根據函式的輸出，從可迭代物件（例如列表）中選取項目： 

```py
filter(my_function, my_list)
```

`filter()` 以函式作為它的第一個引數，以可迭代物件作為它的第二個引數。它傳回一個迭代器，這是一種特殊的物件，能讓你對群集（例如列表）的元素進行迭代。

上述範例的結果是一個迭代器，包含 `my_function` 傳回 `True` 的 `my_list` 元素。 

在 `filter_expenses_by_category` 函式（程式）中，呼叫 `filter()`，將你在前一步撰寫的 `lambda` 函式（程式）作為第一個引數，並將 `expenses` 列表作為第二個引數。

# --hints--

你應該在 `filter_expenses_by_category` 函式（程式）內呼叫 `filter()`。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

你應該將 `lambda expense: expense['category'] == category` 作為第一個引數傳遞給 `filter()` 呼叫。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

你應該將 `expenses` 作為第二個引數傳遞給 `filter()` 呼叫。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
