---
id: 65824561f3478e5371a33ae5
title: 步驟 32
challengeType: 20
dashedName: step-32
---

# --description--

下一步是定義主函數，它將成爲交互式費用跟蹤程序的入口點。

定義一個名爲 `main` 的函數，沒有參數。 使用你在此項目開始時創建的 `expenses` 列表填充函數主體。 你將使用此列表來存儲費用記錄。

# --hints--

你應該定義一個名爲 `main()` 的函數，沒有參數。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

你應該將 `expenses` 列表移到 `main()` 函數體內部。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
