---
id: 66314f2add0a8b57da5d6d91
title: 第 22 步
challengeType: 20
dashedName: step-22
---

# --description--

你應該會在主控台看到類似 `<map object at 0xd273a8>` 的列印內容，這是由 `map()` 傳回的 map 物件的字串表現。

要取得可讀的輸出，你需要將 `map` 物件轉換成列表。透過將 `map()` 呼叫作為引數傳遞給 `list()` 函式來完成。

# --hints--

你應該將 `map()` 呼叫作為引數傳遞給 `list()` 函式（程式）。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
