---
id: 66694bfa435cef241f6f9a68
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在開始專案之前，你將會學習關於列表。

與數字和字串不同，<dfn>列表</dfn>是一種作為其他值容器的資料型別：

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

列表的特徵是在所有值的周圍有方括號，且值之間有逗號，例如：`["A", "happy", "list"]`。如果列表不包含任何值，則為空列表：`[]`。

列表可以包含不同的資料型別：`[1, "Up", ["Down", "Twice"]]`。這包括所有可能的資料型別。它也可以包含另一個列表！

創建一個名為 `my_list` 的變數，並將一個空列表指定給它。

# --hints--

你應該有一個名為 `my_list` 的變數。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

變數 `my_list` 應該具有空列表的值。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
