---
id: 666951d2e4b27e3a3c72ca5f
title: 步驟4
challengeType: 20
dashedName: step-4
---

# --description--

你還可以訪問單個元素來獲取其值。

列表與字符串一樣都是從零開始索引的。 這意味着第一個元素位於索引 `0`，第二個元素位於索引 `1`，依此類推。

要訪問元素，請使用括號表示法。 例如，`example_list[1]` 正在訪問索引 `1` 處的元素，即 `example_list` 的第二個元素。

打印 `my_list` 的第一個元素。

# --hints--

你應該使用 `print(my_list[0])` 打印列表的第一個元素。

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
