---
id: 6669539c1379793f9cb8917c
title: 步驟5
challengeType: 20
dashedName: step-5
---

# --description--

Python 列表是可變的，這意味着列表項的值可以改變。 你可以使用括號符號來更改元素的值。

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

這將使 `example_list` 具有 `[4, 'oh', 6, 7]` 的值。

將 `my_list` 的第一個元素更改爲 `0`，然後打印列表以檢查其值。

# --hints--

不要更改列表聲明。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

`my_list` 的第一個元素應該是 0。

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

你應該使用 `my_list[0] = 0` 更改第一個元素。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

你應該打印這份清單。

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

你應該在更改第一個元素後打印列表。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
