---
id: 66694fc4bba24f33ca01fa5b
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

Python 有少量的列表方法。 例如添加或刪除列表項的方法。

你可以使用 `append()` 方法將項目添加到列表末尾。 例如：

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` 現在是 `[4, 5, 6, 7]`。

嘗試使用 `append()` 方法將 `3` 添加到 `my_list`。 然後打印該列表。

# --hints--

`my_list` 的最後一個元素應該是 `3`。

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` 的長度應爲 3。

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

你不應該更改列表聲明。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

你應該打印這份清單。

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

你應該在使用 `append()` 後打印列表。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
