---
id: 66694eaaeb7d253049fa7883
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

要為你的列表添加項目，你可以將它們寫在方括號之間，並用逗號分隔數值。

為你的列表添加數字 `1` 和 `2`。

# --hints--

`my_list` 的第一個元素應該是 `1`。

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

`my_list` 的第二個元素應該是 `2`。

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` 應該包含兩個元素。

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

你應該直接將 `[1, 2]` 指定到 `my_list`。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
