---
id: 6669539c1379793f9cb8917c
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

Python 列表是可變的，這表示列表項目的值可以被更改。你可以使用括號標記法更改元素的值。

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

這將使 `example_list` 的值為 `[4, 'oh', 6, 7]`。

將 `my_list` 的第一個元素改為 `0`，然後列印該列表以檢查值。

# --hints--

請勿更改列表宣告。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

`my_list` 的第一個元素應該是 0。

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

你應該使用 `my_list[0] = 0` 來更改第一個元素。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

你應該列印該列表。

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

你應該在更改第一個元素後列印該列表。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
