---
id: 666951d2e4b27e3a3c72ca5f
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

你也可以存取單一元素以取得其值。

列表和字串一樣是從零開始編號。這表示第一個元素的索引是 `0`，第二個元素的索引是 `1`，依此類推。

要存取一個元素，你使用括號表示法。例如，`example_list[1]` 是存取 `example_list` 中索引為 `1` 的元素，也就是第二個元素。

列印 `my_list` 的第一個元素。

# --hints--

你應該使用 `print(my_list[0])` 列印列表的第一個元素。

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
