---
id: 666aae6a5d30a71f1fd7749f
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

`insert` 方法可以在列表中的任何位置添加元素。第一個引數是要添加元素的位置，第二個引數是要添加的元素。<dfn>例如</dfn>，以下是如何在 `example_list` 的第三個位置添加新元素：

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

使用 `insert()`，將 `1` 添加到你的 `my_list` 中適當的位置，使其按遞增順序排列，然後列印該列表。

# --hints--

你應該使用 `insert()` 將 `1` 添加到 `my_list` 的索引 `1`。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list` 應該有 4 個元素。

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` 應該相等於 `[0, 1, 2, 3]`。

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

你應該在為列表添加新元素後列印該列表。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
