---
id: 666ab9013cd5563f2b39947a
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

既然你已經探索了列表，請移除你為審查所撰寫的所有程式碼，以便你可以開始專案。

# --hints--

你應該刪除所有內容。

```js
({
    test: () => runPython(`
    assert not _Node(_code)
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)

my_list.pop()
print(my_list)

--fcc-editable-region--
```
