---
id: 646c48df8674cf2b91020ecc
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

أضف عنصر `link` داخل عنصر `head`. بالنسبة لذلك العنصر `link`، عيّن الخاصية `rel` إلى `stylesheet` والخاصية `href` إلى `./styles.css`.

# --hints--

يجب أن يكون لديك عنصر `link`.

```js
assert.match(code, /<link/)
```

يجب أن يحتوي عنصر `link` الخاص بك على خاصيتي `rel` و `href` مع تعيينهما إلى `stylesheet` و `./styles.css`.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
