---
id: 646ceb843412c74edee27a79
title: الخطوة 22
challengeType: 0
dashedName: step-22
---

# --description--

الآن ستتعلم خدعة في CSS لرسم مثلثات.

سيُستخدم هذا لرسم آذان القطة.

باستخدام محدد صنف، أعطِ عنصر `.cat-right-ear` خاصيتي `height` و `width` بقيمة `100px`. اضبط `background-color` على `white`. اضبط الحدود اليسرى واليمنى لتكون `35px solid blue`. اضبط الحدود العلوية والسفلية لتكون `35px solid red`.

# --hints--

يجب أن يكون لديك محدد `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

يجب أن يحتوي محدد `.cat-right-ear` على خاصية `height` بقيمة `100px`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

يجب أن يحتوي محدد `.cat-right-ear` على خاصية `width` بقيمة `100px`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

يجب أن يحتوي محدد `.cat-right-ear` على خاصية `background-color` بقيمة `white`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

يجب أن يحتوي محدد `.cat-right-ear` على خاصية `border-left` بقيمة `35px solid blue`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

يجب أن يحتوي محدد `.cat-right-ear` على خاصية `border-right` بقيمة `35px solid blue`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

يجب أن يحتوي محدد `.cat-right-ear` على خاصية `border-top` بقيمة `35px solid red`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

يجب أن يحتوي محدد `.cat-right-ear` على خاصية `border-bottom` بقيمة `35px solid red`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
