---
id: 646c47867800472a4ed5d2ea
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ابدأ بالهيكل الأساسي لـ HTML. أضف مرجع `DOCTYPE` من نوع `html` وعنصر `html` مع تعيين خاصية `lang` إلى `en`. أضف أيضًا عنصر `head` وعنصر `body` داخل عنصر `html`.

# --hints--

يجب أن يكون لديك إعلان `DOCTYPE` من نوع `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

يجب أن يكون إعلان `DOCTYPE` في بداية ملف HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك وسم `html` مفتوح مع خاصية `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يكون لديك وسم `html` مغلق.

```js
assert(code.match(/<\/html>/i));
```

يجب أن يكون لديك وسم `head` مفتوح.

```js
assert(code.match(/<head>/i));
```

يجب أن يكون لديك وسم `head` مغلق.

```js
assert(code.match(/<\/head>/i));
```

يجب أن يكون لديك وسم `body` مفتوح.

```js
assert(code.match(/<body>/i));
```

يجب أن يكون لديك وسم `body` مغلق.

```js
assert(code.match(/<\/body>/i));
```

يجب أن يأتي عنصر `body` بعد عنصر `head`.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

يجب أن يكون عنصرا `head` و `body` داخل عنصر `html`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
