---
id: 646c5d7057c45f432fcdd46c
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

باستخدام محدِّد صنف، أعطِ عنصر `.cat-head` عرضًا مقداره `205px` وارتفاعًا مقداره `180px`. كذلك، أعطه حدودًا بقيمة `1px solid #000` وانحناء زوايا (`border-radius`) مقداره `46%`.

# --hints--

يجب أن يكون لديك محدِّد `.cat-head`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

يجب أن يحتوي محدِّد `.cat-head` على `width` مضبوطًا على `205px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

يجب أن يحتوي محدِّد `.cat-head` على `height` مضبوطًا على `180px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

يجب أن يحتوي محدِّد `.cat-head` على `border` مضبوطًا على `1px solid #000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

يجب أن يحتوي محدِّد `.cat-head` على `border-radius` مضبوطًا على `46%`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
