---
id: 646c5ffef5598d449b52ec12
title: خطوة 19
challengeType: 0
dashedName: step-19
---

# --description--

الآن يجب أن تعمل على آذان القط. سيكون هناك أذن يمنى وأذن يسرى، وداخل كل واحدة منهما، سيكون هناك أذن داخلية.

داخل عنصر `.cat-head` الخاص بك، أنشئ عنصر `div` بالصنف `cat-ears`.

# --hints--

لا يجب أن تغيّر عنصر `div` الموجود بالصنف `cat-head`.

```js
assert(document.querySelectorAll('div.cat-head').length === 1);
```

يجب أن تنشئ عنصر `div` واحدًا داخل عنصر `.cat-head` الخاص بك.

```js
assert(document.querySelectorAll('.cat-head div').length === 1);
```

يجب أن يكون لعنصر `div` الخاص بك الصنف `cat-ears`.

```js
assert(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class') === 'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">
        
      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
