---
id: 646c5ace05e4be4211407935
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

لرؤية عنصر `.cat-head`، امنحه خلفية بتدرج لوني خطي مع `#5e5e5e` عند `85%` و`#45454f` عند `100%`.

قد لا تلاحظ الفرق بين هذين اللونين، لكنه موجود.

# --hints--

يجب أن يحتوي محدِّد `.cat-head` على خاصية `background`.

```js
assert.match(code, /background:/)
```

يجب أن تستخدم خاصية `background` الدالة `linear-gradient`.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

يجب أن تحدد دالة `linear-gradient` اللون الأول ليكون `#5e5e5e` عند `85%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

يجب أن تحدد دالة `linear-gradient` اللون الثاني ليكون `#45454f` عند `100%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
