---
id: 646c48df8674cf2b91020ecb
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

داخل عنصر `head`، أضف وسم `meta` مع خاصية `charset` بقيمة `utf-8`. أضف أيضًا عنصر `title` بالنص `fCC Cat Painting`.

# --hints--

يجب أن تضيف عنصر `meta` واحد فقط.

```js
assert(document.querySelectorAll('meta').length === 1);
```

يجب أن يحتوي عنصر `meta` على خاصية `charset` مضبوطة على `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

يجب أن تضيف عنصر `title` واحد فقط.

```js
assert(document.querySelectorAll('title').length === 1);
```

يجب أن يحتوي عنصر `title` على النص `fCC Cat Painting`. لاحظ أن التهجئة وحالة الأحرف مهمة.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
