---
id: 646ddd3f9f97a0667b964bdb
title: الخطوة 42
challengeType: 0
dashedName: step-42
---

# --description--

أنشئ عنصرين من نوع `div`، واحد يحمل الصنف `cat-left-inner-eye` داخل العنصر `.cat-left-eye` وآخر يحمل الصنف `cat-right-inner-eye` داخل العنصر `.cat-right-eye`.

# --hints--

لا تغير عنصر `div` الموجود الذي يحمل الصنف `cat-left-eye`.

```js
assert(document.querySelectorAll('div.cat-left-eye').length === 1);
```

لا تغير عنصر `div` الموجود الذي يحمل الصنف `cat-right-eye`.

```js
assert(document.querySelectorAll('div.cat-right-eye').length === 1);
```

يجب أن يكون لديك عنصر `div` واحد داخل العنصر `.cat-left-eye`.

```js
assert(document.querySelectorAll('.cat-left-eye div').length === 1);
```

يجب أن يكون لديك عنصر `div` واحد داخل العنصر `.cat-right-eye`.

```js
assert(document.querySelectorAll('.cat-right-eye div').length === 1);
```

يجب أن يحمل عنصر `div` الجديد داخل `.cat-left-eye` الصنف `cat-left-inner-eye`.

```js
assert(document.querySelectorAll('.cat-left-eye div')[0]?.classList.contains('cat-left-inner-eye'));
```

يجب أن يحمل عنصر `div` الجديد داخل `.cat-right-eye` الصنف `cat-right-inner-eye`.

```js
assert(document.querySelectorAll('.cat-right-eye div')[0]?.classList.contains('cat-right-inner-eye'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>

        <div class="cat-eyes">
          --fcc-editable-region--
          <div class="cat-left-eye">
            
          </div>
          <div class="cat-right-eye">
            
          </div>
          --fcc-editable-region--
        </div>        
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-left-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

.cat-right-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}
```
