---
id: 6476fd4213318f6ee211028a
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

استبدل الآن قيمة خاصية `position` في محدِّد `.cat-head` إلى `fixed`. اترك كل من `top` و `left` كما هما.

بعد ذلك، قم بالتمرير لأعلى ولأسفل لترى كيف تعمل القيمة `fixed`.

# --hints--

يجب أن يحتوي محدِّد `.cat-head` على خاصية `position` مضبوطة على `fixed`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'fixed')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `top` مضبوطة على `100px`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `left` مضبوطة على `100px`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
