---
id: 6476f7a4827bcc61682f2347
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

الخاصية التالية لـ position هي `absolute`. عند استخدام القيمة `absolute` لخاصية `position`، يُخرج العنصر من التدفق الطبيعي للوثيقة، ثم يُحدد موقعه بواسطة خصائص `top` و`right` و`bottom` و`left`.

عيّن خاصية `position` لعنصر `.cat-head` إلى `absolute`، ثم عيّن خصائص `top` و`left` إلى أي قيمة بكسل موجبة.

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

يجب أن يحتوي محدِّد `.cat-head` على خاصية `position` مضبوطة على `absolute`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `top` مضبوطة على أي رقم بوحدة `px`. تأكد من إضافة فاصلة منقوطة.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `left` مضبوطة على أي رقم بوحدة `px`. تأكد من إضافة فاصلة منقوطة.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
