---
id: 6476fc5cf14b276e6d04e82a
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

استخدم محدِّد صنف لإعطاء عنصر `.box` عرضًا مقداره `200px`، وارتفاعًا مقداره `600px`، ولون خلفية `#000`. كذلك، عيّن له خاصية `position` بقيمة `absolute`، و`top` بقيمة `800px`، و`left` بقيمة `650px`.

# --hints--

يجب أن يكون لديك محدِّد `.box`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

يجب أن يحتوي محدِّد `.box` على خاصية `width` مضبوطة على `200px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

يجب أن يحتوي محدِّد `.box` على خاصية `height` مضبوطة على `600px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

يجب أن يحتوي محدِّد `.box` على خاصية `background-color` مضبوطة على `#000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

يجب أن يحتوي محدِّد `.box` على خاصية `position` مضبوطة على `absolute`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

يجب أن يحتوي محدِّد `.box` على خاصية `left` مضبوطة على `650px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

يجب أن يحتوي محدِّد `.box` على خاصية `top` مضبوطة على `800px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

