---
id: 6476f5c17f99146071ee884c
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

يمكنك أن ترى أنه لا يحدث شيء. لم يتحرك عنصر `.cat-head` رغم تعيين `top` و `left` بقيمة `100px` لكل منهما. لكن هذا ليس الحال مع التمركز بـ `relative`.

عند استخدام القيمة `relative`، يبقى العنصر موضوعًا وفقًا للتدفق الطبيعي للوثيقة، لكن قيم `top` و `left` و `bottom` و `right` تصبح فعالة.

بدلًا من `static`، عرّف لمحدد `.cat-head` الخاص بك موقعًا بقيمة `relative`، واترك خاصيتي `top` و `left` كما هما.

# --hints--

يجب أن يحتوي محدد `.cat-head` على خاصية `position` مضبوطة على `relative`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'relative')
```

يجب أن يحتوي محدد `.cat-head` على خاصية `top` مضبوطة على `100px`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

يجب أن يحتوي محدد `.cat-head` على خاصية `left` مضبوطة على `100px`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
