---
id: 6477062778c85972eb648030
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

يجب الآن أن تضع رأس القطة في المنتصف.

أعطِ عنصر `.cat-head` خاصية `position` بقيمة `absolute`. عيّن قيمة `0` للخصائص `right` و`left` و`top` و`bottom`، ثم عيّن خاصية `margin` على جميع الجوانب إلى `auto`. هذه طريقة لوضع عنصر في المنتصف عموديًا وأفقيًا باستخدام تحديد المواقع في CSS.

# --hints--

يجب أن يحتوي محدِّد `.cat-head` على خاصية `position` بقيمة `absolute`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `top` بقيمة `0`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `left` بقيمة `0`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `right` بقيمة `0`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `bottom` بقيمة `0`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `margin` بقيمة `auto`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
