---
id: 64770351e8586671ec0911f0
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

آخر قيمة لخاصية position هي `sticky`. تحديد موقع `sticky` هو مزيج من تحديد الموقع بـ `relative` و `fixed`. يسمح لعنصر أن **يلتصق** بموقع محدد داخل العنصر الحاوي أو إطار العرض، بناءً على موضع التمرير.

غيّر قيمة خاصية `position` في `.cat-head` إلى `sticky`، واضبط `top` على `0`، ثم أزل خاصية `left` وقيمتها.

**ملاحظة**: لرؤية كيفية عمل `sticky`، يجب وضع بعض النصوص قبل وبعد عنصر `div` الخاص بـ `.cat-head`. إذا قمت بالتمرير لأسفل بعد ذلك، سترى أن `.cat-head` يلتصق بالأعلى ويبقى هناك.

# --hints--

يجب أن يحتوي محدد `.cat-head` على خاصية `position` مضبوطة على `sticky`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

يجب أن يحتوي محدد `.cat-head` على خاصية `top` مضبوطة على `0`. تأكد من إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

يجب ألا تحتوي شفرتك على خاصية `left` وقيمتها.

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
