---
id: 647d855ac12cd436059acd39
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

استخدم المحدِّد العام لإضافة `box-sizing: border-box;` إلى CSS الخاص بك. هذا يضمن أن العناصر تشمل الحشو والحدود ضمن العرض والارتفاع المحددين لها.

# --hints--

يجب أن يكون لديك المحدِّد العام (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

يجب أن يحتوي المحدِّد العام على خاصية `box-sizing` مضبوطة على `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
