---
id: 6676a8b01e56ec1a1e07c202
title: الخطوة 23
challengeType: 0
dashedName: step-23
---

# --description--

لاحظ الآن أن لديك مربعًا أبيض بحدود سميكة. تحتوي الحدود على حواف مائلة يمكن استخدامها لإنشاء مثلثات.

داخل نفس محدِّد الصنف عدّل `height` و `width` إلى `0`. اجعل الحدود اليسرى واليمنى والعلوية `transparent`.

أزل خاصية `background-color`، ويجب أن ترى مثلثًا.

# --hints--


يجب أن يكون لديك محدِّد `.cat-right-ear`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

يجب أن يحتوي محدِّد `.cat-right-ear` على خاصية `height` مضبوطة على `0`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height == '0px')
```

يجب أن يحتوي محدِّد `.cat-right-ear` على خاصية `width` مضبوطة على `0`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '0px')
```

يجب أن يحتوي محدِّد `.cat-right-ear` على خاصية `border-left` مضبوطة على `35px solid transparent`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

يجب أن يحتوي محدِّد `.cat-right-ear` على خاصية `border-right` مضبوطة على `35px solid transparent`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

يجب أن يحتوي محدِّد `.cat-right-ear` على خاصية `border-top` مضبوطة على `35px solid transparent`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid transparent')
```


يجب أن يحتوي محدِّد `.cat-right-ear` على خاصية `border-bottom` مضبوطة على `35px solid red`. لا تنسَ إضافة فاصلة منقوطة.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

يجب ألا يحتوي محدِّد `.cat-right-ear` على خاصية `background-color`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
