---
id: 66a9c006e4f58819396605b3
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

الآن بعد أن أكملت قسم الرأس، حان الوقت لإضافة <dfn>جسم الجدول</dfn>، `tbody`. جسم الجدول سيمثل جميع أسماء الطلاب ودرجاتهم.

أضف عنصر جسم جدول إلى جدولك.

# --hints--

يجب أن يكون لديك وسم افتتاحي `tbody`.

```js
assert.match(code, /<tbody>/i);
```

يجب أن يكون لديك وسم إغلاق `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

يجب أن يأتي وسم الافتتاح `tbody` قبل وسم الإغلاق `table`.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
