---
id: 66a9cb26ec6bd41cf6c82bc5
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

آخر قسم تضيفه إلى الجدول سيكون عنصر <dfn>تذييل الجدول</dfn>، `tfoot`. يُستخدم عنصر تذييل الجدول لعرض المعدل المتوسط لجميع الطلاب في الجدول.

أضف عنصر `tfoot` إلى الجدول.

داخل عنصر `tfoot`، أضف عنصر `tr`.

داخل عنصر `tr`، أضف عنصرين `td`.

يجب أن يحتوي أول عنصر `td` على النص `Average Grade`. ويجب أن يحتوي ثاني عنصر `td` على النص `78.8`.

# --hints--

يجب أن يكون لديك وسم افتتاحي لـ `tfoot`.

```js
assert.match(code, /<tfoot>/);
```

يجب أن يكون لديك وسم إغلاق لـ `tfoot`.

```js
assert.match(code, /<\/tfoot>/);
```

يجب أن يكون لديك عنصر `tr` داخل عنصر `tfoot` الخاص بك.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

يجب أن يحتوي أول عنصر `td` على النص `Average Grade`. تأكد من أنه داخل عنصر `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

يجب أن يحتوي ثاني عنصر `td` على النص `78.8`. تأكد من أنه داخل عنصر `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

يجب أن يحتوي عنصر `tr` على عنصرين `td` بداخله.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
