---
id: 66a9b48a211a73155621d0df
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتدرب على العمل مع جداول HTML من خلال بناء جدول اختبار نهائي لمجموعة من الطلاب.

في الدروس السابقة، تعلمت كيفية العمل مع عنصر `table` لتمثيل البيانات الجدولية.

داخل عنصر `body` الخاص بك، قم بتسكين عنصر `table`.

# --hints--

يجب أن يكون لديك وسم افتتاحي `table`.

```js
assert.match(code, /<table>/i);
```

يجب أن يكون لديك وسم إغلاق `table`.

```js
assert.match(code, /<\/table>/i);
```

يجب أن يكون عنصر `table` الخاص بك داخل عنصر `body`.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
