---
id: 66a9ca099e3ecb1be10a2696
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

حان الوقت الآن لإضافة طالبين آخرين إلى الجدول.

باتباع نفس النمط في الخطوات السابقة، أضف صفًا رابعًا لطلاب الجدول. استخدم البيانات التالية لعناصر بيانات الجدول:

- اسم العائلة: `Thompson`
- الاسم الأول: `Jane`
- الدرجة: `77`

لصف الطالب الخامس في الجدول، استخدم البيانات التالية لعناصر بيانات الجدول:

- اسم العائلة: `Williams`
- الاسم الأول: `Natalie`
- الدرجة: `83`

# --hints--

يجب أن يكون لديك عنصر `tr` رابع داخل عنصر `tbody` الخاص بك.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Thompson` داخل عنصر `tr` الرابع.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Jane` داخل عنصر `tr` الرابع.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `77` داخل عنصر `tr` الرابع.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

يجب أن يكون لديك عنصر `tr` خامس داخل عنصر `tbody` الخاص بك.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Williams` داخل عنصر `tr` الخامس.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Natalie` داخل عنصر `tr` الخامس.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `83` داخل عنصر `tr` الخامس.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
