---
id: 66a9b8f14b963916a3baa732
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

لإضافة تسمية توضيحية إلى `table`، يمكنك استخدام عنصر <dfn>تسمية توضيحية الجدول</dfn>.

إليك مثالًا يستخدم عنصر `caption`:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

داخل عنصر `table`، قم بتسكين عنصر `caption` مع النص `Calculus Final Exam Grades`.

# --hints--

يجب أن يكون لديك وسم `caption` افتتاحي.

```js
assert.match(code, /<caption>/i);
```

يجب أن يكون لديك وسم `caption` إغلاق.

```js
assert.match(code, /<\/caption>/i);
```

يجب أن يكون عنصر `caption` داخل عنصر `table`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

يجب أن يحتوي `caption` على النص `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
