---
id: 66a9c7a4521d2b1b1ec6dcf0
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

حان الوقت الآن لإضافة طالبين آخرين إلى الجدول.

باتباع نفس النمط في الخطوة السابقة، أضف صف جدول طالب ثاني. استخدم البيانات التالية لعناصر بيانات الجدول:

- اسم العائلة: `Doe`
- الاسم الأول: `Samantha`
- الدرجة: `92`

لصف جدول الطالب الثالث، استخدم البيانات التالية لعناصر بيانات الجدول:

- اسم العائلة: `Rodriguez`
- الاسم الأول: `Marcus`
- الدرجة: `88`

# --hints--

يجب أن يكون لديك عنصر `tr` ثاني داخل عنصر `tbody` الخاص بك.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Doe` داخل عنصر `tr` الثاني.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Samantha` داخل عنصر `tr` الثاني.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `92` داخل عنصر `tr` الثاني.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

يجب أن يكون لديك عنصر `tr` ثالث داخل عنصر `tbody` الخاص بك.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Rodriguez` داخل عنصر `tr` الثالث.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Marcus` داخل عنصر `tr` الثالث.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `88` داخل عنصر `tr` الثالث.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
