---
id: 66a9bb8578459a175432e7d0
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

في القسم الأول من الجدول، عليك تجميع محتوى القسم العلوي الذي يمثل مسميات الأعمدة للاسم الأول للطالب، واسم العائلة، ودرجة الاختبار النهائي.

يُستخدم عنصر <dfn>القسم العلوي للجدول</dfn>، `thead`، لتجميع محتوى القسم العلوي في الجدول.

إليك مثالًا يستخدم عنصر `thead`:

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

أسفل عنصر `caption` الخاص بك، أضف عنصر القسم العلوي للجدول.

# --hints--

يجب أن يحتوي على وسم افتتاحي `thead`.

```js
assert.match(code, /<thead>/i);
```

يجب أن يحتوي على وسم إغلاق `thead`.

```js
assert.match(code, /<\/thead>/i);
```

يجب أن يكون عنصر `thead` داخل عنصر `table` الخاص بك.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
