---
id: 66a9bcf00f13a418368a272e
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

يتكون عنصر القسم العلوي للجدول من عنصر <dfn>صف الجدول</dfn>، `tr`، الذي يحتوي على عناصر <dfn>خلية رأس الجدول</dfn>، `th`.

إليك مثالًا يستخدم العنصرين `tr` و `th` لجدول رياضي:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

داخل عنصر `thead`، أضف عنصر `tr`.

داخل عنصر `tr`، أضف ثلاثة عناصر `th`.

يجب أن يحتوي أول عنصر `th` على النص `Last Name`. يجب أن يحتوي ثاني عنصر `th` على النص `First Name`. يجب أن يحتوي ثالث عنصر `th` على النص `Grade`.

# --hints--

يجب أن يكون لديك وسم افتتاحي `tr`.

```js
assert.match(code, /<tr>/i);
```

يجب أن يكون لديك وسم إغلاق `tr`.

```js
assert.match(code, /<\/tr>/i);
```

يجب أن يكون عنصر `tr` داخل عنصر `thead`.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

يجب أن يحتوي أول عنصر `th` على النص `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

يجب أن يحتوي ثاني عنصر `th` على النص `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

يجب أن يحتوي ثالث عنصر `th` على النص `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

يجب أن تكون كل عناصر `th` داخل عنصر `tr`.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
