---
id: 66a9c14e3b34c719e34bae20
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

لإضافة بيانات الطلاب إلى الجدول، تحتاج إلى استخدام عناصر صف الجدول وبيانات الجدول.

عنصر <dfn>بيانات الجدول</dfn>، `td`، يُستخدم لإنشاء خلية في الجدول.

إليك مثالًا على استخدام عنصر `td` لجدول لاعبي الرياضة:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

داخل عنصر جسم الجدول الخاص بك، أضف عنصر صف جدول، `tr`، مع ثلاثة عناصر بيانات جدول، `td`.

يجب أن يحتوي أول عنصر بيانات جدول على اسم العائلة `Davis`.

يجب أن يحتوي ثاني عنصر بيانات جدول على الاسم الأول `Alex`.

يجب أن يحتوي ثالث عنصر بيانات جدول على الدرجة `54`.

# --hints--

يجب أن يكون لديك عنصر `tr` داخل عنصر `tbody` الخاص بك.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Davis` داخل عنصر `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `Alex` داخل عنصر `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

يجب أن يكون لديك عنصر `td` يحتوي على النص `54` داخل عنصر `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

يجب أن تكون جميع عناصر `td` الثلاثة داخل عنصر `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
