---
id: 66a9cd4ffa48cf1ef7333640
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

جدولك شبه مكتمل لكن هناك شيء أخير يجب إضافته.

سيكون من الجيد أن يمتد عنصر `td` المستخدم لـ `Average Grade` عبر عمودين بدلًا من عمود واحد فقط.

للقيام بذلك، يمكنك استخدام الخاصية <dfn>colspan</dfn> على عنصر `td`.

تحدد الخاصية `colspan` عدد الأعمدة التي يجب أن يمتد عبرها الخلية.

إليك مثالًا على استخدام الخاصية `colspan` في جدول رياضي:

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

داخل وسم الفتح `td`، أضف الخاصية `colspan` واضبطها على القيمة `"2"`.

وبهذا التغيير، يصبح جدولك مكتملًا!

# --hints--

لا يجب تعديل محتوى عنصر `td` الموجود.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child').textContent.trim(), 'Average Grade');
```

يجب أن يحتوي عنصر `td` على خاصية `colspan`.

```js
assert.isTrue(document.querySelector('tfoot tr td:first-child').hasAttribute('colspan'));
```

يجب أن يحتوي عنصر `td` على خاصية `colspan` بقيمة `"2"`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child').getAttribute('colspan'), '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>

      <tfoot>
        <tr>
          --fcc-editable-region--
          <td>Average Grade</td>
          --fcc-editable-region--
          <td>78.8</td>
        </tr>
      </tfoot>
    </table>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>

      <tfoot>
        <tr>
          <td colspan="2">Average Grade</td>
          <td>78.8</td>
        </tr>
      </tfoot>
    </table>
  </body>
</html>
```
