---
id: 669e2f60e83c011754f711f9
title: 旅行代理店ページを作成する
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `DOCTYPE` 宣言が必要です。
1. `lang` 属性が `en` に設定された `html` 要素が必要です。
1. `charset` 属性が `utf-8` に設定された `meta` 空要素と、テキストが `Travel Agency Page` の `title` を含む `head` 要素が必要です。
1. SEO のためにウェブサイトの簡単な説明を含む `meta` タグが `head` 要素内に必要です。
1. 旅行先を紹介する `h1` 要素が必要です。ページ内で唯一の `h1` 要素でなければなりません。
1. `h1` 要素の下に旅行の機会を紹介する段落が必要です。
1. テキストが `Packages` の `h2` 要素が必要です。
1. さまざまなパッケージを簡単に紹介する `p` 要素が必要です。
1. 2つのリスト項目を持つ番号なしリスト要素が必要です。リスト項目のテキストはそれぞれ `Group Travels` と `Private Tours` で、各リスト項目のテキストはアンカー要素で囲まれている必要があります。
1. テキストが `Top Itineraries` の `h2` 要素が必要です。
1. 少なくとも3つの `figure` 要素が必要で、それぞれにアンカー要素と図の画像を説明する `figcaption` 要素が含まれている必要があります。
1. 3つのアンカー要素には、それぞれ適切な `alt` 属性と有効な画像の `src` 属性を持つ `img` 要素が含まれている必要があります。必要に応じて `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`、`https://cdn.freecodecamp.org/curriculum/labs/alps.jpg`、`https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` を使用できます。
1. 5つすべてのアンカー要素に `href` 属性が `https://www.freecodecamp.org/learn`、`target` 属性が `_blank` に設定されている必要があります。

# --hints--

旅行代理店ページには `!DOCTYPE html` 宣言が必要です。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

`html` に設定された `lang` を持つ `en` 要素が必要です。

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

`html` 要素は終了タグを持つ必要があります。

```js
assert.match(code, /<\/html\s*>/i);
```

`head` 要素内に `html` 要素が必要です。

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

あなたの `head` 要素には終了タグが必要です。

```js
assert.match(code, /<\/head\s*>/i);
```

あなたの `head` 要素内に `title` 要素が必要です。

```js
assert.isNotNull(document.querySelector('head title'));
```

あなたの `title` 要素には終了タグが必要です。

```js
assert.match(code, /<\/title\s*>/i);
```

`meta` 要素内に2つの `head` 要素が必要です。

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

1つの `meta` 要素は、値が `description` の `name` 属性と、空でない `content` 属性を持つ必要があります。

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

1つの `meta` 要素の `charset` 属性は `utf-8` に設定されている必要があります。

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

あなたの `title` 要素には旅行代理店の名前を入れてください。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

`body` 要素内に `html` 要素が必要です。

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

あなたの `body` 要素には終了タグが必要です。

```js
assert.match(code, /<\/body\s*>/i);
```

旅行先を紹介するための `h1` 要素を用意してください。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

`h1` 要素は1つだけにしてください。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`h1` 要素の下に `p` 要素を配置してください。

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

最初の`p`要素は旅行の機会を紹介してください。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

最初の`h2`要素には`Packages`というテキストが含まれている必要があります。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

最初の`h2`要素の下に`p`要素を配置してください。

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

2番目の`p`要素は、さまざまなパッケージを簡単に紹介してください。

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

2つ目の`p`要素の下に番号なしリスト要素を配置してください。

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

番号なしリストには2つの項目が必要です。

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

両方のリスト項目にはアンカー要素を含めてください。

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

最初のリスト項目のアンカー要素は、テキスト `Group Travels` を囲む必要があります。

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

2番目のリスト項目のアンカー要素は、テキスト `Private Tours` を囲む必要があります。

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

番号なしリストの後に `h2` 要素を配置してください。

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

2つ目の `h2` 要素には `Top Itineraries` というテキストを入れてください。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

少なくとも3つの `figure` 要素を用意してください。

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

各 `figure` 要素の最初のネストされた要素としてアンカー要素を含めてください。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

各 `figure` 要素の2番目のネストされた要素として `figcaption` 要素を含めてください。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

各 `figcaption` 要素にはテキストを含めてください。

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

各 `figure` 要素にネストされた `a` 要素には画像を含めてください。

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

各 `img` 要素には有効な `src` 属性を設定してください。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

各 `img` 要素には適切な値の `alt` 属性を設定してください。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

各 `a` 要素には `href` 属性として `https://www.freecodecamp.org/learn` を設定してください。リスト項目内のリンクも忘れないでください。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

各 `a` 要素には `target` 属性として `_blank` を設定してください。リスト項目内のリンクも忘れないでください。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
