---
id: 642df9df4b5216350de7b0d2
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

ستحتاج إلى دالة لتوليد نطاق من الأرقام.

أعلن عن دالة `range` فارغة تأخذ معلمتي `start` و `end`. استخدم دالة البناء `Array()` وأرجع مصفوفة فارغة ضمنيًا.

# --hints--

يجب إعلان متغير `range`.

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

يجب استخدام `const` لإعلان متغير `range`.

```js
assert.match(code, /const\s+range/);
```

يجب أن يكون متغير `range` دالة.

```js
assert.isFunction(range);
```

يجب أن تستخدم دالة `range` صيغة السهم.

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

يجب أن تأخذ دالة `range` المعلمة `start` أولاً.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

يجب أن تأخذ دالة `range` المعلمة `end` ثانيًا.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

يجب أن تستخدم دالة `range` إرجاعًا ضمنيًا. تذكر أن هذا يعني عدم استخدام الأقواس المعقوفة.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

يجب أن تستخدم دالة `range` دالة البناء `Array()`. دوال البناء الأولية لا تحتاج إلى كلمة `new`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

يجب ألا تمرر أي شيء إلى دالة البناء `Array()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
