---
id: 642db8c409d9991d0b3b2f0d
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

يبدأ مشروعك بحاوية HTML أساسية وبعض CSS المقابلة. ستكون مهمتك الأولى هي توليد الخلايا لجدول البيانات الخاص بك برمجيًا.

يمثل الكائن العالمي `window` نافذة المتصفح (أو التبويب). لديه خاصية `onload` التي تسمح لك بتعريف سلوك عند تحميل النافذة للصفحة كاملة، بما في ذلك أوراق الأنماط والسكريبتات.

ابدأ بتعيين خاصية `onload` في `window` إلى دالة سهمية بدون معلمات. في الدالة، أعلن عن متغير `container` وقم بتعيين قيمته بالحصول على العنصر بواسطة `id` الخاص بـ `"container"`.

# --hints--

يجب أن تصل إلى خاصية `onload` في كائن `window`.

```js
assert.match(code, /window\.onload/);
```

يجب أن تعيّن خاصية `onload` إلى دالة.

```js
assert.isFunction(window.onload);
```

يجب أن تستخدم صيغة الدالة السهمية.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

يجب ألا تأخذ دالة `onload` أي معلمات.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

يجب أن تعلن عن متغير `container` داخل دالة `onload` الخاصة بك.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

يجب إعلان متغير `container` باستخدام `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

يجب أن تستخدم `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

يجب أن تحصل على العنصر الذي يحمل `id` بقيمة `"container"`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

يجب أن تعيّن عنصر `#container` إلى المتغير `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
