---
id: 642ddfdea4200e313f80a4b6
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

تذكر أن كائن `document` يحتوي على طريقة `.createElement()` التي تسمح لك بإنشاء عناصر HTML جديدة بشكل ديناميكي.

في دالتك `createLabel`، أعلن عن متغير `label` وقم بتعيين عنصر `div` جديد له.

# --hints--

يجب أن تعلن عن متغير `label` في دالتك `createLabel`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

يجب إعلان متغير `label` باستخدام `const`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

يجب أن تستخدم طريقة `.createElement()` الخاصة بكائن `document`.

```js
assert.match(code, /document\.createElement\(/);
```

يجب أن تمرر السلسلة النصية `"div"` إلى طريقة `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

يجب أن تعين عنصر `div` الجديد إلى المتغير `label`.

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
