---
id: 642e02be7845f13b014cd2b0
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

الآن بعد أن لديك دالة `range`، يمكنك استخدامها لإنشاء نطاق من الحروف أيضًا.

أعلن عن دالة `charRange` باستخدام `const` وصيغة السهم. يجب أن تأخذ معلمتين `start` و `end`. يجب أن تُرجع الدالة ضمنيًا نتيجة استدعاء `range()` مع `start` و `end` كمعلمات.

# --hints--

يجب أن تعلن عن متغير `charRange`.

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

يجب إعلان متغير `charRange` باستخدام `const`.

```js
assert.match(code, /const\s+charRange/);
```

يجب أن يكون متغير `charRange` دالة.

```js
assert.isFunction(charRange);
```

يجب أن تستخدم دالة `charRange` صيغة السهم.

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

يجب أن تأخذ دالة `charRange` المعلمة `start` كأول معلمة.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

يجب أن تأخذ دالة `charRange` المعلمة `end` كالمعلمة الثانية.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

يجب أن تستخدم دالة `charRange` إرجاعًا ضمنيًا.

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

يجب أن تستدعي دالة `charRange` دالتك `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

يجب أن تمرر `start` و `end` كمعلمات لاستدعاء `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
