---
id: 643498755d54c6279ba09078
title: الخطوة 20
challengeType: 0
dashedName: step-20
---

# --description--

تتضمن معظم برامج جداول البيانات دوالًا مدمجة للحساب.

أعلن دالة `sum` التي تأخذ معلمة `nums`، والتي ستكون مصفوفة من الأرقام. يجب أن تُرجع نتيجة استدعاء `reduce` على المصفوفة لجمع كل الأرقام.

# --hints--

يجب أن تعلن متغيرًا باسم `sum`.

```js
assert.match(code, /(?:let|const|var)\s+sum/);
```

يجب أن تستخدم `const` لإعلان متغير `sum` الخاص بك.

```js
assert.match(code, /const\s+sum/);
```

يجب أن يكون متغير `sum` دالة.

```js
assert.isFunction(sum);
```

يجب أن تستخدم دالة `sum` صيغة السهم.

```js
assert.match(code, /const\s+sum\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

يجب أن تحتوي دالة `sum` على معلمة `nums`.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>/);
```

يجب أن تستخدم دالة `sum` إرجاعًا ضمنيًا.

```js
assert.notMatch(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

يجب أن تُرجع دالة `sum` نتيجة استدعاء `.reduce()` على `nums`.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*nums\.reduce\(/);
```

يجب أن تُرجع دالة `sum` مجموع كل الأرقام في `nums`.

```js
const numbers = [1, 2, 3, 4, 5];
assert.equal(sum(numbers), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
