---
id: 6434750c53db16218f41e6e1
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

تذكر أن `range()` تُرجع مصفوفة، لذا يمكنك ربط طرق المصفوفة مباشرة باستدعاء الدالة.

استدعِ `range()` مع `1` و `99` كمعلمتين، واربط طريقة `.forEach()`. مرر إلى `.forEach()` دالة استدعاء راجع فارغة تأخذ `number` كمعلمة.

# --hints--

يجب أن تستدعي دالة `range()` الخاصة بك.

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

يجب أن تمرر `1` كأول معلمة لاستدعاء `range()` الخاص بك.

```js
assert.match(code, /range\(\s*1/);
```

يجب أن تمرر `99` كمعلمة ثانية لاستدعاء `range()` الخاص بك.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

يجب أن تربط طريقة `.forEach()` باستدعاء `range()` الخاص بك.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

يجب أن تمرر دالة استدعاء راجع إلى `.forEach()` باستخدام صيغة السهم.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

يجب أن تحتوي دالة الاستدعاء الراجع على `number` كمعلمة وحيدة.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
