---
id: 6434759f78ec812264ff8f34
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

في استدعاءك الراجع، ستحتاج إلى إجراء استدعائين لدالة. ابدأ باستدعاء `createLabel()` ومرر `number` كمعلمة. يجب أن ترى بعض الأرقام تظهر في جدول البيانات الخاص بك.

ثم استدعِ طريقة `.forEach()` على مصفوفة `letters`. مرر دالة استدعاء راجع فارغة تأخذ معلمة `letter`.

# --hints--

يجب أن تستدعي دالتك `createLabel()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*createLabel\(/);
```

يجب أن تمرر `number` إلى استدعاء `createLabel()` الخاص بك.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*createLabel\(/)
```

يجب أن تستدعي طريقة `.forEach()` على مصفوفة `letters`.

```js
assert.lengthOf(code.match(/letters\.forEach\(/g), 2)
```

يجب أن تمرر دالة استدعاء راجع بصيغة السهم إلى طريقة `.forEach()` الخاصة بك.

```js
assert.match(code, /letters\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>\s*\{/)
```

يجب أن تحتوي دالة الاستدعاء الراجع على معلمة `letter`.

```js
assert.match(code, /letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{/)
```

يجب أن تكون دالة الاستدعاء الراجع فارغة.

```js
assert.match(code, /letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{\s*\}/)
```

يجب أن تكون دالة الاستدعاء الراجع لـ `letters.forEach()` مسكنة داخل دالة الاستدعاء الراجع لـ `range(1, 99).forEach(number => {}`.

```js
assert.match(code, /range\s*\(\s*1\s*,\s*99\s*\)\s*.forEach\s*\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*[^}]*letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{\s*\}\s*\)\s*;?\s*\}\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--
  range(1, 99).forEach(number => {

  })
--fcc-editable-region--
}
```
