---
id: 643715013330824ecaa70442
title: الخطوة 23
challengeType: 0
dashedName: step-23
---

# --description--

ستحسب الدالة التالية القيمة الوسيطة لمصفوفة من الأرقام. ابدأ بإعلان دالة سهمية `median` تأخذ معلمة `nums`.

في الدالة، أعلن متغيرًا `sorted` وعيّنه بقيمة فرز نسخة من مصفوفة `nums`.

يجب استخدام طريقة `slice()` لإنشاء نسخة سطحية من المصفوفة.

# --hints--

يجب إعلان متغير `median`.

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

يجب استخدام `const` لإعلان متغير `median`.

```js
assert.match(code, /const\s+median/);
```

يجب أن يكون متغير `median` دالة.

```js
assert.isFunction(median);
```

يجب أن تستخدم دالة `median` صيغة الدالة السهمية.

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

يجب أن تحتوي دالة `median` على معلمة `nums`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

يجب ألا تستخدم دالة `median` إرجاعًا ضمنيًا.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

يجب أن تحتوي دالة `median` على متغير `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

يجب استخدام `const` لإعلان متغير `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

يجب استخدام `.slice()` لتعيين نسخة من مصفوفة `nums` إلى `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

يجب ربط طريقة `.sort()` بطريقة `.slice()`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

يجب تمرير دالة استدعاء راجع إلى طريقة `sort` لترتيب الأرقام بدقة بترتيب تصاعدي. استخدم الإرجاع الضمني للوضوح.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
