---
id: 64496df724dd3716a71fe971
title: الخطوة 26
challengeType: 0
dashedName: step-26
---

# --description--

تتكون خصائص الكائن من أزواج مفتاح/قيمة. يمكنك استخدام أسماء الخصائص المختصرة عند إعلان كائن حرفي. عند استخدام صيغة الخاصية المختصرة، يصبح اسم المتغير هو مفتاح الخاصية وقيمته هي قيمة الخاصية.

المثال التالي يعلن عن كائن `user` يحتوي على الخصائص `userId` و`firstName` و`loggedIn`.

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

لتتبع جميع دوال جدول البيانات الخاص بك، أعلن عن كائن `spreadsheetFunctions`. باستخدام صيغة الخاصية المختصرة، عيّن `sum` و`average` و`median` كخصائص على كائن `spreadsheetFunctions`.

# --hints--

يجب إعلان متغير `spreadsheetFunctions`.

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

يجب استخدام `const` لإعلان متغير `spreadsheetFunctions`.

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

يجب أن يكون متغير `spreadsheetFunctions` كائنًا.

```js
assert.isObject(spreadsheetFunctions);
```

يجب أن يحتوي كائن `spreadsheetFunctions` على خاصية `sum`.

```js
assert.property(spreadsheetFunctions, "sum");
```

يجب أن تكون خاصية `sum` هي دالة `sum` الخاصة بك.

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

يجب أن يحتوي كائن `spreadsheetFunctions` على خاصية `average`.

```js
assert.property(spreadsheetFunctions, "average");
```

يجب أن تكون خاصية `average` هي دالة `average` الخاصة بك.

```js
assert.equal(spreadsheetFunctions?.average, average);
```

يجب أن يحتوي كائن `spreadsheetFunctions` على خاصية `median`.

```js
assert.property(spreadsheetFunctions, "median");
```

يجب أن تكون خاصية `median` هي دالة `median` الخاصة بك.

```js
assert.equal(spreadsheetFunctions?.median, median);
```

يجب استخدام صيغة التفكيك لتعيين خصائصك.

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
