---
id: a103376db3ba46b2d50db289
title: Implementa un convertidor a Spinal Case
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

El spinal case es un formato de cadena donde todas las palabras están en minúsculas y separadas por guiones. `"this-is-spinal-tap"` es un ejemplo de una cadena en spinal case.

En este laboratorio, crearás una función que convierte una cadena dada a spinal case.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `spinalCase`.
2. La función `spinalCase` debe tomar un solo argumento, una cadena.
3. La función `spinalCase` debería devolver la cadena en formato spinal case. Por ejemplo, si el argumento es `JavaScript is awesome`, la función debería devolver `javascript-is-awesome`.

# --hints--

Debes crear una función llamada `spinalCase`.

```js
assert.isFunction(spinalCase);
```

`spinalCase` debe tomar un solo argumento.

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` debe devolver la cadena `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` debe devolver la cadena `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` debe devolver la cadena `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` debe devolver la cadena `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` debe devolver la cadena `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
