---
id: 686d749d6f32f708fdf2b7ad
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En una lección anterior, aprendiste sobre los elementos `svg` y cómo a menudo se usan como íconos en proyectos. En bases de código reales, normalmente confiarías en bibliotecas de íconos, por lo que no necesitas crear elementos `svg` desde cero. Sin embargo, en este taller, construirás un ícono de corazón para aprender sobre los atributos principales que se usan dentro de un elemento `svg`.

Comienza creando un elemento `svg` en la página.

# --hints--

Debes tener un elemento `svg` en la página.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
