---
id: 60eebd07ea685b0e777b5583
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

欢迎来到注册表单项目！ 首先，在文档的顶部添加 `!DOCTYPE html` 声明，这样浏览器就知道它在解析什么类型的文档。

# --hints--

你的代码应该包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你应在 `DOCTYPE` 后添加一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你应将文档类型定义为 `html` 。

```js
assert(code.match(/html/gi));
```

你应该在 `DOCTYPE` 声明的 `type` 后用 `>` 关闭。

```js
assert(code.match(/html\s*>/gi));
```

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
