---
id: 60f030d388cb74067cf291c3
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

在标题下方，添加如下文本的段落元素来鼓励用户注册：

```md
Please fill out this form with the required information
```

# --hints--

你应该在 `body` 中添加一个 `p` 元素。

```js
assert.exists(document.querySelector('body > p'));
```

你应该在 `h1` 下方添加 `p` 元素。

```js
assert.exists(document.querySelector('h1 + p'));
```

你应该给 `p` 元素添加文本 `Please fill out this form with the required information`。

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>

  </body>
--fcc-editable-region--
</html>
```
