---
id: 5f33071498eb2472b87ddee4
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS is the language used to style an HTML document. It describes how HTML elements should be displayed on the screen.

As you learned in the last few steps of the Cat Photo App, there is a basic structure needed to start building your web page. 每個 HTML 文檔都應包含 `DOCTYPE` 聲明和 `html` 元素。 `DOCTYPE` 告訴瀏覽器文檔的 HTML 版本。 而 `html` 元素代表根元素，它包含所有其他元素。

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Add the `<!DOCTYPE html>` tag, and an `html` element with a `lang` attribute of `en`.

# --hints--

你應該有 `<!DOCTYPE html>` 聲明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你應該有一個 `<html lang="en">` 開始標籤，語言設置爲英語。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你應該有一個 `</html>` 結束標籤。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
