---
id: 69757cc0faae0152c1418aad
title: خطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستتدرب على العمل مع الأعداد والعمليات الرياضية لبناء أداة تقسيم الفاتورة. ستُحسب هذه الأداة مقدار ما يدين به كل شخص بعد إضافة تكاليف الوجبة والبقشيش.

لبداية العمل، تحتاج إلى طريقة لتتبع المجموع الكلي أثناء إضافة التكاليف.

أنشئ متغيرًا باسم `running_total` وقم بتعيين القيمة `0` له.

# --hints--

يجب أن يكون لديك متغير باسم `running_total`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

يجب تعيين العدد الصحيح `0` إلى متغير `running_total` الخاص بك.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
