---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

لكل دورة تكلفة. تحتاج إلى تخزين هذه المبالغ في متغيرات لاستخدامها لاحقًا. بما أن هذه المبالغ تشمل السنتات، ستستخدم نوع float، الذي يُستخدم لتمثيل الأعداد العشرية. إليك مثالًا على متغير بقيمة float:

```py
change = 2.35
```

أنشئ أربعة متغيرات: `appetizers` بقيمة `37.89`، و`main_courses` بقيمة `57.34`، و`desserts` بقيمة `39.39`، و`drinks` بقيمة `64.21`.

# --hints--

يجب تعريف متغير باسم `appetizers` وتعيين القيمة `37.89` له.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

يجب تعريف متغير باسم `main_courses` وتعيين القيمة `57.34` له.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

يجب تعريف متغير باسم `desserts` وتعيين القيمة `39.39` له.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

يجب تعريف متغير باسم `drinks` وتعيين القيمة `64.21` له.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
