---
id: 6976cc0b0c135686a4620b94
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

بعد ذلك، تحتاج إلى احتساب عدد الأشخاص الذين يشاركون الفاتورة. خزّن هذه القيمة في متغير، كما فعلت في الخطوة السابقة.

أنشئ متغيرًا باسم `num_of_friends` وقم بتعيين القيمة `4` له. سيُستخدم هذا لاحقًا في ورشة العمل لحساب التقسيم النهائي.

# --hints--

يجب أن يكون لديك متغير باسم `num_of_friends`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

يجب أن تعيّن العدد الصحيح `4` إلى متغير `num_of_friends`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
