---
id: 69779dcb44016eccc05c15a4
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

كان الخدمة ممتازة، لذلك قرر الفريق ترك بقشيش بنسبة 25%. لحساب النسبة المئوية في Python، يمكنك الضرب في المجموع بالقيمة العشرية المكافئة للنسبة المئوية.

على سبيل المثال، لإيجاد 10% من قيمة، تضربها في `0.10` باستخدام العامل `*`:

```py
tax = total * 0.10
```

أنشئ متغيرًا باسم `tip` وقم بتعيين نتيجة ضرب `running_total` في `0.25` له.

أخيرًا، استخدم `print()` لعرض السلسلة النصية `Tip amount:` متبوعة بمسافة وقيمة المتغير `tip` الخاص بك.

# --hints--

يجب أن يكون لديك متغير باسم `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

يجب أن يكون متغير `tip` هو نتيجة `running_total * 0.25`.

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

يجب أن تطبع السلسلة النصية `Tip amount:` متبوعة بمسافة والمتغير `tip`.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Tip amount: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - ((37.89 + 57.34 + 39.39 + 64.21) * 0.25)) < 1e-6
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
