---
id: 6977a8bbcbdd557886bd0718
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

مع إضافة البقشيش الآن، لديك المبلغ النهائي للمجموعة بأكملها. عليك تحديد مقدار ما يدين به كل شخص بقسمة إجمالي الفاتورة على عدد الأصدقاء.

في Python، تستخدم الشرطة المائلة الأمامية `/` لإجراء القسمة. على سبيل المثال:

```py
half = 10 / 2
```

أنشئ متغيرًا باسم `final_bill` وقم بتعيين نتيجة قسمة `running_total` على `num_of_friends` له.

أخيرًا، استخدم الدالة `print()` لعرض السلسلة النصية `Bill per person:` متبوعة بمسافة وقيمة `final_bill`.

# --hints--

يجب أن تعرف متغيرًا باسم `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

يجب أن تستخدم العامل `/` لقسمة `running_total` على `num_of_friends` وتعيين النتيجة إلى متغير `final_bill` الخاص بك.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

يجب أن تطبع السلسلة النصية `Bill per person:` متبوعة بمسافة ومتغير `final_bill`.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Bill per person: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - (((37.89 + 57.34 + 39.39 + 64.21) * 1.25) / 4)) < 1e-6
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
