---
id: 6982684f3a25f379e195a5fc
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

الآن بعد أن خزنت التكاليف الفردية، يمكنك حساب الإجمالي.

تذكر أن العامل `+=` يضيف قيمة إلى متغير موجود ويُحدّثه في نفس الوقت. على سبيل المثال:

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

استخدم العامل `+=` مرة واحدة لإضافة `appetizers` و`main_courses` و`desserts` و`drinks` إلى `running_total`.

أخيرًا، استخدم `print()` لعرض السلسلة النصية `Total bill so far:` متبوعة بمسافة وقيمة `running_total`.

**ملاحظة:** قد تلاحظ أن الناتج يحتوي على أرقام عشرية أكثر مما توقعت. كما تعلمت في درس سابق، يحدث هذا لأن الأرقام تُخزن بصيغة ثنائية، والعديد من القيم العشرية لا يمكن تمثيلها بدقة في هذا التنسيق، مما يؤدي إلى أخطاء تقريب.

# --hints--

يجب تحديث `running_total` باستخدام العامل `+=` مرة واحدة لإضافة تكاليف الدورات الأربعة كلها.

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total += {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

يجب طباعة السلسلة النصية `Total bill so far:` متبوعة بمسافة وقيمة `running_total`.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Total bill so far: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - (37.89 + 57.34 + 39.39 + 64.21)) < 1e-6
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
