---
id: 6842a6cd9836f0114a5b7a8a
title: 数字パターンジェネレーターを作成する
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

このラボでは、Pythonの基本を練習し、数字のパターンを作成する小さなアプリを作ります。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `number_pattern` という名前の関数を定義し、1つのパラメータ`n`（正の整数を表す）を受け取る必要があります。
1. `number_pattern` は`for` ループを使う必要があります。
1. `number_pattern(n)` は1から`n`（含む）までのすべての整数をスペースで区切った文字列を返す必要があります。例えば、`number_pattern(4)` は文字列`1 2 3 4` を返すべきです。
1. 関数に渡された引数が整数値でない場合、関数は`Argument must be an integer value.` を返す必要があります。
1. 関数に渡された引数が1未満の場合、関数は`Argument must be an integer greater than 0.` を返す必要があります。

# --hints--

`number_pattern` 関数を用意してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

`number_pattern` 関数は`n` という名前の1つのパラメータを持つべきです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` は`1 2 3 4` を返すべきです。

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` は`1 2 3 4 5 6 7 8 9 10 11 12` を返すべきです。

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` は任意の正の整数に対してスペース区切りの数字リストを返すべきです。

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

`number_pattern` は整数でない値が渡された場合に`Argument must be an integer value.` を返すべきです。

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` は正でない整数が渡された場合に`Argument must be an integer greater than 0.` を返すべきです。

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
