---
id: ab6137d4e35944e21037b769
title: Construye un convertidor a Title Case
challengeType: 26
dashedName: build-a-title-case-converter
---

# --description--

En este laboratorio crearás una función que convierte una cadena a formato de título. Formato de título significa que la primera letra de cada palabra está en mayúscula y el resto de la palabra en minúscula.

"Web Development Is Awesome" es un ejemplo de una cadena en formato título.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Deberías tener una función `titleCase` que tome una cadena como argumento.
1. La función `titleCase` debe devolver una cadena con la primera letra de cada palabra en mayúscula y el resto de la palabra en minúscula.
1. `titleCase("I like to code")` debería devolver `"I Like To Code"`.
1. `titleCase("javaScript is fun")` debería devolver `"Javascript Is Fun"`.

# --hints--

Debes crear una función llamada `titleCase`.

```js
assert.isFunction(titleCase);
```

`titleCase` debe tomar un solo argumento.

```js
assert.lengthOf(titleCase, 1);
```


`titleCase("I'm a little tea pot")` debe devolver una cadena.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` debe devolver la cadena `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` debe devolver la cadena `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` debe devolver la cadena `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}
```
