---
id: 66db57ad34c7089b9b41bfd6
title: بناء عارض Lightbox
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

يُستخدم Lightbox في مواقع الويب لعرض المحتوى المتعدد الوسائط بطريقة أكثر جاذبية من خلال نافذة منبثقة أو نافذة مودال فوق المحتوى الرئيسي للصفحة.

في هذا المختبر، ستنشئ معرض Lightbox يعرض الصور بالحجم الكامل عند النقر على صورة مصغرة. لكل صورة، هناك نسختان: صورة مصغرة وصورة بالحجم الكامل. الصورة بالحجم الكامل هي نفسها الصورة المصغرة، لكن بدون اللاحقة `-thumbnail`.

**الهدف:** حقق قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك `div` بصنف `gallery` داخل عنصر `body`.
1. داخل عنصر `.gallery`، يجب أن يكون لديك ثلاث صور مصغرة، كل منها بصنف `gallery-item`. يجب استخدام الروابط التالية لصور المصغرات:
   
    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. يجب أن يحتوي المستند على عنصر `div` بصنف `lightbox` داخل عنصر `body`.
1. يجب أن يحتوي عنصر `.lightbox` على زر `button` يحمل خاصية `id` تساوي `close-btn`. يمكنك استخدام `&times;` كنص للزر إذا رغبت.
1. يجب أن يحتوي عنصر `.lightbox` على صورة `img` تحمل خاصية `id` تساوي `lightbox-image`.
1. يجب أن يكون عنصر `.lightbox` ذا تموضع ثابت بحيث تظهر المعاينة فوق الصور الحالية.
1. يجب أن يغطي عنصر `.lightbox` كامل إطار العرض من خلال ضبط الارتفاع والعرض إلى 100% من الحاوية. تأكد أن عنصر `.lightbox` يبدأ من الزاوية العلوية اليسرى للحاوية.
1. يجب أن يكون لعنصر `.lightbox` لون خلفية. في البداية، يجب ضبط خاصية `display` له إلى `none` لإخفائه.
1. عند نقر أحد عناصر `.gallery-item`، يجب ضبط خاصية `display` لعنصر `.lightbox` إلى `flex` ليظهر عنصر `.lightbox` والعنصران الموجودان بداخله.
1. عند نقر أحد عناصر `.gallery-item`، يجب ضبط خاصية `src` لعنصر `#lightbox-image` إلى نسخة بالحجم الكامل من الصورة التي تم النقر عليها، وذلك بإزالة `-thumbnail` من خاصية `src` للصورة. توجد الصور بالحجم الكامل على الروابط التالية:

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`
  
1. عندما يكون عنصر `.lightbox` مرئيًا وتنقر على زر `#close-btn` أو على عنصر `.lightbox`، يجب إعادة تعيين خاصية `display` لعنصر `.lightbox` إلى `none`.

**ملاحظة:** تأكد من ربط ورقة الأنماط وملف جافاسكريبت في ملف HTML الخاص بك.

# --hints--

يجب أن يكون لديك `div` بصنف `gallery` داخل عنصر `body`.

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

داخل عنصر `.gallery`، يجب أن يكون لديك ثلاث عناصر `img` بصنف `gallery-item`.

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

داخل عنصر `.gallery`، يجب أن يكون لديك عنصر `img` بخاصية `src` مضبوطة على `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

داخل عنصر `.gallery`، يجب أن يكون لديك عنصر `img` بخاصية `src` مضبوطة على `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

داخل عنصر `.gallery`، يجب أن يكون لديك عنصر `img` بخاصية `src` مضبوطة على `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

يجب أن يكون لديك عنصر `div` بصنف `lightbox` داخل عنصر `body`.

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

داخل عنصر `.lightbox` يجب أن يكون لديك عنصر `button` مع تعيين `id` إلى `close-btn`.

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

داخل عنصر `.lightbox` يجب أن يكون لديك عنصر `img` مع تعيين `id` إلى `lightbox-image`.

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

يجب أن يكون عنصر `.lightbox` بموقع ثابت.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

يجب أن يغطي عنصر `.lightbox` كامل إطار العرض.

```js
const style = new __helpers.CSSHelp(document).getStyle(".lightbox");

const width = style?.width;
const height = style?.height;

assert.oneOf(
  width,
  ["100%", "100vw"]
);

assert.oneOf(
  height,
  ["100%", "100vh"]
);
```

يجب أن يكون عنصر `.lightbox` محاذيًا مع الزاوية العلوية اليسرى للحاوية.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

يجب أن يكون لعنصر `.lightbox` لون خلفية.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

يجب أن يكون عنصر `.lightbox` مخفيًا في البداية.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

عند النقر على أحد عناصر `.gallery-item`، يجب تعيين خاصية `display` لعنصر `.lightbox` إلى `flex` لجعل `.lightbox` والعنصرين بداخله مرئيين.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

عند النقر على أحد عناصر `.gallery-item`، يجب تعيين خاصية `src` لعنصر `#lightbox-image` إلى نسخة بالحجم الكامل من الصورة التي تم النقر عليها عن طريق إزالة `-thumbnail` من خاصية `src` للصورة.

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

عندما يكون عنصر `.lightbox` مرئيًا وتنقر على زر `#close-btn`، يجب إعادة تعيين خاصية `display` لعنصر `.lightbox` إلى `none`.

```js
const lightbox = document.querySelector(".lightbox");
const closeBtn = document.getElementById("close-btn"); 
const galleryItem = document.querySelector(".gallery-item");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

closeBtn.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

عندما يكون عنصر `.lightbox` مرئيًا وتنقر على عنصر `.lightbox`، يجب إعادة تعيين خاصية `display` لعنصر `.lightbox` إلى `none`.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

lightbox.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```
