---
id: 5895f70cf9fc0f352b528e65
title: Passport のセットアップ
challengeType: 2
forumTopicId: 301565
dashedName: set-up-passport
---

# --description--

ついに *Passport* をセットアップして、ユーザーがアカウントを登録またはログインできるようにしましょう。Passport に加えて、セッション管理には Express-session を使います。Express-session には多くの高度な機能がありますが、今回は基本的な使い方だけを行います。このミドルウェアはクライアントにセッションIDをクッキーとして保存し、そのIDを使ってサーバー側でセッションデータにアクセスできるようにします。こうすることで、クライアントが使うクッキーに個人のアカウント情報を入れずに済み、サーバー上に保存されたデータにアクセスするための *鍵* を保持できます。

`passport@~0.4.1` と `express-session@~1.17.1` はすでにインストールされており、どちらも `package.json` ファイルの依存関係に記載されています。

セッションの設定を行い、Passport を初期化する必要があります。まず、`session` と `passport` の変数を作成し、それぞれ `express-session` と `passport` を require してください。

次に、以下のオプションを定義して Express アプリにセッションを使うよう設定します。

```javascript
app.use(session({
  secret: process.env.SESSION_SECRET,
  resave: true,
  saveUninitialized: true,
  cookie: { secure: false }
}));
```

`SESSION_SECRET` を `.env` ファイルに追加し、ランダムな値を設定することを忘れないでください。これはクッキーを暗号化するためのハッシュ計算に使われます！

これらを設定したら、Express アプリに `passport.initialize()` と `passport.session()` を **use** するように伝えてください。

正しくできたと思ったらページを提出してください。エラーが出る場合は、<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-passport-3" target="_blank" rel="noopener noreferrer nofollow">ここまで完成したプロジェクトを確認できます</a>。

# --hints--

Passport と Express-session は依存関係に含まれている必要があります。

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport',
    'Your project should list "passport" as a dependency'
  );
  assert.property(
    packJson.dependencies,
    'express-session',
    'Your project should list "express-session" as a dependency'
  );
```

依存関係は正しく require されている必要があります。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport("|')/,
    'You should have required passport'
  );
  assert.match(
    data,
    /require.*("|')express-session("|')/,
    'You should have required express-session'
  );
```

Express アプリは新しい依存関係を使うように設定してください。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /passport\.initialize/, 'Your express app should use "passport.initialize()"');
  assert.match(data, /passport\.session/, 'Your express app should use "passport.session()"');
```

セッションとセッションシークレットは正しく設定してください。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /secret *:\s*process\.env(\.SESSION_SECRET|\[(?<q>"|')SESSION_SECRET\k<q>\])/,
    'Your express app should have express-session set up with your secret as process.env.SESSION_SECRET'
  );
```
