---
id: 5895f70cf9fc0f352b528e67
title: パスポートユーザーのシリアライズを実装する
challengeType: 2
forumTopicId: 301556
dashedName: implement-the-serialization-of-a-passport-user
---

# --description--

実際のユーザーオブジェクトは読み込んでいません。なぜならデータベースが設定されていないからです。サーバー起動時に一度だけデータベースに接続し、アプリのライフサイクル全体で持続的な接続を維持してください。そのために、データベースの接続文字列（例：`mongodb+srv://<username>:<password>@cluster0-jvwxi.mongodb.net/?retryWrites=true&w=majority`）を環境変数`MONGO_URI`に追加します。これは`connection.js`ファイルで使用されます。

*MongoDB Atlasで無料のデータベースを設定する際に問題がある場合は、この<a href="https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/" target="_blank" rel="noopener noreferrer nofollow">チュートリアル</a>を参照してください。*

次に、データベースに接続してからリクエストの待機を開始したいです。これは、データベースが接続されていない場合やエラーがある場合にリクエストを受け付けないようにするためです。これを実現するために、シリアライズ処理とアプリのルートを次のコードで囲んでください：

```javascript
myDB(async client => {
  const myDataBase = await client.db('database').collection('users');

  // Be sure to change the title
  app.route('/').get((req, res) => {
    // Change the response to render the Pug template
    res.render('index', {
      title: 'Connected to Database',
      message: 'Please login'
    });
  });

  // Serialization and deserialization here...

  // Be sure to add this...
}).catch(e => {
  app.route('/').get((req, res) => {
    res.render('index', { title: e, message: 'Unable to connect to database' });
  });
});
// app.listen out here...
```

`myDataBase`内の`deserializeUser`のコードのコメントアウトを解除し、`done(null, null)`に`doc`を含めるよう編集してください。

正しくできたと思ったらページを提出してください。エラーが発生する場合は、<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implement-the-serialization-of-a-passport-user-5" target="_blank" rel="noopener noreferrer nofollow">ここまで完成したプロジェクト</a>を確認できます。

# --hints--

データベース接続が存在している必要があります。

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Connected to Database/gi,
    'You successfully connected to the database!'
  );
```

デシリアライズは現在データベースを正しく使用しており、`done(null, null)`は`doc`を使って呼び出される必要があります。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /null,\s*doc/gi,
    'The callback in deserializeUser of (null, null) should be altered to (null, doc)'
  );
```
