---
id: 5895f70df9fc0f352b528e6a
title: 新しいミドルウェアを作成する
challengeType: 2
forumTopicId: 301551
dashedName: create-new-middleware
---

# --description--

現状では、ユーザーが認証されているかどうかに関わらず、`/profile` にURLを入力すれば誰でもアクセスできます。これを防ぐために、プロフィールページをレンダリングする前にユーザーが認証されているかどうかをチェックしたいです。まさにミドルウェアを作成するのに最適な例です。

ここでのチャレンジは、`ensureAuthenticated(req, res, next)` というミドルウェア関数を作成することです。この関数は、Passportの`isAuthenticated` メソッドを`request` に対して呼び出し、`req.user` が定義されているかどうかをチェックします。定義されていれば`next()` を呼び出します。そうでなければ、ログインのためにホームページへリダイレクトするレスポンスを返せばよいです。

このミドルウェアの実装例は以下の通りです：

```javascript
function ensureAuthenticated(req, res, next) {
  if (req.isAuthenticated()) {
    return next();
  }
  res.redirect('/');
};
```

上記のミドルウェア関数を作成し、GETリクエストの引数の前に`ensureAuthenticated` をプロフィールページへのリクエストのミドルウェアとして渡してください：

```javascript
app
 .route('/profile')
 .get(ensureAuthenticated, (req,res) => {
    res.render('profile');
 });
```

正しくできたと思ったらページを提出してください。エラーが出る場合は、<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#create-new-middleware-8" target="_blank" rel="noopener noreferrer nofollow">ここまで完成したプロジェクトを確認できます</a>。

# --hints--

ミドルウェア`ensureAuthenticated` を実装し、`/profile` ルートにアタッチしてください。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /ensureAuthenticated[^]*req.isAuthenticated/,
    'Your ensureAuthenticated middleware should be defined and utilize the req.isAuthenticated function'
  );
  assert.match(
    data,
    /profile[^]*get[^]*ensureAuthenticated/,
    'Your ensureAuthenticated middleware should be attached to the /profile route'
  );
```

認証されていないGETリクエストが`/profile` に来た場合、正しく`/` にリダイレクトされる必要があります。

```js
  const url = new URL("/profile", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/,
    'An attempt to go to the profile at this point should redirect to the homepage since we are not logged in'
  );
```
