---
id: 5895f70ef9fc0f352b528e6b
title: プロフィールの作成方法
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

ユーザーが`/profile`にアクセスしていることを認証できるようになったので、ページ上の`req.user`に含まれる情報を使えます。

プロフィールビューの`username`メソッドの第2引数として、プロパティ`req.user.username`と値`render`を含むオブジェクトを渡してください。

次に、`profile.pug`ビューに移動し、既存の`h1`要素の下、同じインデントレベルに次の行を追加してください。

```pug
h2.center#welcome Welcome, #{username}!
```

これにより、クラス`h2`、id `center`を持つ`welcome`要素が作成され、テキスト`Welcome, `とユーザー名が続きます。

また、`profile.pug`には、ユーザーの認証解除のロジックをホストする`/logout`ルートへのリンクを追加してください。

```pug
a(href='/logout') Logout
```

正しくできたと思ったらページを提出してください。エラーが出る場合は、<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">ここまで完成したプロジェクトを確認できます</a>。

# --hints--

`/profile`にPugのレンダリング変数を正しく追加してください。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );
```
