---
id: 58966a17f9fc0f352b528e6d
title: 新規ユーザー登録
challengeType: 2
forumTopicId: 301561
dashedName: registration-of-new-users
---

# --description--

これでサイトに新しいユーザーがアカウントを登録できるようにします。ホームページの`res.render`で渡されるオブジェクトに新しい変数`showRegistration: true`を追加してください。ページを更新すると、すでに`index.pug`ファイルで作成されている登録フォームが表示されるはずです。このフォームは`/register`に対して**POST**するよう設定されているので、そのルートを作成し、以下のロジックに従ってユーザーオブジェクトをデータベースに追加してください。

登録ルートのロジックは次の通りです：

1. 新しいユーザーを登録する
2. 新しいユーザーを認証する
3. `/profile`へリダイレクトする

ステップ1のロジックは次の通りです：

1. `findOne`でデータベースをクエリする
2. エラーがあれば、エラーを引数にして`next`を呼び出す
3. ユーザーが返された場合はホームにリダイレクトする
4. ユーザーが見つからずエラーもなければ、ユーザー名とパスワードでデータベースに`insertOne`する。ここでエラーがなければ、`next`を呼び出してステップ2の新しいユーザー認証に進む。認証のロジックはすでに`POST /login`ルートで書いてあります。

```js
app.route('/register')
  .post((req, res, next) => {
    myDataBase.findOne({ username: req.body.username }, (err, user) => {
      if (err) {
        next(err);
      } else if (user) {
        res.redirect('/');
      } else {
        myDataBase.insertOne({
          username: req.body.username,
          password: req.body.password
        },
          (err, doc) => {
            if (err) {
              res.redirect('/');
            } else {
              // The inserted document is held within
              // the ops property of the doc
              next(null, doc.ops[0]);
            }
          }
        )
      }
    })
  },
    passport.authenticate('local', { failureRedirect: '/' }),
    (req, res, next) => {
      res.redirect('/profile');
    }
  );
```

正しくできたと思ったらページを提出してください。エラーが出る場合は<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#registration-of-new-users-11" target="_blank" rel="noopener noreferrer nofollow">ここまで完成したプロジェクトを確認できます</a>。

**注意：** ここから先は*ピクチャーインピクチャー*ブラウザの使用に関連した問題が発生する場合があります。エディター内でアプリのプレビューを提供するオンラインIDEを使っている場合は、このプレビューを新しいタブで開くことをお勧めします。

# --hints--

`/register`ルートがあり、ホームページに登録フォームが表示されているはずです。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /showRegistration:( |)true/gi,
    'You should be passing the variable showRegistration as true to your render function for the homepage'
  );
  assert.match(
    data,
    /register[^]*post[^]*findOne[^]*username:( |)req.body.username/gi,
    'You should have a route that accepts a POST request on /register that queries the db with findOne and the query being username: req.body.username'
  );
```

登録が機能しているはずです。

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttp = new XMLHttpRequest();
  xhttp.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      test(this);
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttp.open('POST', url + '/register', true);
  xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
  xhttp.send(`username=${user}&password=${user}`);
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Register should work, and redirect successfully to the profile.'
    );
  }
```

ログインが機能しているはずです。

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttpReg = new XMLHttpRequest();
  xhttpReg.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      login();
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttpReg.open('POST', url + '/register', true);
  xhttpReg.setRequestHeader(
    'Content-type',
    'application/x-www-form-urlencoded'
  );
  xhttpReg.send(`username=${user}&password=${user}`);
  function login() {
    const xhttp = new XMLHttpRequest();
    xhttp.onreadystatechange = function () {
      if (this.readyState == 4 && this.status == 200) {
        test(this);
      } else {
        throw new Error(`${this.status} ${this.statusText}`);
      }
    };
    xhttp.open('POST', url + '/login', true);
    xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
    xhttp.send(`username=${user}&password=${user}`);
  }
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Login should work if previous test was done successfully and redirect successfully to the profile.'
    );
    assert.match(
      data,
      new RegExp(user, 'g'),
      'The profile should properly display the welcome to the user logged in'
    );
  }
```

ログアウトが機能しているはずです。

```js
  $.ajax({
    url: code + '/logout',
    type: 'GET',
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(data, /Home/gi, 'Logout should redirect to home');
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

ログアウト後はプロフィールが使えなくなっているはずです。

```js
  $.ajax({
    url: code + '/profile',
    type: 'GET',
    crossDomain: true,
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(
        data,
        /Home/gi,
        'Profile should redirect to home when we are logged out now again'
      );
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```
