---
id: 58965611f9fc0f352b528e6c
title: ユーザーのログアウト処理
challengeType: 2
forumTopicId: 301560
dashedName: logging-a-user-out
---

# --description--

ログアウトのロジックを作成するのは簡単です。ルートはユーザーの認証を解除し、ビューをレンダリングする代わりにホームページへリダイレクトするだけでよいです。

passportでは、ユーザーの認証を解除するのは`req.logout()`を呼び出してからリダイレクトするだけで簡単にできます。これを行うために`/logout`ルートを追加してください。

```js
app.route('/logout')
  .get((req, res) => {
    req.logout();
    res.redirect('/');
});
```

404エラー（ページが見つからない場合）を処理していないことに気づいたかもしれません。Nodeでこれを処理する一般的な方法は次のミドルウェアです。すべての他のルートの後にこれを追加してください。

```js
app.use((req, res, next) => {
  res.status(404)
    .type('text')
    .send('Not Found');
});
```

正しくできたと思ったらページを提出してください。エラーが出る場合は、<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#logging-a-user-out-10" target="_blank" rel="noopener noreferrer nofollow">ここまで完成したプロジェクトを確認できます</a>。

# --hints--

`req.logout()`は`/logout`ルート内で呼び出す必要があります。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /req.logout/gi,
    'You should be calling req.logout() in your /logout route'
  );
```

`/logout`はホームページへリダイレクトする必要があります。

```js
  const url = new URL("/logout", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/gi,
    'When a user logs out they should be redirected to the homepage'
  );
```
