---
id: 589fc831f9fc0f352b528e76
title: 切断の処理
challengeType: 2
forumTopicId: 301552
dashedName: handle-a-disconnect
---

# --description--

これまでユーザー数を増やす処理しかしていなかったことに気づくかもしれません。ユーザーの切断を処理するのは、最初の接続を処理するのと同じくらい簡単ですが、サーバー全体ではなく各ソケットで切断を監視する必要があります。

これを行うには、既存の`'connect'`リスナーの中に、データを渡さずにソケットで`'disconnect'`を監視する別のリスナーを追加してください。この機能は、ユーザーが切断したことをコンソールにログ出力するだけでテストできます。

```js
socket.on('disconnect', () => {
  /*anything you want to do on disconnect*/
});
```

クライアントが常に現在のユーザー数の最新情報を持つようにするため、切断が発生したときに`currentUsers`を1減らし、更新されたカウントを含む`'user count'`イベントを送信してください。

**注意:** `'disconnect'`と同様に、ソケットがサーバーに送信できる他のすべてのイベントも、'socket'が定義されている接続リスナー内で処理する必要があります。

正しくできたと思ったらページを提出してください。エラーが発生する場合は、<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#handle-a-disconnect-8" target="_blank" rel="noopener noreferrer nofollow">ここまで完成したプロジェクトを確認できます</a>。

# --hints--

サーバーはソケットからのdisconnectイベントを処理する必要があります。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /socket.on.*('|")disconnect('|")/s, '');
```

クライアントは`'user count'`イベントを監視している必要があります。

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
