---
id: 589fc832f9fc0f352b528e78
title: 新しいユーザーの接続を通知する
challengeType: 2
forumTopicId: 301546
dashedName: announce-new-users
---

# --description--

多くのチャットルームでは、ユーザーが接続または切断したときにその情報を通知し、チャットに接続している全ユーザーに表示します。すでに接続時と切断時にイベントを発火しているので、そのイベントを修正してこの機能をサポートするだけで済みます。最も論理的な方法は、イベントに3つのデータを送ることです：接続または切断したユーザーのユーザー名、現在のユーザー数、そしてそのユーザー名が接続したのか切断したのか。

イベント名を `'user'` に変更し、`username`、`currentUsers`、`connected` のフィールドを含むオブジェクトを渡してください（接続の場合は `true`、切断の場合は `false` になります）。`'user count'` の両方のイベントを変更し、切断時のイベントでは `false` フィールドに `connected` を送るようにし、接続時に送る `true` とは異なるようにしてください。

```js
io.emit('user', {
  username: socket.request.user.username,
  currentUsers,
  connected: true
});
```

これでクライアントは現在のユーザー数を正しく表示し、ユーザーの接続や切断を通知するために必要な情報をすべて受け取れます！クライアント側でこのイベントを処理するには `'user'` をリッスンし、jQueryを使って `#num-users` のテキストを `'{NUMBER} users online'` に更新し、さらに id が `<li>` の番号なしリストに `messages` を追加して `'{NAME} has {joined/left} the chat.'` を表示してください。

この実装例は以下のようになります：

```js
socket.on('user', data => {
  $('#num-users').text(data.currentUsers + ' users online');
  let message =
    data.username +
    (data.connected ? ' has joined the chat.' : ' has left the chat.');
  $('#messages').append($('<li>').html('<b>' + message + '</b>'));
});
```

正しくできたと思ったらページを提出してください。エラーが出る場合は、<a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135/3#announce-new-users-10" target="_blank" rel="noopener noreferrer nofollow">ここまで完成したプロジェクト</a>を確認できます。

# --hints--

イベント `'user'` は `username`、`currentUsers`、`connected` を含めて発火する必要があります。

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  // Regex is lenient to match both `username` and `name` as the key on purpose.
  assert.match(
    data,
    /io.emit.*('|")user\1.*name.*currentUsers.*connected/s,
    'You should have an event emitted named user sending name, currentUsers, and connected'
  );
```

クライアントはイベント `'user'` からの新しいデータを正しく処理し表示する必要があります。

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*num-users/s,
    'You should change the text of "#num-users" within on your client within the "user" event listener to show the current users connected'
  );
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*messages.*li/s,
    'You should append a list item to "#messages" on your client within the "user" event listener to announce a user came or went'
  );
```
