---
id: 68e51d8f163d3f8b70f9c4e5
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, practicarás trabajando con los métodos `includes()` y `slice()`.

Para comenzar, crea una variable llamada `fccSentence` y asígnale la cadena `"freeCodeCamp is a great place to learn web development."`.

# --hints--

Debes declarar una variable llamada `fccSentence`.

```js
assert.exists(fccSentence);
```

Tu variable `fccSentence` debe ser una cadena.

```js
assert.isString(fccSentence);
```

Tu variable `fccSentence` debe ser asignada con la cadena `"freeCodeCamp is a great place to learn web development."`.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
