---
id: 68e6172975008fea742a63d0
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Puedes usar `slice()` para extraer desde el inicio de una cadena.

Crea una variable llamada `greetingWord`. Luego asígnale el resultado de usar `message.slice(0, 7)` para extraer la primera palabra `"Welcome"`.

Luego usa una plantilla literal o concatenación de cadenas para registrar el mensaje `The first word is "<greetingWord>".` en la consola. Reemplaza `<greetingWord>` con el valor real de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Debes declarar una variable llamada `greetingWord`.

```js
assert.exists(greetingWord);
```

Tu variable `greetingWord` debe ser una cadena.

```js
assert.isString(greetingWord);
```

Debes usar `message.slice(0, 7)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Debes asignar el resultado de `message.slice(0, 7)` a tu variable `greetingWord`.

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Tu variable `greetingWord` debería ser igual a `"Welcome"`.

```js
assert.strictEqual(greetingWord, "Welcome");
```

Debes registrar el mensaje `The first word is "<greetingWord>".`, donde `<greetingWord>` debe ser reemplazado por el valor real de la variable.

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
