---
id: 68e6172975008fea742a63d1
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

El método `slice()` puede usar índices negativos para contar desde el final de la cadena. `-1` se refiere al último carácter.

Aquí hay un ejemplo:

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

Crea una variable llamada `endPunctuation`. Luego asígnale el resultado de usar `message.slice(-1)` para extraer el último carácter.

Luego usa una plantilla literal o concatenación de cadenas para registrar el mensaje `The ending punctuation mark is a "<endPunctuation>"` en la consola. Reemplaza `<endPunctuation>` con el valor real de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Debes declarar una variable llamada `endPunctuation`.

```js
assert.exists(endPunctuation);
```

Tu variable `endPunctuation` debería ser una cadena.

```js
assert.isString(endPunctuation);
```

Debes usar `message.slice(-1)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

Debes asignar el resultado de `message.slice(-1)` a tu variable `endPunctuation`.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

Tu variable `endPunctuation` debería ser igual a `"!"`.

```js
assert.strictEqual(endPunctuation, "!");
```

Debes registrar el mensaje `The ending punctuation mark is a "<endPunctuation>"`, donde `<endPunctuation>` debe ser reemplazado con el valor real de la variable.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
