---
id: 68e6172875008fea742a63c9
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Como recuerdas de las lecciones anteriores, el método `includes()` verifica si una cadena contiene una subcadena específica y devuelve `true` o `false`.

Aquí tienes un ejemplo usando el método `includes()`:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Crea una variable llamada `hasFreeCodeCamp`. Luego asígnale el resultado de usar el método `includes()` en `fccSentence` para verificar si contiene `"freeCodeCamp"`.

# --hints--

Debes declarar una variable llamada `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

Debes usar `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Debes asignar el resultado de `fccSentence.includes("freeCodeCamp")` a tu variable `hasFreeCodeCamp`.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

El resultado de `fccSentence.includes("freeCodeCamp")` debería ser `true`. 

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
