---
id: 68e6172975008fea742a63cd
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora explorarás el método `slice()`, que extrae una porción de una cadena.

Crea una variable llamada `message` y asígnale la cadena `"Welcome to freeCodeCamp!"`.

Luego usa `console.log()` para registrar `"Here are some examples of the slice() method:"` en la consola.

# --hints--

Debes declarar una variable llamada `message`.

```js
assert.exists(message);
```

Tu variable `message` debe ser una cadena.

```js
assert.isString(message);
```

Tu variable `message` debería ser `"Welcome to freeCodeCamp!"`.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

Debes registrar `"Here are some examples of the slice() method:"` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
