---
id: 68e6172975008fea742a63ce
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

El método `slice()` extrae una sección de una cadena. Toma dos argumentos: el índice de inicio (incluido) y el índice de fin (no incluido).

Aquí hay un ejemplo:

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

Los índices de cadenas comienzan en 0. En la cadena `message`, el carácter en el índice 11 es `f` (el inicio de "freeCodeCamp").

Crea una variable llamada `platform`. Luego asígnale el resultado de usar `message.slice(11, 23)` para extraer `"freeCodeCamp"`.

# --hints--

Debes declarar una variable llamada `platform`.

```js
assert.exists(platform);
```

Tu variable `platform` debe ser una cadena.

```js
assert.isString(platform);
```

Debes usar `message.slice(11, 23)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Debes asignar el resultado de `message.slice(11, 23)` a tu variable `platform`.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Tu variable `platform` debería ser igual a `"freeCodeCamp"`.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
