---
id: 68e6172975008fea742a63ca
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Ahora usa una plantilla literal o concatenación de cadenas para registrar el mensaje `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` en la consola. Reemplaza `<hasFreeCodeCamp>` con el valor real de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Debes registrar el mensaje `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`, donde `<hasFreeCodeCamp>` debe ser reemplazado con el valor real de la variable.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
