---
id: 68e6172975008fea742a63cc
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

El método `includes()` distingue entre mayúsculas y minúsculas, lo que significa que diferencia entre letras en mayúscula y minúscula.

Crea una variable llamada `hasLowercaseFCC`. Luego asígnale el resultado de usar el método `includes()` para verificar si `fccSentence` incluye `"freecodecamp"` (todo en minúsculas).

Luego usa una plantilla literal o concatenación de cadenas para registrar el mensaje `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` en la consola. Reemplaza `<hasLowercaseFCC>` con el valor real de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Debes declarar una variable llamada `hasLowercaseFCC`.

```js
assert.exists(hasLowercaseFCC);
```

Debes usar `fccSentence.includes("freecodecamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Debes asignar el resultado de `fccSentence.includes("freecodecamp")` a tu variable `hasLowercaseFCC`.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

El resultado de `fccSentence.includes("freecodecamp")` debería ser `false`. 

```js
assert.isFalse(hasLowercaseFCC);
```

Debes registrar el mensaje que incluye `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`, donde `<hasLowercaseFCC>` debe ser reemplazado con el valor real de la variable.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
