---
id: 68e6172975008fea742a63cb
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora es momento de ver qué sucede cuando la subcadena NO se encuentra.

Crea una variable llamada `hasJavaScript`. Luego asígnale el resultado de usar el método `includes()` para verificar si `fccSentence` incluye `"JavaScript"`.

Luego usa una plantilla literal o concatenación de cadenas para registrar el mensaje `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` en la consola. Reemplaza `<hasJavaScript>` con el valor real de la variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Debes declarar una variable llamada `hasJavaScript`.

```js
assert.exists(hasJavaScript);
```

Debes usar `fccSentence.includes("JavaScript")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Debes asignar el resultado de `fccSentence.includes("JavaScript")` a tu variable `hasJavaScript`.

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

El resultado de `fccSentence.includes("JavaScript")` debería ser `false`. 

```js
assert.isFalse(hasJavaScript);
```

Debes registrar el mensaje `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`, donde `<hasJavaScript>` debe ser reemplazado por el valor real de la variable.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
