---
id: 5b7d72c338cd7e35b63f3e14
title: تحسين التوافق مع أنظمة المتصفح البديلة
challengeType: 0
forumTopicId: 301087
dashedName: improve-compatibility-with-browser-fallbacks
---

# --description--

عند العمل مع CSS، من المحتمل أن تواجه مشاكل في التوافق مع المتصفحات في مرحلة ما. لهذا من المهم توفير نظام بديل للمتصفح لتجنب المشاكل المحتملة.

عندما يقوم متصفحك بتحليل CSS لصفحة ويب، يتجاهل أي خاصيات لا يتعرف عليها أو لا يدعمها. على سبيل المثال، إذا استخدمت متغير CSS لتعيين لون الخلفية في موقع، فإن Internet Explorer سيتجاهل لون الخلفية لأنه لا يدعم متغيرات CSS. في هذه الحالة، يستخدم المتصفح أي قيمة موجودة لتلك الخاصية. وإذا لم يجد أي قيمة أخرى محددة لتلك الخاصية، فسيرجع إلى القيمة الافتراضية، والتي عادةً ما تكون غير مثالية.

هذا يعني أنه إذا أردت توفير نظام بديل للمتصفح، فالأمر بسيط مثل تعيين قيمة أخرى مدعومة على نطاق أوسع مباشرة قبل إعلانك. بهذه الطريقة، سيكون لدى المتصفح القديم شيء يعتمد عليه، بينما يفسر المتصفح الأحدث الإعلان الذي يأتي لاحقًا في التتابع.

# --instructions--

يبدو أن متغيرًا يُستخدم لتعيين لون الخلفية لصنف `.red-box`. لنحسن توافق المتصفح لدينا بإضافة إعلان `background` آخر مباشرة قبل الإعلان الحالي وتعيين قيمته إلى `red`.

# --hints--

يجب أن تتضمن قاعدة `.red-box` نظامًا بديلاً مع تعيين `background` إلى `red` مباشرة قبل إعلان `background` الحالي.

```js
const spacelessCode = __helpers.removeWhiteSpace(__helpers.removeCssComments(code));
assert.match(
  spacelessCode,
  /\.red-box{background:(red|#ff0000|#f00|rgb\(255,0,0\)|rgb\(100%,0%,0%\)|hsl\(0,100%,50%\));background:var\(--red-color\);height:200px;width:200px;}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  :root {
    --red-color: red;
  }
  .red-box {

    background: var(--red-color);
    height: 200px;
    width:200px;
  }
</style>
<div class="red-box"></div>
```

# --solutions--

```html
<style>
  :root {
    --red-color: red;
  }
  .red-box {
    background: red;
    background: var(--red-color);
    height: 200px;
    width:200px;
  }
</style>
<div class="red-box"></div>
```
