---
id: bad87fee1248bd9aedf08824
title: إضافة هوامش مختلفة لكل جانب من عناصر الصفحة
challengeType: 0
forumTopicId: 16633
dashedName: add-different-margins-to-each-side-of-an-element
---

# --description--

أحيانًا تريد تخصيص عنصر بحيث يكون له `margin` مختلف على كل جانب من جوانبه.

تسمح لك CSS بالتحكم في `margin` لكل من الجوانب الأربعة الفردية للعنصر باستخدام الخصائص `margin-top` و`margin-right` و`margin-bottom` و`margin-left`.

# --instructions--

أعطِ الصندوق الأزرق `margin` بقيمة `40px` في الأعلى وعلى الجانب الأيسر، و`20px` فقط في الأسفل وعلى الجانب الأيمن.

# --hints--

يجب أن يمنح صنف `blue-box` أعلى العناصر هامشًا بقيمة `40px` من `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '40px');
```

يجب أن يمنح صنف `blue-box` الجانب الأيمن للعناصر هامشًا بقيمة `20px` من نوع `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginRight = window.getComputedStyle(blueBox)['margin-right'];
assert.strictEqual(marginRight, '20px');
```

يجب أن يمنح صنف `blue-box` أسفل العناصر هامشًا بقيمة `20px` من `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginBottom = window.getComputedStyle(blueBox)['margin-bottom'];
assert.strictEqual(marginBottom, '20px');
```

يجب أن يمنح صنف `blue-box` الجهة اليسرى للعناصر هامشًا بقيمة `40px` من نوع `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginLeft = window.getComputedStyle(blueBox)['margin-left'];
assert.strictEqual(marginLeft,'40px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin-top: 40px;
    margin-right: 20px;
    margin-bottom: 20px;
    margin-left: 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin-top: 40px;
    margin-right: 20px;
    margin-bottom: 20px;
    margin-left: 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    margin-top: 40px;
    margin-right: 20px;
    margin-bottom: 20px;
    margin-left: 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
