---
id: bad87fee1348bd9aedf08807
title: استيراد خط من Google
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9MRsJ'
forumTopicId: 18200
dashedName: import-a-google-font
---

# --description--

بالإضافة إلى تحديد الخطوط الشائعة الموجودة في معظم أنظمة التشغيل، يمكننا أيضًا تحديد خطوط ويب مخصصة وغير قياسية لاستخدامها في موقعنا. هناك مصادر عديدة لخطوط الويب على الإنترنت. في هذا المثال، سنركز على مكتبة خطوط Google.

خطوط Google هي مكتبة مجانية لخطوط الويب يمكنك استخدامها في CSS عبر الإشارة إلى عنوان URL الخاص بالخط.

لذا، لنبدأ باستيراد وتطبيق خط من Google (لاحظ أنه إذا كان Google محجوبًا في بلدك، ستحتاج إلى تخطي هذا التحدٍّ).

لاستيراد خط من Google، يمكنك نسخ عنوان URL الخاص بالخط من مكتبة خطوط Google ثم لصقه في HTML الخاص بك. في هذا التحدٍّ، سنستورد خط `Lobster`. للقيام بذلك، انسخ مقتطف الشفرة التالي والصقه في أعلى محرر الشفرة الخاص بك (قبل عنصر `style` الافتتاحي):

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
```

الآن يمكنك استخدام خط `Lobster` في CSS الخاص بك باستخدام `Lobster` كاسم العائلة FAMILY_NAME كما في المثال التالي:

```css
font-family: FAMILY_NAME, GENERIC_NAME;
```

الاسم العام GENERIC_NAME اختياري، وهو خط النظام البديل في حال عدم توفر الخط المحدد الآخر. هذا مغطى في التحدٍّ التالي.

أسماء العائلات حساسة لحالة الأحرف ويجب أن تُحاط بعلامات اقتباس إذا كان هناك فراغ في الاسم. على سبيل المثال، تحتاج إلى علامات اقتباس لاستخدام الخط `"Open Sans"`، لكن لا تحتاجها لاستخدام الخط `Lobster`.

# --instructions--

استورد خط `Lobster` إلى صفحة الويب الخاصة بك. ثم استخدم محدد عنصر لتعيين `Lobster` كـ `font-family` لعنصر `h2` الخاص بك.

# --hints--

يجب أن تستورد خط `Lobster`.

```js
assert.exists(document.querySelector('link[href*="googleapis" i]'));
```

يجب أن يستخدم عنصر `h2` الخاص بك الخط `Lobster`.

```js
const h2 = document.querySelector('h2'); 
const fontFamily = window.getComputedStyle(h2)['font-family']; 
assert.match(fontFamily, /lobster/i);
```

يجب أن تستخدم فقط محدد عنصر `h2` لتغيير الخط.

```js
assert.match(__helpers.removeHtmlComments(code), /\s*[^\.]h2\s*\{\s*font-family\s*:\s*('|"|)Lobster\1\s*(,\s*('|"|)[a-z -]+\3\s*)?(;\s*\}|\})/gi);
```

يجب أن يظل عنصر `p` يستخدم الخط `monospace`.

```js
const paragraphElement = document.querySelector('p');
const fontFamily = window.getComputedStyle(paragraphElement)['font-family']; 
assert.match(fontFamily, /monospace/i);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }

  h2 {
    font-family: Lobster;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
