---
id: bad88fee1348bd9aedf08825
title: ضبط حشو عنصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cED8ZC2'
forumTopicId: 301083
dashedName: adjust-the-padding-of-an-element
---

# --description--

لنضع تطبيق صور القطط جانبًا لفترة قصيرة ونتعلم المزيد عن تنسيق HTML.

ربما لاحظت هذا بالفعل، لكن كل عناصر HTML هي في الأساس مستطيلات صغيرة.

ثلاث خصائص مهمة تتحكم في المساحة التي تحيط بكل عنصر HTML: `padding`، و`border`، و`margin`.

حشو (`padding`) العنصر يتحكم في مقدار المساحة بين محتوى العنصر و`border` الخاص به.

هنا، نرى أن الصندوق الأزرق والصندوق الأحمر مسكنان داخل الصندوق الأصفر. لاحظ أن الصندوق الأحمر يحتوي على حشو (`padding`) أكثر من الصندوق الأزرق.

عندما تزيد حشو (`padding`) الصندوق الأزرق، ستزداد المسافة (`padding`) بين النص والحدود المحيطة به.

# --instructions--

غيّر حشو (`padding`) الصندوق الأزرق ليطابق حشو الصندوق الأحمر.

# --hints--

يجب أن يمنح صنف `blue-box` العناصر مقدار `20px` من `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
