---
id: bad87fee1348bd9aedf08726
title: استخدام شفرة Hex للألوان المحددة
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

هل تعلم أن هناك طرقًا أخرى لتمثيل الألوان في CSS؟ إحدى هذه الطرق تُسمى شفرة سداسية عشرية، أو شفرة hex للاختصار.

عادةً نستخدم <dfn>الأعداد العشرية</dfn>، أو أرقام الأساس 10، التي تستخدم الرموز من 0 إلى 9 لكل رقم. <dfn>السداسيات عشرية</dfn> (أو <dfn>hex</dfn>) هي أرقام أساسها 16. هذا يعني أنها تستخدم ستة عشر رمزًا مميزًا. مثل الأعداد العشرية، تمثل الرموز من 0 إلى 9 القيم من صفر إلى تسعة. ثم تمثل A وB وC وD وE وF القيم من عشرة إلى خمسة عشر. معًا، يمكن للرموز من 0 إلى F تمثيل رقم في النظام السداسي عشر، مما يعطينا 16 قيمة ممكنة إجمالًا. يمكنك العثور على مزيد من المعلومات عن <a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">الأعداد السداسية عشرية هنا</a>.

في CSS، يمكننا استخدام 6 أرقام سداسية عشرية لتمثيل الألوان، رقمين لكل من مكوِّنات الأحمر (R) والأخضر (G) والأزرق (B). على سبيل المثال، `#000000` هو اللون الأسود وهو أيضًا أدنى قيمة ممكنة. يمكنك العثور على مزيد من المعلومات عن <a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">نظام ألوان RGB هنا</a>.

```css
body {
  color: #000000;
}
```

# --instructions--

استبدل كلمة `black` في خاصية background-color لعنصر `body` بتمثيلها باستخدام شفرة hex، `#000000`.

# --hints--

يجب أن يحتوي عنصر `body` على خاصية `background-color` باللون الأسود.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

يجب استخدام شفرة hex للون الأسود بدلًا من كلمة `black`.

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
