---
id: bad87fee1348bd9aedf08736
title: تنسيق عنصر body في HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB77PHW'
forumTopicId: 18313
dashedName: style-the-html-body-element
---

# --description--

لنبدأ من جديد ونتحدث عن وراثة CSS.

كل صفحة HTML تحتوي على عنصر `body`.

# --instructions--

يمكننا إثبات وجود عنصر `body` هنا بإعطائه خاصية `background-color` باللون الأسود.

يمكننا فعل ذلك بإضافة التالي إلى عنصر `style` الخاص بنا:

```css
body {
  background-color: black;
}
```

# --hints--

يجب أن يحتوي عنصر `body` على خاصية `background-color` باللون الأسود.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color'];

assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

يجب أن تكون قاعدة CSS الخاصة بك منسقة بشكل صحيح مع أقواس معقوفة مفتوحة ومغلقة.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

يجب أن تنتهي قاعدة CSS الخاصة بك بفاصلة منقوطة.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

# --seed--

## --seed-contents--

```html
<style>

</style>
```

# --solutions--

```html
<style>
body {
  background-color: black;
}
</style>
```
