---
id: bad82fee1322bd9aedf08721
title: فهم الوحدات المطلقة مقابل الوحدات النسبية
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN66JSL'
forumTopicId: 301089
dashedName: understand-absolute-versus-relative-units
---

# --description--

جميع التحديات الأخيرة ضبطت هامش أو حشو عنصر بوحدة البكسل (`px`). البكسل هو نوع من وحدات الطول، وهي التي تُخبر المتصفح كيف يحدد حجم العنصر أو يفرغه. بالإضافة إلى `px`، يحتوي CSS على عدة خيارات مختلفة لوحدات الطول يمكنك استخدامها.

النوعان الرئيسيان لوحدات الطول هما المطلقة والنسبية. الوحدات المطلقة ترتبط بوحدات الطول الفيزيائية. على سبيل المثال، `in` و `mm` تشير إلى البوصة والمليمتر على التوالي. وحدات الطول المطلقة تقارب القياس الفعلي على الشاشة، لكن هناك بعض الاختلافات حسب دقة وضوح الشاشة.

الوحدات النسبية، مثل `em` أو `rem`، تعتمد على قيمة طول أخرى. على سبيل المثال، `em` يعتمد على حجم خط العنصر. إذا استخدمته لضبط خاصية `font-size` نفسها، فهو يعتمد على `font-size` الخاص بالعنصر الأب.

**ملاحظة:** هناك عدة خيارات لوحدات نسبية مرتبطة بحجم إطار العرض. يتم تغطيتها في قسم مبادئ تصميم الويب المتجاوب.

# --instructions--

أضف خاصية `padding` إلى العنصر الذي يحمل الصنف `red-box` واضبطها على `1.5em`.

# --hints--

يجب أن يحتوي صنف `red-box` على خاصية `padding`.

```js
const redBox =document.querySelector('.red-box'); 
const style = window.getComputedStyle(redBox); 

assert.notEqual(style['padding-top'], '0px');
assert.notEqual(style['padding-right'], '0px');
assert.notEqual(style['padding-bottom'], '0px');
assert.notEqual(style['padding-left'], '0px');
```

يجب أن يمنح صنف `red-box` العناصر حشو `padding` بمقدار 1.5em.

```js
assert.match(__helpers.removeCssComments(code), /\.red-box\s*?{[\s\S]*padding\s*:\s*?1\.5em/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;

  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;
    padding: 1.5em;
  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```
