---
id: bad87fee1348bd9aede08718
title: استخدام قيم RGB لتلوين العناصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

طريقة أخرى لتمثيل الألوان في CSS هي باستخدام قيم `RGB`.

قيمة `RGB` للون الأسود تبدو هكذا:

```css
rgb(0, 0, 0)
```

قيمة `RGB` للون الأبيض تبدو هكذا:

```css
rgb(255, 255, 255)
```

بدلًا من استخدام ستة أرقام عشرية سداسية كما تفعل مع شفرة hex، مع `RGB` تحدد سطوع كل لون برقم بين 0 و255.

إذا حسبت، الرقمان لكل لون يساويان 16 في 16، مما يعطينا 256 قيمة إجمالية. لذا `RGB`، التي تبدأ العد من الصفر، لديها نفس عدد القيم الممكنة تمامًا مثل شفرة hex.

إليك مثالًا على كيفية تغيير خلفية عنصر `body` إلى اللون البرتقالي باستخدام شفرة RGB الخاصة به.

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

لنستبدل شفرة hex في لون خلفية عنصر `body` بقيمة RGB للون الأسود: `rgb(0, 0, 0)`

# --hints--

يجب أن يكون لعنصر `body` خلفية سوداء.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

يجب استخدام `rgb` لإعطاء عنصر `body` خلفية سوداء.

```js
assert.match(__helpers.removeCssComments(code), /rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
