---
id: bad87fee1348bd9aedf06756
title: تجاوز إعلانات الصنف باستخدام الأنماط المضمَّنة
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDRha'
forumTopicId: 18252
dashedName: override-class-declarations-with-inline-styles
---

# --description--

لقد أثبتنا أن إعلانات المعرف تتجاوز إعلانات الصنف، بغض النظر عن مكان إعلانها في CSS داخل عنصر `style` الخاص بك.

هناك طرق أخرى لتجاوز CSS. هل تتذكر الأنماط المضمَّنة؟

# --instructions--

استخدم نمطًا مضمّنًا لتحاول جعل عنصر `h1` الخاص بنا أبيض. تذكر، تبدو الأنماط المضمَّنة هكذا:

```html
<h1 style="color: green;">
```

اترك صنفي `blue-text` و `pink-text` على عنصر `h1` الخاص بك.

# --hints--

يجب أن يحتوي عنصر `h1` على الصنف `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

يجب أن يحتوي عنصر `h1` على الصنف `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

يجب أن يحتوي عنصر `h1` على المعرف `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

يجب أن يحتوي عنصر `h1` على نمط مضمّن.

```js
assert.exists(document.querySelector('h1[style]'));
```

يجب أن يكون عنصر `h1` أبيض.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
