---
id: bad87fee1348bd9aedf08823
title: إضافة هامش سلبي إلى عنصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpyGs3'
forumTopicId: 16166
dashedName: add-a-negative-margin-to-an-element
---

# --description--

يتحكم `margin` الخاص بالعنصر في مقدار المساحة بين `border` الخاص بالعنصر والعناصر المحيطة به.

إذا ضبطت `margin` الخاص بالعنصر على قيمة سالبة، سينمو حجم العنصر.

# --instructions--

حاول ضبط `margin` على قيمة سالبة مثل القيمة الخاصة بالمربع الأحمر.

غيّر `margin` الخاص بالمربع الأزرق إلى `-15px`، ليملأ العرض الأفقي الكامل للمربع الأصفر المحيط به.

# --hints--

يجب أن يمنح صنف `blue-box` العناصر `-15px` من `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)["margin-top"]; 
assert.strictEqual(marginTop, "-15px"); 
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: -15px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }
</style>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: -15px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
    margin-top: -15px;
  }
</style>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
