---
id: bad87fee1348bd9aedf08808
title: تحديد كيفية تدهور الخطوط
challengeType: 0
videoUrl: 'https://scrimba.com/c/cpVKBfQ'
forumTopicId: 18304
dashedName: specify-how-fonts-should-degrade
---

# --description--

هناك عدة خطوط افتراضية متوفرة في جميع المتصفحات. تشمل عائلات الخطوط العامة هذه `monospace` و`serif` و`sans-serif`.

عندما لا يتوفر خط معين، يمكنك إخبار المتصفح بـ"التدهور" إلى خط آخر.

على سبيل المثال، إذا أردت أن يستخدم عنصر ما الخط `Helvetica`، لكنه يتدهور إلى الخط `sans-serif` عندما لا يتوفر `Helvetica`، فستحدده كما يلي:

```css
p {
  font-family: Helvetica, sans-serif;
}
```

أسماء عائلات الخطوط العامة غير حساسة لحالة الأحرف. كما أنها لا تحتاج إلى علامات اقتباس لأنها كلمات مفتاحية في CSS.

# --instructions--

ابدأ بتطبيق الخط `monospace` على عنصر `h2`، ليصبح لديه الآن خطان - `Lobster` و`monospace`.

في التحدي السابق، استوردت الخط `Lobster` باستخدام وسم `link`. الآن أضف تعليقًا على استيراد الخط `Lobster` (باستخدام تعليقات HTML التي تعلمتها سابقًا) من Google Fonts حتى لا يصبح متوفرًا بعد الآن. لاحظ كيف يتدهور عنصر `h2` إلى الخط `monospace`.

**ملاحظة:** إذا كان الخط `Lobster` مثبتًا على جهازك، فلن ترى التدهور لأن المتصفح قادر على العثور على الخط.

# --hints--

يجب أن يستخدم عنصر h2 الخط `Lobster`.

```js
const h2Element = document.querySelector('h2');
const fontFamily = window.getComputedStyle(h2Element)['font-family']; 
assert.match(fontFamily, /^"?lobster/i);
```

يجب أن يتدهور عنصر h2 إلى الخط `monospace` عندما لا يتوفر `Lobster`.

```js
assert.match(__helpers.removeCssComments(code), /\s*h2\s*\{\s*font-family\s*\:\s*(\'|"|)Lobster\1\s*,\s*monospace\s*;?\s*\}/gi);
```

يجب عليك إضافة تعليق على استدعائك لـ Google للخط `Lobster` بوضع `<!--` أمامه.

```js
assert.match(code, /<!--[^fc]/gi);
```

يجب إغلاق التعليق بإضافة `-->`.

```js
assert.match(code, /[^fc]-->/gi);
```

# --seed--

## --seed-contents--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<!--<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">-->
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster, monospace;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
