---
id: bad87fee1348bd9aedf08756
title: أعطِ أولوية لنمط واحد على آخر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ8wnHv'
forumTopicId: 18258
dashedName: prioritize-one-style-over-another
---

# --description--

أحيانًا تتلقى عناصر HTML الخاصة بك أنماطًا متعددة تتعارض مع بعضها البعض.

على سبيل المثال، لا يمكن لعنصر `h1` أن يكون أخضر ووردي في الوقت نفسه.

لنرَ ماذا يحدث عندما ننشئ صنفًا يجعل النص ورديًا، ثم نطبقه على عنصر. هل صنفنا سيُتجاوز خاصية CSS `color: green;` لعنصر `body`؟

# --instructions--

أنشئ صنف CSS باسم `pink-text` يمنح العنصر اللون الوردي.

أعطِ عنصر `h1` الخاص بك الصنف `pink-text`.

# --hints--

يجب أن يحتوي عنصر `h1` على الصنف `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

يجب أن يحتوي `<style>` على صنف CSS باسم `pink-text` يغير الخاصية `color`.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*\{\s*color\s*:\s*.+\s*;?\s*\}/g);
```

يجب أن يكون عنصر `h1` ورديًا.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
</style>
<h1>Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```
