---
id: bad87fee1348bd9aedf08746
title: وراثة الأنماط من عنصر body
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bmdtR'
forumTopicId: 18204
dashedName: inherit-styles-from-the-body-element
---

# --description--

لقد أثبتنا الآن أن كل صفحة HTML تحتوي على عنصر `body`، وأنه يمكن تنسيق عنصر `body` أيضًا باستخدام CSS.

تذكر، يمكنك تنسيق عنصر `body` تمامًا مثل أي عنصر HTML آخر، وجميع العناصر الأخرى سترث أنماط عنصر `body`.

# --instructions--

أولًا، أنشئ عنصر `h1` يحتوي على النص `Hello World`.

بعدها، امنح جميع العناصر في صفحتك اللون `green` بإضافة `color: green;` إلى إعلان نمط عنصر `body`.

أخيرًا، امنح عنصر `body` خاصية الخط `monospace` بإضافة `font-family: monospace;` إلى إعلان نمط عنصر `body`.

# --hints--

يجب أن تنشئ عنصر `h1`.

```js
assert.isNotEmpty(document.querySelectorAll('h1'));
```

يجب أن يحتوي عنصر `h1` على النص `Hello World`.

```js
assert.match(
  document.querySelector('h1').textContent,
  /hello world/i
);
```

يجب أن يحتوي عنصر `h1` على وسم الإغلاق.

```js
const commentlessCode = __helpers.removeHtmlComments(code);
assert.match(commentlessCode, /<\/h1>/g);
assert.match(commentlessCode, /<h1/g);
assert.lengthOf(commentlessCode.match(/<\/h1>/g), commentlessCode.match(/<h1/g).length);
```

يجب أن يحتوي عنصر `body` على الخاصية `color` بالقيمة `green`.

```js
const bodyElement = document.querySelector('body');
const color = window.getComputedStyle(bodyElement)['color']; 
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

يجب أن يحتوي عنصر `body` على الخاصية `font-family` بالقيمة `monospace`.

```js
const bodyElement = document.querySelector('body');
const fontFamily = window.getComputedStyle(bodyElement)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

يجب أن يرث عنصر `h1` الخط `monospace` من عنصر `body`.

```js
const h1Element = document.querySelector('h1');
const fontFamily = window.getComputedStyle(h1Element)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

يجب أن يرث عنصر `h1` اللون `green` من عنصر `body`.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }

</style>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }

</style>
<h1>Hello World!</h1>
```
