---
id: bad87fee1348bd9aedf07756
title: تجاوز جميع الأنماط الأخرى باستخدام Important
challengeType: 0
videoUrl: 'https://scrimba.com/c/cm24rcp'
forumTopicId: 18249
dashedName: override-all-other-styles-by-using-important
---

# --description--

رائع! لقد أثبتنا للتو أن الأنماط المضمَّنة ستتجاوز جميع إعلانات CSS في عنصر `style` الخاص بك.

لكن انتظر. هناك طريقة أخيرة لتجاوز CSS. هذه هي الطريقة الأقوى على الإطلاق. لكن قبل أن نفعل ذلك، دعنا نتحدث عن سبب رغبتك في تجاوز CSS أصلاً.

في كثير من الحالات، ستستخدم مكتبات CSS. قد تتجاوز هذه المكتبات CSS الخاص بك عن طريق الخطأ. لذلك عندما تحتاج إلى التأكد تمامًا من أن عنصرًا ما يحتوي على CSS محدد، يمكنك استخدام `!important`.

لنعد إلى إعلان الصنف `pink-text`. تذكر أن صنف `pink-text` تم تجاوزه بواسطة إعلانات الأصناف التالية، وإعلانات المعرفات، والأنماط المضمَّنة.

# --instructions--

لنضيف الكلمة المفتاحية `!important` إلى إعلان لون عنصر pink-text لتتأكد 100% أن عنصر `h1` سيكون ورديًا.

مثال على كيفية القيام بذلك هو:

```css
color: red !important;
```

# --hints--

يجب أن يحتوي عنصر `h1` على الصنف `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

يجب أن يحتوي عنصر `h1` على الصنف `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

يجب أن يحتوي عنصر `h1` على `id` بقيمة `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

يجب أن يحتوي عنصر `h1` على النمط المضمَّن `color: white`.

```js
const commentessCode = __helpers.removeHtmlComments(code);
assert.match(commentessCode, /<h1.*style/gi);
assert.match(commentessCode, /<h1.*style.*color\s*?:/gi);
```

يجب أن يحتوي إعلان صنف `pink-text` على الكلمة المفتاحية `!important` لتجاوز جميع الإعلانات الأخرى.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*?\{[\s\S]*?color:.*pink.*!important\s*;?[^\.]*\}/g);
```

يجب أن يكون عنصر `h1` ورديًا.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink !important;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
