---
id: bad87fee1348bd9aedf08826
title: استخدام تدوين عقارب الساعة لتحديد حشو العنصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB7mBS9'
forumTopicId: 18346
dashedName: use-clockwise-notation-to-specify-the-padding-of-an-element
---

# --description--

بدلًا من تحديد خاصيات `padding-top` و`padding-right` و`padding-bottom` و`padding-left` لعنصر بشكل منفصل، يمكنك تحديدها جميعًا في سطر واحد، مثل هذا:

```css
padding: 10px 20px 10px 20px;
```

تعمل هذه القيم الأربع مثل عقارب الساعة: الأعلى، اليمين، الأسفل، اليسار، وستنتج نفس النتيجة تمامًا كما لو استخدمت تعليمات الحشو الخاصة بكل جانب.

# --instructions--

استخدم تدوين عقارب الساعة لإعطاء الصنف `.blue-box` حشو `padding` بمقدار `40px` في الأعلى وعلى الجانب الأيسر، ولكن فقط `20px` في الأسفل وعلى الجانب الأيمن.

# --hints--

يجب أن يمنح صنف `blue-box` الحشو (`padding`) بمقدار `40px` في أعلى العناصر.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '40px');
```

يجب أن يمنح صنف `blue-box` الجانب الأيمن للعناصر `20px` من `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingRight = window.getComputedStyle(blueBox)['padding-right'];
assert.strictEqual(paddingRight, '20px');
```

يجب أن يمنح صنف `blue-box` أسفل العناصر مقدار `20px` من `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingBottom = window.getComputedStyle(blueBox)['padding-bottom'];
assert.strictEqual(paddingBottom, '20px');
```

يجب أن يمنح صنف `blue-box` الجهة اليسرى للعناصر `40px` من `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingLeft = window.getComputedStyle(blueBox)['padding-left'];
assert.strictEqual(paddingLeft, '40px');
```

يجب أن تستخدم تدوين عقارب الساعة لتعيين الحشو لصنف `blue-box`.

```js
const css =  __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(css, /\.blue-box\s*{[\s\S]*padding\s*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
