---
id: bad87fee1348bd9aedf08719
title: استخدام رمز اللون السداسي المختصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpKAm'
forumTopicId: 18338
dashedName: use-abbreviated-hex-code
---

# --description--

يشعر كثير من الناس بالإرهاق من احتمالات أكثر من 16 مليون لون. ومن الصعب تذكر رمز اللون السداسي. لحسن الحظ، يمكنك اختصاره.

على سبيل المثال، يمكن اختصار رمز اللون الأحمر `#FF0000` إلى `#F00`. هذا الشكل المختصر يعطي رقمًا واحدًا للأحمر، ورقمًا واحدًا للأخضر، ورقمًا واحدًا للأزرق.

هذا يقلل العدد الإجمالي للألوان الممكنة إلى حوالي 4,000 لون. لكن المتصفحات ستفسر `#FF0000` و `#F00` على أنهما نفس اللون تمامًا.

# --instructions--

تابع، جرّب استخدام رموز اللون السداسية المختصرة لتلوين العناصر الصحيحة.

<table><tbody><tr><th>اللون</th><th>رمز اللون السداسي المختصر</th></tr><tr><td>سماوي</td><td><code>#0FF</code></td></tr><tr><td>أخضر</td><td><code>#0F0</code></td></tr><tr><td>أحمر</td><td><code>#F00</code></td></tr><tr><td>فوشيا</td><td><code>#F0F</code></td></tr></tbody></table>

# --hints--

يجب أن يُعطى عنصر `h1` الذي يحتوي على النص `I am red!` خاصية `color` باللون الأحمر.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

يجب استخدام رمز اللون السداسي المختصر للون الأحمر بدلاً من رمز اللون `#FF0000`.

```js
assert.match(code,/\.red-text\s*?{\s*?color\s*:\s*?#F00\s*?;?\s*?}/gi);
```

يجب أن يُعطى عنصر `h1` الذي يحتوي على النص `I am green!` خاصية `color` باللون الأخضر.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

يجب استخدام رمز اللون السداسي المختصر للون الأخضر بدلاً من رمز اللون `#00FF00`.

```js
assert.match(__helpers.removeCssComments(code), /\.green-text\s*?{\s*?color\s*:\s*?#0F0\s*?;?\s*?}/gi);
```

يجب أن يُعطى عنصر `h1` الذي يحتوي على النص `I am cyan!` خاصية `color` باللون السماوي.

```js
const cyanText = document.querySelector('.cyan-text');
const color = window.getComputedStyle(cyanText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 255)');
```

يجب استخدام رمز اللون السداسي المختصر للون السماوي بدلاً من رمز اللون `#00FFFF`.

```js
assert.match(__helpers.removeCssComments(code), /\.cyan-text\s*?{\s*?color\s*:\s*?#0FF\s*?;?\s*?}/gi);
```

يجب أن يُعطى عنصر `h1` الذي يحتوي على النص `I am fuchsia!` خاصية `color` باللون الفوشيا.

```js
const fuchsiaText = document.querySelector('.fuchsia-text');
const color = window.getComputedStyle(fuchsiaText)['color']; 
assert.strictEqual(color,'rgb(255, 0, 255)');
```

يجب استخدام رمز اللون السداسي المختصر للون الفوشيا بدلاً من رمز اللون `#FF00FF`.

```js
assert.match(__helpers.removeCssComments(code), /\.fuchsia-text\s*?{\s*?color\s*:\s*?#F0F\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .fuchsia-text {
    color: #000000;
  }
  .cyan-text {
    color: #000000;
  }
  .green-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #F00;
  }
  .fuchsia-text {
    color: #F0F;
  }
  .cyan-text {
    color: #0FF;
  }
  .green-text {
    color: #0F0;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```
