---
id: bad87fee1348bd9aedf08721
title: استخدام شفرة Hex لخلط الألوان
challengeType: 0
videoUrl: 'https://scrimba.com/c/cK89PhP'
forumTopicId: 18359
dashedName: use-hex-code-to-mix-colors
---

# --description--

للمراجعة، تستخدم شفرة hex ستة أرقام عشرية سداسية لتمثيل الألوان، رقمين لكل من مكوِّنات الأحمر (R)، الأخضر (G)، والأزرق (B).

من هذه الألوان الثلاثة النقية (الأحمر، الأخضر، والأزرق)، يمكننا تغيير كميات كل منها لإنشاء أكثر من 16 مليون لون آخر!

على سبيل المثال، البرتقالي هو أحمر نقي ممزوج ببعض الأخضر، ولا يحتوي على أزرق. في شفرة hex، يُترجم هذا إلى `#FFA500`.

الرقم `0` هو أدنى رقم في شفرة hex، ويمثل غياب اللون تمامًا.

الرقم `F` هو أعلى رقم في شفرة hex، ويمثل أقصى سطوع ممكن.

# --instructions--

استبدل كلمات الألوان في عنصر `style` الخاص بنا بشفرات hex الصحيحة الخاصة بها.

<table><tbody><tr><th>اللون</th><th>شفرة Hex</th></tr><tr><td>Dodger Blue</td><td><code>#1E90FF</code></td></tr><tr><td>Green</td><td><code>#00FF00</code></td></tr><tr><td>Orange</td><td><code>#FFA500</code></td></tr><tr><td>Red</td><td><code>#FF0000</code></td></tr></tbody></table>

# --hints--

يجب أن يُعطى عنصر `h1` الخاص بك مع النص `I am red!` خاصية `color` باللون الأحمر.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

يجب استخدام `hex code` للون الأحمر بدلًا من كلمة `red`.

```js
assert.match(code, /\.red-text\s*?{\s*?color\s*:\s*?(#FF0000|#F00)\s*?;?\s*?}/gi);
```

يجب أن يُعطى عنصر `h1` الخاص بك مع النص `I am green!` خاصية `color` باللون الأخضر.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

يجب استخدام `hex code` للون الأخضر بدلًا من كلمة `green`.

```js
assert.match(code, /\.green-text\s*?{\s*?color\s*:\s*?(#00FF00|#0F0)\s*?;?\s*?}/gi);
```

يجب أن يُعطى عنصر `h1` الخاص بك مع النص `I am dodger blue!` خاصية `color` بلون dodger blue.

```js
const blueText = document.querySelector('.dodger-blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(30, 144, 255)');
```

يجب استخدام `hex code` للون dodger blue بدلًا من كلمة `dodgerblue`.

```js
assert.match(code, /\.dodger-blue-text\s*?{\s*?color\s*:\s*?#1E90FF\s*?;?\s*?}/gi);
```

يجب أن يُعطى عنصر `h1` الخاص بك مع النص `I am orange!` خاصية `color` باللون البرتقالي.

```js
const orangeText = document.querySelector('.orange-text');
const color = window.getComputedStyle(orangeText)['color']; 

assert.strictEqual(color, 'rgb(255, 165, 0)');
```

يجب استخدام `hex code` للون البرتقالي بدلًا من كلمة `orange`.

```js
assert.match(code, /\.orange-text\s*?{\s*?color\s*:\s*?#FFA500\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: black;
  }
  .green-text {
    color: black;
  }
  .dodger-blue-text {
    color: black;
  }
  .orange-text {
    color: black;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #FF0000;
  }
  .green-text {
    color: #00FF00;
  }
  .dodger-blue-text {
    color: #1E90FF;
  }
  .orange-text {
    color: #FFA500;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```
