---
id: bad87fee1348bd9aedf08822
title: ضبط هامش عنصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJarHW'
forumTopicId: 16654
dashedName: adjust-the-margin-of-an-element
---

# --description--

يتحكم `margin` الخاص بالعنصر في مقدار المساحة بين `border` الخاص بالعنصر والعناصر المحيطة به.

هنا، نرى أن الصندوق الأزرق والصندوق الأحمر مسكنان داخل الصندوق الأصفر. لاحظ أن الصندوق الأحمر له `margin` أكبر من الصندوق الأزرق، مما يجعله يبدو أصغر.

عندما تزيد `margin` الخاص بالصندوق الأزرق، ستزداد المسافة بين حدوده والعناصر المحيطة به.

# --instructions--

غيّر قيمة `margin` للمربع الأزرق لتطابق قيمة المربع الأحمر.

# --hints--

يجب أن يمنح صنف `blue-box` العناصر `20px` من `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
