---
id: bad82fee1348bd9aedf08721
title: استخدام RGB لخلط الألوان
challengeType: 0
videoUrl: 'https://scrimba.com/c/cm24JU6'
forumTopicId: 18368
dashedName: use-rgb-to-mix-colors
---

# --description--

تمامًا كما في شفرة hex، يمكنك خلط الألوان في RGB باستخدام تركيبات من قيم مختلفة.

# --instructions--

استبدل رموز hex في عنصر `style` بالقيم الصحيحة لـ RGB.

<table><tbody><tr><th>اللون</th><th>RGB</th></tr><tr><td>أزرق</td><td><code>rgb(0, 0, 255)</code></td></tr><tr><td>أحمر</td><td><code>rgb(255, 0, 0)</code></td></tr><tr><td>أوركيد</td><td><code>rgb(218, 112, 214)</code></td></tr><tr><td>سيينا</td><td><code>rgb(160, 82, 45)</code></td></tr></tbody></table>

# --hints--

يجب أن يحتوي عنصر `h1` الذي يحمل النص `I am red!` على `color` أحمر.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

يجب استخدام `rgb` للون الأحمر.

```js
assert.match(code, /\.red-text\s*{\s*color\s*:\s*rgb\(\s*255\s*,\s*0\s*,\s*0\s*\)\s*;?\s*}/gi);
```

يجب أن يحتوي عنصر `h1` الذي يحمل النص `I am orchid!` على `color` أوركيد.

```js
const orchidText = document.querySelector('.orchid-text');
const color = window.getComputedStyle(orchidText)['color']; 
assert.strictEqual(color, 'rgb(218, 112, 214)');
```

يجب استخدام `rgb` للون الأوركيد.

```js
assert.match(__helpers.removeCssComments(code), /\.orchid-text\s*{\s*color\s*:\s*rgb\(\s*218\s*,\s*112\s*,\s*214\s*\)\s*;?\s*}/gi);
```

يجب أن يحتوي عنصر `h1` الذي يحمل النص `I am blue!` على `color` أزرق.

```js
const blueText = document.querySelector('.blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

يجب استخدام `rgb` للون الأزرق.

```js
assert.match(__helpers.removeCssComments(code), /\.blue-text\s*{\s*color\s*:\s*rgb\(\s*0\s*,\s*0\s*,\s*255\s*\)\s*;?\s*}/gi);
```

يجب أن يحتوي عنصر `h1` الذي يحمل النص `I am sienna!` على `color` سيينا.

```js
const siennaText = document.querySelector('.sienna-text');
const color = window.getComputedStyle(siennaText)['color']; 
assert.strictEqual(color, 'rgb(160, 82, 45)');
```

يجب استخدام `rgb` للون السيينا.

```js
assert.match(__helpers.removeCssComments(code), /\.sienna-text\s*{\s*color\s*:\s*rgb\(\s*160\s*,\s*82\s*,\s*45\s*\)\s*;?\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .orchid-text {
    color: #000000;
  }
  .sienna-text {
    color: #000000;
  }
  .blue-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="orchid-text">I am orchid!</h1>

<h1 class="sienna-text">I am sienna!</h1>

<h1 class="blue-text">I am blue!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: rgb(255, 0, 0);
  }
  .orchid-text {
    color: rgb(218, 112, 214);
  }
  .sienna-text {
    color: rgb(160, 82, 45);
  }
  .blue-text {
    color:rgb(0, 0, 255);
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="orchid-text">I am orchid!</h1>

<h1 class="sienna-text">I am sienna!</h1>

<h1 class="blue-text">I am blue!</h1>
```
