---
id: bad87fee1348bd8aedf06756
title: تجاوز إعلانات الأصناف بتنسيق خصائص المعرف
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpDhB'
forumTopicId: 18251
dashedName: override-class-declarations-by-styling-id-attributes
---

# --description--

أثبتنا للتو أن المتصفحات تقرأ CSS من الأعلى إلى الأسفل حسب ترتيب الإعلان. هذا يعني أنه في حالة وجود تعارض، يستخدم المتصفح إعلان CSS الذي جاء أخيرًا. لاحظ أنه حتى لو وضعنا `blue-text` قبل `pink-text` في أصناف عنصر `h1`، فإنه سيظل ينظر إلى ترتيب الإعلان وليس ترتيب الاستخدام!

لكننا لم ننتهِ بعد. هناك طرق أخرى لتجاوز CSS. هل تتذكر خصائص المعرف؟

دعنا نتجاوز أصناف `pink-text` و `blue-text`، ونجعل عنصر `h1` برتقاليًا، عن طريق إعطاء عنصر `h1` معرفًا ثم تنسيق ذلك المعرف.

# --instructions--

أعطِ عنصر `h1` الخاص بك خاصية `id` بقيمة `orange-text`. تذكر، أنماط المعرف تبدو هكذا:

```html
<h1 id="orange-text">
```

اترك أصناف `blue-text` و `pink-text` على عنصر `h1` الخاص بك.

أنشئ إعلان CSS لمعرف `orange-text` داخل عنصر `style` الخاص بك. إليك مثالًا على ما يبدو عليه هذا:

```css
#brown-text {
  color: brown;
}
```

**ملاحظة:** لا يهم إذا أعلنت هذا الـ CSS فوق أو تحت صنف `pink-text`، لأن الخاصية `id` ستأخذ الأولوية دائمًا.

# --hints--

يجب أن يحتوي عنصر `h1` على الصنف `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

يجب أن يحتوي عنصر `h1` على الصنف `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

يجب أن يحتوي عنصر `h1` على المعرف `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'),'orange-text');
```

يجب أن يكون هناك عنصر `h1` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

يجب أن يحتوي معرف `orange-text` على إعلان CSS.

```js
assert.match(__helpers.removeCssComments(code), /#orange-text\s*{/gi);
```

يجب ألا يحتوي عنصر `h1` على أي خصائص `style`.

```js
assert.notMatch(__helpers.removeHtmlComments(code), /<h1.*style.*>/gi);
```

يجب أن يكون عنصر `h1` برتقالي اللون.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 165, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
  #orange-text {
    color: orange;
  }  
</style>
<h1 id="orange-text"  class="pink-text blue-text">Hello World!</h1>
```
