---
id: bad87fee1348bd9aecf08806
title: استخدام صنف CSS لتنسيق عنصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MvDtV'
forumTopicId: 18337
dashedName: use-a-css-class-to-style-an-element
---

# --description--

الصنف هو أنماط قابلة لإعادة الاستخدام يمكن إضافتها إلى عناصر HTML.

إليك مثال على إعلان صنف CSS:

```html
<style>
  .blue-text {
    color: blue;
  }
</style>
```

يمكنك أن ترى أننا أنشأنا صنف CSS يسمى `blue-text` داخل وسم `<style>`. يمكنك تطبيق الصنف على عنصر HTML بهذا الشكل: `<h2 class="blue-text">CatPhotoApp</h2>`. لاحظ أنه في عنصر `style` الخاص بـ CSS، تبدأ أسماء الأصناف بنقطة. أما في خاصية الصنف في عناصر HTML، فلا تتضمن اسم الصنف النقطة.

# --instructions--

داخل عنصر `style` الخاص بك، غيّر المحدد `h2` إلى `.red-text` وحدث قيمة اللون من `blue` إلى `red`.

أضف إلى عنصر `h2` الخاص بك الخاصية `class` بقيمة `red-text`.

# --hints--

يجب أن يكون عنصر `h2` الخاص بك باللون الأحمر.

```js
const h2Element = document.querySelector('h2'); 
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

يجب أن يحتوي عنصر `h2` الخاص بك على الصنف `red-text`.

```js
assert.isTrue(document.querySelector('h2').classList.contains('red-text'));
```

يجب أن يعلن ملف الأنماط الخاص بك عن صنف `red-text` ويضبط لونه على `red`.

```js
assert.match(__helpers.removeCssComments(code), /\.red-text\s*\{\s*color\s*:\s*red;?\s*\}/g);
```

لا يجب أن تستخدم إعلانات الأنماط المضمّنة مثل `style="color: red"` في عنصر `h2` الخاص بك.

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

# --seed--

## --seed-contents--

```html
<style>
  h2 {
    color: blue;
  }
</style>

<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  .red-text {
    color: red;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
