---
id: bad87fee1348bd9afdf08726
title: استخدام تدوين عقارب الساعة لتحديد هامش عنصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpybAd'
forumTopicId: 18345
dashedName: use-clockwise-notation-to-specify-the-margin-of-an-element
---

# --description--

لنجرّب هذا مرة أخرى، لكن مع `margin` هذه المرة.

بدلاً من تحديد خاصيات `margin-top` و`margin-right` و`margin-bottom` و`margin-left` لعنصر بشكل منفصل، يمكنك تحديدها كلها في سطر واحد، مثل هذا:

```css
margin: 10px 20px 10px 20px;
```

تعمل هذه القيم الأربع مثل عقارب الساعة: الأعلى، اليمين، الأسفل، اليسار، وستنتج نفس النتيجة تمامًا كما لو استخدمت تعليمات الهامش الخاصة بكل جانب.

# --instructions--

استخدم تدوين عقارب الساعة لإعطاء العنصر الذي يحمل صنف `blue-box` هامشًا قدره `40px` في الأعلى وعلى الجانب الأيسر، ولكن فقط `20px` في الأسفل وعلى الجانب الأيمن.

# --hints--

يجب أن يمنح صنف `blue-box` أعلى العناصر هامشًا بقيمة `40px` من `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '40px');
```

يجب أن يمنح صنف `blue-box` الجانب الأيمن للعناصر هامشًا بقيمة `20px` من نوع `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginRight = window.getComputedStyle(blueBox)['margin-right'];
assert.strictEqual(marginRight, '20px');
```

يجب أن يمنح صنف `blue-box` أسفل العناصر هامشًا بقيمة `20px` من `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginBottom = window.getComputedStyle(blueBox)['margin-bottom'];
assert.strictEqual(marginBottom, '20px');
```

يجب أن يمنح صنف `blue-box` الجهة اليسرى للعناصر هامشًا بقيمة `40px` من نوع `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginLeft = window.getComputedStyle(blueBox)['margin-left'];
assert.strictEqual(marginLeft, '40px');
```

يجب أن تستخدم تدوين عقارب الساعة لتعيين هامش صنف `blue-box`.

```js
const cssCode = __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(cssCode, /\.blue-box\s*{[\s\S]*margin[\s]*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    margin: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
