---
id: bad87fee1348bd9aedf08805
title: استخدام محددات CSS لتنسيق العناصر
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKMBT2'
forumTopicId: 18349
dashedName: use-css-selectors-to-style-elements
---

# --description--

مع CSS، هناك المئات من خصائص CSS التي يمكنك استخدامها لتغيير مظهر العنصر على صفحتك.

عندما أدخلت `<h2 style="color: red;">CatPhotoApp</h2>`، كنت تُنسق ذلك العنصر الفردي `h2` باستخدام CSS مضمَّن، وهو اختصار لـ Cascading Style Sheets.

هذه طريقة واحدة لتحديد نمط عنصر، لكن هناك طريقة أفضل لتطبيق CSS.

في أعلى شفرتك، أنشئ كتلة `style` مثل هذه:

```html
<style>
</style>
```

داخل تلك الكتلة، يمكنك إنشاء <dfn>محدِّد CSS</dfn> لجميع عناصر `h2`. على سبيل المثال، إذا أردت أن تكون كل عناصر `h2` باللون الأحمر، تضيف قاعدة نمط تبدو هكذا:

```html
<style>
  h2 {
    color: red;
  }
</style>
```

لاحظ أنه من المهم وجود كل من أقواس متعرجة (`{` و `}`) الافتتاحية والختامية حول قاعدة (أو قواعد) نمط كل عنصر. كما يجب التأكد من أن تعريف نمط العنصر موجود بين وسمَي النمط الافتتاحي والختامي. وأخيرًا، تأكد من إضافة فاصلة منقوطة في نهاية كل قاعدة نمط للعنصر.

# --instructions--

احذف خاصية النمط من عنصر `h2`، وأنشئ بدلًا من ذلك كتلة `style` في CSS. أضف قواعد CSS اللازمة لجعل جميع عناصر `h2` زرقاء.

# --hints--

يجب حذف خاصية `style` من عنصر `h2` الخاص بك.

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

يجب أن تنشئ عنصر `style`.

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
assert.isAtLeast(document.querySelectorAll('style:not(.fcc-hide-header)').length, 1);
```

يجب أن يكون عنصر `h2` الخاص بك أزرق.

```js
const h2Element = document.querySelector('h2');
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

يجب أن يكون إعلان `h2` في ورقة الأنماط صالحًا مع فاصلة منقوطة وقوس إغلاق.

```js
assert.match(__helpers.removeCssComments(code), /h2\s*\{\s*color\s*:.*;\s*\}/g);
```

يجب أن تكون كل عناصر `style` صالحة ولها وسم إغلاق.

```js
assert.match(__helpers.removeHtmlComments(code), /<\/style>/g);

const closingElementLength = __helpers.removeHtmlComments(code).match(/<\/style>/g).length;
const openingElementsLength = __helpers.removeHtmlComments(code).match(/<style((\s)*((type|media|scoped|title|disabled)="[^"]*")?(\s)*)*>/g).length;
assert.strictEqual(closingElementLength, openingElementsLength);
```

# --seed--

## --seed-contents--

```html
<h2 style="color: red;">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  h2 {
    color: blue;
  }
</style>
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
