---
id: bad87fee1348bd9aedf08824
title: إضافة حشو مختلف لكل جانب من عناصر الصفحة
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB7mwUw'
forumTopicId: 16634
dashedName: add-different-padding-to-each-side-of-an-element
---

# --description--

أحيانًا تريد تخصيص عنصر بحيث يحتوي على كميات مختلفة من `padding` على كل جانب من جوانبه.

تسمح CSS بالتحكم في `padding` لكل من الجوانب الأربعة الفردية للعنصر باستخدام الخاصيات `padding-top`، و`padding-right`، و`padding-bottom`، و`padding-left`.

# --instructions--

أعطِ الصندوق الأزرق `padding` بقيمة `40px` في الأعلى وعلى الجانب الأيسر، ولكن فقط `20px` في الأسفل وعلى الجانب الأيمن.

# --hints--

يجب أن تعطي صنف `blue-box` الجزء العلوي من العناصر `padding` بقيمة `40px`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '40px');
```

يجب أن تعطي صنف `blue-box` الجانب الأيمن من العناصر `padding` بقيمة `20px`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingRight = window.getComputedStyle(blueBox)['padding-right'];
assert.strictEqual(paddingRight, '20px');
```

يجب أن تعطي صنف `blue-box` الجزء السفلي من العناصر `padding` بقيمة `20px`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingBottom = window.getComputedStyle(blueBox)['padding-bottom'];
assert.strictEqual(paddingBottom, '20px');
```

يجب أن تعطي صنف `blue-box` الجانب الأيسر من العناصر `padding` بقيمة `40px`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingLeft = window.getComputedStyle(blueBox)['padding-left'];
assert.strictEqual(paddingLeft, '40px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding-top: 40px;
    padding-right: 20px;
    padding-bottom: 20px;
    padding-left: 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding-top: 40px;
    padding-right: 20px;
    padding-bottom: 20px;
    padding-left: 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding-top: 40px;
    padding-right: 20px;
    padding-bottom: 20px;
    padding-left: 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
