---
id: bad87fee1348bd9aedf04756
title: تجاوز الأنماط في CSS اللاحقة
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDQug'
forumTopicId: 18253
dashedName: override-styles-in-subsequent-css
---

# --description--

صنف `pink-text` تجاوز إعلان CSS لعنصر `body`!

لقد أثبتنا للتو أن أصنافنا ستتجاوز CSS لعنصر `body`. إذًا السؤال المنطقي التالي هو، ماذا يمكننا أن نفعل لتجاوز صنف `pink-text`؟

# --instructions--

أنشئ صنف CSS إضافي يسمى `blue-text` يمنح العنصر اللون الأزرق. تأكد من أنه تحت إعلان صنف `pink-text` الخاص بك.

طبّق صنف `blue-text` على عنصر `h1` بالإضافة إلى صنف `pink-text`، ولنرَ أيهما ينتصر.

تطبيق عدة خصائص صنف على عنصر HTML يتم بوضع مسافة بينها هكذا:

```html
class="class1 class2"
```

**ملاحظة:** لا يهم ترتيب الأصناف في عنصر HTML.

لكن ترتيب إعلانات `class` في قسم `<style>` هو المهم. الإعلان الثاني دائمًا له الأسبقية على الأول. وبما أن `.blue-text` معلن ثانيًا، فإنه يتجاوز خصائص `.pink-text`.

# --hints--

يجب أن يحتوي عنصر `h1` على صنف `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

يجب أن يحتوي عنصر `h1` على صنف `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

يجب أن ينتمي كل من `blue-text` و `pink-text` إلى نفس عنصر `h1`.

```js
assert.isTrue(document.querySelector('.pink-text').classList.contains('blue-text'));
```

يجب أن يكون عنصر `h1` أزرق.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }

  .blue-text {
    color: blue;
  }  
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```
