---
id: 698c55e7aaacfae48a227459
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Fügen Sie unter dem `h2`-Element ein neues `audio`-Element mit dem Text `Cruising for a Musing` hinzu.

Unter dem `h2` fügen Sie ein `p`-Element mit dem Text `Artist: Quincy Larson` hinzu.

# --hints--

Sie sollten ein zweites `h2`-Element haben.

```js
assert.exists(document.querySelector('h2:nth-of-type(2)'));
```

Ihr zweites `h2`-Element sollte den Text `Cruising for a Musing` enthalten.

```js
const h2 = document.querySelector('h2:nth-of-type(2)');
assert.strictEqual(h2?.textContent.trim(), 'Cruising for a Musing');
```

Sie sollten ein zweites `p`-Element haben.

```js
assert.exists(document.querySelector('p:nth-of-type(2)'));
```

Ihr zweites `p`-Element sollte den Text `Artist: Quincy Larson` enthalten.

```js
const p = document.querySelector('p:nth-of-type(2)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
