---
id: 698c4fe24f1f1a17be74a23e
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Erstellen Sie als Nächstes ein `audio`-Element unterhalb des `p`-Elements. In den nächsten Schritten fügen Sie die notwendigen Attribute hinzu, damit das Audioelement Musik abspielt.

# --hints--

Sie sollten ein `audio`-Element haben.

```js
assert.exists(document.querySelector('audio'));
```

Das neue `audio`-Element sollte unterhalb des `p`-Elements sein.

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
