---
id: 698c58745b73289bb6c89ece
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Das letzte Lied heißt `Scratching the Surface`, der Künstler ist Quincy Larson, und die zu verwendende Datei ist `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3`.

Fügen Sie das letzte Lied hinzu, um den Musikplayer zu vervollständigen, und folgen Sie dabei dem gleichen Muster, das Sie für die ersten beiden Lieder verwendet haben.

# --hints--

Sie sollten ein drittes `h2`-Element haben.

```js
assert.exists(document.querySelector('h2:nth-of-type(3)'));
```

Ihr drittes `h2`-Element sollte den Text `Scratching the Surface` enthalten.

```js
const h2 = document.querySelector('h2:nth-of-type(3)');
assert.strictEqual(h2?.textContent.trim(), 'Scratching the Surface');
```

Sie sollten ein drittes `p`-Element haben.

```js
assert.exists(document.querySelector('p:nth-of-type(3)'));
```

Ihr drittes `p`-Element sollte den Text `Artist: Quincy Larson` enthalten.

```js
const p = document.querySelector('p:nth-of-type(3)');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

Sie sollten ein drittes `audio`-Element haben.

```js
assert.exists(document.querySelector('audio:nth-of-type(3)'));
```

Das neue `audio`-Element sollte unter dem `p`-Element sein.

```js
assert.exists(document.querySelector('p:nth-of-type(3) + audio:nth-of-type(3)'));
```

Ihr `audio`-Element sollte ein `src`-Attribut haben.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('src'));
```

Ihr `audio`-Element sollte ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3` haben.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

Ihr `audio`-Element sollte das `loop`-Attribut haben.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Ihr `audio`-Element sollte das `controls`-Attribut haben.

```js
const audio = document.querySelector('audio:nth-of-type(3)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>

  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3" loop controls></audio>

  <h2>Scratching the Surface</h2>
  <p>Artist: Quincy Larson</p>

  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3" loop controls></audio>
  
</body>
</html>
```
