---
id: 698c4e69abe14a9746c72312
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop bauen Sie einen HTML-Musikplayer. Der HTML-Boilerplate-Code wurde für Sie bereitgestellt.

Erstellen Sie ein `h1`-Element und geben Sie ihm den Text `freeCodeCamp Tunes`.

# --hints--

Sie sollten ein `h1`-Element haben.

```js
assert.exists(document.querySelector('h1'));
```

Ihr `h1`-Element sollte den Text `freeCodeCamp Tunes` enthalten.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'freeCodeCamp Tunes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
