---
id: 698c5229e47d57c6dfaa91bd
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Das `controls`-Attribut stellt Wiedergabesteuerungen bereit, darunter Pause/Fortsetzen der Wiedergabe, Springen und Lautstärkeregelung für das `audio`-Element.

Das `controls`-Attribut ist ein Boolean-Attribut und benötigt keinen Wert.

Fügen Sie das `controls`-Attribut zum `audio`-Element hinzu.

Jetzt sollten Sie das `audio`-Element auf der Seite sehen.

# --hints--

Ihr `audio`-Element sollte das `controls`-Attribut haben.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
