---
id: 698c50535b54b2d75fa985fb
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Das Attribut `loop` startet die Audio-Wiedergabe neu, sobald sie abgeschlossen ist. Denken Sie an ein Internet-Meme, das die Wiedergabe wiederholt. Wenn Sie das Attribut `loop` weglassen, wird die Audio nur einmal abgespielt.

Das Attribut `loop` ist ein Boolean-Attribut und benötigt keinen Wert.

Fügen Sie das Attribut `loop` zum Element `audio` hinzu.

# --hints--

Ihr Element `audio` sollte das Attribut `loop` haben.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
