---
id: 698c4f4e965bd853ad5084e1
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie unter dem `h1` ein `h2`-Element mit dem Text `Can't Stay Down` hinzu. Dies ist der Titel des ersten Songs.

Fügen Sie unter dem `h2`-Element ein `p`-Element mit dem Text `Artist: Quincy Larson` hinzu.

# --hints--

Sie sollten ein `h2`-Element haben.

```js
assert.exists(document.querySelector('h2'));
```

Ihr `h2`-Element sollte unter dem `h1`-Element sein.

```js
assert.exists(document.querySelector('h1 + h2'));
```

Ihr `h2`-Element sollte den Text `Can't Stay Down` enthalten.

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

Sie sollten ein `p`-Element haben.

```js
assert.exists(document.querySelector('p'));
```

Ihr `p`-Element sollte den Text `Artist: Quincy Larson` enthalten.

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
