---
id: 698c5710136539797f296098
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Fügen Sie unter dem `p`-Element ein `audio`-Element hinzu, geben Sie ihm ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3` sowie die Attribute `loop` und `controls`.

# --hints--

Sie sollten ein zweites `audio`-Element haben.

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

Das neue `audio`-Element sollte unter dem `p`-Element sein.

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

Ihr `audio`-Element sollte ein `src`-Attribut haben.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

Ihr `audio`-Element sollte ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3` haben.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

Ihr `audio`-Element sollte das `loop`-Attribut haben.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Ihr `audio`-Element sollte das `controls`-Attribut haben.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
