---
id: 68c2fb153760d075321c5da3
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Camperbot створив сторінку блогу про подорож у програмуванні, але здається, що на сторінці є проблеми з доступністю та погані практики. Ваше завдання в цьому практичному занятті — виправити ці проблеми для Camperbot.

Перша множина помилок пов’язана з використанням заголовків. У попередньому уроці ви дізналися, що найкраща практика — використовувати лише один `h1` на сторінку. Але здається, Camperbot використовує кілька `h1`.

Залиште `<h1>Welcome to Camperbot's Blog</h1>` без змін і змініть інші `h1` елементи на `h2` елементи.

# --hints--

Ви не повинні змінювати `<h1>Welcome to Camperbot's Blog</h1>`.

```js
const heading = document.querySelectorAll("h1")[0];
assert.equal(heading?.textContent, "Welcome to Camperbot's Blog");
```

Ви повинні змінити `h1` елемент із текстом `My Journey Learning to Code` на `h2` елемент.

```js
const firstPost = document.getElementById("post1");
assert.equal(firstPost?.tagName, "H2");
assert.equal(firstPost?.textContent, "My Journey Learning to Code");
```

Ви повинні змінити `h1` елемент із текстом `Accessibility Matters` на `h2` елемент.

```js
const secondPost = document.getElementById("post2");
assert.equal(secondPost?.tagName, "H2");
assert.equal(secondPost?.textContent, "Accessibility Matters");
```

Ви повинні змінити `h1` елемент із текстом `What's Next?` на `h2` елемент.

```js
const thirdPost = document.getElementById("post3");
assert.equal(thirdPost?.tagName, "H2");
assert.equal(thirdPost?.textContent, "What's Next?");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <h2>Navigation</h2>
  <ul>
    <li><a href="#post1">My Journey</a></li>
    <li><a href="#post2">Accessibility</a></li>
    <li><a href="#post3">Next Steps</a></li>
  </ul>
--fcc-editable-region--
  <div>
    <h1 id="post1">My Journey Learning to Code</h1>
    <p>I started learning to code a few months ago and it's been a wild ride!</p>
    
    <h4>Early Challenges</h4>
    <p>At first, syntax was really confusing.</p>
    
    <h4>Breakthroughs</h4>
    <p>Eventually things started to click.</p>
  </div>

  <div>
    <h1 id="post2">Accessibility Matters</h1>
    <p>Today I learned that not everyone uses the web the same way I do.</p>

    <h5>Screen Readers</h5>
    <p>These tools help visually impaired users browse websites.</p>
  </div>

  <div>
    <h1 id="post3">What's Next?</h1>
    <p>I'm excited to dive into JavaScript and build interactive features!</p>

    <h3>Coming soon: My first JavaScript project!</h3>
    <p>Stay tuned for some exciting interactive blog features.</p>
  </div>
--fcc-editable-region--
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
</body>
</html>
```
