---
id: 685239f0cad43c7a909cbd98
title: налагодження сторінки з усиновленням тварин
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Селлі, власниця магазину з усиновлення тварин, створила свою першу вебсторінку, але там є деякі проблеми.

Ваше завдання — виправити всі помилки, щоб Селлі могла продовжувати створювати свою сторінку.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Селлі хоче використати зображення котів, але воно відображається неправильно. Вам потрібно виправити наступне в елементі `img`:
   - Замініть атрибут `href` на правильний атрибут для джерела зображення.
   - Замініть атрибут `att` на правильний атрибут, який представляє короткий опис зображення.
   - Видаліть кінцевий тег `</img>`, бо елементи `img` є порожніми і не мають кінцевих тегів.
2. Селлі хоче використати посилання, щоб направляти користувачів на сторінки з собаками та котами. Але посилання працюють неправильно. Вам потрібно виправити наступне в елементах `a`:
   - Замініть обидва атрибути `src` на правильні атрибути, які використовуються для вказання URL-адрес.

# --hints--

Ваш елемент `img` має мати атрибут `src` замість атрибута `href`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

Ваш елемент `img` має мати атрибут `alt` замість неіснуючого атрибута `att`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

Ваш елемент `img` не повинен мати кінцевий тег `</img>`.

```js
assert.notMatch(code, /<\/img>/);
```

Ваш елемент `a` з текстом `Visit cats page` має мати атрибут `href` замість атрибута `src`.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

Ваш елемент `a` з текстом `Visit dogs page` має мати атрибут `href` замість атрибута `src`.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
