---
id: 67cd525b0e575b1e5abf4c03
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Kama ilivyotajwa awali, kitendakazi cha set kinahusika na kusahihisha hali. Njia ya kawaida wakati wa kubadilisha kati ya thamani za kweli au siyo kweli ni kusahihisha hali kwa njia hii:

```js
const [isUserLoggedIn, setIsUserLoggedIn] = useState(false);

// updating state
setIsUserLoggedIn(!isUserLoggedIn);
```

Mfano ufuatao unatumia kitendakazi cha set kinachoitwa `setIsUserLoggedIn` kusahihisha hali ya kuingia kwa mtumizi. Ni desturi ya kawaida kutumia kielezi cha mantiki cha NOT (`!`) kubadilisha kati ya `true` na `false`. Kwa kuwa thamani ya awali ni `false`, kutumia kielezi cha `!` kutabadilisha kuwa `true`.

Ndani ya kitendakazi cha `handleToggleVisibility`, tumia `setIsVisible` kusahihisha hali ya uonekano kwa kutumia kielezi cha mantiki cha NOT (`!`).

# --hints--

Unapaswa kutumia `setIsVisible` kusahihisha uonekano wa maandishi. Rejelea mfano ikiwa unahitaji msaada.

```js
assert.match(code, /setIsVisible\s*\(\s*!isVisible\s*\)/);  
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(false);
  
  const handleToggleVisibility = () => {
  --fcc-editable-region--
    
  --fcc-editable-region--
  }
  
  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      {isVisible && <p id="message">I love freeCodeCamp!</p>}
    </div>
  );
};
```
