---
id: 67cd49e05693111a41771f29
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kuunda vigezo vya hali kwa kutumia `useState` hook na ufumbuzi wa muundo wa safu ya data:

```js
const [user, setUser] = useState("Jessica");
```

Mkataba wa kuitaja vigezo vya hali ni kutumia muundo wa `example` na `setExample`.

`useState` inarudisha safu ya data yenye thamani mbili. Thamani ya kwanza ni hali ya sasa. Thamani ya pili ni kitendakazi cha kuweka thamani kitakachotumika kusahisha hali.

Unda kigezo kipya cha hali na kitendakazi cha kuweka thamani kinachoitwa `isVisible` na `setIsVisible`. Kumbuka kutumia sintaksia ya ufumbuzi wa muundo wa safu ya data kama ilivyoonyeshwa katika mfano.

Kisha weka hook ya `useState` na thamani ya awali ya `false`.

# --hints--

Unapaswa kutumia sintaksia ya ufumbuzi wa muundo wa safu ya data kwa vigezo vya `isVisible` na `setIsVisible`. Rejelea tena mfano.

```js
assert.match(code, /(const|let|var)\s+\[\s*(isVisible)\s*,\s*(setIsVisible)\s*\]\s*/);
```

Unapaswa kutumia hook ya `useState`.

```js
  const abuseState = __helpers.spyOn(React, "useState");

  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.ToggleApp);
  assert.isAtLeast(abuseState.calls.length, 1);
```

Hook yako ya `useState` inapaswa kuwa na thamani ya awali ya `false`.

```js
  const abuseState = __helpers.spyOn(React, "useState");

  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.ToggleApp);
  assert.equal(abuseState.calls[0]?.[0], false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  --fcc-editable-region--

  --fcc-editable-region--
  
  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      <p id="message">I love freeCodeCamp!</p>
    </div>
  );
};
```
