---
id: 587d7fb5367417b2b2512c04
title: rimuovere un pacchetto dalle tue dipendenze
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

Hai già provato diversi modi per gestire le dipendenze del tuo progetto usando la sezione dependencies del package.json. Hai anche incluso pacchetti esterni aggiungendoli al file e hai indicato a npm quali tipi di versioni vuoi, usando caratteri speciali come la tilde o l'accento circonflesso.

Ma cosa succede se vuoi rimuovere un pacchetto esterno che non ti serve più? Probabilmente l'hai già capito: basta eliminare la coppia chiave-valore corrispondente a quel pacchetto dalle tue dipendenze.

Questo stesso metodo vale anche per rimuovere altri campi nel tuo package.json.

# --instructions--

Rimuovi il pacchetto `@freecodecamp/example` dalle tue dipendenze.

**Nota:** assicurati di avere il giusto numero di virgole dopo averlo rimosso.

# --hints--

`"dependencies"` non dovrebbe includere `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
