---
id: 587d7fb4367417b2b2512bfd
title: aggiungere parole chiave al tuo package.json
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

Il campo `keywords` è dove puoi descrivere il tuo progetto usando parole chiave correlate. Ecco un esempio:

```json
"keywords": [ "descriptive", "related", "words" ],
```

Come puoi vedere, questo campo è strutturato come un array di stringhe racchiuse tra doppi apici.

# --instructions--

Aggiungi un array di stringhe adatte al campo `keywords` nel file package.json del tuo progetto.

Una delle parole chiave dovrebbe essere "freecodecamp".

# --hints--

package.json dovrebbe avere una chiave "keywords" valida

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Il campo "keywords" dovrebbe essere un array

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

"keywords" dovrebbe includere "freecodecamp"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
