---
id: 587d7fb4367417b2b2512c00
title: espandi il tuo progetto con pacchetti esterni da npm
challengeType: 2
forumTopicId: 301527
dashedName: expand-your-project-with-external-packages-from-npm
---

# --description--

Uno dei motivi principali per usare un gestore di pacchetti è la sua potente gestione delle dipendenze. Invece di dover controllare manualmente di avere tutte le dipendenze ogni volta che configuri un progetto su un nuovo computer, npm installa automaticamente tutto per te. Ma come fa npm a sapere esattamente di cosa ha bisogno il tuo progetto? Ecco la sezione `dependencies` del tuo file package.json.

In questa sezione, i pacchetti richiesti dal tuo progetto sono archiviati usando il seguente formato:

```json
"dependencies": {
  "package-name": "version",
  "express": "4.14.0"
}

```

# --instructions--

Aggiungi la versione `1.1.0` del pacchetto `@freecodecamp/example` al campo `dependencies` del tuo file `package.json`.

**Nota:** `@freecodecamp/example` è un pacchetto fittizio usato come strumento didattico.

# --hints--

`"dependencies"` dovrebbe includere `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

La versione `"@freecodecamp/example"` dovrebbe essere `"1.1.0"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^[\^\~]?1\.1\.0/,
        'Wrong version of "@freecodecamp/example" installed. It should be 1.1.0'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
