---
id: 587d7fb5367417b2b2512c02
title: Usa il carattere tilde per usare sempre l’ultima versione patch di una dipendenza
challengeType: 2
forumTopicId: 301532
dashedName: use-the-tilde-character-to-always-use-the-latest-patch-version-of-a-dependency
---

# --description--

Nell’ultima sfida, hai detto a npm di includere solo una versione specifica di un pacchetto. È un modo utile per bloccare le dipendenze se devi assicurarti che le diverse parti del progetto restino compatibili tra loro. Ma nella maggior parte dei casi non vuoi perdere le correzioni di bug, dato che spesso includono importanti patch di sicurezza e (si spera) non rompono nulla nel farlo.

Per permettere a una dipendenza npm di aggiornarsi all’ultima versione PATCH, puoi anteporre alla versione della dipendenza il carattere tilde (`~`). Ecco un esempio di come consentire aggiornamenti a qualsiasi versione `1.3.x`.

```json
"package": "~1.3.8"
```

# --instructions--

Nel file package.json, la tua regola attuale su come npm può aggiornare `@freecodecamp/example` è usare una versione specifica (`1.2.13`). Ma ora vuoi permettere l’ultima versione `1.2.x`.

Usa il carattere tilde (`~`) per anteporre la versione di `@freecodecamp/example` nelle tue dipendenze e permetti a npm di aggiornarla a qualsiasi nuova release _patch_.

**Nota:** I numeri di versione non devono essere modificati.

# --hints--

`"dependencies"` dovrebbe includere `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

La versione `"@freecodecamp/example"` dovrebbe corrispondere a `"~1.2.13"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\~1\.2\.13/,
        'Wrong version of "@freecodecamp/example". It should be ~1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
