---
id: 587d7fb3367417b2b2512bfc
title: Aggiungi una descrizione al tuo package.json
challengeType: 2
forumTopicId: 301522
dashedName: add-a-description-to-your-package-json
---

# --description--

La parte successiva di un buon file package.json è il campo `description`; qui va inserita una descrizione breve ma informativa del tuo progetto.

Se un giorno prevedi di pubblicare un pacchetto su npm, questa è la stringa che dovrebbe convincere l’utente a installare o meno il tuo pacchetto. Tuttavia, non è l’unico caso d’uso della descrizione: è un ottimo modo per riassumere cosa fa un progetto. È altrettanto importante in qualsiasi progetto Node.js per aiutare altri sviluppatori, i futuri manutentori o anche te stesso in futuro a capire rapidamente il progetto.

Indipendentemente da cosa prevedi per il tuo progetto, una descrizione è sicuramente consigliata. Ecco un esempio:

```json
"description": "A project that does something awesome",
```

# --instructions--

Aggiungi un `description` al file package.json del tuo progetto.

**Nota:** Ricorda di usare le virgolette doppie per i nomi dei campi (") e le virgole (,) per separare i campi.

# --hints--

package.json dovrebbe avere una chiave "description" valida

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.description, '"description" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
