---
id: 587d7fb4367417b2b2512bfe
title: aggiungi una licenza al tuo package.json
challengeType: 2
forumTopicId: 301523
dashedName: add-a-license-to-your-package-json
---

# --description--

Il campo `license` è dove informi gli utenti su cosa possono fare con il tuo progetto.

Alcune licenze comuni per progetti open source includono MIT e BSD. Le informazioni sulla licenza non sono obbligatorie, e le leggi sul diritto d'autore nella maggior parte dei paesi ti garantiscono la proprietà di ciò che crei per impostazione predefinita. Tuttavia, è sempre buona pratica dichiarare esplicitamente cosa gli utenti possono e non possono fare. Ecco un esempio del campo license:

```json
"license": "MIT",
```

# --instructions--

Compila il campo `license` nel file package.json del tuo progetto come ritieni più opportuno.

# --hints--

package.json dovrebbe avere una chiave "license" valida

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.license, '"license" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
