---
id: 686e64be3cd9db25282a956d
title: Projetar uma Exposição de Fotografia
challengeType: 25
dashedName: lab-photography-exhibit
demoType: onClick
---

# --description--

Neste laboratório, você criará um layout de exposição fotográfica usando as classes utilitárias do Tailwind CSS. 

Para as imagens, você pode usar os seguintes URLs se desejar:

- https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg
- https://cdn.freecodecamp.org/curriculum/labs/alps.jpg
- https://cdn.freecodecamp.org/curriculum/labs/sea.jpg

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um elemento com uma classe chamada `main-container`.
2. Seu elemento `.main-container` também deve conter as seguintes classes utilitárias do Tailwind CSS:
   - a classe utilitária correta para criar um contêiner CSS Grid.
   - uma classe utilitária para definir o número de colunas dentro da grade. Você deve usar a classe utilitária que utiliza o número fixo de colunas. Ex. `grid-cols-<number>`.
   - uma classe utilitária para definir o espaçamento entre linhas e colunas dos itens da grade.
3. Dentro do seu elemento `.main-container`, deve haver pelo menos três elementos, cada um com uma classe chamada `card`.
4. Cada um dos seus elementos `.card` também deve conter as seguintes classes utilitárias do Tailwind CSS:
   - uma classe utilitária para definir uma borda arredondada de sua escolha.
   - uma classe utilitária para definir o padding de sua escolha. Você deve usar este formato `p-<number>`.
5. Dentro de cada um dos seus elementos `.card`, você deve ter os seguintes elementos:
   - um elemento de imagem com os atributos `src` e `alt`.
   - um elemento com a classe chamada `subheading`.
   - um elemento com a classe chamada `description`.
6. Cada um dos seus elementos de imagem deve ter uma classe utilitária para definir uma borda arredondada de sua escolha.
7. Cada um dos seus elementos `.subheading` deve usar classes utilitárias para definir o peso e o tamanho da fonte de sua escolha.
8. Cada um dos seus elementos `.description` deve usar uma classe utilitária para definir o tamanho da fonte de sua escolha.

# --hints--

Você deve ter um elemento com uma classe chamada `main-container`.

```js
assert.exists(document.querySelector(".main-container"));
```

Seu elemento `.main-container` deve usar a classe utilitária correta do Tailwind CSS para criar um contêiner CSS Grid.

```js
const mainContainer = document.querySelector(".main-container");
assert.isTrue(mainContainer.classList.contains("grid"));
```

Seu elemento `.main-container` deve usar uma classe utilitária para definir o número de colunas dentro da grade.

```js
const mainContainer = document.querySelector(".main-container");
// Grid classes can vary, so this check is for the presence of any grid-cols class
const hasGridCols = [...mainContainer.classList].some(className => className.startsWith("grid-cols-"));
assert.isTrue(hasGridCols);
```

Seu elemento `.main-container` deve usar uma classe utilitária para definir o espaçamento entre linhas e colunas dos itens da grade.

```js
const mainContainer = document.querySelector(".main-container");
// Check for any gap class including gap-x and gap-y

const hasGap = [...mainContainer.classList].some(c => c.startsWith("gap-")) ||
  (classList.some(c => c.startsWith("gap-x-")) &&
   classList.some(c => c.startsWith("gap-y-")));

assert.isTrue(hasGap);
```

Seu elemento `.main-container` deve conter pelo menos três elementos, cada um com uma classe chamada `card`.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isAtLeast(cards.length, 3);
```

Cada um dos seus elementos `.card` deve ter uma classe utilitária para definir uma borda arredondada.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);

const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

cards.forEach(card => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => card.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
})
```

Cada um dos seus elementos `.card` deve ter uma classe utilitária para definir o padding.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);
const hasPadding = [...cards].every(card => {
  return [...card.classList].some(c => c.startsWith("p-"));
});
assert.isTrue(hasPadding);
```

Cada um dos seus elementos `.card` deve conter um elemento de imagem.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
```

Cada um dos seus elementos `img` deve ter um atributo `src`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("src")));
images.forEach(img => assert.isTrue(img.getAttribute("src").length > 0));
```

Cada um dos seus elementos `img` deve ter um atributo `alt`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("alt")));
images.forEach(img => assert.isTrue(img.getAttribute("alt").length > 0));
```

Cada um dos seus elementos `.card` deve conter um elemento com a classe chamada `subheading`.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isAtLeast(subheadings.length, 3);
```

Cada um dos seus elementos `.subheading` deve conter algum conteúdo de texto.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(subheadings);

subheadings.forEach(subheading => {
  assert.isNotEmpty(subheading.textContent.trim());
});
```

Cada um dos seus elementos `.card` deve conter um elemento com a classe chamada `description`.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
assert.isAtLeast(descriptions.length, 3);
```

Cada um dos seus elementos `.description` deve conter algum texto.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(descriptions);

descriptions.forEach(description => {
  assert.isNotEmpty(description.textContent.trim());
});
```

Cada um dos seus elementos `img` deve ter uma classe utilitária para definir uma borda arredondada.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isNotEmpty(images);
const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

images.forEach(img => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => img.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
});
```

Cada um dos seus elementos `.subheading` deve usar classes utilitárias para definir o peso da fonte.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedFontWeights = [
  "font-thin",
  "font-extralight",
  "font-light",
  "font-normal",
  "font-medium",
  "font-semibold",
  "font-bold",
  "font-extrabold",
  "font-black"
];

subheadings.forEach(subheading => {
  const hasFontWeight = predefinedFontWeights.some(weight => subheading.classList.contains(weight));
  assert.isTrue(hasFontWeight);
});
```

Cada um dos seus elementos `.subheading` deve usar classes utilitárias para definir o tamanho da fonte.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

subheadings.forEach(subheading => {
  const hasFontSize = predefinedSizes.some(size => subheading.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

Cada um dos seus elementos `.description` deve usar uma classe utilitária para definir o tamanho da fonte.

```js
const descriptions = document.querySelectorAll(".description");
assert.isNotEmpty(descriptions);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

descriptions.forEach(description => {
  const hasFontSize = predefinedSizes.some(size => description.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
 
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-purple-800 text-yellow-400 min-h-screen">
    <h1 class="text-center text-4xl font-bold my-12">Photography Exhibit</h1>

    <main class="main-container grid grid-cols-1 md:grid-cols-3 gap-8 px-4">
      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="Ancient Colosseum"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Colosseum, Rome</h2>
          <p class="description text-sm mb-4">
            There is so much history and beauty. The light hits it just right.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="Alpine Mountains"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">The Alps</h2>
          <p class="description text-sm mb-4">
            The mountains go on forever and the view is breathtaking.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="Calm Sea"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Seascape</h2>
          <p class="description text-sm mb-4">
            The water was still, the air was cool and the feeling is pure peace.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>
    </main>
  </body>
</html>
```
