---
id: 689f7996426c1534fa5bea66
title: Construir um Tradutor de Pontuação de Golfe
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

No jogo de Golf, cada buraco tem um `par`, que significa o número médio de `strokes` que um golfista espera dar para encaçapar a bola no buraco e completar a jogada. Dependendo da distância acima ou abaixo de `par` que seu número de `strokes` estiver, há diferentes apelidos.

Neste laboratório, você escreverá uma função que converte o `par` e as `strokes` para seus apelidos. 

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `golfScore`.
1. `golfScore` deve receber dois argumentos numéricos, que são o par do campo e a quantidade de tacadas realizadas.
1. `golfScore` deve retornar uma string.
1. `golfScore` deve retornar `"Hole-in-one!"` se `strokes` for `1`.
1. `golfScore` deve retornar `"Águia"` se `strokes` for menor ou igual a `par` menos `2`.
1. `golfScore` deve retornar `"Birdie"` se `strokes` for igual a `par` menos `1`.
1. `golfScore` deve retornar `"Par"` se `strokes` for igual a `par`.
1. `golfScore` deve retornar `"Bogey"` se `strokes` for igual a `par` mais `1`.
1. `golfScore` deve retornar `"Double Bogey"` se `strokes` for igual a `par` mais `2`.
1. `golfScore` deve retornar `"Vai para casa!"` se `strokes` for maior ou igual a `par` mais `3`.

# --hints--

Você deve criar uma função chamada `golfScore`.

```js
assert.isFunction(golfScore);
```

`golfScore` deve receber dois parâmetros.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` deve retornar uma string

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` deve retornar a string `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` deve retornar a string `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` deve retornar a string `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` deve retornar a string `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` deve retornar a string `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` deve retornar a string `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` deve retornar a string `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` deve retornar a string `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` deve retornar a string `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` deve retornar a string `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` deve retornar a string `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` deve retornar a string `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` deve retornar a string `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` deve retornar a string `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` deve retornar a string `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` deve retornar a string `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` deve retornar a string `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` deve retornar a string `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` deve retornar a string `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` deve retornar a string `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` deve retornar a string `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
