---
id: 66325a250690a3612c1db0f6
title: Hatua ya 5
challengeType: 0
dashedName: step-5
---

# --description--

CamperBot ameunda kigezo kipya kinachoitwa `bgHexCodeSpanElement` kuhifadhi rejea kwa kipengee cha `span` chenye `id` ya `bg-hex-code`. Hata hivyo, wanapojaribu kuandika kigezo hicho kwenye konsoli, wanapata `null`. 

`null` ni thamani maalum katika JavaScript inayowakilisha kutokuwepo kwa thamani. Hii inaweza kutokea unapojaribu kupata sifa ya kitu ambacho hakipo. 

Katika kesi hii, CamperBot haipitishi kichaguzi sahihi kwa njia ya `document.querySelector`.

Rekebisha mstari wa `document.querySelector("bg-hex-code")` ili uchague kipengee chenye `id` ya `bg-hex-code` kwa usahihi.

# --hints--

Unapaswa kurekebisha `"bg-hex-code"` ili i-rejea kwa usahihi `id` ya `bg-hex-code`. Kumbuka kuwa wachaguzi wa `id` katika CSS huanza na alama ya hash (`#`).

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

--fcc-editable-region--
const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("bg-hex-code");

console.log(bgHexCodeSpanElement);
--fcc-editable-region--
```
