---
id: 66a84111965a0c46df6bbd0a
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

In previous lessons, you learned how to associate a `label` element with an `input` like this:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

`label` 元素的 `for` 属性应与对应的 `input` 元素的 `id` 属性匹配。 这称为显式关联。

在你的 `legend` 元素下方，添加一个 `label`元素，其文本为 `Name (required):`。 将其 `for` 属性设置为 `"full-name"`。

然后在你的 `label` 元素下方，添加一个无属性的 `input` 元素。 在接下来的步骤中，你将添加必要的属性。

# --hints--

你应该有一个 `label` 开始标签。

```js
assert.match(code, /<label.*>/);
```

你应该有一个 `label` 结束标签。

```js
assert.match(code, /<\/label>/);
```

你的 `label` 元素应放在 `fieldset` 元素中 `legend` 的下方。

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

你的 `label` 元素应该有一个 `for` 属性，值为 `"full-name"`。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

你的 `label` 元素的文本应为 `Name (required):`。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

你应该在 `label` 元素下方有一个 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
