---
id: 66a8380d911e3f4270d5cadc
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

现在，需要添加 `main` 元素，它代表页面的主要内容。

# --hints--

你应该有一个 `main` 开始标签。

```js
assert.match(code, /<main>/i);
```

你应该有一个 `main` 结束标签。

```js
assert.match(code, /<\/main>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
