---
id: 66a83601cd819e37f0dccd14
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

对于介绍性文字，你需要显示主标题，然后是关于留下反馈意见的简短说明。

在 `body` 元素内，添加一个 `header` 元素。

在 `header` 元素中，添加一个 `h1` 元素。 `h1` 元素的文本应为 `Hotel Feedback Form`。

在你的 `h1` 元素下方，添加一个 `p` 元素。 `p` 元素的文本应为 `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

你应该有一个 `header` 开始标签。

```js
assert.match(code, /<header>/i);
```

你应该有一个 `header` 结束标签。

```js
assert.match(code, /<\/header>/i);
```

你应该有一个 `h1` 开始标签。

```js
assert.match(code, /<h1>/i);
```

你应该有一个 `h1` 结束标签。

```js
assert.match(code, /<\/h1>/i);
```

你的 `h1` 元素应包含文本 `Hotel Feedback Form`。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

你的 `h1` 元素应该在 `header` 元素中。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

你应该有一个 `p` 开始标签。

```js
assert.match(code, /<p>/i);
```

你应该有一个 `p` 结束标签。

```js
assert.match(code, /<\/p>/i);
```

你的段落元素应包含以下文本 `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

你的段落元素应位于 `header` 内。

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

你的段落元素应位于 `h1` 之后。

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
