---
id: 66a83e5e491625454b6f62c3
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

表单由 `inputs`组成，用户可在其中输入数据。 你可以使用 `fieldset` 元素将相关输入分组。

下面是使用 `fieldset` 元素的示例：

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

在你的 `form` 元素中，添加一个 `fieldset` 元素。

# --hints--

你应该有一个 `fieldset` 开始标签。

```js
assert.match(code, /<fieldset>/i);
```

你应该有一个 `fieldset` 结束标签。

```js
assert.match(code, /<\/fieldset>/i);
```

你的 `fieldset` 元素应该在 `form` 元素中。

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">

      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
