---
id: 66a83fec026a7a4631e084d2
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

在使用 `fieldset` 元素时，通常使用一个标题来描述输入组。 为此，你可以使用 `legend` 元素。

下面是使用 `legend` 元素的示例：

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

在你的 `fieldset` 元素中，添加一个 `legend` 元素，文本为 `Personal Information`。

# --hints--

你应该有一个 `legend` 开始标签。

```js
assert.match(code, /<legend>/);
```

你应该有一个 `legend` 结束标签。

```js
assert.match(code, /<\/legend>/);
```

你的 `legend` 元素应该在 `fieldset` 元素中。

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

你的 `legend` 元素应包含文本 `Personal Information`。

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
