---
id: 66a8290a27c2c625e2355042
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本讲座中，你将通过制作酒店反馈表来练习使用 HTML 表单。

首先添加 `<!DOCTYPE html>`，然后添加具有值为 `en` 的 `lang` 属性的 `html` 元素。

在 `html` 元素中，添加一个 `head` 元素。


# --hints--

你应该有 `<!DOCTYPE html>`。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你应该有一个 `html` 开始标签，语言设置为英语。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

你应该有一个 `html` 结束标签。

```js
assert.match(code, /<\/html>/i);
```

你的 `DOCTYPE` 应放在 `html` 元素之前。

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

你应该有一个 `head` 开始标签。

```js
assert.match(code, /<head>/i);
```

你应该有一个 `head` 结束标签。

```js
assert.match(code, /<\/head>/i);
```

你的 `head` 开始标签应在 `head` 结束标签之前。

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

你的 `head` 元素应该在 `html` 元素中。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
