---
id: 66a8347f97a33e36ffc81b9f
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

为了准备创建一些实际的内容，在 `head` 元素下面添加一个 `body` 元素。

# --hints--

你应该有一个 `body` 开始标签。

```js
assert.match(code, /<body>/i);
```

你应该有一个 `body` 结束标签。

```js
assert.match(code, /<\/body>/i);
```

你不应该改变你的 `head` 元素。 确认你没有删除结束标签。

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

你的 `body` 元素应该在你的 `head` 元素之后。

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>

  --fcc-editable-region--
</html>
```
