---
id: 66a954b2bcddba72076c1857
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

在你的 `label` 元素下方，添加一个 `radio` 按钮，`id` 设置为 `"no-option"`，名字属性设置为 `"hotel-stay"`。

在你的新 `radio` 按钮下方，添加另一个 `label` 元素，其 `for` 属性设置为 `"no-option"`。 `label` 的文本应为 `No`。

完成后，你就可以尝试使用单选按钮，每次选择一个选项。

# --hints--

你应该在第一个 `label` 元素下方有一个 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

你的 `input` 元素应该有一个 `type` 属性，值为 `radio`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

你的 `input` 元素应该有一个 `id` 属性，值为 `"no-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

你的 `input` 元素应该有一个 `name` 属性，值为 `"hotel-stay"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

你的第二个 `input` 元素下方应该有第二个 `label` 元素。

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

你的 `label` 元素应该有一个 `for` 属性，设置为 `"no-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

你的 `label` 元素应该包含文本 `No`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          --fcc-editable-region--
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
