---
id: 66a937e74920ba68ebe5e86d
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

在之前的课程中，你学习了如何使用 `placeholder` 和 `required` 属性，如下所示：

```html
<input type="text" placeholder="e.g., John Doe" required>
```

为你现有的 `input` 元素添加一个值为 `"e.g., John Doe"` 的 `placeholder` 属性。

也在 `input` 元素中添加 `required` 属性。

# --hints--

你不应删除现有的 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

你的 `input` 应该有一个 `placeholder` 属性，值为 `"e.g., John Doe"`。

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

你的 `input` 应该具有 `required` 属性。

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
