---
id: 66a93bbe65a26169dbf3bc39
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

酒店反馈表还应收集用户的电子邮件地址。

首先，在表单中添加一个新的 `label` 元素，文本为 `Email address (required):`。 你的 `label` 元素应该有一个 `for` 属性，值设置为 `"email"`。

# --hints--

你应该有一个 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

你的 `label` 元素应包含文本 `Email address (required):`。

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

你的 `label` 元素应该有一个 `for` 属性，值设置为 `"email"`。

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--

--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
