---
id: 66a93c95bc58e26a8fe95818
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

接下来，在你的邮箱 `label` 下方，添加一个 `input`，类型为 `"email"`。 该 `input` 元素应该有一个 `id` 属性，值设置为 `"email"`。 将 `name` 属性也设置为 `"email"`。

该 `input`也是必填的，因此，请确保添加 `required` 属性。

最后，添加一个 `placeholder` 属性，设置为 `"example@email.com"`。

# --hints--

你的邮箱 `label` 元素下面应该有一个 `input` 元素。

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

你的 `input` 元素应该有一个 `type`，值为 `"email"`。

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

你的 `input` 元素应有一个 `id` 属性，设置为 `"email"`。

```js
assert.isNotNull(document.querySelector('input#email'));
```

你的 `input` 元素应有一个 `name` 属性，设置为 `"email"`。

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

你的 `input` 元素应该具有 `required` 属性。

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

您的 `输入的` 元素应该有一个 `占位符` 属性设置为 `"example@email.com"`

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>
        --fcc-editable-region--
          <label for="email">Email address (required):</label>

        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
