---
id: 66a972137acd1179fa3fe8a0
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

在你的 `select` 元素中，添加以下五个 `option` 元素，并为 `option` 文本和 `value` 属性添加相应的值：

**值属性：**

- poor
- satisfactory
- good
- very-good
- excellent

**选项元素文本：**

- Poor
- Satisfactory
- Good
- Very Good
- Excellent

# --hints--

你应该有一个值设置为 `"poor"` 的 `option` 元素。

```js
assert.exists(document.querySelector('option[value="poor"]'));
```

`value` 为 `"poor"` 的 `option` 元素应包含文本 `Poor`。

```js
assert.strictEqual(document.querySelector('option[value="poor"]')?.textContent.trim(), 'Poor');
```

你应该有一个 `value` 设置为 `"satisfactory"` 的 `option` 元素。

```js
assert.exists(document.querySelector('option[value="satisfactory"]'));
```

`value` 为 `"satisfactory"` 的 `option` 元素应包含文本 `Satisfactory`。

```js
assert.strictEqual(document.querySelector('option[value="satisfactory"]')?.textContent.trim(), 'Satisfactory');
```

你应该有一个 `value` 设置为 `"good"` 的 `option` 元素。

```js
assert.exists(document.querySelector('option[value="good"]'));
```

`value` 为 `"good"` 的 `option` 元素应包含文本 `Good`。

```js

assert.strictEqual(document.querySelector('option[value="good"]')?.textContent.trim(), 'Good');
```

你应该有一个值设置为 `"very-good"` 的 `option` 元素。

```js
assert.exists(document.querySelector('option[value="very-good"]'));
```

`value` 为 `"very-good"` 的 `option` 元素应包含文本 `Very Good`。

```js
assert.strictEqual(document.querySelector('option[value="very-good"]')?.textContent.trim(), 'Very Good');
```

你应该有一个值设置为 `"excellent"` 的 `option` 元素。

```js
assert.exists(document.querySelector('option[value="excellent"]'));
```

`value` 为 `"excellent"` 的 `option` 元素应包含文本 `Excellent`。

```js
assert.strictEqual(document.querySelector('option[value="excellent"]')?.textContent.trim(), 'Excellent');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>

          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          --fcc-editable-region--
          <select name="service" id="service">

          </select>
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
