---
id: 66a94ea5df66236ebed260e8
title: 步骤18
challengeType: 0
dashedName: step-18
---

# --description--

表单的下一部分将负责询问用户是否曾入住过该酒店。

首先添加一个 `fieldset` 元素。

在 `fieldset` 元素内部，添加一个 `legend` 元素，其文本为 `Was this your first time at our hotel?`。

# --hints--

你应该有一个 `fieldset` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 2);
```

你应该在 `fieldset` 元素中包含一个 `legend` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset legend'), 2);
```

你的 `legend` 元素的文本应为 `Was this your first time at our hotel?`。

```js
assert.strictEqual(document.querySelectorAll('fieldset legend')[1]?.textContent.trim(), 'Was this your first time at our hotel?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
      --fcc-editable-region--

      --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
