---
id: 66a969951120be7818d8ee49
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

表单的下一部分将为用户提供为酒店评分的功能。

首先添加一个新的 `fieldset` 元素，并在其中嵌套一个 `legend` 元素。 `legend` 的文本应为 `Ratings`。

在 `legend` 元素下方，添加一个 `label` 元素，其文本为 `How was the service?`。 `for` 属性应设置为 `"service"`。

# --hints--

你应该有一个 `fieldset` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 4);
```

你的 `fieldset` 应该有一个 `legend` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend'), 1);
```

你的 legend 应包含文本 `Ratings` 。

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend')[0]?.textContent.trim(), 'Ratings');
```

你的 `fieldset` 应该在 `legend` 下方有一个 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend + label'), 1);
```

你的 `label` 应包含文本 `How was the service?`。

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.textContent.trim(), 'How was the service?');
```

你的 `label` 应将 `for` 属性设置为 `"service"`。

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.getAttribute('for'), 'service');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>

          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
