---
id: 66a93730719e1f68410cce54
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

`name` 属性用于在表单数据提交到服务器后对其进行标识。

下面是一个如何使用 `name` 属性的示例：

```html
<input type="email" name="email">
```

为 `input` 元素添加一个 `name` 属性，值为 `"name"`。

# --hints--

你不应删除现有的 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

你不应创建新的 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

你的 `input` 元素应该有一个 `name` 属性，值为 `"name"`。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
