---
id: 66a930b20f589b6664c51cb0
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

当用户提供全名时，`input` 将接受纯文本。

在之前的课程中，你学习了如何这样使用 `type` 属性：

```html
<input type="text">
```

对于你现有的 `input` 元素，添加一个 `type` 属性，设置为 `"text"`。 

也为 `input` 元素赋予一个 `id` 属性，其值为 `"full-name"`。

# --hints--

你不应删除现有的 `input` 元素。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

你不应创建新的 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

你的 `input` 元素应该有一个 `type` 属性，值为 `"text"`。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

你的 `input` 元素应该有一个 `id` 属性，值为 `"full-name"`。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
