---
id: 66a975c259525b7bc2d5c776
title: Step 28
challengeType: 0
dashedName: step-28
---

# --description--

你的酒店反馈表还应允许用户对食物进行评分。

首先添加一个 `label` 元素，其文本为 `How was the food?`。 该 `label` 元素应将 `for` 属性设置为 `"food"`。

在你的 `label` 元素下方，添加一个 `select` 元素，`id` 和 `name` 设置为 `"food"`。

# --hints--

你应将 `label` 元素的 `for` 属性设置为 `"food"`。

```js
assert.exists(document.querySelector('label[for="food"]'));
```

你的 `label` 应包含文本 `How was the food?`。

```js
assert.strictEqual(document.querySelector('label[for="food"]')?.textContent.trim(), 'How was the food?');
```

你应将 `label` 下方的 `select` 元素的 `id` 设置为 `"food"`。

```js
assert.exists(document.querySelector('label + select#food'));
```

你的 `select` 元素应将 `name` 属性设置为 `"food"`。

```js
assert.exists(document.querySelector('label + select[name="food"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>

          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option selected value="excellent">Excellent</option>
          </select>

          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
