---
id: 66a9419e2d18476c645ce693
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

数字 `input` 用于创建数字输入框。

下面是一个数字输入框的示例：

```html
<input type="number" id="age" name="age" min="18" max="100">
```

`min` 和 `max` 属性用于设置输入字段中可输入的最小值和最大值。

在你的 `label` 元素下方，添加一个输入，`type` 属性设置为 `"number"`，`id` 为 `"age"`。

`name` 属性应设置为 `"age"`，`min` 属性应设置为 `"3"`，`max` 属性应设置为 `"100"`。

# --hints--

你的 `label` 元素下方应该有一个数字 `input`。

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

你的数字 `input` 应该有一个 `id`，值为 `age`。

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

你的数字 `input` 应有一个 `name` 属性，设置为 `age`。

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

你的数字 `input` 应有一个 `min` 属性，设置为 `"3"`。

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

你的数字 `input` 应有一个 `max` 属性，设置为 `"100"`。

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
      --fcc-editable-region--
          <label for="age">Age (optional):</label>

      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
