---
id: 66a9401c9d660d6bb15993e2
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

你的酒店反馈表应为用户添加年龄选项。

首先，在表单中添加一个 `label` 元素，文本为 `Age (optional):`。

`for` 属性应设置为 `"age"`。

# --hints--

你应该有一个 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

你的 `label` 元素应将 `for` 属性设置为 `"age"`。

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

你的 `label` 元素应包含文本 `Age (optional):`。

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--

--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
