---
id: 660ef19b95d3308e7dd31bb6
title: 步驟5
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>console</dfn>（控制檯指令）允許你打印並查看 JavaScript 的輸出。 你可以使用 `console.log()` 來向控制檯發送信息。 舉個例子，下面這串代碼將會在控制檯打印 `"Naomi"`。

```js
let developer = "Naomi";
console.log(developer);
```

上面的代碼通過在 `console.log()` 中的名稱讀取到了變量 `developer` 。 請注意，括號內的值是將要被打印的內容。

在控制檯打印變量 `character` 的值。 然後點擊 “Console”（控制檯）按鈕來查看 JavaScript 控制檯。

# --hints--

`character` 變量初始化不應被更改。

```js
assert.strictEqual(character, 'Hello');
```

You should access the `console` in your code.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

You should access the `log` property of the `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

You should use parentheses to call the `.log()` method.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

You should print the `character` variable to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
