---
id: 660ef31a5be625914a0102cd
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

當使用 `let` 關鍵字宣告變數時，你可以在之後 <dfn>重新指定、設值</dfn>（或更改該變數的值）。在此範例中，`programmer` 的值從 `"Naomi"` 更改為 `"CamperChan"`。

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

請注意，當重新指定變數時，你**不**會再次使用 `let` 關鍵字。

在你的 `console.log` 之後，將值 `"World"` 指定給你的 `character` 變數。

# --hints--

你應該在你的程式碼中使用兩次 `character`。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

你應該使用指定運算子來重新指定 `character`。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

你應該在你的程式碼中使用字串 `"World"`。

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

你的 `character` 變數應該有值 `"World"`。

```js
assert.equal(character, "World");
```

你的重新指派不應該使用 `let`。

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
