---
id: 660ef6355e8f5a9e67fe5f46
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

使用 `number` 資料型別，你可以執行數學操作，例如加法。試著將 `count + 1` 列印到主控台。

# --hints--

你應該在你的程式碼中存取 `console`。

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

你應該存取 `console` 的 `.log()`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

你應該將 `count + 1` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
