---
id: 660ef0f7c4b8e68ccd1f0786
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript 有七種原始資料型別，其中 `String` 是其中之一。在 JavaScript 中，<dfn>字串</dfn> 表現為一連串字元，可以用單引號（`'`）或雙引號（`"`）括起來。

請注意，字串是<dfn>不可變的</dfn>，這表示一旦它們被創建，就無法更改。變數仍然可以被重新指定另一個值。

將你的 `"Hello"` 字串改為使用單引號。

# --hints--

你的程式碼中不應該有雙引號。

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

你應該為你的 `"Hello"` 字串使用單引號。

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

你仍然應該在你的程式碼中使用 `let`。

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

你仍然應該宣告一個 `character` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

你的 `character` 變數的值仍應該是字串 `"Hello"`。

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
