---
id: 660ef857f2806aa626d29d17
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

你也可以執行減法（`-`）、乘法（`*`）和除法（`/`）。隨意在你的 `console.log` 中試驗這些運算子和數字。當你準備好繼續時，請移除 `console.log`。

# --hints--

你的程式碼中不應該有 `console.log()`。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

你的 `character` 變數應該仍然具有值 `"Hello"`。

```js
assert.equal(character, "Hello");
```

你的 `count` 變數應該仍然有值 `8`。

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
