---
id: 660ef5105b8ba095307a0e50
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

現在再次將你的 `character` 變數記錄到主控台。你應該會在主控台看到字串 `"Hello"`，接著是字串 `"World"`。

# --hints--

你應該第二次使用 `console.log`。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

你應該第二次將 `character` 記錄到主控台。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


你的新 `console.log()` 應該放在你重新指派之後。

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
