---
id: 660ef5c1904955978a986a5c
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

使用 `let` 宣告一個 `count` 變數。將 <dfn>數字</dfn> `8` 指定給它。當使用數字值時，不要使用引號。例如：

```js
let money = 100;
```

# --hints--

你的程式碼中應該有第二個 `let` 關鍵字。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

你應該使用 `let` 來宣告一個 `count` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

你應該將數字 `8` 指定到你的 `count` 變數。

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
