---
id: 660ee6e3a242da6bd579de69
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

程式設計中最重要的概念之一是變數。<dfn>變數</dfn> 指向儲存值的特定記憶體位址。變數會被賦予一個名稱，該名稱可在你的程式碼中用來存取該值。

宣告一個變數就是給它一個名稱。在 JavaScript 中，這通常是使用 `let` 關鍵字來完成。 例如，這是你如何宣告一個 `hello` 變數：

```js
let hello;
```

變數命名遵循特定規則：名稱可以包含字母、數字、美元符號和底線，但不能包含空格且不得以數字開頭。

使用 `let` 關鍵字宣告一個名為 `character` 的變數。

*注意*：在 JavaScript 中以分號 `;` 結束述語是常見的做法。

# --hints--

你應該在你的程式碼中使用 `let`。

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

你應該在你的程式碼中使用 `character`。

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

你應該使用 `let` 來宣告一個 `character` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

你的宣告應該以分號結尾。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
