---
id: 660ef55dd468079679ee0092
title: 步驟 12
challengeType: 1
dashedName: step-12
---

# --description--

你現在已經準備好宣告你的下一個變數。移除兩個 `console.log` 述語，以及 `character` 的重新指定。

同時移除你的 `secondCharacter` 變數，但保持 `character` 初始化不變。

# --hints--

你的程式碼中不應該有任何 `console.log()` 述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

你不應該重新指定 `character` 的值。

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

你的 `character` 變數應該有值 `"Hello"`。

```js
assert.equal(character, "Hello");
```

你不應該有 `secondCharacter` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
