---
id: 660ef19b95d3308e7dd31bb6
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>主控台</dfn>允許你列印並檢視 JavaScript 輸出。你可以使用 `console.log()` 將資訊傳送到主控台。 例如，這段程式碼會在主控台列印 `"Naomi"`：

```js
let developer = "Naomi";
console.log(developer);
```

上述程式碼在 `console.log()` 中以名稱存取 `developer` 變數。請注意，小括弧中的值是將被列印的值。

將 `character` 變數的值列印到主控台。然後，點擊「Console」按鈕以檢視 JavaScript 主控台。

# --hints--

`character` 變數的初始化不應該被更改。

```js
assert.strictEqual(character, 'Hello');
```

你應該在你的程式碼中存取 `console`。

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

你應該存取 `console` 的 `log` 屬性。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

你應該使用小括弧來呼叫 `.log()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

你應該將 `character` 變數列印到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
