---
id: 660eebd83100d37862268781
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

你的 `character` 變數目前沒有值。你可以使用 <dfn>指派</dfn> 運算子 `=` 來指定值。例如：

```js
let hello = "Hello";
```

在宣告變數的同時指定值稱為<dfn>初始化</dfn>。

在宣告 `character` 變數時，透過指定其值為 `"Hello"` 來初始化它。

# --hints--

你應該使用指定運算子 `=`。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

你應該使用字串 `"Hello"`。

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

你應該為你的 `"Hello"` 字串使用雙引號。

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

你應該將 `"Hello"` 指定到你的 `character` 變數。

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
