---
id: 660f2b6fd54ac1fc142804dd
title: 步驟 47
challengeType: 1
dashedName: step-47
---

# --description--

格式化這個金字塔的邏輯很可能會變得複雜，這表示現在是將該程式碼抽取到函式中的好時機。

<dfn>函式（程式）</dfn>是一個可以在你的應用程式中重複使用的程式碼區塊。函式使用以下語法宣告：

```js
function name(parameter) {

}
```

`function` 關鍵字告訴 JavaScript `name` 變數將會是一個函式（程式）。`parameter` 是一個變數，用來表述傳入函式的值。函式可以有任意多或任意少的 <dfn>參數</dfn>。像 `for` 迴圈一樣，大括號之間的空間是 <dfn>函式主體</dfn>。

宣告一個 `padRow` 函式（程式）。暫時不要建立任何參數變數。函式（程式）本體應該是空的。請記得你需要使用駝峰式命名法。

# --hints--

你應該使用 `function` 關鍵字。

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

你應該宣告一個 `padRow` 函式（程式）。

```js
assert.isFunction(padRow);
```

你的 `padRow()` 函式（程式）不應該有任何參數。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

你的 `padRow()` 函式（程式）應該有一個空的主體。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
