---
id: 660f5179b3b0ca558f6b4d4f
title: 步驟 113
challengeType: 1
dashedName: step-113
---

# --description--

使用 `if` 述語來檢查 `inverted` 是否為 true。請記住這裡不需要使用等號運算子。

# --hints--

你應該使用 `if` 述語。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/);
```

你的 `if` 述語應該檢查 `inverted` 是否為 `true`。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted/);
```

你的 `if` 條件不應該使用任何比較運算子。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {

  rows.unshift(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
