---
id: 660f061d259bbebc37461080
title: 第 20 步
challengeType: 1
dashedName: step-20
---

# --description--

你可以使用值的 <dfn>索引</dfn> 來存取陣列內的值。索引是一個數字，表述該值在陣列中的位置，第一個值的索引從 `0` 開始。

你可以使用 <dfn>括號表示法</dfn> 存取值，例如 `array[0]`。

使用 `console.log` 和括號表示法列印你的 `rows` 陣列中的第一個值。

# --hints--

你應該在你的程式碼中有一個 `console.log()` 述語。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

你應該存取你的 `rows` 陣列。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

你應該對你的 `rows` 陣列使用括號表示法。

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

你應該使用括號表示法來存取你的 `rows` 陣列的第一個元素。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

你應該記錄你的 `rows` 陣列的第一個元素。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
