---
id: 660f4de78f775e480ba2e451
title: 步驟 107
challengeType: 1
dashedName: step-107
---

# --description--

陣列的 `.unshift()` 方法允許你將一個值添加到陣列的**開頭**，不同於 `.push()`，它會將值添加到陣列的末端。`.unshift()` 會傳回被呼叫陣列的新長度。

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

使用 `const` 宣告一個 `unshifted` 變數，並將呼叫你的 `numbers` 陣列的 `.unshift()` 結果設值給它。將 `5` 作為引數傳入。然後列印你的 `unshifted` 變數。

# --hints--

你應該宣告你的 `unshifted` 變數。

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

你應該在你的 `numbers` 陣列上呼叫 `.unshift()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

你應該將 `5` 作為引數傳遞給你的 `.unshift()` 呼叫。

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

你應該將 `.unshift()` 呼叫的結果指定給你的 `unshifted` 變數。

```js
assert.equal(unshifted, 4);
```

你應該記錄你的 `unshifted` 變數。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
