---
id: 660f1b6e60bd9edf902c81fd
title: 第 36 步
challengeType: 1
dashedName: step-36
---

# --description--

你的 <dfn>iteration</dfn> 述語會告訴你的迴圈在每次執行後如何處理迭代器。

當你重新指定一個變數時，你可以使用該變數來引用重新指定前的先前值。這讓你能夠執行像是將三加到現有數字的操作。<dfn>例如</dfn>，`bees = bees + 3` 會將 `bees` 的值增加三。

使用該語法將你的 `"iteration"` 字串取代為一個將 `i` 增加一的重新指定述語。

# --hints--

你應該為你的 `i` 變數添加 1。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

你應該將 `i + 1` 指定回你的 `i` 變數。

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

你的 `for` 迴圈應該在每次迭代時將 `i` 增加 `1`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
