---
id: 660f0c34aad72dc712b97624
title: 步驟 27
challengeType: 1
dashedName: step-27
---

# --description--

另一個對這個專案至關重要的方法是 `.pop()` 方法。它會從陣列中移除最後一個元素並 <dfn>傳回</dfn> 該元素。

當一個方法傳回一個值時，你可以將它視為將該值還給你，使其可用於你程式碼的其他零件。

創建一個名為 `popped` 的新變數，並將 `rows.pop()` 的結果指定給它。然後，將 `popped` 輸出到主控台。

# --hints--

你應該宣告一個名為 `popped` 的變數。

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

你應該使用 `let` 來宣告名為 `popped` 的變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

你應該呼叫 `.pop()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

你應該在你的 `rows` 陣列上呼叫 `.pop()` 方法。 

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

你應該記錄你的 `popped` 變數。 

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
