---
id: 660f20473aef47e9b8c9afc6
title: 步驟 40
challengeType: 1
dashedName: step-40
---

# --description--

為你的程式碼添加一個列印 `result` 值的日誌述語。根據你使用的主控台，可能不會看到任何列印內容。

# --hints--

你應該為你的程式碼添加一個 `console.log` 呼叫。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

你應該記錄你的 `result` 變數。 

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
