---
id: 660f07d231941bc11719f664
title: 第 21 步
challengeType: 1
dashedName: step-21
---

# --description--

陣列的特殊之處在於它們被視為 <dfn>mutable</dfn>。這表示你可以直接更改索引處的值。

例如，這段程式碼會將數字 `25` 指定給陣列中的第二個元素：

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

將你的 `rows` 陣列的**第三**個元素更新為數字 `10`。然後將 `rows` 陣列列印到你的主控台。

# --hints--

你應該再次對 `rows` 陣列使用括號表示法。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

你應該存取 `rows` 陣列的第三個元素。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

你應該對 `rows` 陣列的第三個元素使用指定運算子。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

你應該將值 `10` 指定到你的 `rows` 陣列的第三個元素。

```js
assert.equal(rows[2], 10);
```

你的程式碼中應該有第二個 `console.log` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

你應該記錄 `rows` 陣列。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
