---
id: 660f3b664421471aa595170f
title: 步驟 74
challengeType: 1
dashedName: step-74
---

# --description--

金字塔現在看起來有點奇怪。因為你是從 `1` 而不是 `0` 開始迴圈，所以在將 `i` 傳遞給 `padRow` 時不需要加一。

將你對 `padRow` 的第一次呼叫的第一個引數更新為 `i`。

# --hints--

你不應該將 `i + 1` 傳遞給 `padRow`。

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

你應該將 `i` 傳遞給 `padRow`。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
