---
id: 660f487dc0c8fa38084f9754
title: 步驟 92
challengeType: 1
dashedName: step-92
---

# --description--

既然你已將比較移到 `while` 條件中，你可以移除整個 `if` 述語。

# --hints--

你不應該再有 `if` 述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(\s*done\s*===\s*count\s*\)/);
```

你不應該再將 `continueLoop` 設為 `false`。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/continueLoop\s*=\s*false/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (done !== count) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
