---
id: 660f540c2176ea5dec01306d
title: 步驟 118
challengeType: 1
dashedName: step-118
---

# --description--

做得好！嘗試為你的 `character`、`count` 和 `inverted` 變數使用不同的值。

當你準備好進入下一個專案時，將 `character` 設為 `"!"`，`count` 設為 `10`，並將 `inverted` 設為 `false` 以繼續。

恭喜你完成你的第一個 JavaScript 專案！

# --hints--

你應該將 `character` 設定為 `"!"`。

```js
assert.equal(character, "!");
```

你應該將 `count` 設定為 `10`。

```js
assert.equal(count, 10);
```

你應該將 `inverted` 設為 `false`。

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
