---
id: 660f48e1d3682f39e81843c4
title: 步驟 94
challengeType: 1
dashedName: step-94
---

# --description--

你的金字塔產生器仍在運作中。然而，仍有可能再次出現無窮迴圈。

因為你只是在檢查 `done` 是否不相等於 `count`，如果 `done` 大於 `count`，你的迴圈將會無限執行。

更新你的迴圈條件以檢查 `done` 是否小於或相等於 `count`。

# --hints--

你的 `while` 迴圈應該檢查 `done` 是否小於或相等於 `count`。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
