---
id: 660f415b76859a2736771607
title: 步驟 78
challengeType: 1
dashedName: step-78
---

# --description--

你的金字塔又消失了。沒關係－這是預期中的情況。

在你創建新的迴圈之前，你需要了解 `if` 述語。<dfn>`if` 述語</dfn>允許你僅在條件符合時執行一個區塊的程式碼。它們使用以下語法：

```js
if (condition) {
  logic
}
```

使用布林值 `true` 作為條件，創建一個 `if` 述語。在主體中，列印字串 `"Condition is true"`。

# --hints--

你應該創建一個 `if` 述語。

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

你的 `if` 述語應該以 `true` 作為條件。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

你的 `if` 主體應該記錄 `"Condition is true"`。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
