---
id: 660f2a70ad6225fa503e71c3
title: 步驟 46
challengeType: 1
dashedName: step-46
---

# --description--

你越來越接近了！此時，你遇到的是所謂的 <dfn>off-by-one error</dfn>，這是在像 JavaScript 這類以零為基底的索引語言中常見的問題。

你的 `rows` 陣列的第一個索引是 `0`，這就是為什麼你從 `i = 0` 開始你的 `for` 迴圈。但將字串重複零次會導致沒有任何東西可列印。

為了解決這個問題，將 `1` 添加到你在 `.repeat()` 呼叫中 `i` 的值。不要像在你的迴圈條件中那樣將它重新指定給 `i`。

# --hints--

你應該在你的 `.repeat()` 方法中將 `1` 添加到 `i`。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
