---
id: 660f3ba3cceef11b6ba08b59
title: 步驟 75
challengeType: 1
dashedName: step-75
---

# --description--

不幸的是，現在金字塔的底部已經消失。這是因為你產生了另一個 <dfn>off-by-one error</dfn>。

你的原始迴圈中 `i` 的值從 `0` 到 `7`，因為 `count` 是 `8`，且你的條件要求 `i` 必須小於 `count`。你的迴圈現在是對 `i` 的值從 `1` 到 `7` 進行執行。

你的迴圈也需要更新，以便在 `i` 為 `8` 時執行。根據你的邏輯，這表示你的迴圈應該在 `i` <dfn>小於或相等於</dfn> `count` 時執行。你可以使用小於或相等於運算子 `<=` 來達成這個目的。

更新你的迴圈條件，使其在 `i` 小於或相等於 `count` 時執行。

# --hints--

你的 `for` 迴圈不應該檢查 `i` 是否小於 `count`。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

你的 `for` 迴圈應該檢查 `i` 是否小於或相等於 `count`。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
