---
id: 660f4377a359972c521d3f4b
title: 步驟 79
challengeType: 1
dashedName: step-79
---

# --description--

你會在主控台看到字串被列印，因為 `true` 事實上是 true。

將你的 `if` 述語條件改為布林值 `false`。

# --hints--

你的 `if` 條件應該以 `false` 作為條件。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (true) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
