---
id: 660f280dda5040f707c76b4a
title: 步驟 45
challengeType: 1
dashedName: step-45
---

# --description--

現在你有一串 `#` 字元，但金字塔形狀仍然缺失。幸運的是，`i` 變數表述你迴圈中目前的「列」號碼，使你能夠用它來製作類似金字塔的結構。

為了達成此目的，你將使用字串可用的 `.repeat()` 方法。此方法接受一個數字作為引數，指定要重複標的字串的次數。例如，使用 `.repeat()` 來產生字串 `"Code! Code! Code!"`：

```js
const activity = "Code! ";
activity.repeat(3);
```

對你的 `character` 使用 `.repeat()` 方法，並給它 `i` 作為數量。

# --hints--

你應該使用 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

你應該在你的 `character` 變數上使用 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

你應該將 `i` 傳遞給你的 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

你應該在 `.push()` 方法中使用 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
