---
id: 660f4cffb1459d45e34902d1
title: 步驟 106
challengeType: 1
dashedName: step-106
---

# --description--

你其實可以建置倒金字塔，而不需要像你那樣「反向」迴圈。

為了做到這點，你需要學習幾個新的陣列方法。首先使用 `const` 宣告一個 `numbers` 變數。將它設值為包含元素 `1`、`2` 和 `3` 的陣列。然後記錄 `numbers` 陣列。

# --hints--

你應該使用 `const` 來宣告一個 `numbers` 變數。

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

你的 `numbers` 變數應該是一個陣列。

```js
assert.isArray(numbers);
```

你的 `numbers` 陣列應該依序包含元素 `1`、`2` 和 `3`。

```js
assert.deepEqual(numbers, [1,2,3]);
```

你應該記錄你的 `numbers` 陣列。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
