---
id: 660f17d4e9f227d86e834abd
title: 步驟 33
challengeType: 1
dashedName: step-33
---

# --description--

要產生金字塔，你需要創建多個列。當你必須重複執行任務直到條件達成時，你會使用 <dfn>loop</dfn>。有許多方法可以撰寫迴圈。

你將從基本的 `for` 迴圈開始。`for` 迴圈使用以下語法：

```js
for (iterator; condition; iteration) {
  logic;
}
```

在接下來的步驟中，你將詳細探索迴圈的每個組件。現在，構建一個包含 `"iterator"`、`"condition"` 和 `"iteration"` 三個組件的 `for` 迴圈。保持迴圈的 <dfn>body</dfn>，即大括號 `{}` 內的部分，為空。


# --hints--

你應該有一個 `for` 迴圈。

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

你的 `for` 迴圈的第一個組件應該是字串 `"iterator"`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

你的 `for` 迴圈的第二個組件應該是字串 `"condition"`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

你的 `for` 迴圈的第三個組件應該是字串 `"iteration"`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

你的 `for` 迴圈主體應該是空的。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
