---
id: 660f455b044d3230ed971e98
title: 步驟 85
challengeType: 1
dashedName: step-85
---

# --description--

<dfn>`while`</dfn> 迴圈會不斷執行，直到指定的 `condition` 不再為真。它具有以下語法：

```js
while (condition) {
  logic;
}
```

使用該語法宣告一個以 `continueLoop` 為條件的 `while` 迴圈。主體應該是空的。

# --hints--

你應該使用 `while` 迴圈。

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

你的 `while` 迴圈應該使用 `continueLoop` 作為條件。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
