---
id: 660f374d532dc41189cc9cc2
title: 步驟 69
challengeType: 1
dashedName: step-69
---

# --description--

現在是時候進行一些數學計算。考慮一個三列金字塔。如果我們想要將它置中，它看起來會像這樣：

```md
··#··
·###·
#####
```

空白已被中點符號取代，以提升可讀性。如果你推斷該樣式，你可以看到列的開頭和結尾的空白遵循一個樣式。

將你的空白字串更新為重複 `rowCount - rowNumber` 次。

打開主控台以查看結果。

# --hints--

你應該對你的 `" "` 字串呼叫 `.repeat()`，以重複 `rowCount - rowNumber` 次。

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
