---
id: 660f1cedf3676fe26122ebf6
title: 第 38 步
challengeType: 1
dashedName: step-38
---

# --description--

你應該會在你的主控台中看到從零到七的數字列印出來，每行一個。這將作為生成你的金字塔的基礎。

將你的記錄述語取代為將 `i` 推入你的 `rows` 陣列的述語。

# --hints--

你不應該有 `console.log` 呼叫。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

你應該對你的 `rows` 陣列呼叫 `.push()`。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

你應該將 `i` 推入你的 `rows` 陣列。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

你的 `.push()` 應該發生在你的 `for` 迴圈中。 

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
