---
id: 660f44f10ea40f300b896a5e
title: 步驟 84
challengeType: 1
dashedName: step-84
---

# --description--

既然你已經練習過使用 `if...else if...else` 述語，就可以將它們從你的程式碼中移除了。

完成後，使用 `let` 宣告一個 `continueLoop` 變數並指定其為布林值 `false`。接著使用 `let` 宣告一個 `done` 變數並指定其值為 `0`。

# --hints--

你不應該有 `if` 述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

你不應該有 `else if` 述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

你不應該有 `else` 述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

你應該使用 `let` 來宣告 `continueLoop` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

你的 `continueLoop` 變數應該有 `false` 的值。

```js
assert.isFalse(continueLoop);
```

你應該使用 `let` 來宣告一個 `done` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

你的 `done` 變數應該有值 `0`。

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
