---
id: 660f0a55847d6cc485f29ba5
title: 步驟 26
challengeType: 1
dashedName: step-26
---

# --description--

JavaScript 中的 <dfn>方法</dfn> 是與特定值或物件相關的函式（程式）。你已經遇過的一個範例是 `.log()` 方法，它是 `console` 物件的部分。

陣列有自己的方法，首先你將探索的是 `.push()` 方法。這個方法允許你將一個值「推入」陣列的末端。以下是一個將數字 `12` 添加到陣列末端的範例：

```js
array.push(12);
```

使用 `.push()` 將字串 `"freeCodeCamp"` 添加到你的 `rows` 陣列的末端。將此程式碼添加到你的 `console.log` 之前，這樣你就可以看到對陣列所做的更改。

# --hints--

你應該在你的程式碼中使用 `.push()`。

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

你應該使用你的 `rows` 陣列的 `.push()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

你應該將字串 `"freeCodeCamp"` 傳遞給你的 `.push()` 方法。 

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
