---
id: 660f229d2dbe09ef2954a4a1
title: 步驟 42
challengeType: 1
dashedName: step-42
---

# --description--

請記得在你之前的迴圈中，你使用了加法運算子 `+` 將 `i` 的值增加了 `1`。

你可以對字串值做類似的事情，方法是將新的字串附加到現有的字串。 例如，`hello = hello + " World";` 將會為儲存在 `hello` 變數中的現有字串添加字串 `" World"`。這稱為 <dfn>串接</dfn>。

在你的 `for...of` 迴圈中，使用加法運算子將 `row` 值串接到 `result` 值。

# --hints--

你應該在你的 `result` 變數上使用串接運算子。

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

你應該將 `row` 串接到你的 `result` 變數。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

你應該將串接的結果指定回 `result` 變數。

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
