---
id: 660f4f79e2a82a4e92290f44
title: 步驟 110
challengeType: 1
dashedName: step-110
---

# --description--

有時你可能想要恢復之前你註解掉的程式碼。你可以透過移除包圍該程式碼的 `/*` 和 `*/` 來達成。這稱為 <dfn>取消註解</dfn>。

只取消註解你的第一個 `for` 迴圈。保留單行註解和另外兩個多行註解。

# --hints--

你應該取消註解你的第一個 `for` 迴圈。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

你不應該移除你的單行註解。

```js
assert.match(code, /\/\//);
```

你不應該取消註解你的 `while` 迴圈。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

你不應該取消註解你的第二個 `for` 迴圈。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
