---
id: 660f4efcb8068e4cb470dca1
title: 步驟 109
challengeType: 1
dashedName: step-109
---

# --description--

既然你已經嘗試過這些方法，你可以進行另一個倒金字塔方法。但首先你需要清理你的實驗過程。

移除你的 `numbers` 陣列，以及方法呼叫和日誌呼叫。

# --hints--

你不應該有 `numbers` 陣列。

```js
assert.notMatch(__helpers.removeJSComments(code), /numbers/);
```

你不應該有 `unshifted` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /unshifted/);
```

你不應該有 `shifted` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /shifted/);
```

你不應該有你的 `console.log` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];
const shifted = numbers.shift();
console.log(shifted);
const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
