---
id: 660f46460f9c36330ebc07d8
title: 步驟 88
challengeType: 1
dashedName: step-88
---

# --description--

等號運算子在 JavaScript 中可能導致一些奇怪的行為。<dfn>例如</dfn>，`"0" == 0` 是 true，儘管一個是字串，一個是數字。

<dfn>嚴格相等性</dfn>運算子 `===` 用於檢查兩個值是否相等且具有相同的型別。一般規則是，你應該永遠使用這個等號運算子。使用嚴格相等性運算子時，`"0" === 0` 會變成 false，因為雖然它們的值可能都是零，但它們的型別並不相等。

將你的 `done == count` 條件更新為使用嚴格等號運算子。

# --hints--

你的 `if` 條件應該使用嚴格相等性。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
