---
id: 660f2fbd45b520046cac68e8
title: 步驟 65
challengeType: 1
dashedName: step-65
---

# --description--

請記得在先前的步驟中，你已經學習過返回值。函式（程式）可以<dfn>return</dfn>一個值，供你的應用程式單獨使用。

在函式（程式）中，`return` 關鍵字用來指定返回值。例如，這個函式會傳回給第一個參數的值：

```js
function name(parameter) {
  return parameter;
}
```

使用 `return` 關鍵字傳回 `character` 變數的值，重複 `rowNumber` 次。

# --hints--

你應該使用 `.repeat()` 方法。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

你應該在你的 `character` 變數上使用 `.repeat()` 方法。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

你應該將 `rowNumber` 傳遞給你的 `.repeat()` 呼叫。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

你應該使用 `return` 關鍵字。

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

你應該傳回 `.repeat()` 呼叫的結果。

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
