---
id: 660f23b53db70af0f2620e78
title: 步驟 43
challengeType: 1
dashedName: step-43
---

# --description--

現在你所有的數字都出現在同一行。這樣無法用來建立金字塔。

你需要為每一列添加一個新行。然而，在 JavaScript 中按下 return 鍵以在引號之間插入換行符會導致解析錯誤。相反，你需要使用特殊的 <dfn>escape sequence</dfn> `\n`，當字串被記錄時，它會被解釋為新行。例如：

```js
lineOne = lineOne + "\n" + lineTwo;
```

使用第二個加法運算子將新列附加在 `result` 和 `row` 值之後。

# --hints--

你應該使用 `\n` 跳脫序列。請記得它需要是一個字串，因此會被包在引號中。

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

你應該將你的 `row` 變數串接到你的 `result` 變數。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

你應該將 `\n` 跳脫序列串接到你的 `row` 變數。

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

你應該將整個串接結果指定回你的 `result` 變數。

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
