---
id: 660f033cf051ebb50ea3bf48
title: 步驟 18
challengeType: 1
dashedName: step-18
---

# --description--

在程式設計中，你經常需要處理大量的資料。有許多資料結構可以幫助你組織和管理你的資料。其中最基本的資料結構之一是 <dfn>array</dfn>。

<dfn>陣列</dfn>是一種非原始型別的資料型別，可以儲存一系列的值。非原始資料型別與原始資料型別不同，因為它們可以儲存更複雜的資料。像字串和數字這類原始資料型別一次只能儲存一個值。

陣列使用方括號（`[]`）表示。以下是一個值為空陣列的變數範例：

```js
let array = [];
```

宣告一個 `rows` 變數並指定一個空的陣列。

# --hints--

你應該有一個 `rows` 變數。

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

你應該使用 `let` 來宣告你的 `rows` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

你應該將一個空的陣列指定給你的 `rows` 變數。

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
