---
id: 660f48a419b40238e2b8b4d5
title: 步驟 93
challengeType: 1
dashedName: step-93
---

# --description--

你的迴圈不再依賴 `continueLoop` 變數。這使該變數成為一個 <dfn>未使用的宣告</dfn>。通常，你會想避免未使用的宣告以防止未來的混淆。

移除你的 `continueLoop` 變數。

# --hints--

你不應該再有 `continueLoop` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /continueLoop/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
