---
id: 660f17294346b7d69e79db3d
title: 步驟 32
challengeType: 1
dashedName: step-32
---

# --description--

你現在已經準備好開始建置你的金字塔產生器。你的 `character` 變數將作為金字塔的建造區塊。

「Hello」可能不太適合那個用途。將 `character` 的值改為雜湊字元 `"#"`。

# --hints--

你的 `character` 變數應該是一個雜湊符號。

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
