---
id: 660f49e32001983c90b75850
title: 步驟 98
challengeType: 1
dashedName: step-98
---

# --description--

那是一個非常乾淨且功能完整的迴圈。做得好！但還有更多可以探索的地方。

使用多行註解來註解你的 `while` 迴圈。

# --hints--

你的 `while` 迴圈應該被註解掉。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

你的 `while` 迴圈主體應該被註解掉。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
