---
id: 660f4774e3e0df35a68bb5f2
title: 步驟 90
challengeType: 1
dashedName: step-90
---

# --description--

為了讓你的金字塔再次產生，將使用 `done` 和 `count` 作為引數呼叫 `padRow` 的結果推入你的 `rows` 陣列，就像你在第一個迴圈中所做的那樣。

# --hints--

你的迴圈應該對你的 `rows` 呼叫 `.push()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(/);
```

你應該在你的 `.push()` 方法中呼叫你的 `padRow` 函式。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(/);
```

你應該將 `done` 作為第一個引數傳遞給你的 `padRow` 呼叫。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done/);
```

你應該將 `count` 作為第二個引數傳遞給你的 `padRow` 呼叫。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--

--fcc-editable-region--
  if (done === count) {
    continueLoop = false;
  } 
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
