---
id: 660f34e99571070d56d2f231
title: 步驟 67
challengeType: 1
dashedName: step-67
---

# --description--

你的 `padRow` 函式有兩個你定義的參數。當函式被呼叫時，值會提供給那些參數。

你提供給函式呼叫的值稱為<dfn>引數</dfn>，你會將引數<dfn>傳遞</dfn>給函式呼叫。這裡有一個函式呼叫，傳遞了 `"Hello"` 作為引數：

```js
test("Hello");
```

將 `i + 1` 和 `count` 作為引數傳遞給你的 `padRow` 呼叫。像參數一樣，引數以逗號分隔。

# --hints--

你應該將 `i + 1` 傳遞給你的 `padRow()` 呼叫。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

你應該在你的 `i + 1` 引數後面加上逗號。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

你應該將 `count` 作為你的第二個引數傳遞。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
