---
id: 660f45ccf4ca5c31f253005a
title: 步驟 86
challengeType: 1
dashedName: step-86
---

# --description--

現在，如果你將 `continueLoop` 改為 true，你的 `while` 迴圈將會無限執行。這稱為 <dfn>無窮迴圈</dfn>，你應該小心避免這種情況。無窮迴圈可能會鎖住你的系統，必須重新啟動才能脫離。

為避免這種情況，請先在你的迴圈中使用累加運算子來增加 `done` 變數的值。

# --hints--

你的 `while` 迴圈應該遞增 `done` 變數。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
