---
id: 660f4a83373de83ea101685f
title: 步驟 100
challengeType: 1
dashedName: step-100
---

# --description--

因為你將要反方向迴圈，所以你的迴圈需要在 `i` 大於 `0` 時執行。你可以使用 <dfn>greater than</dfn> 運算子 `>` 來達成這個目的。

將你的迴圈條件設定為當 `i` 大於 `0` 時執行。

# --hints--

當 `i` 大於 0 時，你的 `for` 迴圈應該執行。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
