---
id: 660f3ce51f70571e1c5227c8
title: 步驟 76
challengeType: 1
dashedName: step-76
---

# --description--

註解可以幫助說明為什麼你的程式碼採取某種方法，或為你未來的自己留下待辦事項。

在 JavaScript 中，你可以使用 `//` 在你的程式碼中留下單行註解。

在你的 `for` 迴圈上方添加單行註解，以提醒自己將程式碼更改為不同種類的迴圈。

# --hints--

你應該以 `//` 開始單行註解。

```js
assert.match(code, /\/\//);
```

你的單行註解應該至少有五個字元長。

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
