---
id: 660f2eccfe3f820304af1b39
title: 步驟 64
challengeType: 1
dashedName: step-64
---

# --description--

為了知道如何格式化一列，你的 `padRow` 函式（程式）需要知道你目前在第幾列，以及總共產生了多少列。

最好的方法是為它們建立函式（程式）參數。為你的 `padRow` 函式（程式）添加 `rowNumber` 和 `rowCount` 參數。多個參數以逗號分隔：

```js
function name(first, second) {

}
```

# --hints--

你的 `padRow` 函式（程式）應該有一個 `rowNumber` 參數。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

你應該在你的 `rowNumber` 參數後面添加逗號。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

你的 `padRow` 函式應該有一個 `rowCount` 參數。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
