---
id: 660f4ae5b3924c3fc3373973
title: 步驟 101
challengeType: 1
dashedName: step-101
---

# --description--

你的迭代述語也會不同。你需要在每個迴圈中從 `i` 減去 `1`，而不是加上 `1`。

像你之前對 `i = i + 1` 所做的那樣，更新你的迭代述語，使 `i` 的值為從自身減去 `1`。

# --hints--

你的 `for` 迴圈應該使用 `i = i - 1` 作為迭代。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*=\s*i\s*-\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
