---
id: 660f383d4c772c12ff59904b
title: 步驟 70
challengeType: 1
dashedName: step-70
---

# --description--

你可以將完整的表達式作為引數傳遞。函式會接收評估該表達式的結果。<dfn>例如</dfn>，下面這兩個函式呼叫會產生相同的結果。（請注意，運算優先順序規則 PEMDAS——括號、指數、乘法、除法、加法、減法——適用）：

```js
test(2 * 3 + 1);
test(7);
```

再次查看這個樣式：

```md
··#··
·###·
#####
```

將 `character` 值更新為重複 `2 * rowNumber - 1` 次。

再次打開主控台以查看更新後的結果。

# --hints--

你的函式（程式）應該將 `2 * rowNumber - 1` 傳遞給你的 `.repeat()` 呼叫。

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
