---
id: 660f34626216270c682e2f7b
title: 步驟 66
challengeType: 1
dashedName: step-66
---

# --description--

<dfn>函式呼叫</dfn>讓你實際使用一個函式。你可能沒有注意到，但你一直在使用的像 `.push()` 這樣的方法就是函式呼叫。

函式（程式）是透過取用函式（程式）名稱並添加 `()` 來呼叫的。以下是如何呼叫 `test` 函式（程式）：

```js
test();
```

將 `.push()` 呼叫中的 `character.repeat(i + 1)` 取代為對你的 `padRow` 函式（程式）的呼叫。暫時不要為它添加任何引數。

# --hints--

你不應該在你的 `push` 呼叫中使用 `i + 1`。

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

你不應該在你的 `.push()` 呼叫中使用 `character.repeat`。

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

你應該在你的 `.push()` 呼叫中呼叫 `padRow`。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

你不應該在你的 `padRow` 呼叫中有引數。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
