---
id: 660f505d02b2bd513a1c3468
title: 步驟 111
challengeType: 1
dashedName: step-111
---

# --description--

你的金字塔不再是倒置的。這是因為你正在將新的列添加到陣列的**末端**。

更新你的迴圈主體，將新列添加到陣列的開頭。

# --hints--

你應該使用 `rows` 的 `unshift` 方法。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

你應該將 `padRow()` 呼叫作為 `.unshift()` 方法的引數傳入。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

你應該將 `i` 作為第一個引數傳遞給你的 `padRow()` 呼叫。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

你應該將 `count` 作為第二個引數傳遞給你的 `padRow()` 呼叫。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
