---
id: 660f1e3f047bf4e403268713
title: 步驟 39
challengeType: 1
dashedName: step-39
---

# --description--

不幸的是，現在你無法看到你的迴圈正在做什麼。

使用 `let` 宣告一個 `result` 變數，並指定其為空字串。空字串是由中間沒有任何內容的引號表述，例如 `""`。

# --hints--

你應該宣告一個 `result` 變數。

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

你應該使用 `let` 來宣告你的 `result` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

你的 `result` 變數應該是一個空字串。

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
