---
id: 660f38c34a4de6141c0c369f
title: 步驟 71
challengeType: 1
dashedName: step-71
---

# --description--

你的金字塔產生器現在運作如預期。但這是進一步探索你所撰寫程式碼的絕佳機會。

加法運算子並不是將值添加到變數的唯一方式。<dfn>加法指派</dfn>運算子可以用作簡寫，表示「取變數的原始值，加上此值，然後將結果指定回變數」。例如，這兩個述語會產生相同的結果：

```js
test = test + 1;
test += 1;
```

更新你在 `for` 迴圈中的迭代述語以使用加法指派。

# --hints--

你的 `for` 迴圈不應該使用 `i = i + 1`；

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

你的 `for` 迴圈應該使用加法指派搭配 `i`。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
