---
id: 660f4a1472f8e63d76162ce5
title: 步驟 99
challengeType: 1
dashedName: step-99
---

# --description--

如果你把你的金字塔做成倒置的，或是<dfn>反轉</dfn>的呢？是時候試試看了！

從建立一個新的 `for` 迴圈開始。宣告你的迭代器 `i` 並將它指定為 `count` 的值，然後在條件和迭代述語中使用布林值 `false`。

# --hints--

你的程式碼應該有一個 `for` 迴圈。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

你的 `for` 迴圈應該用 `count` 的值來初始化 `i`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count/);
```

你的 `for` 迴圈應該使用 `false` 作為條件。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false/);
```


你的 `for` 迴圈應該使用 `false` 作為迭代條件。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
