---
id: 660f39b444fd6f16d1e49c1f
title: 步驟 73
challengeType: 1
dashedName: step-73
---

# --description--

與其必須將 `i + 1` 傳遞給你的 `padRow` 呼叫，不如將你的迴圈從 `1` 開始。這將允許你創建一個以 1 為起點的迴圈。

將你的迭代器更新為從 `1` 開始，而不是 `0`。

# --hints--

你的 `for` 迴圈應該將 `i` 初始化為 `1`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
