---
id: 660f359af3e32e0f1a6880b7
title: 步驟 68
challengeType: 1
dashedName: step-68
---

# --description--

你現在應該會在你的主控台看到相同的一堆字元。你的 `padRow` 函式正在做你之前所做的完全相同的事情，但現在它是在一個可重複使用的獨立區段中。

使用加法運算子將單一空格 `" "` 串接到你重複的 `character` 字串的開頭和結尾。

請記住，你可以使用 `+` 運算子來串接字串，如此：

```js
" " + "string"
```

# --hints--

你應該在傳回值的開頭串接一個空格。

```js
assert.match(padRow(1, 1), /^\s/);
```

你應該在傳回的值末尾串接一個空格。

```js
assert.match(padRow(1, 1), /\s$/);
```

你的 `padRow()` 函式應該傳回在字元序列前後各有一個空格的重複 `character` 序列。

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
