---
id: 660f0da9bf1035c9097af20a
title: 步驟 28
challengeType: 1
dashedName: step-28
---

# --description--

你應該已經看到 `"freeCodeCamp"` 被列印到主控台。這是因為 `.pop()` 傳回從陣列中移除的值－而你之前將 `"freeCodeCamp"` 推入陣列的末端。

但是 `.push()` 傳回什麼？將你現有的 `rows.push()` 指定給一個新的 `pushed` 變數，並將它記錄下來。

# --hints--

你應該宣告一個 `pushed` 變數。

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

你應該使用 `let` 來宣告你的 `pushed` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

你應該將 `rows.push("freeCodeCamp")` 指定給你的 `pushed` 變數。

```js
assert.equal(pushed, 4);
```

你應該記錄你的 `pushed` 變數。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
