---
id: 660f47afe4c98536715d5fa4
title: 步驟 91
challengeType: 1
dashedName: step-91
---

# --description--

<dfn>嚴格不相等</dfn>運算子 `!==` 允許你檢查兩個值是否不相等，或型別不相同。語法與等號運算子類似：`value !== 4`。

目前 `while` 迴圈僅在 `continueLoop` 為 true 時執行。更新 `while` 迴圈條件以檢查 `done` 是否不相等於 `count`。

# --hints--

你的 `while` 迴圈應該檢查 `done` 和 `count` 是否不相等。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
