---
id: 660f3dd626be3a1ffe27e5d1
title: 步驟 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript 也支援多行註解。多行註解以 `/*` 開始，以 `*/` 結束。

與單行註解不同，多行註解會包覆多行內容。

使用 `/*` 和 `*/` 將你目前的 `for` 迴圈（包含主體）變成多行註解。

# --hints--

你應該以 `/*` 開始多行註解。

```js
assert.match(code, /\/\*/);
```

你應該以 `*/` 結束多行註解。

```js
assert.match(code, /\*\//);
```

你的整個 `for` 迴圈應該被註解掉。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
