---
id: 660f1a00ac619ddc1e259a66
title: 步驟 35
challengeType: 1
dashedName: step-35
---

# --description--

`for` 迴圈的 <dfn>條件</dfn> 告訴迴圈應該迭代多少次。當 `condition` 變為 false 時，迴圈將停止。

在 JavaScript 中，布林值可以是 `true` 或 `false`。這些不是字串－你將在稍後學習更多關於它們之間的差異。

目前，你將使用 <dfn>小於</dfn> 運算子（`<`）。這讓你可以檢查左邊的值是否小於右邊的值。 例如，`count < 3` 如果 `count` 是 `2`，則會計算為 `true`，如果 `count` 是 `4`，則會計算為 `false`。

將你的 `"condition"` 字串取代為檢查 `i` 是否小於 `count` 的條件。

# --hints--

你應該使用小於運算子。

```js
assert.match(__helpers.removeJSComments(code), /</);
```

你應該使用小於運算子來檢查 `i` 是否小於 `count`。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

你的 `for` 迴圈應該使用 `i < count` 作為條件。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
