---
id: 660f0ee51d7460ce88cd248d
title: 第 29 步
challengeType: 1
dashedName: step-29
---

# --description--

你是否預期在主控台看到 `4`？`.push()` 會在加入你給它的值後，傳回陣列的新長度。

了解一個方法傳回哪些值非常重要。花點時間試驗 `.push()` 和 `.pop()`。當你準備好時，移除所有的 `.push()` 和 `.pop()` 呼叫，以及你的 `console.log` 述語。

# --hints--

你不應該有 `.push()` 呼叫。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

你不應該有 `.pop()` 呼叫。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

你不應該有任何日誌述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

你不應該有 `popped` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

你不應該有 `pushed` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
