---
id: 660f0f980e98e8cf77f1ce31
title: 第 30 步
challengeType: 1
dashedName: step-30
---

# --description--

將你的 `rows` 宣告重新指定為空陣列。

此外，將你的 `'Hello'` 字串改回使用雙引號。一般來說，兩者中你偏好哪一種並不重要，但你會想在整個專案中保持該選擇的一致性。

# --hints--

你的 `rows` 陣列應該是空的。

```js
assert.empty(rows);
```

你的 `"Hello"` 字串應該使用雙引號。

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
