---
id: 660f530d6e33d159e1bf4947
title: 步驟 115
challengeType: 1
dashedName: step-115
---

# --description--

如果你的金字塔不是倒置的，則你會想要有一個 `else` 區塊，以正常順序建造金字塔。

在先前的步驟中，你學會了如何像這樣使用 `else` 述語：

```js
if (condition) {
  // if condition is true, run this code
} else {
  // if condition is false, run this code
}
```

為你的 `if` 區塊添加一個 `else` 區塊。

# --hints--

你應該為你的程式碼添加一個 `else` 區塊。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)\s*\{\s*rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)\s*\);\s*\}\s*else\s*\{/);
```

你的 `else` 區塊應該是空的。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)\s*\{\s*rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)\s*\);\s*\}\s*else\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  }

}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
