---
id: 660f09a2694b59c3a10ee304
title: 第 23 步
challengeType: 1
dashedName: step-23
---

# --description--

目前，移除你的第一個主控台日誌和你的 `rows[rows.length - 1]` 指派。第二個 `rows` 日誌述語留待稍後使用。

# --hints--

你應該移除你的 `console.log(rows[0])` 述語。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

你應該移除對 `rows[rows.length - 1]` 的重新指派。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

你不應該移除你的 `console.log(rows)` 述語。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
