---
id: 660f447efc0e722f016c1be0
title: 步驟 81
challengeType: 1
dashedName: step-81
---

# --description--

文字又出現了！這是因為 `"false"` 是一個字串，當被計算成布林值時會變成 `true`。這表示 `"false"` 是一個真值。

<dfn>truthy value</dfn> 是在以布林值計算時被視為 true 的值。你在 JavaScript 中遇到的大多數值都是 truthy。

<dfn>假值</dfn>是相反的——在作為布林值計算時被視為假的值。JavaScript 有一個定義好的假值列表。其中一些包括 `false`、`0`、`""`、`null`、`undefined` 和 `NaN`。

嘗試將你的 `if` 條件改為空字串 `""`，這是一個假值。

# --hints--

你的 `if` 述語應該以 `""` 作為條件。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
