---
id: 660f4934fb48f63abd5ae371
title: 步驟 95
challengeType: 1
dashedName: step-95
---

# --description--

使用 `done` 來追蹤已產生的列數是可行的，但你其實可以進一步簡化邏輯。

陣列有一個特殊的 `length` 屬性，允許你查看陣列中有多少個值或<dfn>元素</dfn>。你可以使用像 `myArray.length` 這樣的語法來存取這個屬性。

請注意，`padRow` 呼叫中的 `rows.length` 會導致一個偏差一的錯誤，因為 `done` 是在呼叫 *之前* 被遞增的。

更新你的條件以檢查 `rows.length` 是否小於 `count`。

# --hints--

你的 `while` 迴圈應該檢查 `rows.length` 是否小於 `count`。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
