---
id: 660f4990b1caa03b9dc97a43
title: 步驟 97
challengeType: 1
dashedName: step-97
---

# --description--

現在你不再需要你的 `done` 變數。從你的迴圈中移除遞增操作，並移除 `done` 的變數宣告。

# --hints--

你不應該遞增 `done` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /done\+\+/);
```

你不應該再有 `done` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /done/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (rows.length < count) {
  done++;
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
