---
id: 660f18f059fe0fda192ce394
title: 步驟 34
challengeType: 1
dashedName: step-34
---

# --description--

你的迴圈現在需要一個適當的迭代器。<dfn>iterator</dfn> 是你可以在 `for` 迴圈中專門宣告的變數，用來控制迴圈如何迭代或執行你的邏輯。

在迴圈中使用 `i` 作為你的迭代器變數是一個常見的慣例。`for` 迴圈允許你在小括弧 `()` 中宣告這個變數。<dfn>例如</dfn>，這裡有一個 `for` 迴圈，它宣告了一個 `index` 變數並將其設值為 `100`。

```js
for (let index = 100; "second"; "third") {

}
```

將字串 `"iterator"` 取代為變數 `i` 的 `let` 宣告。指定它的值為 `0` 作為起始。這將使 `i` 變數在你的迴圈**第一次**執行時擁有值 `0`。

# --hints--

你應該使用 `let` 來宣告一個 `i` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

你應該將 `0` 指定到你的 `i` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

你的 `for` 迴圈應該從 `0` 開始一個 `i` 迭代器。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
