---
id: 660f165270622fd4ec0da3f7
title: 步驟 31
challengeType: 1
dashedName: step-31
---

# --description--

使用 `let` 關鍵字宣告變數允許它被重新指定、設值。這表示你之後可以將 `character` 改成完全不同的值。

對於這個專案，你不會想要更改這些變數值。因此，你應該使用 `const` 來宣告它們。`const` 變數是特殊的。

首先，`const` 變數無法像 `let` 變數那樣被重新指定。這段程式碼會發出錯誤：

```js
const firstName = "Naomi";
firstName = "Jessica";
```

`const` 變數也不能未初始化。這段程式碼會發出錯誤：

```js
const firstName;
```

將你的 `let` 關鍵字取代為 `const`。

# --hints--

你應該使用 `const` 來宣告你的 `character` 變數。

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

你應該使用 `const` 來宣告你的 `count` 變數。

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

你應該使用 `const` 來宣告你的 `rows` 變數。

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

你不應該在你的程式碼中使用 `let`。

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
