---
id: 660f50a21fe7645252804f2b
title: 步驟 112
challengeType: 1
dashedName: step-112
---

# --description--

如果你有一種方法可以在倒金字塔和標準金字塔之間切換，會怎麼樣？

首先宣告一個 `inverted` 變數，並指定其值為 `true`。你在程式碼中不會更改這個變數，但你需要使用 `let`，以便我們的測試稍後可以修改它。

# --hints--

你應該使用 `let` 宣告一個 `inverted` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

你應該將 `inverted` 初始化為值 `true`。

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
