---
id: 660f255022991ef34ed0ee88
title: 步驟 44
challengeType: 1
dashedName: step-44
---

# --description--

列印數字不會產生視覺上吸引人的金字塔。既然你正在輸出 `rows` 陣列的格式化內容，現在是時候更新你原本的迴圈了。

不要將 `i` 推入陣列，而是將你的 `character` 變數的值推入陣列。

# --hints--

你不應該再推送你的 `i` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

你應該推送你的 `character` 變數。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
