---
id: 660f4cde8dd305450514a1cb
title: 步驟 105
challengeType: 1
dashedName: step-105
---

# --description--

也使用多行註解來註解這個迴圈，以準備下一個方法。

# --hints--

你的 `for` 迴圈應該使用多行註解註解起來。

```js
assert.lengthOf(code.match(/\/\*\s*for\s*\(/g), 2);
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 1);
```

你的 `for` 迴圈主體應該作為多行註解的一部分被註解掉。

```js
assert.notMatch(code, /\/\/\s*rows\.push/);
assert.lengthOf(code.match(/rows\.push/g), 3)
assert.notMatch(__helpers.removeJSComments(code), /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
