---
id: 660f3915b41a441537ec9f5e
title: 步驟 72
challengeType: 1
dashedName: step-72
---

# --description--

因為你只將 `i` 增加 1，所以你可以使用 <dfn>累加運算子</dfn> `++`。這個運算子會將變數的值增加 1，並更新該變數的指派。<dfn>例如</dfn>，`test` 在這裡會變成 `8`：

```js
let test = 7;
test++;
```

將你的加法指派取代為累加運算子以進行你的迴圈疊代。

# --hints--

你的 `for` 迴圈不應該使用加法指派與 `i`。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

你的 `for` 迴圈應該對 `i` 使用累加運算子。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
