---
id: 660f039ff313dbb696b007ca
title: 第 19 步
challengeType: 1
dashedName: step-19
---

# --description--

當一個陣列包含值，或<dfn>元素</dfn>時，這些值會以逗號分隔。這是一個包含兩個字串的陣列：

```js
let array = ["first", "second"];
```

將你的 `rows` 宣告改為包含字串 `"Naomi"`、`"Quincy"` 和 `"CamperChan"` 的陣列。陣列中值的順序很重要，所以請遵循該順序。請記得字串是區分大小寫的。

# --hints--

你的陣列中的第一個元素應該是字串 `"Naomi"`。

```js
assert.equal(rows[0], "Naomi");
```

你的陣列中的第二個元素應該是字串 `"Quincy"`。

```js
assert.equal(rows[1], "Quincy");
```

你的陣列中的第三個元素應該是字串 `"CamperChan"`。

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
