---
id: 660f207334fabaeac3269c38
title: 步驟 41
challengeType: 1
dashedName: step-41
---

# --description--

要操作 `result` 字串，你將使用不同類型的迴圈。具體來說，是 `for...of` 迴圈，它會迭代可迭代物件中的每個項目，並暫時將其指定給一個變數。

`for...of` 迴圈的語法如下：

```js
for (const value of iterable) {

}
```

請注意，你可以使用 `const`，因為該變數只存在於單次迴圈中，而非整個迴圈期間。

創建一個 `for...of` 迴圈來迭代你的 `rows` 陣列，並將每個值指定給一個 `row` 變數。

# --hints--

你應該使用另一個 `for` 關鍵字。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

你應該宣告一個 `row` 變數。

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

你應該使用 `const` 來宣告你的 `row` 變數。

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

你的 `for...of` 迴圈應該宣告你的 `row` 變數。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

你的 `row` 變數應該使用 `of` 關鍵字從 `rows` 中擷取。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

你的 `for...of` 迴圈主體應該是空的。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
