---
id: 660f53ad3d39175c5d4335ac
title: 第 117 步
challengeType: 1
dashedName: step-117
---

# --description--

你的金字塔產生器現在已完成，擁有比你原本計劃更多的功能！下一步是清理你的程式碼。

從你的程式碼中移除所有註解，包括單行和多行註解。

# --hints--

你的程式碼中不應該有任何單行註解。

```js
assert.notMatch(code, /\/\//);
```

你的程式碼中不應該有任何多行註解。

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

你的程式碼中不應該有任何註解。

```js
assert.equal(code, __helpers.removeJSComments(code));
```

你應該移除被多行註解註解掉的程式碼。

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
