---
id: 6610b9f7619764fad5fd516d
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

你也可以將一個變數的值指定給另一個變數。例如：

```js
let first = "One";
let second = "Two";
second = first;
```

`second` 變數現在會有值 `"One"`。

要看到這個效果，請將你的 `secondCharacter` 指派從 `"Test"` 改為你的 `character` 變數。

然後打開主控台查看記錄了什麼。

# --hints--

你不應該將值 `"Test"` 指定給你的 `secondCharacter` 變數。

```js
assert.notEqual(secondCharacter, "Test");
```

你應該將 `character` 變數的值指定給你的 `secondCharacter` 變數。別忘了你的分號。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

你的 `secondCharacter` 變數現在應該有值 `"World"`。


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
