---
id: 6610b8017d1671f2814e8c77
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

當你宣告一個變數卻沒有初始化時，它會被視為 <dfn>未初始化</dfn>。目前，你的 `secondCharacter` 變數是未初始化的。

為你的程式碼添加 `console.log()` 以查看你的 `secondCharacter` 變數的值。

# --hints--

你應該添加第二個 `console.log`。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

你應該記錄你的 `secondCharacter` 變數。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
