---
id: 6610c21b3ef82015573ffbbe
title: 步驟 51
challengeType: 1
dashedName: step-51
---

# --description--

你的 `call` 變數有一個 `undefined` 值，儘管你已經定義它！這是因為你的 `padRow` 函式（程式）目前沒有傳回任何值。函式（程式）預設會傳回 `undefined` 作為它們的值。

為了傳回其他東西，你需要使用 `return` 關鍵字。以下是一個傳回字串 `"Functions are cool!"` 的函式範例：

```js
function demo() {
  return "Functions are cool!";
}
```

使用 `return` 關鍵字讓你的函式傳回字串 `"Hello!"`。

# --hints--

你的 `padRow` 函式應該使用 `return` 關鍵字。

```js
assert.match(padRow.toString(), /return/);
```

你的 `padRow` 函式應該傳回字串 `"Hello!"`。

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
