---
id: 6610c71600966a2191d3a64a
title: 步驟 59
challengeType: 1
dashedName: step-59
---

# --description--

從函式傳回的值是透過呼叫該函式來使用的。你可以直接將函式呼叫當作它傳回的值使用，或是將傳回的值指定給一個變數。這樣，你就可以在函式創建的範圍外使用指定給區域變數的值。

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

要使用你的 `"Testing"` 值，請透過更新你的 `return` 述語，使 `padRow` 函式只傳回 `test` 變數。


# --hints--

你的 `padRow` 函式應該傳回 `test` 變數。

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
