---
id: 6610c48c4ea0891afa7c4696
title: 步驟 54
challengeType: 1
dashedName: step-54
---

# --description--

如果你再次打開你的主控台，你會看到你的 `padRow` 函式正在傳回 `undefined`，即使你定義了返回值！這是因為在你**呼叫**函式時，參數需要被賦予值。

當你將一個值傳遞給函式（程式）的呼叫時，該值稱為一個 <dfn>引數</dfn>。以下是一個呼叫 `demo` 函式（程式）並將 `"Naomi"` 作為 `name` 參數的引數的範例。

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

將你自己的名字作為引數傳遞給 `padRow` 呼叫中的 `name` 參數。請記得你的名字是字串，因此你需要使用引號。

# --hints--

你應該將字串傳遞給你的 `padRow()` 呼叫。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
