---
id: 6610c1d97b1671140f95cfbb
title: 步驟 50
challengeType: 1
dashedName: step-50
---

# --description--

現在為你的 `call` 變數添加一個列印其值的日誌述語。

# --hints--

你應該為你的程式碼添加一個 `console.log` 呼叫。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

你應該記錄你的 `call` 變數。別忘了分號。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
