---
id: 6610c16c4fa0df12c0e30675
title: 步驟 49
challengeType: 1
dashedName: step-49
---

# --description--

你正在呼叫你的 `padRow` 函式，但沒有對該函式呼叫做任何事情。JavaScript 中所有函式都會 <dfn>傳回</dfn>一個值，這表示它們會提供呼叫時定義的結果，讓你在其他地方使用。

要查看呼叫你的 `padRow` 函式（程式）的結果，請宣告一個 `call` 變數並將你現有的 `padRow` 呼叫指定給該變數。

# --hints--

你應該宣告一個 `call` 變數。

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

你應該使用 `const` 來宣告你的 `call` 變數。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

你應該將 `call` 指定為你的 `padRow` 呼叫的結果。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
