---
id: 6610c77d50636722e5b6be17
title: 步驟 60
challengeType: 1
dashedName: step-60
---

# --description--

在 `return` 述語下方，將字串 `"This works!"` 記錄到主控台。

完成後，你會看到字串 `"This works!"` 不會顯示在主控台中，且 `console.log("This works!")` 這一行會呈現灰色。

複製主控台日誌並將它貼到 `return` 述語上方。現在，字串 `"This works!"` 應該會出現在主控台中。

關於 `return` 關鍵字，一個重要的事情是它不僅定義要從你的函式傳回的值，它也會停止你在函式或區段述語內的程式碼執行。這表示在 `return` 述語之後的任何程式碼都不會執行。

# --hints--

你的 `padRow` 函式應該傳回 `test` 變數。

```js
assert.equal(padRow("Naomi"), "Testing");
```

你的第一個 `console.log` 應該放在你的 `return` 關鍵字之後。

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

你的第二個 `console.log` 應該放在你的 `return` 關鍵字之前。

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
