---
id: 6610c8cfe4cf4d278e35c156
title: 步驟 63
challengeType: 1
dashedName: step-63
---

# --description--

如預期，你的函式（程式）現在再次傳回 `undefined`。你的 `call` 變數不再必要，因此請移除 `call` 宣告式以及 `call` 變數的 `console.log`。

# --hints--

你不應該有 `call` 宣告式。

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

你不應該記錄你的 `call` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
