---
id: 6610c105bbdacc114d6cdc44
title: 第 48 步
challengeType: 1
dashedName: step-48
---

# --description--

為了使用函式（程式），你需要呼叫它。<dfn>函式呼叫</dfn> 告訴你的應用程式在你選擇呼叫的任何地方執行該函式的程式碼。函式呼叫的語法是函式名稱後面接小括弧。 例如，這段程式碼定義並呼叫了 `test` 函式。

```js
function test() {

}
test();
```

呼叫你的 `padRow` 函式。

# --hints--

你應該呼叫 `padRow` 函式（程式）。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

你的 `padRow` 函式（程式）主體應該是空的。

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
