---
id: 6610c424b7119919b62932f4
title: 步驟 53
challengeType: 1
dashedName: step-53
---

# --description--

函式不必傳回寫死的值。它可以傳回儲存在變數中的值。參數是函式的特殊變數，因此它們也可以被傳回。

將你的 `padRow` 函式直接傳回 `name` 參數。

# --hints--

你的 `padRow` 函式應該傳回 `name` 參數的值。

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
