---
id: 6610b8f6a98d25f4d485a94d
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

未初始化變數的預設值是 `undefined`。這是一種特殊的資料型別，表述尚未有定義的值。

你仍然可以為未初始化的變數指定值。以下是一個範例：

```js
let uninitialized;
uninitialized = "assigned";
```

在你的宣告式下方，將字串 `"Test"` 指定給你的 `secondCharacter` 變數。打開主控台以查看你的日誌如何變化。

# --hints--

你不應該初始化 `secondCharacter`。請記住，初始化是指在宣告變數時指定值。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

你應該對 `secondCharacter` 使用指定運算子。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

你應該將字串 `"Test"` 指定給你的 `secondCharacter` 變數。

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
