---
id: 6610bbed59bc2a0194d85533
title: 第 22 步
challengeType: 1
dashedName: step-22
---

# --description--

注意你的 `rows` 陣列內的值是如何被直接更改的？這稱為 <dfn>mutation</dfn>。當你進一步學習陣列時，你會了解何時應該變異陣列，何時不應該。

在繼續之前，這是一個學習常見陣列用法的好機會。目前，你的程式碼是用 `rows[2]` 存取陣列中的最後一個元素。但當你想要最後一個元素時，可能不知道陣列中有多少個元素。

你可以使用陣列的 `.length` 屬性－這會傳回陣列中元素的數量。要取得任何陣列的最後一個元素，你可以使用以下語法：

```js
array[array.length - 1]
```

`array.length` 傳回陣列中元素的數量。透過減去 `1`，你可以取得陣列中最後一個元素的索引。你可以將這個相同的概念套用到你的 `rows` 陣列。

更新你的 `rows[2]` 以動態存取 `rows` 陣列中的最後一個元素。請參考上方的範例以協助你。

你不應該在你的主控台中看到任何變化。

# --hints--

你應該使用你的 `rows` 陣列的 `.length` 屬性。

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

你應該在你的括號表示法中使用 `rows.length`。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

你應該從你的括號表示法中的長度減去 `1`。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
