---
id: 6610c87eac0f0b256d7b037e
title: 步驟 62
challengeType: 1
dashedName: step-62
---

# --description--

因為你的函式（程式）不再使用該參數，改變引數並未影響它。

請繼續移除你 `padRow` 函式中的 `test` 宣告和 `return` 述語，讓該函式再次變成空的。

# --hints--

你的 `padRow` 函式不應該有 `test` 變數。

```js
assert.notMatch(padRow.toString(), /test/);
```

你的 `padRow` 函式不應該傳回任何值。

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
