---
id: 6610b741b54b90f0c0fb3d58
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

當變數名稱超過一個字時，對於如何將字母大寫有特定的命名慣例。在 JavaScript 中，慣例是使用 <dfn>camel case</dfn>。

駝峰式大小寫表示名稱中的第一個字完全是小寫，但後續的字都是首字大寫。以下是一些駝峰式大小寫的範例：

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

使用駝峰式大小寫宣告新的 `secondCharacter` 變數。

# --hints--

你應該宣告一個 `secondCharacter` 變數。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

你應該使用 `let` 來宣告你的 `secondCharacter` 變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

你不應該為你的 `secondCharacter` 變數指定值。不要忘記你的分號。

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
