---
id: 6610c538372aa61cc0f5b122
title: 步驟 57
challengeType: 1
dashedName: step-57
---

# --description--

JavaScript 中的變數存在於特定的 <dfn>範圍</dfn>。換句話說，變數宣告的位置決定了它在你的程式碼中可以使用的範圍。

第一個範圍是全域的範圍。宣告在任何像函式或 `for` 迴圈這類「區塊」之外的變數，都屬於<dfn>全域的範圍</dfn>。你的 `character`、`count` 和 `rows` 變數都在全域的範圍內。

當變數位於全域的範圍時，函式可以在其定義中存取該變數。以下是一個使用全域的 `title` 變數的函式範例：

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

此範例會傳回 `"Professor Naomi"`。更新你的 `padRow` 函式（程式），以傳回將你的 `character` 變數串接到 `name` 參數開頭的值。

# --hints--

你的 `padRow` 函式應該將 `character` 串接到 `name` 的開頭。

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

你的 `padRow` 函式應該傳回 `character + name` 的結果。

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
