---
id: 6610c6541c82551f95e765ab
title: 步驟 58
challengeType: 1
dashedName: step-58
---

# --description--

變數也可以在函式（程式）內宣告。這些變數被視為在<dfn>區域的範圍</dfn>或<dfn>區塊範圍</dfn>內。在函式（程式）內宣告的變數只能在該函式（程式）內使用。如果你嘗試在函式（程式）外存取它，會得到引用錯誤。

要看到這個效果，請在你的 `padRow` 函式中使用 `const` 宣告一個 `test` 變數。並以值 `"Testing"` 初始化它。

然後，在你的函式下方，嘗試將 `test` 記錄到主控台。你會看到錯誤，因為它未在函式的區域作用域之外定義。移除該 `console.log` 以通過測試並繼續。

# --hints--

你的函式（程式）應該宣告一個 `test` 變數。

```js
assert.match(padRow.toString(), /var\s+test/);
```

你應該用值 `"Testing"` 初始化 `test`。別忘了分號。

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

你的 `test` 變數應該放在你的 `return` 關鍵字之前。

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
