---
id: 6610c83b52583e245a079217
title: 步驟 61
challengeType: 1
dashedName: step-61
---

# --description--

現在你的 `call` 變數具有值 `"Testing"`。但你的函式（程式）不再使用 `name` 參數。

從你的函式宣告中移除 `name` 參數，然後從 `padRow` 呼叫中移除你的 `"CamperChan"` 字串。

此外，從 `padRow` 函式中移除兩個 `console.log`。
# --hints--

你的 `padRow` 函式不應該有 `name` 參數。

```js
assert.notMatch(padRow.toString(), /name/);
```

你不應該將 `"CamperChan"` 傳遞給你的 `padRow` 呼叫。

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

你仍然應該呼叫你的 `padRow` 函式（程式）。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

你不應該在 `return` 關鍵字之前有 `console.log`。

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

你不應該在你的 `return` 關鍵字之後有 `console.log`。

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
