---
id: 6610bf6fa14d700beed1b109
title: 步驟 87
challengeType: 1
dashedName: step-87
---

# --description--

<dfn>相等性</dfn>運算子 `==` 用於檢查兩個值是否相等。要比較兩個值，你會使用像 `value == 8` 這樣的述語。

在你的迴圈中 `done++` 下面，為你的程式碼添加一個 `if` 述語。該述語應該使用等號運算子檢查 `done` 是否相等於 `count`。

# --hints--

你應該在你的迴圈中使用 `if` 述語。它應該在 `done++` 之後添加。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

你的 `if` 述語應該使用等號運算子來比較條件中的 `done` 和 `count`。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
