---
id: 6610c2d8d67563174fcf96dc
title: 步驟 52
challengeType: 1
dashedName: step-52
---

# --description--

當你在程式碼中明確寫入一個值，例如你函式中的 `"Hello!"` 字串時，該值即被視為 <dfn>寫死的</dfn>。在函式內寫死一個值可能會讓它沒有你期望的那麼可重複使用。

相反地，你可以為函式定義<dfn>參數</dfn>。參數是特殊的變數，當你呼叫函式時會被賦予值，並且可以在你的函式中使用，以動態改變函式程式碼的結果。

為你的函式（程式）添加參數，你需要在小括弧內添加變數名稱。 例如，這個 `demo` 函式（程式）有一個 `name` 參數：

```js
function demo(name) {

}
```

`name` 聽起來像是一個有用的參數，所以請繼續將它添加到你的 `padRow` 函式中。

# --hints--

你的 `padRow` 函式應該有一個 `name` 參數。

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
