---
id: 664599653fcd6e97104f9261
title: 步驟 83
challengeType: 1
dashedName: step-83
---

# --description--

有時候，當所有的 `if...else if` 條件都是 `false` 時，你會想要執行不同的程式碼。你可以透過添加一個 `else` 區塊來做到這一點。 

只有在 `if` 和 `else if` 區塊中的條件不符合時，`else` 區塊才會被計算。

這裡將 `else` 區塊添加到 `else if` 區塊。 

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

為你的 `else if` 區塊添加一個 `else` 區塊。在 `else` 區塊內，將字串 `"This is the else block"` 記錄到主控台。

要在主控台中查看結果，你可以手動將 `else if` 述語中的 `<` 改成 `>`。這樣條件會變成 `false`，`else` 區塊就會被執行。

# --hints--

你應該有一個 `else` 區塊。

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

你的 `else` 區塊應該將字串 `"This is the else block"` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

