---
id: 66458f0a05df478aa627629e
title: 步驟 82
challengeType: 1
dashedName: step-82
---

# --description--

文字又不見了！空字串會被計算為 `false`，使它們成為一種 <dfn>falsy value</dfn>。你將在未來的專案中學習更多關於 truthy 和 falsy 值的知識。

除了 `if` 述語之外，JavaScript 也有 <dfn>else if</dfn> 述語。`else if` 述語允許你在單一區塊的程式碼中檢查多個條件。

這是 `else if` 述語的語法：

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

如果第一個條件是 `false`，JavaScript 將會檢查串鏈中的下一個條件。如果第二個條件是 `false`，JavaScript 將會檢查第三個條件，依此類推。

在你的 `if` 述語下方，添加一個 `else if` 述語來檢查 `5` 是否小於 `10`。然後在 `else if` 述語的主體內，將字串 `"5 is less than 10"` 輸出到主控台。

檢查主控台以查看結果。

# --hints--

你應該有一個 `else if` 述語。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

你的 `else if` 述語應該檢查 `5` 是否小於 `10`。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

你應該將字串 `"5 is less than 10"` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

