---
id: 664ed9344ba3495e1f67f8ea
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

現在你已經檢視過宣告式和初始化，請移除你為檢視所撰寫的程式碼。

請勿移除你的 `character` 變數。

# --hints--

你不應該有 `profession` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

你不應該有 `age` 變數。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

你不應該有 `console.log` 述語用於 `age`。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

你不應該有 `console.log` 述語用於 `profession`。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

你應該仍然有你的 `character` 變數。

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
