---
id: 664ed97c55a99e5ffab759dc
title: 步驟 13
challengeType: 1
dashedName: step-13
---

# --description--

在繼續之前，你應該花點時間複習你已經學過的概念。

使用 `let` 關鍵字宣告一個 `profession` 變數和一個 `age` 變數。將 `profession` 初始化為字串 `"teacher"`，但不要為 `age` 初始化任何值。

將你的兩個變數都記錄到主控台以查看結果。

# --hints--

你應該宣告一個 `profession` 變數。

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

你應該使用 `let` 來宣告 `profession` 變數。

```js
assert.match(code, /let\s+profession/);
```

你應該將字串 `"teacher"` 指定到 `profession` 變數。

```js
assert.strictEqual(profession, "teacher");
```

你應該宣告一個 `age` 變數。

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

你應該使用 `let` 來宣告 `age` 變數。

```js
assert.match(code, /let\s+age/);
```

你不應該為 `age` 變數指定值。

```js
assert.isUndefined(age);
```

你應該將 `age` 記錄到主控台。

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

你應該將 `profession` 記錄到主控台。

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
