---
id: 665732da4815b70bb083915e
title: 第 24 步
challengeType: 1
dashedName: step-24
---

# --description--

在最後幾個步驟中，你已經學會了如何操作陣列。花點時間回顧你所學到的內容。

首先宣告一個 `cities` 變數並將其初始化為包含字串 `"London"`、`"New York"` 和 `"Mumbai"` 的陣列。然後將該變數記錄到主控台。

記錄後，將 `cities` 的最後一個元素改為字串 `"Mexico City"`，然後再次記錄 `cities` 變數。

當正確完成時，你應該會在主控台看到此輸出。

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

你應該使用 `let` 來宣告一個 `cities` 變數。

```js
assert.match(code, /let\s+cities/);
```

你應該將包含字串 `"London"`、`"New York"` 和 `"Mumbai"` 的陣列指定給 `cities` 變數。

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

你應該使用 `console.log()` 將整個 `cities` 陣列記錄到主控台。

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

你應該將 `cities` 陣列的最後一個成員更新為字串 `"Mexico City"`。請記得你可以使用 `array[array.length - 1]` 來存取陣列的最後一個成員。

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

你的程式碼中應該有兩個 `console.log(cities)` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
