---
id: 665732e62eb8b30c249b5c2c
title: 第 25 步
challengeType: 1
dashedName: step-25
---

# --description--

現在你已經準備好進入下一組陣列課程。

移除你先前步驟中的所有程式碼。

# --hints--

你應該移除所有對你的 `cities` 變數的引用。

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
