---
id: 66643c9a9e0ad1c787f5fe2b
title: 步驟 56
challengeType: 1
dashedName: step-56
---

# --description--

完成快速回顧後，你應該移除你的 `addTwoNumbers` 函式、`sum` 變數和記錄述語。

# --hints--

你的程式碼中不應該有 `addTwoNumbers` 函式。

```js
assert.notMatch(code, /addTwoNumbers/);
```

你的程式碼中不應該有 `sum` 變數。

```js
assert.notMatch(code, /sum/);
```

你不應該記錄 `sum` 變數。

```js
assert.notMatch(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
function addTwoNumbers(num1, num2) {
  return num1 + num2;
}

const sum = addTwoNumbers(5, 10);
console.log(sum)

--fcc-editable-region--

const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
