---
id: 66643c93e05093c728abdbe9
title: 步驟 55
challengeType: 1
dashedName: step-55
---

# --description--

在繼續之前，花點時間回顧函式（程式）如何運作。

宣告一個名為 `addTwoNumbers` 的函式（程式）。此函式應該接受兩個引數並傳回這兩個引數的總和。

你的函式不應該使用寫死的值。寫死函式的範例可能是：

```js
function sayName(firstName, lastName) {
  return "John Doe";
}

sayName("Camper", "Cat");
```

此函式會傳回 `"John Doe"`，不論傳遞給參數 `firstName` 和 `lastName` 的引數為何，因此 `"John Doe"` 被視為寫死的值。

宣告一個 `sum` 變數，並將呼叫你的 `addTwoNumbers` 函式（程式）且以 `5` 和 `10` 作為引數的結果設值給它。將 `sum` 變數記錄到主控台。

# --hints--

你應該有一個名為 `addTwoNumbers` 的函式（程式）。

```js
assert.isFunction(addTwoNumbers);
```

你的函式（程式）`addTwoNumbers`應該有兩個參數。

```js
assert.lengthOf(addTwoNumbers, 2);
```

你的函式（程式）應該傳回兩個參數的總和。

```js
assert.strictEqual(addTwoNumbers(5,10), 15);
```

你的函式不應該傳回寫死的值。也就是說，它應該能處理任意兩個數字引數。

```js
assert.strictEqual(addTwoNumbers(3, 5), 8);
```

你應該宣告一個 `sum` 變數。

```js
assert.isDefined(sum);
```

你的 `sum` 變數應該有值 `15`。

```js
assert.strictEqual(sum, 15);
```

你應該將 `sum` 指定為呼叫 `addTwoNumbers` 函式（程式）並以 `5` 和 `10` 作為引數的回傳值。

```js
assert.match(code, /sum\s*=\s*addTwoNumbers\s*\(/);
```

你應該記錄你的 `sum` 變數。

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
