---
id: 66b205e6eacba4c4e54ea434
title: Створіть маску для е-пошти
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

У цій лабораторній роботі ви замаскуєте частину імені користувача в адресі е-пошти за допомогою зірочок. Маскування — це термін, який означає приховування або заміну конфіденційної інформації зірочками чи іншими символами.

Наприклад, якщо адреса е-пошти була `myEmail@email.com`, то замаскована адреса буде `m*****l@email.com`.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Створіть функцію з іменем `maskEmail`, яка приймає `email` як аргумент.
2. Усередині функції потрібно замаскувати `email` і додати до нього доменне ім’я. Пам’ятайте, що ви можете використовувати методи, такі як `slice`, `repeat`, `indexOf` або навіть `replace`, щоб допомогти собі.
3. Поза межами функції оголосіть змінну з іменем `email`, щоб зберегти адресу е-пошти, яку потрібно замаскувати.
4. Викличте функцію `maskEmail` із змінною `email` і виведіть результат у консоль.
5. `maskEmail("apple.pie@example.com")` має повертати `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` має повертати `"f**********p@example.com"`.
7. `maskEmail("info@test.dev")` має повертати `"i**o@test.dev"`.
8. `maskEmail("user@domain.org")` має повертати `"u**r@domain.org"`.

# --hints--

Ви повинні визначити функцію з іменем `maskEmail`.

```js
assert.isFunction(maskEmail);
```

Функція `maskEmail` має приймати рядок `email` як аргумент.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

Поза межами функції у вас має бути змінна `email`.

```js
assert.isDefined(email);
```

Ви повинні призначити дійсну адресу е-пошти змінній `email`.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` має повертати `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` має повертати `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` має повертати `"i**o@test.dev"`.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` має повертати `"u**r@domain.org"`.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

Ваш `maskEmail` має давати правильний результат.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

Ви повинні вивести у консоль результат виклику `maskEmail` з аргументом `email`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
