---
id: 615f2d4150fe0d4cbd0f2628
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Below your `h1` element, add a `p` element with the text `8 servings per container`.

# --hints--

Du solltest ein neues `p`-Element hinzufügen.

```js
assert.exists(document.querySelector('p'));
```

Dein `p`-Element sollte sich innerhalb deines `body`-Elements befinden.

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

Dein `p`-Element sollte nach deinem `h1`-Element stehen.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

Dein `p`-Element sollte den Text `8 servings per container` enthalten.

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>

</body>
</html>
--fcc-editable-region--
```

```css

```
