---
id: 615f34ecc1091b4fd5a8a484
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Within your `head` element, add a `link` element with the `rel` attribute set to `stylesheet` and the `href` attribute set to `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

This will import the `Open Sans` font family, with the font weight values `400`, `700`, and `800`.

Also add a `link` element to link your `styles.css` file.

# --hints--

Dein Code sollte zwei `link`-Elemente enthalten.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Dein`link`-Element sollte ein leeres Element sein, es sollte kein schließendes Tag besitzen `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Deine beiden `link`-Elemente sollten innerhalb des `head`-Elements liegen.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Both of your `link` elements should have the `rel` attribute set to `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Eines deiner `link`-Elemente sollte ein `href`-Attribut mit dem Wert `styles.css` enthalten.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

One of your `link` elements should have an `href` attribute set to `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
