---
id: 615f34ecc1091b4fd5a8a484
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Füge innerhalb des `head`-Elements ein `link`-Element hinzu, bei dem das `rel`-Attribut auf `stylesheet` und das `href`-Attribut auf `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800` gesetzt sind.

Auf diese Weise wird die `Open Sans`-Schriftart-Familie mit den Schriftstärke-Einstellungen `400`, `700` und `800` importiert.

Füge außerdem ein `link`-Element hinzu, um deine `styles.css`-Datei zu verknüpfen.

# --hints--

Dein Code sollte zwei `link`-Elemente enthalten.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Dein`link`-Element sollte ein leeres Element sein, es sollte kein schließendes Tag besitzen `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Deine beiden `link`-Elemente sollten innerhalb des `head`-Elements liegen.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Beide `link`-Elemente benötigen ein `rel`-Attribut mit dem Wert `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Eines deiner `link`-Elemente sollte ein `href`-Attribut mit dem Wert `styles.css` enthalten.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

Eines der `link`-Elemente sollte ein `href`-Attribut mit `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800` als Wert enthalten.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
