---
id: 587d7fb7367417b2b2512c0a
title: 使用 model.create() 创建多条记录
challengeType: 2
forumTopicId: 301537
dashedName: create-many-records-with-model-create
---

# --description--

有时你需要创建许多模型实例，例如在用初始数据填充数据库时。`Model.create()` 以类似 `[{name: 'John', ...}, {...}, ...]` 的对象数组作为第一个参数，并将它们全部储存在数据库中。

# --instructions--

修改 `createManyPeople` 方法，使用 `arrayOfPeople` 作为 `Model.create()` 的参数来创建多个 people 实例。

**注意：** 你可以使用在上一个挑战中创建的 model 来完成当前挑战。

# --hints--

应当成功地一次性创建多条数据

```js
  const response = await fetch(code + '/_api/create-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'John', age: 24, favoriteFoods: ['pizza', 'salad'] },
      { name: 'Mary', age: 21, favoriteFoods: ['onions', 'chicken'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an array');
  assert.equal(
    data.length,
    2,
    'the response does not contain the expected number of items'
  );
  assert.equal(data[0].name, 'John', 'The first item is not correct');
  assert.equal(
    data[0].__v,
    0,
    'The first item should be not previously edited'
  );
  assert.equal(data[1].name, 'Mary', 'The second item is not correct');
  assert.equal(
    data[1].__v,
    0,
    'The second item should be not previously edited'
  );
```

