---
id: 587d7fb8367417b2b2512c0f
title: 在 document 中执行新的更新方式——使用 model.findOneAndUpdate()
challengeType: 2
forumTopicId: 301542
dashedName: perform-new-updates-on-a-document-using-model-findoneandupdate
---

# --description--

Mongoose 的最新版本有简化文档更新的方法。一些更高级的特色（例如 pre/post 钩子、验证）在这种方法下表现不同，因此经典方法在许多情况下仍然有用。`findByIdAndUpdate()` 可用于按 id 查找时。

# --instructions--

修改 `findAndUpdate` 函数，通过 `Name` 查询人，并将查到的人的年龄设为 `20` 岁， 将函数参数 `personName` 作为查询关键字。

**提示：** 你需要返回更新后的 document。 你可以把 `findOneAndUpdate()` 的第三个参数设置为 `{ new: true }` 。 默认情况下，这个方法会返回修改前的数据。

# --hints--

应成功地使用 findOneAndUpdate 更新数据

```js
  const response = await fetch(code + '/_api/find-one-update', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Dorian Gray',
      age: 35,
      favoriteFoods: ['unknown']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Dorian Gray', 'item.name is not what expected');
  assert.equal(data.age, 20, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['unknown'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(
    data.__v,
    0,
    'findOneAndUpdate does not increment version by design!'
  );
```

