---
id: 669e2f60e83c011754f711f9
title: Crea una página de agencia de viajes
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

Cumple con las historias de usuario a continuación y haz que todas las pruebas pasen para completar el laboratorio.

**Instrucciones:**

1. Debe tener una declaración `DOCTYPE`.
1. Debe tener un elemento `html` con `lang` establecido en `en`.
1. Debe tener un elemento `head` que contenga un elemento vacío `meta` con `charset` establecido a `utf-8` y un `title` con el texto `Travel Agency Page`.
1. Debería tener una etiqueta `meta` en su elemento `head` que contenga una descripción corta de su sitio web para SEO.
1. Debe tener un elemento `h1` para presentar los destinos de viaje.
1. Debe tener un párrafo debajo del elemento `h1` que introduzca las oportunidades de viaje.
1. Debe tener un elemento `h2` con el texto `Packages`.
1. Debe tener un elemento `p` que introduzca brevemente los diversos paquetes.
1. Debe tener un elemento de lista desordenada con dos elementos de lista. Los dos elementos de la lista deben tener el texto `Group Travels` y `Private Tours`, respectivamente. El texto de cada elemento de la lista debe estar encerrado en un elemento de anclaje.
1. Debe tener un elemento `h2` con el texto `Top Itineraries`.
1. Debe tener al menos tres elementos `figure`, cada uno con un elemento de anclaje y un elemento `figcaption`.
1. Los tres elementos de anclaje deben tener un elemento `img` con un atributo `alt` apropiado y un atributo `src` configurado a una imagen válida como su contenido. Puede usar `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` y `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` si desea.
1. Todos sus cinco elementos de anclaje deben tener un atributo `href` con el valor de `https://www.freecodecamp.org/learn` y un atributo `target` con el valor de `_blank`.

# --hints--

Su página de agencia de viajes debería tener una declaración `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Debería tener un elemento `html` con `lang` configurado a `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

Debería tener un elemento `head` dentro del elemento `html`.

```js
assert.match(
  code,
  /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Debería tener dos elementos `meta` dentro de su elemento `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*(<\s*meta[\s\S]*>[\s\S]*){2,}<\/\s*head\s*>/i);
```

Un elemento `meta` debería tener un atributo `name` con el valor de `description` y un atributo `content` no vacío.

 ```js
 assert.match(code,
  /<\s*meta\s+[^>]*name\s*=\s*["']description["'][^>]*content\s*=\s*["'][^"']+?["'][^>]*>/i,
);
```

Un elemento `meta` debería tener su atributo de `charset` configurado a `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Debería tener un elemento `title` dentro de su elemento `head`.

```js
assert.match(
  code,
  /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i
);
```

Su elemento `title` debería tener el nombre de su agencia de viajes.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Debería tener un elemento `body` dentro de su elemento `html`.

```js
assert.match(
  code,
  /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Debería tener un elemento `h1` para presentar sus destinos de viaje.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Solo Debes tener un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Debería tener un elemento `p` debajo del elemento `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Su primer elemento `p` debería presentar las oportunidades de viaje.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Su primer elemento `h2` debería tener el texto `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Packages');
```

Debería tener un elemento `p` debajo de su primer elemento `h2`.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Su segundo elemento `p` debería presentar brevemente los varios paquetes.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Debería tener un elemento de lista no ordenada debajo de su segundo elemento `p`.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

Debería tener dos elementos en su lista no ordenada.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Ambos elementos de lista deberían contener un elemento ancla.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

El elemento ancla de su primer elemento de lista debería envolver el texto `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText, 'Group Travels');
```

El elemento ancla de su segundo elemento de lista debería envolver el texto `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText, 'Private Tours');
```

Debería tener un elemento `h2` después de su lista no ordenada.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Su segundo elemento `h2` debería tener el texto `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Top Itineraries');
```

Debería tener al menos tres elementos `figure`.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Cada elemento `figure` debería contener un elemento ancla como su primer hijo.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Cada elemento `figure` debería contener un elemento `figcaption` como su segundo hijo.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Cada `figcaption` debería contener algo de texto.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Cada uno de los elementos `a` que son hijos de sus elementos `figure` debería contener una imagen.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Cada elemento `img` debería tener un atributo `src` válido.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Cada elemento `img` debería tener un atributo `alt` con un valor apropiado.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Cada elemento `a` debería tener un atributo `href` con el valor de `https://www.freecodecamp.org/learn`. No olvide los enlaces en los elementos de la lista.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Cada elemento `a` debería tener un atributo `target` con el valor de `_blank`. No olvide los enlaces en los elementos de la lista.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
