---
id: 642def66e6a60432c9a0371e
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

将 `label` 元素的 `className` 设为`"label"`，并将 `textContent` 设置为 `name` 参数。

# --hints--

你应该访问 `label` 元素的 `className` 属性。

```js
assert.match(code, /label\.className/);
```

你应该将 `className` 属性设置为字符串 `"label"`。

```js
assert.match(code, /label\.className\s*=\s*('|"|`)label\1/);
```

你应该访问 `label` 元素的 `textContent` 属性。

```js
assert.match(code, /label\.textContent/);
```

你应该将 `textContent` 属性设置为 `name` 的值。

```js
assert.match(code, /label\.textContent\s*=\s*name/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");

  }
}
--fcc-editable-region--
```
