---
id: 643475e13dc727231acd0f72
title: Step 17
challengeType: 0
dashedName: step-25
---

# --description--

现在在嵌套的 `.forEach()` 调用中，声明一个 `input` 变量。 使用 `document` 对象的 `.createElement()` 方法创建 `input` 元素。 将 `type` 属性设置为 `"text"`，将 `id` 属性设置为 `letter + number`。

# --hints--

你应该声明一个 `input` 变量。

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

你应该使用 `const` 来声明你的 `input` 变量。

```js
assert.match(code, /const\s+input/)
```

你应该调用 `document` 对象的 `.createElement()` 方法。

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

你应该将字符串 `"input"` 传递给 `.createElement()` 方法。

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

你应该将新的 `input` 元素分配给你的 `input` 变量。

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

你应该访问 `input` 元素的 `type` 属性。

```js
assert.match(code, /input\.type/);
```

你应该将 `input` 元素的 `type` 属性设置为 `"text"`。

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

你应该访问 `input` 元素的 `id` 属性。

```js
assert.match(code, /input\.id/);
```

你应该将 `input` 元素的 `id` 属性设置为 `letter + number`。

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
