---
id: 6434750c53db16218f41e6e1
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

请记住，`range()` 返回一个数组，因此你可以将数组方法直接链接到函数调用。

使用 `1` 和 `99` 作为参数调用 `range()`，并链接 `.forEach()` 方法。 向 `.forEach()` 方法传递一个以 `number` 作为参数的空回调。

# --hints--

你应该调用 `range()` 函数。

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

你应该将 `1` 作为第一个参数传递给 `range()` 调用。

```js
assert.match(code, /range\(\s*1/);
```

你应该将 `99` 作为第二个参数传递给 `range()` 调用。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

你应该将 `.forEach()` 方法链接到你的 `range()` 调用。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

你应该使用箭头语法将回调函数传递给 `.forEach()`。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

你的回调函数应该以 `number` 作为唯一参数。

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
