---
id: 6437124c4c03dd4c8fb35d56
title: 步骤21
challengeType: 0
dashedName: step-21
---

# --description--

声明一个 `isEven` 函数，该函数接受 `num` 参数，如果数字为偶数则返回 `true`，否则返回 `false`。 使用模运算符 `%` 来确定数字是偶数还是奇数。

# --hints--

你应该声明一个 `isEven` 变量。

```js
assert.match(code, /(?:let|const|var)\s+isEven/);
```

你应该使用 `const` 来声明你的 `isEven` 变量。

```js
assert.match(code, /const\s+isEven/);
```

你的 `isEven` 变量应该是一个函数。

```js
assert.isFunction(isEven);
```

你的 `isEven` 函数应该使用箭头语法。

```js
assert.match(code, /const\s+isEven\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

你的 `isEven` 函数应该有一个 `num` 参数。

```js
assert.match(code, /const\s+isEven\s*=\s*(\(\s*num\s*\)|num)\s*=>/);
```

你的 `isEven` 函数应该使用模运算符 `%`。

```js
assert.match(isEven.toString(),  /%/);
```

对于偶数，你的 `isEven` 函数应该返回 `true`。

```js
assert.isTrue(isEven(2));
assert.isTrue(isEven(1000));
assert.isTrue(isEven(42));
```

对于奇数，`isEven` 函数应该返回 `false`。

```js
assert.isFalse(isEven(1));
assert.isFalse(isEven(333));
assert.isFalse(isEven(777777777));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
const sum = nums => nums.reduce((acc, el) => acc + el, 0);

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
