---
id: 64496df724dd3716a71fe971
title: 步骤 26
challengeType: 0
dashedName: 步骤-26
---

# --description--

对象属性由键/值对组成。 声明对象文字时，可以使用简写属性名称。 当使用简写属性名称语法时，变量的名称将成为属性键，其值将成为属性值。

下面的示例声明了一个 `user` 对象，其属性为 `userId`、`firstName` 和 `loggedIn`。

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

为了跟踪电子表格的所有功能，请声明一个 `spreadsheetFunctions` 对象。 使用简写符号语法，将 `sum`、`average` 和 `median` 设置为 `spreadsheetFunctions` 对象上的属性。

# --hints--

你应该声明一个 `spreadsheetFunctions` 变量。

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

你应该使用 `const` 来声明你的 `spreadsheetFunctions` 变量。

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

你的 `spreadsheetFunctions` 变量应该是一个对象。

```js
assert.isObject(spreadsheetFunctions);
```

你的 `spreadsheetFunctions` 对象应该具有 `sum` 属性。

```js
assert.property(spreadsheetFunctions, "sum");
```

你的 `sum` 属性应该是你的 `sum` 函数。

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

你的 `spreadsheetFunctions` 对象应该具有 `average` 属性。

```js
assert.property(spreadsheetFunctions, "average");
```

你的 `average` 属性应该是你的 `average` 函数。

```js
assert.equal(spreadsheetFunctions?.average, average);
```

你的 `spreadsheetFunctions` 对象应该具有 `median` 属性。

```js
assert.property(spreadsheetFunctions, "median");
```

你的 `median` 属性应该是你的 `median` 函数。

```js
assert.equal(spreadsheetFunctions?.median, median);
```

你应该使用解构语法来分配你的属性。

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
