---
id: 641d9a19bff38d34d5a5edb8
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 Fetch API를 사용해 외부 소스에서 데이터를 불러와 freeCodeCamp 저자들로 가득한 페이지에 표시하는 방법을 배우세요.

이 워크숍에 필요한 모든 HTML과 CSS가 제공되어 있습니다. 두 파일을 살펴보며 익숙해지세요.

준비가 되면 `#author-container`와 `#load-more-btn` 요소에 접근하는 것부터 시작하세요. 이때 `getElementById` 또는 `querySelector` 메서드를 사용할 수 있다는 점을 기억하세요.

그다음 이 요소들을 각각 `authorContainer`와 `loadMoreBtn`라는 변수에 할당하세요.

# --hints--

`authorContainer`라는 변수가 있어야 합니다.

```js
assert.isDefined(authorContainer);
```

`#author-container` 요소를 `authorContainer` 변수에 할당해야 합니다.

```js
assert.strictEqual(authorContainer, document.querySelector('#author-container'));
```

`loadMoreBtn`라는 변수가 있어야 합니다.

```js
assert.isDefined(loadMoreBtn);
```

`#load-more-btn` 요소를 `loadMoreBtn` 변수에 할당해야 합니다.

```js
assert.strictEqual(loadMoreBtn, document.querySelector('#load-more-btn'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
