---
id: 641da5abaac81844a54adb03
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

이제 페이지에 저자를 표시할 함수를 만들 차례입니다.

`displayAuthors` 매개변수를 받는 `authors`라는 함수를 먼저 만드세요. 그런 다음 `authors` 배열을 순회하며 각 저자마다 카드 하나를 `authorContainer` 요소에 추가해야 합니다.

카드는 다음 구조를 가져야 합니다:

```html
<div id="${index}" class="user-card">
  <h2 class="author-name">${author}</h2>
</div>
```

**팁:**

- `authors` 배열을 순회하는 데 어떤 종류의 반복문도 사용할 수 있습니다. 예: `for`, `forEach`, `map` 등.
- 반복문 안에서 `innerHTML`에 카드를 추가하려면 `authorContainer` 속성을 사용해야 합니다. 마크업을 추가할 때는 템플릿 리터럴 구문도 사용해야 합니다.

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

`displayAuthors`은 함수여야 합니다.

```js
assert.isFunction(displayAuthors)
```

`displayAuthors` 함수는 `authors` 매개변수를 받아야 합니다.

```js
const parameters = __helpers.getFunctionParams(displayAuthors.toString());
assert.deepEqual(parameters[0].name,"authors"); 
```

`displayAuthors` 함수가 저자 배열과 함께 호출되면 각 저자마다 카드를 `#authorContainer` 요소에 추가해야 합니다.

```js
const exampleAuthors = [
    [{
        "author": "Quincy Larson",
    }, {
        "author": "Jessica Wilkins",
    }, ],
    [{
        "author": "Naomi Carrigan",
    }, {
        "author": "Beau Carnes",
    }, {
        "author": "David Clinton",
    }],
];

for (const authorSet of exampleAuthors) {
    displayAuthors(authorSet);

    const authorCards = document.querySelectorAll("#author-container div");
    const authorTitles = document.querySelectorAll("#author-container div h2");
    assert.deepEqual(authorCards.length, authorTitles.length);
    for (let i = 0; i < authorCards.length; i++) {
        assert.deepEqual(authorCards[i].className, "user-card");
        assert.deepEqual(authorCards[i].id, i.toString());
        assert.deepEqual(authorTitles[i].className, "author-name");
        assert.deepEqual(authorTitles[i].innerText, authorSet[i].author);
    }
    authorContainer.innerHTML = "";
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });

--fcc-editable-region--

--fcc-editable-region--
```
