---
id: 681cb1afdab50c87ddb2e516
title: "Challenge 5: Jbelmud Text"
challengeType: 29
dashedName: challenge-5
---

# --description--

Ukipata mfuatano wa herufi, rudisha toleo lililochanganywa la mfuatano huo ambapo kila neno hubadilishwa kwa kutumia vizingiti vifuatavyo:

- Herufi ya kwanza na ya mwisho ya maneno hubaki mahali pake
- Herufi zote kati ya herufi ya kwanza na ya mwisho hupangwa kwa mpangilio wa alfabeti.
- Mfuatano wa herufi wa ingizo hautakuwa na alama za uandishi, na utakuwa mzima kwa herufi ndogo.

# --hints--

`jbelmu("hello world")` should return `"hello wlord"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("hello world"), "hello wlord")`)
}})
```

`jbelmu("i love jumbled text")` should return `"i love jbelmud text"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("i love jumbled text"), "i love jbelmud text")`)
}})
```

`jbelmu("freecodecamp is my favorite place to learn to code")` should return `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe")`)
}})
```

`jbelmu("the quick brown fox jumps over the lazy dog")` should return `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def jbelmu(text):

    return text
```

# --solutions--

```py
def jbelmu(text):
    words = text.split()
    jumbled = []

    for word in words:
        if len(word) <= 3:
            jumbled.append(word)
        else:
            first = word[0]
            last = word[-1]
            middle = ''.join(sorted(word[1:-1]))
            jumbled.append(first + middle + last)
    
    return ' '.join(jumbled)
```
