---
id: 6821ebf8237de8297eaee798
title: "Challenge 22: Tribonacci Sequence"
challengeType: 29
dashedName: challenge-22
---

# --description--

Mfululizo wa Tribonacci ni mfululizo wa nambari ambapo kila nambari ni jumla ya nambari tatu zilizotangulia. Unapoanza na `0`, `0` na `1`, nambari kumi za kwanza katika mfululizo ni `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24`, `44`.

Ukitoa safu ya data yenye nambari tatu za kwanza za mfululizo wa Tribonacci, na nambari kamili inayowakilisha urefu wa mfululizo, rudisha safu ya data yenye mfululizo wa urefu uliotolewa.

- Kitendakazi chako kinapaswa kushughulikia mfululizo wa urefu wowote mkubwa au sawa na sifuri.
- Ikiwa urefu ni sifuri, rudisha safu ya data tupu.
- Kumbuka kuwa nambari za kuanzia ni sehemu ya mfululizo.

# --hints--

`tribonacci_sequence([0, 0, 1], 20)` should return `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513])`)
}})
```

`tribonacci_sequence([21, 32, 43], 1)` should return `[21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([21, 32, 43], 1), [21])`)
}})
```

`tribonacci_sequence([0, 0, 1], 0)` should return `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([0, 0, 1], 0), [])`)
}})
```

`tribonacci_sequence([10, 20, 30], 2)` should return `[10, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 2), [10, 20])`)
}})
```

`tribonacci_sequence([10, 20, 30], 3)` should return `[10, 20, 30]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([10, 20, 30], 3), [10, 20, 30])`)
}})
```

`tribonacci_sequence([123, 456, 789], 8)` should return `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tribonacci_sequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134])`)
}})
```

# --seed--

## --seed-contents--

```py
def tribonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def tribonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return [start_sequence[0], start_sequence[1]]
    if length == 3:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2] + sequence[-3]
        sequence.append(next_value)

    return sequence
```
