---
id: 6821ebce237de8297eaee790
title: "Challenge 14: Character Battle"
challengeType: 29
dashedName: challenge-14
---

# --description--

Kutolewa mfuatano wa herufi miwili unaowakilisha jeshi lako na jeshi la mpinzani, kila herufi kutoka jeshi lako inapigana na herufi iliyopo katika nafasi sawa kutoka jeshi la mpinzani kwa kutumia sheria zifuatazo:

- Herufi `a-z` zina nguvu ya `1-26`, mtawaliwa.
- Herufi `A-Z` zina nguvu ya `27-52`, mtawaliwa.
- Nambari `0-9` zina nguvu sawa na thamani yao ya uso.
- Herufi zote nyingine zina thamani ya sifuri.
- Kila herufi inaweza kupigana vita moja tu.

Kwa kila vita, herufi yenye nguvu zaidi hushinda. Jeshi lenye ushindi mwingi zaidi hushinda vita. Rudisha thamani zifuatazo:

- `"Opponent retreated"` ikiwa jeshi lako lina herufi nyingi zaidi kuliko jeshi la mpinzani.
- `"We retreated"` ikiwa jeshi la mpinzani lina herufi nyingi zaidi kuliko lako.
- `"We won"` ikiwa jeshi lako limeshinda vita nyingi zaidi.
- `"We lost"` ikiwa jeshi la mpinzani limeshinda vita nyingi zaidi.
- `"It was a tie"` ikiwa majeshi yote mawili yameshinda vita idadi sawa.

# --hints--

`battle("Hello", "World")` should return `"We lost"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` should return `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` should return `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` should return `"Opponent retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` should return `"We retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` should return `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` should return `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
