---
id: 68adce01c0e1144d0a90295e
title: "Challenge 29: Acronym Builder"
challengeType: 29
dashedName: challenge-29
---

# --description--

Ikiwa umepewa mfuatano wa herufi unao kuwa na neno moja au zaidi, rudisha herufi za awali za maneno hayo kwa kutumia masharti yafuatayo:

- Herufi za awali zinapaswa kuwa herufi kubwa, isipokuwa kama imetajwa vinginevyo.
- Herufi za awali za maneno haya yafuatayo zisitiliwe isipokuwa kama ni neno la kwanza katika mfuatano uliotolewa: `a`, `for`, `an`, `and`, `by`, na `of`.
- Herufi za awali zinapaswa kurudishwa kwa mpangilio zilivyopewa.
- Herufi za awali hazipaswi kuwa na nafasi yoyote kati yao.

# --hints--

`build_acronym("Search Engine Optimization")` should return `"SEO"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Search Engine Optimization"), "SEO")`)
}})
```

`build_acronym("Frequently Asked Questions")` should return `"FAQ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Frequently Asked Questions"), "FAQ")`)
}})
```

`build_acronym("National Aeronautics and Space Administration")` should return `"NASA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("National Aeronautics and Space Administration"), "NASA")`)
}})
```

`build_acronym("Federal Bureau of Investigation")` should return `"FBI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Federal Bureau of Investigation"), "FBI")`)
}})
```

`build_acronym("For your information")` should return `"FYI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("For your information"), "FYI")`)
}})
```

`build_acronym("By the way")` should return `"BTW"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("By the way"), "BTW")`)
}})
```

`build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` should return `"AUHWPOTIMSH"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH")`)
}})
```

# --seed--

## --seed-contents--

```py
def build_acronym(s):

    return s
```

# --solutions--

```py
def build_acronym(s):
    small_words = {"a", "for", "an", "and", "by", "of"}
    words = s.split()
    acronym = ""

    for i, word in enumerate(words):
        if i == 0 or word.lower() not in small_words:
            acronym += word[0].upper()

    return acronym
```
