---
id: 68adce01c0e1144d0a902956
title: "Challenge 25: Vowel Repeater"
challengeType: 29
dashedName: challenge-25
---

# --description--

Ukitoa mfuatano wa herufi, rudisha toleo jipya la mfuatano huo ambapo kila vokali inarudiwa mara moja zaidi kuliko vokali iliyotangulia uliyokutana nayo. Kwa mfano, vokali ya kwanza katika sentensi haibadiliki. Vokali ya pili inaonekana mara mbili mfululizo. Vokali ya tatu inaonekana mara tatu mfululizo, na kadhalika.

- Herufi `a`, `e`, `i`, `o`, na `u`, iwe ni herufi kubwa au herufi ndogo, zinachukuliwa kuwa vokali.
- Vokali ya asili inapaswa kuhifadhi aina yake ya herufi.
- Vokali zinazorudiwa zinapaswa kuwa herufi ndogo.
- Herufi zote zisizo za vokali zinapaswa kuhifadhi aina yao ya herufi kama zilivyo asili.

# --hints--

`repeat_vowels("hello world")` should return `"helloo wooorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` should return `"freeeCooodeeeeCaaaaamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` should return `"AEeIiiOoooUuuuu"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` should return `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
