---
id: 68b06e589bf2273243814777
title: "Zoezi la 39: Jaza Tanki"
challengeType: 29
dashedName: challenge-39
---

# --description--

Kutolewa ukubwa wa tanki la mafuta, kiwango cha sasa cha mafuta, na bei kwa galoni, rudisha gharama ya kujaza tanki lote.

- Hoja ya kwanza ni uwezo wa jumla wa tanki kwa galoni.
- Hoja ya pili ni kiasi cha sasa cha mafuta katika tanki kwa galoni.
- Hoja ya tatu ni gharama ya galoni moja ya mafuta.
- Thamani inayorudishwa inapaswa kuzungushwa hadi sehemu mbili za desimali kwa muundo: `"$d.dd"`.

# --hints--

`cost_to_fill(20, 0, 4.00)` inapaswa kurudisha `"$80.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` inapaswa kurudisha `"$17.50"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` inapaswa kurudisha `"$29.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` inapaswa kurudisha `"$0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` inapaswa kurudisha `"$21.89"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
