---
id: 68b06e589bf227324381476f
title: "Challenge 35: Word Frequency"
challengeType: 29
dashedName: challenge-35
---

# --description--

Ukichukulia aya, rudisha safu ya data yenye maneno matatu yanayotokea mara nyingi zaidi.

- Maneno katika aya yatatengwa na nafasi.
- Puuza tofauti za herufi ndogo na kubwa katika aya iliyotolewa. Kwa mfano, chukulia `Hello` na `hello` kama neno moja.
- Puuza alama za uandishi katika aya iliyotolewa. Alama za uandishi ni kama vile koma (`,`), nukta (`.`), na alama za mshangao (`!`).
- Safu ya data inayorudishwa inapaswa kuwa na maneno yote yenye herufi ndogo.
- Safu ya data inayorudishwa inapaswa kuwa katika mpangilio wa kushuka kwa mara ya kutokea, neno linalotokea mara nyingi zaidi likiwa la kwanza.

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` should return `["coding", "python", "in"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` should return `["i", "like", "coding"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` should return `["debug", "test", "deploy"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
