---
id: 68b1f72371a5ac895ac70a06
title: "Challenge 42: Video Storage"
challengeType: 29
dashedName: challenge-42
---

# --description--

Ukizingatia ukubwa wa video, kitengo cha ukubwa wa video, uwezo wa diski kuu, na kitengo cha diski kuu, rudisha idadi ya video ambazo diski kuu inaweza kuhifadhi kwa kutumia vizingiti vifuatavyo:

- Kitengo cha ukubwa wa video kinaweza kuwa bytes (`"B"`), kilobytes (`"KB"`), megabytes (`"MB"`), au gigabytes (`"GB"`).
- Ikiwa hakutolewa mojawapo ya vitengo vya video vilivyotajwa hapo juu, rudisha `"Invalid video unit"`.
- Kitengo cha uwezo wa diski kuu kinaweza kuwa gigabytes (`"GB"`) au terabytes (`"TB"`).
- Ikiwa hakutolewa mojawapo ya vitengo vya diski kuu vilivyotajwa hapo juu, rudisha `"Invalid drive unit"`.
- Rudisha idadi ya video kamili ambazo diski inaweza kuhifadhi.
- Tumia mabadiliko yafuatayo:

| Kitengo | Sawia na |
|:-------:|:--------:|
| 1 B     |   1 B    |
| 1 KB    |  1000 B  |
| 1 MB    |  1000 KB |
| 1 GB    |  1000 MB |
| 1 TB    |  1000 GB |

Kwa mfano, ukipokea `500`, `"MB"`, `100`, na `"GB"` kama hoja, amua ni video ngapi za 500 MB zinaweza kuhifadhiwa kwenye diski kuu ya 100 GB.

# --hints--

`number_of_videos(500, "MB", 100, "GB")` should return `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` should return `"Invalid video unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` should return `"Invalid drive unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` should return `4000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` should return `1466`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```
