---
id: 68b1f72371a5ac895ac70a04
title: "Challenge 41: File Storage"
challengeType: 29
dashedName: challenge-41
---

# --description--

Ukizingatia ukubwa wa faili, kitengo cha ukubwa wa faili, na uwezo wa diski kuu kwa gigabaiti (GB), rudisha idadi ya mafaili ambayo diski kuu inaweza kuhifadhi kwa kutumia vizingiti vifuatavyo:

- Kitengo cha ukubwa wa faili kinaweza kuwa bait (`"B"`), kilobaiti (`"KB"`), au megabaiti (`"MB"`).
- Rudisha idadi ya mafaili kamili ambayo diski inaweza kuingiza.
- Tumia mabadiliko yafuatayo:

| Kitengo | Sawia na |
|:-------:|:--------:|
| 1 B     |   1 B    |
| 1 KB    |  1000 B  |
| 1 MB    |  1000 KB |
| 1 GB    |  1000 MB |

Kwa mfano, ukipata `500`, `"KB"`, na `1` kama hoja, amua ni mafaili mangapi ya 500 KB yanaweza kuhifadhiwa kwenye diski kuu ya 1 GB.

# --hints--

`number_of_files(500, "KB", 1)` should return `2000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(500, "KB", 1), 2000)`)
}})
```

`number_of_files(50000, "B", 1)` should return `20000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(50000, "B", 1), 20000)`)
}})
```

`number_of_files(5, "MB", 1)` should return `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(5, "MB", 1), 200)`)
}})
```

`number_of_files(4096, "B", 1.5)` should return `366210`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4096, "B", 1.5), 366210)`)
}})
```

`number_of_files(220.5, "KB", 100)` should return `453514`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(220.5, "KB", 100), 453514)`)
}})
```

`number_of_files(4.5, "MB", 750)` should return `166666`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4.5, "MB", 750), 166666)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_files(file_size, file_unit, drive_size_gb):

    return file_size
```

# --solutions--

```py
def number_of_files(file_size, file_unit, drive_size_gb):
    drive_size_bytes = drive_size_gb * 1000 * 1000 * 1000

    if file_unit == "B":
        file_size_bytes = file_size
    elif file_unit == "KB":
        file_size_bytes = file_size * 1000
    else:
        file_size_bytes = file_size * 1000 * 1000

    return int(drive_size_bytes // file_size_bytes)
```
