---
id: 68b7cadffed0e75a517da673
title: "Challenge 52: Binary to Decimal"
challengeType: 29
dashedName: challenge-52
---

# --description--

Ikiwa umepewa mfuatano wa herufi unaowakilisha nambari ya binary, rudisha sawa yake katika mfumo wa desimali kama nambari.

Nambari ya binary hutumia tu tarakimu `0` na `1` kuwakilisha nambari yoyote. Ili kubadilisha binary kuwa desimali, zidisha kila tarakimu kwa nguvu ya `2` kisha ziongeze pamoja. Anza kwa kuzidisha tarakimu ya kulia kabisa kwa `2^0`, tarakimu inayofuata kushoto kwa `2^1`, na kadhalika. Mara zote tarakimu zote zitakapozidishwa kwa nguvu ya `2`, ongeza matokeo hayo pamoja.

Kwa mfano, nambari ya binary `101` ni sawa na `5` katika mfumo wa desimali kwa sababu:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` should return `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```
