---
id: 68b7687dded630607aceccb1
title: "Challenge 48: Spam Detector"
challengeType: 29
dashedName: challenge-48
---

# --description--

Ikiwa nambari ya simu iko katika muundo wa `"+A (BBB) CCC-DDDD"`, ambapo kila herufi inawakilisha nambari kama ifuatavyo:

- `A` inawakilisha msimbo wa nchi na inaweza kuwa na nambari yoyote ya tarakimu.
- `BBB` inawakilisha msimbo wa eneo na daima itakuwa na tarakimu tatu.
- `CCC` na `DDDD` zinawakilisha nambari ya ndani na daima zitakuwa na tarakimu tatu na nne mtawalia.

Amua kama ni nambari ya spam kulingana na vigezo vifuatavyo:

- Msimbo wa nchi ni wa tarakimu zaidi ya 2 au hauanza na sifuri (`0`).
- Msimbo wa eneo ni zaidi ya 900 au chini ya 200.
- Jumla ya tarakimu tatu za kwanza za nambari ya ndani inaonekana ndani ya tarakimu nne za mwisho za nambari ya ndani.
- Nambari ina tarakimu ile ile mara nne au zaidi mfululizo (bila kuzingatia herufi za muundo).

# --hints--

`is_spam("+0 (200) 234-0182")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (200) 234-0182"), False)`)
}})
```

`is_spam("+091 (555) 309-1922")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+091 (555) 309-1922"), True)`)
}})
```

`is_spam("+1 (555) 435-4792")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+1 (555) 435-4792"), True)`)
}})
```

`is_spam("+0 (955) 234-4364")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (955) 234-4364"), True)`)
}})
```

`is_spam("+0 (155) 131-6943")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (155) 131-6943"), True)`)
}})
```

`is_spam("+0 (555) 135-0192")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 135-0192"), True)`)
}})
```

`is_spam("+0 (555) 564-1987")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 564-1987"), True)`)
}})
```

`is_spam("+00 (555) 234-0182")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+00 (555) 234-0182"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_spam(number):

    return number
```

# --solutions--

```py
import re
def is_spam(number):
    digits = re.sub(r"\D", "", number)
    match = re.match(r"^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$", number)
    country_code, area_code, ccc, dddd = match.groups()

    if len(country_code) > 2 or not country_code.startswith("0"):
        return True

    area_num = int(area_code)
    if area_num > 900 or area_num < 200:
        return True

    sum_ccc = sum(int(d) for d in ccc)
    if str(sum_ccc) in dddd:
        return True

    if re.search(r"(\d)\1\1\1", digits):
        return True

    return False
```
