---
id: 68c1a929005bf54d342aa8d4
title: "Challenge 57: Space Week Day 3: Phone Home"
challengeType: 29
dashedName: challenge-57
---

# --description--

Kwa siku ya tatu ya Wiki ya Anga, unapewa safu ya data ya nambari zinazowakilisha umbali (kwa kilomita) kati yako, satelaiti, na sayari yako ya nyumbani katika njia ya data ya mawasiliano. Tambua ni muda gani ujumbe utakaopelekwa kupitia njia hiyo utachukua kufikia sayari lengwa kwa kutumia vizingiti vifuatavyo:

- Thamani ya kwanza katika safu ya data ni umbali kutoka mahali ulipo hadi satelaiti ya kwanza.
- Kila thamani inayofuata, isipokuwa ya mwisho, ni umbali hadi satelaiti inayofuata.
- Thamani ya mwisho katika safu ya data ni umbali kutoka satelaiti ya mwisho hadi sayari yako ya nyumbani.
- Ujumbe husafiri kwa kasi ya 300,000 km/s.
- Kila satelaiti ambayo ujumbe unapitisha huongeza ucheleweshaji wa usambazaji wa sekunde 0.5.
- Rudisha nambari iliyozungushwa hadi sehemu 4 za desimali, ukiondoa sifuri za mwisho.

# --hints--

`send_message([300000, 300000])` should return `2.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` should return `3.0627`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` should return `364.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` should return `1674.3333`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` should return `2.4086`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` should return `21.1597`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
