---
id: 68c1a929005bf54d342aa8d6
title: "Challenge 59: Space Week Day 5: Goldilocks Zone"
challengeType: 29
dashedName: challenge-59
---

# --description--

Kwa siku ya tano ya Wiki ya Anga, utahesabu "eneo la Goldilocks" la nyota - eneo linalozunguka nyota ambapo hali ni "sawa kabisa" kwa maji ya kioevu kuwepo.

Ukizingatia uzito wa nyota, rudisha safu ya data yenye umbali wa kuanzia na wa mwisho wa Eneo lake la Goldilocks kwa Vipimo vya Anga.

Ili kuhesabu Eneo la Goldilocks:

1. Tafuta mwangaza wa nyota kwa kuinua uzito wake kwa nguvu ya 3.5.
2. Mwanzo wa eneo ni mara 0.95 ya mzizi wa mraba wa mwangaza wake.
3. Mwisho wa eneo ni mara 1.37 ya mzizi wa mraba wa mwangaza wake.

- Rudisha umbali uliopunguzwa hadi sehemu mbili za desimali.

Kwa mfano, ukipata `1` kama uzito, rudisha `[0.95, 1.37]`.

# --hints--

`goldilocks_zone(1)` should return `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` should return `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` should return `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` should return `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` should return `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
