---
id: 68c497f3aaefc9fd9f1b0e26
title: "Challenge 62: Hex to Decimal"
challengeType: 29
dashedName: challenge-62
---

# --description--

Imetolewa mfuatano wa herufi unaowakilisha nambari ya hexadecimal (msingi 16), rudisha thamani yake ya desimali (msingi 10) kama nambari kamili.

Hexadecimal ni mfumo wa nambari unaotumia tarakimu 16:

- `0-9` wakilisha thamani za `0` hadi `9`.
- `A-F` wakilisha thamani za `10` hadi `15`.

Hapa kuna jedwali la mabadiliko ya sehemu:

| Hexadecimal | Desimali |
|:-----------:|:--------:|
| 0           | 0        |
| 1           | 1        |
| ...         | ...      |
| 9           | 9        |
| A           | 10       |
| ...         | ...      |
| F           | 15       |
| 10          | 16       |
| ...         | ...      |
| 9F          | 159      |
| A0          | 160      |
| ...         | ...      |
| FF          | 255      |
| 100         | 256      |

- Mfuatano wa herufi utakuwa na herufi `0–9` na `A–F` tu.

# --hints--

`hex_to_decimal("A")` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A"), 10)`)
}})
```

`hex_to_decimal("15")` should return `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("15"), 21)`)
}})
```

`hex_to_decimal("2E")` should return `46`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("2E"), 46)`)
}})
```

`hex_to_decimal("FF")` should return `255`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("FF"), 255)`)
}})
```

`hex_to_decimal("A3F")` should return `2623`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A3F"), 2623)`)
}})
```

# --seed--

## --seed-contents--

```py
def hex_to_decimal(hex):

    return hex
```

# --solutions--

```py
def hex_to_decimal(hex):

    return int(hex, 16)
```
