---
id: 68cae5b538ff798bbd4da005
title: "Challenge 67: Email Validator"
challengeType: 29
dashedName: challenge-67
---

# --description--

Kutolewa mfuatano wa herufi, amua kama ni anwani halali ya barua pepe kwa kutumia vizingiti vifuatavyo:

- Lazima iwe na alama moja tu ya `@`.
- Sehemu ya ndani (kabla ya `@`):
  - Inaweza kuwa na herufi tu (`a-z`, `A-Z`), nambari (`0-9`), nukta (`.`), alama za mstari chini (`_`), au alama za mkato (`-`).
  - Haiwezi kuanza au kumalizika na nukta.
- Sehemu ya kikoa (baada ya `@`):
  - Lazima iwe na angalau nukta moja.
  - Lazima imalizike na nukta ikifuatiwa na angalau herufi mbili.
- Sehemu ya ndani wala sehemu ya kikoa haiwezi kuwa na nukta mbili mfululizo.

# --hints--

`validate("a@b.cd")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("a@b.cd"), True)`)
}})
```

`validate("hell.-w.rld@example.com")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hell.-w.rld@example.com"), True)`)
}})
```

`validate(".b@sh.rc")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate(".b@sh.rc"), False)`)
}})
```

`validate("example@test.c0")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("example@test.c0"), False)`)
}})
```

`validate("freecodecamp.org")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("freecodecamp.org"), False)`)
}})
```

`validate("develop.ment_user@c0D!NG.R.CKS")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("develop.ment_user@c0D!NG.R.CKS"), True)`)
}})
```

`validate("hello.@wo.rld")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hello.@wo.rld"), False)`)
}})
```

`validate("hello@world..com")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("hello@world..com"), False)`)
}})
```

`validate("develop..ment_user@c0D!NG.R.CKS")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("develop..ment_user@c0D!NG.R.CKS"), False)`)
}})
```

`validate("git@commit@push.io")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(validate("git@commit@push.io"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def validate(email):

    return email
```

# --solutions--

```py
import re
def validate(email):
    if '..' in email:
        return False

    parts = email.split('@')
    if len(parts) != 2:
        return False

    local, domain = parts

    if local.startswith('.') or local.endswith('.'):
        return False
    if not re.match(r'^[a-zA-Z0-9._-]+$', local):
        return False

    if '.' not in domain:
        return False

    tld = domain.split('.')[-1]
    if len(tld) < 2 or not tld.isalpha():
        return False

    return True
```
