---
id: 68cae5b538ff798bbd4da00a
title: "Challenge 72: Thermostat Adjuster 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

Ukizingatia joto la sasa la chumba kwa Fahrenheit na joto lengwa kwa Celsius, rudisha mfuatano wa herufi unaoelezea jinsi ya kurekebisha joto la chumba kulingana na masharti haya:

- Rudisha `"Heat: X degrees Fahrenheit"` ikiwa joto la sasa liko chini ya lengwa. Ambapo `X` ni idadi ya nyuzi za Fahrenheit za kupasha joto chumba kufikia lengwa, zilikwisha kuzungushwa hadi sehemu moja ya desimali.
- Rudisha `"Cool: X degrees Fahrenheit"` ikiwa joto la sasa liko juu ya lengwa. Ambapo `X` ni idadi ya nyuzi za Fahrenheit za kupooza chumba kufikia lengwa, zilikwisha kuzungushwa hadi sehemu moja ya desimali.
- Rudisha `"Hold"` ikiwa joto la sasa ni sawa na lengwa.

Ili kubadilisha Celsius kuwa Fahrenheit, zidisha joto la Celsius kwa 1.8 kisha ongeza 32 kwa matokeo (`F = (C * 1.8) + 32`).

# --hints--

`adjust_thermostat(32, 0)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` should return `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` should return `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` should return `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
