---
id: 68cae5b538ff798bbd4da007
title: "Challenge 69: Missing Socks"
challengeType: 29
dashedName: challenge-69
---

# --description--

Ukizingatia nambari kamili inayowakilisha idadi ya jozi za soksi ulizozianza nazo, na nambari nyingine inayowakilisha ni mizunguko mingapi ya kuosha umefanya, rudisha idadi ya jozi kamili za soksi ulizonazo kwa sasa ukitumia masharti yafuatayo:

- Kila mizunguko 2 ya kuosha, hupoteza soksi moja.
- Kila mizunguko 3 ya kuosha, unapata soksi moja iliyokosekana.
- Kila mizunguko 5 ya kuosha, soksi moja huharibika na lazima itupwe.
- Kila mizunguko 10 ya kuosha, unanunua jozi mpya ya soksi.
- Huwezi kuwa na soksi chini ya sifuri kabisa.
- Sheria zinaweza kuungana. Kwa mfano, katika mzunguko wa kuosha wa 10, utapoteza soksi moja, utatupa soksi moja, na ununue jozi mpya ya soksi.
- Rudisha idadi ya jozi kamili za soksi.

# --hints--

`sock_pairs(2, 5)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(2, 5), 1)`)
}})
```

`sock_pairs(1, 2)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 2), 0)`)
}})
```

`sock_pairs(5, 11)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(5, 11), 4)`)
}})
```

`sock_pairs(6, 25)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(6, 25), 3)`)
}})
```

`sock_pairs(1, 8)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 8), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sock_pairs(pairs, cycles):

    return pairs
```

# --solutions--

```py
def sock_pairs(pairs, cycles):
    socks = pairs * 2

    for i in range(1, cycles + 1):
        if i % 2 == 0:
            socks -= 1
        if i % 3 == 0:
            socks += 1
        if i % 5 == 0:
            socks -= 1
        if i % 10 == 0:
            socks += 2

        if socks < 0:
            socks = 0

    return socks // 2
```
