---
id: 68ee9e3066cfd4eb2328e8a8
title: "Challenge 89: Counting Cards"
challengeType: 29
dashedName: challenge-89
---

# --description--

Kifurushi cha kawaida cha kadi za kuchezea kina kadi 13 za kipekee katika kila rangi. Ukipata nambari kamili inayowakilisha idadi ya kadi za kuchagua kutoka kwenye kifurushi, rudisha idadi ya mchanganyiko wa kipekee wa kadi unazoweza kuchagua.

- Mpangilio hauna umuhimu. Kuchagua kadi A kisha kadi B ni sawa na kuchagua kadi B kisha kadi A.

Kwa mfano, ukipata `52`, rudisha `1`. Kuna mchanganyiko mmoja tu wa kadi 52 kuchagua kutoka kwenye kifurushi cha kadi 52. Na ukipata `2`, rudisha `1326`. Kuna mchanganyiko 1326 ya kadi unayoweza kupata unapochagua kadi 2 kutoka kwenye kifurushi.

# --hints--

`combinations(52)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` should return `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` should return `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` should return `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
