---
id: 68f6587287ad1f4ad39b0c7f
title: "Zoezi la 93: Vowels na Consonants"
challengeType: 29
dashedName: challenge-93
---

# --description--

Ukipata mfuatano wa herufi, rudisha safu ya data yenye idadi ya vowels na idadi ya consonants katika mfuatano huo.

- Vowels ni `a`, `e`, `i`, `o`, `u` katika hali yoyote.
- Consonants ni herufi zote nyingine katika hali yoyote.
- Puuza herufi zisizo za alfabeti.

Kwa mfano, ukipata `"Hello World"`, rudisha `[3, 7]`.

# --hints--

`count("Hello World")` inapaswa kurudisha `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("JavaScript")` inapaswa kurudisha `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("JavaScript"), [3, 7])`)
}})
```

`count("Python")` inapaswa kurudisha `[1, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Python"), [1, 5])`)
}})
```

`count("freeCodeCamp")` inapaswa kurudisha `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("freeCodeCamp"), [5, 7])`)
}})
```

`count("Hello, World!")` inapaswa kurudisha `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello, World!"), [3, 7])`)
}})
```

`count("The quick brown fox jumps over the lazy dog.")` inapaswa kurudisha `[11, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24])`)
}})
```

# --seed--

## --seed-contents--

```py
def count(s):

    return s
```

# --solutions--

```py
def count(s):
    vowels = "aeiou"
    consonants = "bcdfghjklmnpqrstvwxyz"
    v = 0
    c = 0

    for char in s:
        ch = char.lower()
        if ch in vowels:
            v += 1
        if ch in consonants:
            c += 1

    return [v, c]
```
