---
id: 68f6587287ad1f4ad39b0c83
title: "Challenge 97: GCD"
challengeType: 29
dashedName: challenge-97
---

# --description--

Ukizingatia nambari kamili chanya mbili, rudisha mgawanyiko wao mkubwa zaidi wa pamoja (GCD).

- GCD ya nambari kamili mbili ni nambari kubwa zaidi inayogawanya nambari zote mbili kwa usawa bila kuacha mabaki.

Kwa mfano, wagawanyiko wa `4` ni `1`, `2`, na `4`. Wagawanyiko wa `6` ni `1`, `2`, `3`, na `6`. Kwa hivyo ukizingatia `4` na `6`, rudisha `2`, nambari kubwa zaidi inayojitokeza katika seti zote mbili za wagawanyiko.

# --hints--

`gcd(4, 6)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(4, 6), 2)`)
}})
```

`gcd(20, 15)` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(20, 15), 5)`)
}})
```

`gcd(13, 17)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(13, 17), 1)`)
}})
```

`gcd(654, 456)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(654, 456), 6)`)
}})
```

`gcd(3456, 4320)` should return `864`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(3456, 4320), 864)`)
}})
```

# --seed--

## --seed-contents--

```py
def gcd(x, y):

    return x
```

# --solutions--

```py
def gcd(x, y):
    while y != 0:
        x, y = y, x % y
    return x
```
