---
id: 68f6587287ad1f4ad39b0c84
title: "Challenge 98: Rectangle Count"
challengeType: 29
dashedName: challenge-98
---

# --description--

Ukitoa nambari kamili chanya mbili zinazowakilisha upana na urefu wa mstatili, amua ni mistari mingapi ya mstatili inaweza kuingizwa ndani ya ule uliotolewa.

- Hesabu tu mistari ya mstatili yenye upana na urefu wa nambari kamili.

Kwa mfano, ukitoa `1` na `3`, rudisha `6`. Mistari mitatu ya 1x1, mistari miwili ya 1x2, na mstari mmoja wa 1x3.

# --hints--

`count_rectangles(1, 3)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 3), 6)`)
}})
```

`count_rectangles(3, 2)` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(3, 2), 18)`)
}})
```

`count_rectangles(1, 2)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 2), 3)`)
}})
```

`count_rectangles(5, 4)` should return `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(5, 4), 150)`)
}})
```

`count_rectangles(11, 19)` should return `12540`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(11, 19), 12540)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_rectangles(width, height):

    return width
```

# --solutions--

```py
def count_rectangles(width, height):
    total = 0
    for w in range(1, width + 1):
        for h in range(1, height + 1):
            total += (width - w + 1) * (height - h + 1)
    return total
```
