---
id: 68f6587287ad1f4ad39b0c81
title: "Challenge 95: Array Shift"
challengeType: 29
dashedName: challenge-95
---

# --description--

Ikiwa umepewa safu ya data na nambari kamili inayowakilisha ni nafasi ngapi za kusogeza safu hiyo, rudisha safu ya data iliyosogezwa.

- Nambari kamili chanya husogeza safu ya data kushoto.
- Nambari kamili hasi husogeza safu ya data kulia.
- Kusogeza kunazunguka ndani ya safu ya data.

Kwa mfano, ukipata `[1, 2, 3]` na `1`, sogeza safu ya data mara 1 kushoto, ukirudisha `[2, 3, 1]`.

# --hints--

`shift_array([1, 2, 3], 1)` should return `[2, 3, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], 1), [2, 3, 1])`)
}})
```

`shift_array([1, 2, 3], -1)` should return `[3, 1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([1, 2, 3], -1), [3, 1, 2])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], 5)` should return `["charlie", "alpha", "bravo"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"])`)
}})
```

`shift_array(["alpha", "bravo", "charlie"], -11)` should return `["bravo", "charlie", "alpha"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"])`)
}})
```

`shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` should return `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_array([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_array(arr, n):

    return arr
```

# --solutions--

```py
def shift_array(arr, n):
    length = len(arr)
    n = n % length
    if n < 0:
        n += length

    return arr[n:] + arr[:n]
```
