---
id: 69162d64f96574d9bb629efb
title: "Challenge 110: Word Guesser"
challengeType: 29
dashedName: challenge-110
---

# --description--

Ukizingatia mfuatano wa herufi mbili zenye urefu sawa, neno la siri na neno la kubahatisha, linganisha neno la kubahatisha na neno la siri kwa kutumia sheria zifuatazo:

- Neno la siri na neno la kubahatisha litakuwa na herufi kubwa tu (`"A"` hadi `"Z"`);
- Kwa kila herufi katika neno la kubahatisha, ibadilishe na nambari kulingana na jinsi inavyolingana na neno la siri:
  - `"2"` ikiwa herufi iko katika neno la siri na ipo mahali sahihi.
  - `"1"` ikiwa herufi iko katika neno la siri lakini ipo mahali potofu.
  - `"0"` ikiwa herufi haipo katika neno la siri.
- Kila herufi katika neno la siri inaweza kutumika mara moja tu.
- Mlingano kamili (`"2"`) unapelekwa kwanza, kisha mlingano wa sehemu (`"1"`) unapelekwa kutoka kushoto kwenda kulia kwa herufi zilizobaki.
- Ikiwa herufi inatokea mara nyingi katika neno la kubahatisha, inaweza kulingana mara ngapi tu inavyotokea katika neno la siri.

Kwa mfano, ukizingatia neno la siri la `"APPLE"` na neno la kubahatisha la `"POPPA"`, rudisha `"10201"`:

Herufi ya kwanza ya `"P"` haipo mahali sahihi (`"1"`), `"O"` haipo katika neno la siri (`"0"`), herufi ya pili ya `"P"` iko mahali sahihi (`"2"`), herufi ya tatu ya `"P"` ni sifuri (`"0"`) kwa sababu herufi mbili za `"P"` katika neno la siri zimetumika, na `"A"` haipo mahali sahihi (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` should return `"10201"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("APPLE", "POPPA"), "10201")`)
}})
```

`compare("REACT", "TRACE")` should return `"11221"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("REACT", "TRACE"), "11221")`)
}})
```

`compare("DEBUGS", "PYTHON")` should return `"000000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("DEBUGS", "PYTHON"), "000000")`)
}})
```

`compare("JAVASCRIPT", "TYPESCRIPT")` should return `"0000222222"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222")`)
}})
```

`compare("ORANGE", "ROUNDS")` should return `"110200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("ORANGE", "ROUNDS"), "110200")`)
}})
```

`compare("WIRELESS", "ETHERNET")` should return `"10021000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("WIRELESS", "ETHERNET"), "10021000")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare(word, guess):

    return word
```

# --solutions--

```py
def compare(word, guess):
    n = len(word)
    result = ["0"] * n
    used_secret = [False] * n

    for i in range(n):
        if guess[i] == word[i]:
            result[i] = "2"
            used_secret[i] = True

    for i in range(n):
        if result[i] == "0":
            for j in range(n):
                if not used_secret[j] and guess[i] == word[j]:
                    result[i] = "1"
                    used_secret[j] = True
                    break

    return "".join(result)
```
