---
id: 69162d64f96574d9bb629eff
title: "Challenge 114: Camel to Snake"
challengeType: 29
dashedName: challenge-114
---

# --description--

Ikiwa umepewa mfuatano wa herufi katika camel case, rudisha toleo la snake case la mfuatano huo kwa kutumia sheria zifuatazo:

- Mfuatano wa herufi wa ingizo utakuwa na herufi tu (`A-Z` na `a-z`) na daima utaanza na herufi ndogo.
- Kila herufi kubwa katika mfuatano wa camel case huanzisha neno jipya.
- Badilisha herufi zote kuwa herufi ndogo.
- Tenganisha maneno kwa mstari chini (underscore) (`_`).

# --hints--

`to_snake("helloWorld")` should return `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` should return `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` should return `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
