---
id: 69162d64f96574d9bb629f00
title: "Challenge 115: Markdown Ordered List Item Converter"
challengeType: 29
dashedName: challenge-115
---

# --description--

Ikiwa umepewa mfuatano wa herufi unaowakilisha kipengee cha orodha iliyopangwa katika Markdown, rudisha mfuatano wa herufi wa HTML unaolingana.

Kipengee halali cha orodha iliyopangwa katika Markdown lazima:

- Kuanza na nafasi sifuri au zaidi, ikifuatiwa na
- Nambari (1 au zaidi) na nukta (`.`), ikifuatiwa na
- Angalau nafasi moja, kisha
- Maandishi ya kipengee cha orodha.

Kama mfuatano wa herufi hauendani kabisa na muundo huo hapo juu, rudisha `"Invalid format"`. Vinginevyo, zungusha maandishi ya kipengee cha orodha ndani ya lebo za `li` na rudisha mfuatano wa herufi.

Kwa mfano, ukipokea `"1. My item"`, rudisha `"<li>My item</li>"`.

Kumbuka: Konsoli huenda isionyeshe lebo za HTML ndani ya mfuatano wakati wa kurekodi ujumbe. Angalia konsoli ya kivinjari kuona rekodi zikiwa na lebo zilizoingizwa.

# --hints--

`convert_list_item("1. My item")` should return `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` should return `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` should return `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
