---
id: 691b559495c5cb5a37b9b480
title: "Challenge 120: Pounds to Kilograms"
challengeType: 29
dashedName: challenge-120
---

# --description--

Ukitoa uzito kwa pauni kama nambari, rudisha mfuatano wa herufi `"(lbs) pounds equals (kgs) kilograms."`.

- Badilisha `"(lbs)"` na nambari ya ingizo.
- Badilisha `"(kgs)"` na ingizo lililobadilishwa kuwa kilogramu, zungusha hadi desimali mbili na daima jumuisha desimali mbili katika thamani.
- Pauni 1 ni sawa na kilogramu 0.453592.
- Ikiwa ingizo ni `1`, tumia `"pound"` badala ya `"pounds"`.
- Ikiwa thamani iliyobadilishwa ni `1`, tumia `"kilogram"` badala ya `"kilograms"`.

# --hints--

`convert_to_kgs(1)` should return `"1 pound equals 0.45 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` should return `"0 pounds equals 0.00 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` should return `"100 pounds equals 45.36 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` should return `"2.5 pounds equals 1.13 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` should return `"2.20462 pounds equals 1.00 kilogram."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
