---
id: 691b559495c5cb5a37b9b487
title: "Challenge 127: Speed Check"
challengeType: 29
dashedName: challenge-127
---

# --description--

Kutokana na kasi unayovuka kwa maili kwa saa (MPH), na kikomo cha kasi kwa kilomita kwa saa (KPH), amua kama unavuka kasi na kama utapokea onyo au tiketi.

- Maili 1 ni sawa na kilomita 1.60934.
- Ikiwa unasafiri kwa kasi isiyozidi au sawa na kikomo cha kasi, rudisha `"Not Speeding"`.
- Ikiwa unasafiri kwa kasi ya KPH 5 au chini zaidi ya kikomo cha kasi, rudisha `"Warning"`.
- Ikiwa unasafiri kwa kasi zaidi ya KPH 5 juu ya kikomo cha kasi, rudisha `"Ticket"`.

# --hints--

`speed_check(30, 70)` should return `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` should return `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` should return `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` should return `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` should return `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` should return `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
