---
id: 691b559495c5cb5a37b9b485
title: "Challenge 125: Game of Life"
challengeType: 29
dashedName: challenge-125
---

# --description--

Ikiwa umepewa matriki (safu ya safu) inayowakilisha hali ya sasa katika Mchezo wa Maisha wa Conway, rudisha hali inayofuata ya matriki kwa kutumia sheria hizi:

- Kila seli ni `1` (hai) au `0` (kifo).
- Majirani wa seli ni seli hadi nane zinazozunguka (wima, mlalo, na mwinuko).
- Seli zilizokaribu na mipaka zina majirani wachache kuliko nane.

Sheria za kusasisha kila seli:

- Seli yoyote hai yenye majirani hai chini ya wawili hufa (ukosefu wa watu).
- Seli yoyote hai yenye majirani hai wawili au watatu huendelea kuishi.
- Seli yoyote hai yenye majirani hai zaidi ya watatu hufa (watu wengi sana).
- Seli yoyote iliyokufa yenye majirani hai watatu hasa huamka (uzazi).

Kwa mfano, ukipata:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

return:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Kila seli husasishwa kulingana na idadi ya majirani hai. Kwa mfano, `[0][0]` hubaki amekufa (majirani hai 2), `[0][1]` hubaki hai (majirani hai 2), `[0][2]` hufa (majirani hai 3), na kadhalika.

# --hints--

`game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` should return `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]])`)
}})
```

`game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` should return `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]])`)
}})
```

`game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` should return `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]])`)
}})
```

`game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` should return `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def game_of_life(grid):

    return grid
```

# --solutions--

```py
def game_of_life(grid):
    rows, cols = len(grid), len(grid[0])

    def count_live_neighbors(r, c):
        count = 0
        for i in range(r-1, r+2):
            for j in range(c-1, c+2):
                if 0 <= i < rows and 0 <= j < cols and (i != r or j != c):
                    count += grid[i][j]
        return count

    next_state = [row[:] for row in grid]

    for r in range(rows):
        for c in range(cols):
            live_neighbors = count_live_neighbors(r, c)

            if grid[r][c] == 1:
                if live_neighbors < 2 or live_neighbors > 3:
                    next_state[r][c] = 0
            else:
                if live_neighbors == 3:
                    next_state[r][c] = 1

    return next_state
```
