---
id: 691b559495c5cb5a37b9b482
title: "Challenge 122: Markdown Bold Parser"
challengeType: 29
dashedName: challenge-122
---

# --description--

Ikiwa umepewa mfuatano wa herufi ambao unaweza kuwa na maandishi ya herufi nzito katika Markdown, rudisha mfuatano wa herufi sawa wa HTML.

- Maandishi ya herufi nzito katika Markdown ni maandishi yoyote yanayoanza na kumalizika na nyota mbili (`**`) au mstari chini miwili (`__`).
- Hakuna nafasi yoyote kati ya maandishi na nyota au mistari chini, lakini nafasi zinaweza kuwepo ndani ya maandishi yenyewe.
- Badilisha matukio yote ya herufi nzito kuwa lebo za HTML `b` na rudisha mfuatano wa herufi.

Kwa mfano, ukipokea `"**This is bold**"`, rudisha `"<b>This is bold</b>"`.

Kumbuka: Konsoli huenda isionyeshe lebo za HTML ndani ya mfuatano wakati wa kurekodi ujumbe. Angalia konsoli ya kivinjari kuona rekodi zikiwa na lebo zilizoingizwa.

# --hints--

`parse_bold("**This is bold**")` should return `"<b>This is bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is bold**"), "<b>This is bold</b>")`)
}})
```

`parse_bold("__This is also bold__")` should return `"<b>This is also bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__This is also bold__"), "<b>This is also bold</b>")`)
}})
```

`parse_bold("**This is not bold **")` should return `"**This is not bold **"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is not bold **"), "**This is not bold **")`)
}})
```

`parse_bold("__ This is also not bold__")` should return `"__ This is also not bold__"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__ This is also not bold__"), "__ This is also not bold__")`)
}})
```

`parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog.")` should return `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_bold(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_bold(markdown):
    markdown = re.sub(r'\*\*(\S(?:.*?\S)?)\*\*', r'<b>\1</b>', markdown)
    markdown = re.sub(r'__(\S(?:.*?\S)?)__', r'<b>\1</b>', markdown)

    return markdown
```
