---
id: 691b559495c5cb5a37b9b486
title: "Challenge 126: Capitalize It"
challengeType: 29
dashedName: challenge-126
---

# --description--

Ukipata mfuatano wa herufi unaoitwa title, rudisha mfuatano mpya uliopangwa kwa mtindo wa herufi kubwa kwa kutumia sheria zifuatazo:

- Andika herufi ya kwanza ya kila neno kwa herufi kubwa.
- Fanya herufi zote nyingine katika kila neno kuwa herufi ndogo.
- Maneno hubaki kutenganishwa na nafasi moja tu kila mara.

# --hints--

`title_case("hello world")` should return `"Hello World"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("hello world"), "Hello World")`)
}})
```

`title_case("the quick brown fox")` should return `"The Quick Brown Fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("the quick brown fox"), "The Quick Brown Fox")`)
}})
```

`title_case("JAVASCRIPT AND PYTHON")` should return `"Javascript And Python"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("JAVASCRIPT AND PYTHON"), "Javascript And Python")`)
}})
```

`title_case("AvOcAdO tOAst fOr brEAkfAst")` should return `"Avocado Toast For Breakfast"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast")`)
}})
```

# --seed--

## --seed-contents--

```py
def title_case(title):

    return title
```

# --solutions--

```py
def title_case(title):
    return " ".join(
        w[:1].upper() + w[1:].lower()
        for w in title.split(" ")
    )
```
