---
id: 6925e2068081f40f549ced1b
title: "Challenge 137: Snowflake Generator"
challengeType: 29
dashedName: challenge-137
---

# --description--

Ikiwa umepewa mfuatano wa herufi wenye mistari mingi unaotumia tabia za newline (`\n`) kuwakilisha mapumziko ya mistari, rudisha mfuatano mpya ambapo kila mstari umeakisi kwa usawa na kuambatanishwa mwishoni mwa mstari wa asili.

- Akisi mstari kwa kubadilisha mfuatano wake wote wa herufi, ikijumuisha nafasi.

Kwa mfano, ukipokea `"* \n *\n* "`, ambayo inaandika kwenye konsoli kama:

```sh
* 
 *
* 
```

Rudisha `"*  *\n ** \n*  *"`, ambayo inaandika kwenye konsoli kama:

```sh
*  *
 ** 
*  *
```

Chukua tahadhari maalum kuhusu nafasi katika mfuatano uliotolewa na ule uliorudishwa. Hakikisha haukatagi nafasi yoyote.

# --hints--

`generate_snowflake("* \n *\n* ")` should return `"*  *\n ** \n*  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` should return `"X=~~=X"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` should return `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` should return `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` should return `"*  --  *\n * -- * \n*  --  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
