---
id: 6925e2068081f40f549ced1a
title: "Challenge 136: Markdown Image Parser"
challengeType: 29
dashedName: challenge-136
---

# --description--

Ukipata mfuatano wa herufi wa picha katika Markdown, rudisha mfuatano wa herufi sawa wa HTML.

Picha ya Markdown ina muundo ufuatao: `"![alt text](image_url)"`. Ambapo:

- `alt text` ni maelezo ya picha (thamani ya sifa ya kipengele ya `alt`).
- `image_url` ni URL chanzo cha picha (thamani ya sifa ya kipengele ya `src`).

Rudisha mfuatano wa herufi wa lebo ya HTML `img` yenye `src` imewekwa kuwa URL ya picha na `alt` imewekwa kuwa maandishi ya alt.

Kwa mfano, ukipata `"![Cute cat](cat.png)"` rudisha `'<img src="cat.png" alt="Cute cat">'`;

- Hakikisha lebo, mpangilio wa sifa za kipengele, nafasi, na matumizi ya nukuu ni sawa na hapo juu.

Kumbuka: Konsoli inaweza isionyeshe lebo za HTML katika mfuatano wa herufi wakati wa kurekodi ujumbe — angalia konsoli ya kivinjari ili kuona rekodi zikiwa na lebo zilizoambatanishwa.

# --hints--

`parse_image("![Cute cat](cat.png)")` should return `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` should return `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` should return `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
