---
id: 69373793f5a867f769cde137
title: "Challenge 152: Circular Prime"
challengeType: 29
dashedName: challenge-152
---

# --description--

Ikiwa namba kamili, amua kama ni namba msingi mduara.

Namba msingi mduara ni namba kamili ambapo mizunguko yote ya tarakimu zake ni nambari msingi pia.

Kwa mfano, `197` ni namba msingi mduara kwa sababu mizunguko yote ya tarakimu zake: `197`, `971`, na `719`, ni namba msingi.

# --hints--

`is_circular_prime(197)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(197), True)`)
}})
```

`is_circular_prime(23)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(23), False)`)
}})
```

`is_circular_prime(13)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(13), True)`)
}})
```

`is_circular_prime(89)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(89), False)`)
}})
```

`is_circular_prime(1193)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(1193), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_circular_prime(n):

    return n
```

# --solutions--

```py
import math
def is_prime(n):
    if n < 2:
        return False
    for i in range(2, int(math.isqrt(n)) + 1):
        if n % i == 0:
            return False
    return True

def rotations(n):
    s = str(n)
    return [int(s[i:] + s[:i]) for i in range(len(s))]

def is_circular_prime(n):
    return all(is_prime(rot) for rot in rotations(n))
```
