---
id: 69738771fb5a7b8b24cca2a0
title: "Challenge 174: Zodiac Finder"
challengeType: 29
dashedName: challenge-174
---

# --description--

Kutolewa mfuatano wa herufi wa tarehe katika muundo `"YYYY-MM-DD"`, rudisha alama ya zodiac kwa tarehe hiyo ukitumia chati ifuatayo:

| Kipindi cha tarehe | Alama ya zodiac |
| - | - |
| Machi 21 - Aprili 19 | `"Aries"` |
| Aprili 20 - Mei 20 | `"Taurus"` |
| Mei 21 - Juni 20 | `"Gemini"` |
| Juni 21 - Julai 22 | `"Cancer"` |
| Julai 23 - Agosti 22 | `"Leo"` |
| Agosti 23 - Septemba 22 | `"Virgo"` |
| Septemba 23 - Oktoba 22 | `"Libra"` |
| Oktoba 23 - Novemba 21 | `"Scorpio"` |
| Novemba 22 - Desemba 21 | `"Sagittarius"` |
| Desemba 22 - Januari 19 | `"Capricorn"` |
| Januari 20 - Februari 18 | `"Aquarius"` |
| Februari 19 - Machi 20 | `"Pisces"` |

- Alama za zodiac zinategemea tu mwezi na siku, unaweza kupuuza mwaka.

# --hints--

`get_sign("2026-01-31")` should return `"Aquarius"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2026-01-31"), "Aquarius")`)
}})
```

`get_sign("2001-06-10")` should return `"Gemini"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2001-06-10"), "Gemini")`)
}})
```

`get_sign("1985-09-07")` should return `"Virgo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1985-09-07"), "Virgo")`)
}})
```

`get_sign("2023-03-19")` should return `"Pisces"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2023-03-19"), "Pisces")`)
}})
```

`get_sign("2045-11-05")` should return `"Scorpio"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2045-11-05"), "Scorpio")`)
}})
```


`get_sign("1985-12-06")` should return `"Sagittarius"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1985-12-06"), "Sagittarius")`)
}})
```

`get_sign("2025-12-30")` should return `"Capricorn"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2025-12-30"), "Capricorn")`)
}})
```

`get_sign("2018-10-08")` should return `"Libra"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("2018-10-08"), "Libra")`)
}})
```

`get_sign("1958-05-04")` should return `"Taurus"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_sign("1958-05-04"), "Taurus")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_sign(date_str):

    return date_str
```

# --solutions--

```py
def get_sign(date_str):
    _, month_str, day_str = date_str.split("-")
    month = int(month_str)
    day = int(day_str)

    if (month == 3 and day >= 21) or (month == 4 and day <= 19):
        return "Aries"
    if (month == 4 and day >= 20) or (month == 5 and day <= 20):
        return "Taurus"
    if (month == 5 and day >= 21) or (month == 6 and day <= 20):
        return "Gemini"
    if (month == 6 and day >= 21) or (month == 7 and day <= 22):
        return "Cancer"
    if (month == 7 and day >= 23) or (month == 8 and day <= 22):
        return "Leo"
    if (month == 8 and day >= 23) or (month == 9 and day <= 22):
        return "Virgo"
    if (month == 9 and day >= 23) or (month == 10 and day <= 22):
        return "Libra"
    if (month == 10 and day >= 23) or (month == 11 and day <= 21):
        return "Scorpio"
    if (month == 11 and day >= 22) or (month == 12 and day <= 21):
        return "Sagittarius"
    if (month == 12 and day >= 22) or (month == 1 and day <= 19):
        return "Capricorn"
    if (month == 1 and day >= 20) or (month == 2 and day <= 18):
        return "Aquarius"
    if (month == 2 and day >= 19) or (month == 3 and day <= 20):
        return "Pisces"
```
