---
id: 69738771fb5a7b8b24cca2a2
title: "Challenge 176: Groundhog Day"
challengeType: 29
dashedName: challenge-176
---

# --description--

Leo ni Siku ya Groundhog, ambapo groundhog hutabiri hali ya hewa kulingana na kama inaona kivuli chake au la.

Kutolewa thamani inayowakilisha muonekano wa groundhog, rudisha utabiri sahihi:

- Ikiwa thamani iliyotolewa ni thamani ya kweli au siyo kweli `true` (groundhog aliona kivuli chake), rudisha `"Looks like we'll have six more weeks of winter."`.
- Ikiwa thamani ni thamani ya kweli au siyo kweli `false` (groundhog hakuona kivuli chake), rudisha `"It's going to be an early spring."`.
- Ikiwa thamani ni kitu kingine chochote (groundhog hakuonekana), rudisha `"No prediction this year."`.

# --hints--

`groundhog_day_prediction(True)` should return `"Looks like we'll have six more weeks of winter."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

`groundhog_day_prediction(False)` should return `"It's going to be an early spring."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

`groundhog_day_prediction(None)` should return `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` should return `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

`groundhog_day_prediction("True")` should return `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
