---
id: 697a49e6ff50d756c9b6935e
title: "Zoezi la 181: Michezo ya Majira ya Baridi 2026 Siku ya 2: Snowboarding"
challengeType: 29
dashedName: challenge-181
---

# --description--

Kutolewa mkao wa mwanariadha wa snowboarding anapoanza na mzunguko kwa digrii, amua mkao atakapoangukia.

- Mkao wa mwanariadha wa snowboarding ni `"Regular"` au `"Goofy"`.
- Mizunguko ya mbinu ni mara nyingi ya digrii 90. Chanya inaonyesha mzunguko wa saa, na hasi inaonyesha mzunguko wa kinyume cha saa.
- Mkao wa kuangukia hubadilika kila mzunguko wa digrii 180.

Kwa mfano, ukitolewa `"Regular"` na `90`, rudisha `"Regular"`. Ukitolewa `"Regular"` na digrii `180`, rudisha `"Goofy"`.

# --hints--

`get_landing_stance("Regular", 90)` inapaswa kurudisha `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` inapaswa kurudisha `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` inapaswa kurudisha `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` inapaswa kurudisha `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` inapaswa kurudisha `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` inapaswa kurudisha `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

`get_landing_stance("Goofy", 90)` inapaswa kurudisha `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 90), "Goofy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
