---
id: 697a49e9860d24853adef67d
title: "Challenge 192: 2026 Winter Games Day 13: Nordic Combined"
challengeType: 29
dashedName: challenge-192
---

# --description--

Kutolewa safu ya data ya alama za kuruka kwa wanariadha, hesabu nyakati zao za kuchelewa kuanza kwa sehemu ya mbio za msalaba wa Nordic Combined.

Mwanariadha mwenye alama ya juu zaidi ya kuruka huanza kwanza (hakuna kuchelewa kwa sekunde 0). Wanariadha wengine wote huanza baadaye kulingana na umbali wa alama zao za kuruka kutoka kwa alama bora zaidi.

Ili kuhesabu kuchelewa kwa kila mwanariadha, toa alama ya kuruka ya mwanariadha kutoka kwa alama bora kabisa ya kuruka kisha zidisha matokeo kwa 1.5. Punguza kuchelewa hadi nambari kamili inayofuata.

# --hints--

`calculate_start_delays([120, 110, 125])` should return `[8, 23, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([120, 110, 125]), [8, 23, 0])`)
}})
```

`calculate_start_delays([118, 125, 122, 120])` should return `[11, 0, 5, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([118, 125, 122, 120]), [11, 0, 5, 8])`)
}})
```

`calculate_start_delays([100, 105, 95, 110, 120, 115, 108])` should return `[30, 23, 38, 15, 0, 8, 18]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18])`)
}})
```

`calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` should return `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_start_delays(jump_scores):

    return jump_scores
```

# --solutions--

```py
import math
def calculate_start_delays(jump_scores):
    best_jump = max(jump_scores)
    return [math.ceil((best_jump - score) * 1.5) for score in jump_scores]
```
