---
id: 697a49e6ff50d756c9b6935f
title: "Challenge 182: 2026 Winter Games Day 3: Biathlon"
challengeType: 29
dashedName: challenge-182
---

# --description--

Imetolewa safu ya data ya nambari kamili, ambapo kila thamani inawakilisha idadi ya malengo yaliyopigwa katika mzunguko mmoja wa biathlon, rudisha jumla ya umbali wa adhabu ambayo mwanariadha lazima apige ski.

- Kila mzunguko una malengo 5.
- Kila lengo lililokosewa husababisha mzunguko wa adhabu wa mita 150.

# --hints--

`calculate_penalty_distance([4, 4])` should return `300`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 4]), 300)`)
}})
```

`calculate_penalty_distance([5, 5])` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 5]), 0)`)
}})
```

`calculate_penalty_distance([4, 5, 3, 5])` should return `450`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 5, 3, 5]), 450)`)
}})
```

`calculate_penalty_distance([5, 4, 5, 5])` should return `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 4, 5, 5]), 150)`)
}})
```

`calculate_penalty_distance([4, 3, 0, 3])` should return `1500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 3, 0, 3]), 1500)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_penalty_distance(rounds):

    return rounds
```

# --solutions--

```py
def calculate_penalty_distance(rounds):
    total_penalty = 0

    for hits in rounds:
        misses = 5 - hits
        total_penalty += misses * 150

    return total_penalty
```
