---
id: 697a49e9860d24853adef67c
title: "Challenge 191: 2026 Winter Games Day 12: Bobsled"
challengeType: 29
dashedName: challenge-191
---

# --description--

Imetolewa safu ya data inayowakilisha uzito wa wanariadha kwenye timu ya bobsled na nambari inayowakilisha uzito wa bobsled, amua kama timu inastahili kushindana.

- Urefu wa safu ya data unaamua ukubwa wa timu: timu za mtu 1, 2 au 4.
- Thamani zote za uzito zilizotolewa ni kwa kilo (kg).

Bobsled (sled pekee) lazima iwe na kiwango cha chini cha uzito cha:

- 162 kg kwa timu ya mtu 1
- 170 kg kwa timu ya watu 2
- 210 kg kwa timu ya watu 4

Jumla ya uzito wa bobsled (wanariadha pamoja na sled) haipaswi kuzidi:

- 247 kg kwa timu ya mtu 1
- 390 kg kwa timu ya watu 2
- 630 kg kwa timu ya watu 4

Rudisha `"Eligible"` ikiwa timu inakidhi mahitaji yote, au `"Not Eligible"` ikiwa timu haikutimiza moja au zaidi ya mahitaji hayo.

# --hints--

`check_eligibility([78], 165)` should return `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([78], 165), "Eligible")`)
}})
```

`check_eligibility([80], 160)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 160), "Not Eligible")`)
}})
```

`check_eligibility([80], 170)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([80], 170), "Not Eligible")`)
}})
```

`check_eligibility([85, 90], 170)` should return `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 90], 170), "Eligible")`)
}})
```

`check_eligibility([85, 95], 168)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([85, 95], 168), "Not Eligible")`)
}})
```

`check_eligibility([112, 97], 185)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([112, 97], 185), "Not Eligible")`)
}})
```

`check_eligibility([110, 102, 90, 106], 222)` should return `"Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([110, 102, 90, 106], 222), "Eligible")`)
}})
```

`check_eligibility([106, 99, 90, 88], 205)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 90, 88], 205), "Not Eligible")`)
}})
```

`check_eligibility([106, 99, 103, 96], 227)` should return `"Not Eligible"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_eligibility([106, 99, 103, 96], 227), "Not Eligible")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_eligibility(athlete_weights, sled_weight):

    return athlete_weights
```

# --solutions--

```py
def check_eligibility(athlete_weights, sled_weight):
    team_size = len(athlete_weights)

    rules = {
        1: {"min_sled": 162, "max_total": 247},
        2: {"min_sled": 170, "max_total": 390},
        4: {"min_sled": 210, "max_total": 630}
    }

    min_sled = rules[team_size]["min_sled"]
    max_total = rules[team_size]["max_total"]

    total_weight = sum(athlete_weights) + sled_weight

    if sled_weight < min_sled:
        return "Not Eligible"

    if total_weight > max_total:
        return "Not Eligible"

    return "Eligible"
```
