---
id: 697a49e6ff50d756c9b6935d
title: "Zoezi la 180: Michezo ya Majira ya Baridi Siku ya 1: Siku ya Ufunguzi"
challengeType: 29
dashedName: challenge-180
---

# --description--

Siku 17 zijazo zitakuletea changamoto za uandishi wa programu zilizoongozwa na Michezo ya Majira ya Baridi.

Kwa changamoto ya kwanza, unapewa msimbo wa nchi wenye herufi mbili na unahitaji kurudisha emoji ya bendera ya nchi hiyo.

Tumia orodha hii:

| Nchi | Msimbo | Bendera |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"🇦🇱"` |
| Andorra                    |            `"AD"` | `"🇦🇩"` |
| Argentina                  |            `"AR"` | `"🇦🇷"` |
| Armenia                    |            `"AM"` | `"🇦🇲"` |
| Australia                  |            `"AU"` | `"🇦🇺"` |
| Austria                    |            `"AT"` | `"🇦🇹"` |
| Azerbaijan                 |            `"AZ"` | `"🇦🇿"` |
| Belgium                    |            `"BE"` | `"🇧🇪"` |
| Benin                      |            `"BJ"` | `"🇧🇯"` |
| Bolivia                    |            `"BO"` | `"🇧🇴"` |
| Bosnia and Herzegovina     |            `"BA"` | `"🇧🇦"` |
| Brazil                     |            `"BR"` | `"🇧🇷"` |
| Bulgaria                   |            `"BG"` | `"🇧🇬"` |
| Canada                     |            `"CA"` | `"🇨🇦"` |
| Chile                      |            `"CL"` | `"🇨🇱"` |
| China                      |            `"CN"` | `"🇨🇳"` |
| Colombia                   |            `"CO"` | `"🇨🇴"` |
| Croatia                    |            `"HR"` | `"🇭🇷"` |
| Cyprus                     |            `"CY"` | `"🇨🇾"` |
| Czech Republic             |            `"CZ"` | `"🇨🇿"` |
| Denmark                    |            `"DK"` | `"🇩🇰"` |
| Ecuador                    |            `"EC"` | `"🇪🇨"` |
| Eritrea                    |            `"ER"` | `"🇪🇷"` |
| Estonia                    |            `"EE"` | `"🇪🇪"` |
| Finland                    |            `"FI"` | `"🇫🇮"` |
| France                     |            `"FR"` | `"🇫🇷"` |
| Georgia                    |            `"GE"` | `"🇬🇪"` |
| Germany                    |            `"DE"` | `"🇩🇪"` |
| Great Britain              |            `"GB"` | `"🇬🇧"` |
| Greece                     |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau              |            `"GW"` | `"🇬🇼"` |
| Haiti                      |            `"HT"` | `"🇭🇹"` |
| Hong Kong                  |            `"HK"` | `"🇭🇰"` |
| Hungary                    |            `"HU"` | `"🇭🇺"` |
| Iceland                    |            `"IS"` | `"🇮🇸"` |
| India                      |            `"IN"` | `"🇮🇳"` |
| Iran                       |            `"IR"` | `"🇮🇷"` |
| Ireland                    |            `"IE"` | `"🇮🇪"` |
| Israel                     |            `"IL"` | `"🇮🇱"` |
| Italy                      |            `"IT"` | `"🇮🇹"` |
| Jamaica                    |            `"JM"` | `"🇯🇲"` |
| Japan                      |            `"JP"` | `"🇯🇵"` |
| Kazakhstan                 |            `"KZ"` | `"🇰🇿"` |
| Kenya                      |            `"KE"` | `"🇰🇪"` |
| Kosovo                     |            `"XK"` | `"🇽🇰"` |
| Kyrgyzstan                 |            `"KG"` | `"🇰🇬"` |
| Latvia                     |            `"LV"` | `"🇱🇻"` |
| Lebanon                    |            `"LB"` | `"🇱🇧"` |
| Liechtenstein              |            `"LI"` | `"🇱🇮"` |
| Lithuania                  |            `"LT"` | `"🇱🇹"` |
| Luxembourg                 |            `"LU"` | `"🇱🇺"` |
| Madagascar                 |            `"MG"` | `"🇲🇬"` |
| Malaysia                   |            `"MY"` | `"🇲🇾"` |
| Malta                      |            `"MT"` | `"🇲🇹"` |
| Mexico                     |            `"MX"` | `"🇲🇽"` |
| Moldova                    |            `"MD"` | `"🇲🇩"` |
| Monaco                     |            `"MC"` | `"🇲🇨"` |
| Mongolia                   |            `"MN"` | `"🇲🇳"` |
| Montenegro                 |            `"ME"` | `"🇲🇪"` |
| Morocco                    |            `"MA"` | `"🇲🇦"` |
| Netherlands                |            `"NL"` | `"🇳🇱"` |
| New Zealand                |            `"NZ"` | `"🇳🇿"` |
| Nigeria                    |            `"NG"` | `"🇳🇬"` |
| North Macedonia            |            `"MK"` | `"🇲🇰"` |
| Norway                     |            `"NO"` | `"🇳🇴"` |
| Pakistan                   |            `"PK"` | `"🇵🇰"` |
| Philippines                |            `"PH"` | `"🇵🇭"` |
| Poland                     |            `"PL"` | `"🇵🇱"` |
| Portugal                   |            `"PT"` | `"🇵🇹"` |
| Puerto Rico                |            `"PR"` | `"🇵🇷"` |
| Romania                    |            `"RO"` | `"🇷🇴"` |
| San Marino                 |            `"SM"` | `"🇸🇲"` |
| Saudi Arabia               |            `"SA"` | `"🇸🇦"` |
| Serbia                     |            `"RS"` | `"🇷🇸"` |
| Singapore                  |            `"SG"` | `"🇸🇬"` |
| Slovakia                   |            `"SK"` | `"🇸🇰"` |
| Slovenia                   |            `"SI"` | `"🇸🇮"` |
| South Africa               |            `"ZA"` | `"🇿🇦"` |
| South Korea                |            `"KR"` | `"🇰🇷"` |
| Spain                      |            `"ES"` | `"🇪🇸"` |
| Sweden                     |            `"SE"` | `"🇸🇪"` |
| Switzerland                |            `"CH"` | `"🇨🇭"` |
| Thailand                   |            `"TH"` | `"🇹🇭"` |
| Trinidad & Tobago          |            `"TT"` | `"🇹🇹"` |
| Turkey                     |            `"TR"` | `"🇹🇷"` |
| Ukraine                    |            `"UA"` | `"🇺🇦"` |
| United Arab Emirates       |            `"AE"` | `"🇦🇪"` |
| United States              |            `"US"` | `"🇺🇸"` |
| Uruguay                    |            `"UY"` | `"🇺🇾"` |
| Uzbekistan                 |            `"UZ"` | `"🇺🇿"` |
| Venezuela                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`get_flag("AL")` inapaswa kurudisha `"🇦🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "🇦🇱")`)
}})
```

`get_flag("AD")` inapaswa kurudisha `"🇦🇩"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "🇦🇩")`)
}})
```

`get_flag("AR")` inapaswa kurudisha `"🇦🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "🇦🇷")`)
}})
```

`get_flag("AM")` inapaswa kurudisha `"🇦🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "🇦🇲")`)
}})
```

`get_flag("AU")` inapaswa kurudisha `"🇦🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "🇦🇺")`)
}})
```

`get_flag("AT")` inapaswa kurudisha `"🇦🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "🇦🇹")`)
}})
```

`get_flag("AZ")` inapaswa kurudisha `"🇦🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "🇦🇿")`)
}})
```

`get_flag("BE")` inapaswa kurudisha `"🇧🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "🇧🇪")`)
}})
```

`get_flag("BJ")` inapaswa kurudisha `"🇧🇯"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "🇧🇯")`)
}})
```

`get_flag("BO")` inapaswa kurudisha `"🇧🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "🇧🇴")`)
}})
```

`get_flag("BA")` inapaswa kurudisha `"🇧🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "🇧🇦")`)
}})
```

`get_flag("BR")` inapaswa kurudisha `"🇧🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "🇧🇷")`)
}})
```

`get_flag("BG")` inapaswa kurudisha `"🇧🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "🇧🇬")`)
}})
```

`get_flag("CA")` inapaswa kurudisha `"🇨🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "🇨🇦")`)
}})
```

`get_flag("CL")` inapaswa kurudisha `"🇨🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "🇨🇱")`)
}})
```

`get_flag("CN")` inapaswa kurudisha `"🇨🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "🇨🇳")`)
}})
```

`get_flag("CO")` inapaswa kurudisha `"🇨🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "🇨🇴")`)
}})
```

`get_flag("HR")` inapaswa kurudisha `"🇭🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "🇭🇷")`)
}})
```

`get_flag("CY")` inapaswa kurudisha `"🇨🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "🇨🇾")`)
}})
```

`get_flag("CZ")` inapaswa kurudisha `"🇨🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "🇨🇿")`)
}})
```

`get_flag("DK")` inapaswa kurudisha `"🇩🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "🇩🇰")`)
}})
```

`get_flag("EC")` inapaswa kurudisha `"🇪🇨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "🇪🇨")`)
}})
```

`get_flag("ER")` inapaswa kurudisha `"🇪🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "🇪🇷")`)
}})
```

`get_flag("EE")` inapaswa kurudisha `"🇪🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "🇪🇪")`)
}})
```

`get_flag("FI")` inapaswa kurudisha `"🇫🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "🇫🇮")`)
}})
```

`get_flag("FR")` inapaswa kurudisha `"🇫🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "🇫🇷")`)
}})
```

`get_flag("GE")` inapaswa kurudisha `"🇬🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "🇬🇪")`)
}})
```

`get_flag("DE")` inapaswa kurudisha `"🇩🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "🇩🇪")`)
}})
```

`get_flag("GB")` inapaswa kurudisha `"🇬🇧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "🇬🇧")`)
}})
```

`get_flag("GR")` inapaswa kurudisha `"🇬🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "🇬🇷")`)
}})
```

`get_flag("GW")` inapaswa kurudisha `"🇬🇼"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "🇬🇼")`)
}})
```

`get_flag("HT")` inapaswa kurudisha `"🇭🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "🇭🇹")`)
}})
```

`get_flag("HK")` inapaswa kurudisha `"🇭🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "🇭🇰")`)
}})
```

`get_flag("HU")` inapaswa kurudisha `"🇭🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "🇭🇺")`)
}})
```

`get_flag("IS")` inapaswa kurudisha `"🇮🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "🇮🇸")`)
}})
```

`get_flag("IN")` inapaswa kurudisha `"🇮🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "🇮🇳")`)
}})
```

`get_flag("IR")` inapaswa kurudisha `"🇮🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "🇮🇷")`)
}})
```

`get_flag("IE")` inapaswa kurudisha `"🇮🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "🇮🇪")`)
}})
```

`get_flag("IL")` inapaswa kurudisha `"🇮🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "🇮🇱")`)
}})
```

`get_flag("IT")` inapaswa kurudisha `"🇮🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "🇮🇹")`)
}})
```

`get_flag("JM")` inapaswa kurudisha `"🇯🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "🇯🇲")`)
}})
```

`get_flag("JP")` inapaswa kurudisha `"🇯🇵"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "🇯🇵")`)
}})
```

`get_flag("KZ")` inapaswa kurudisha `"🇰🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "🇰🇿")`)
}})
```

`get_flag("KE")` inapaswa kurudisha `"🇰🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "🇰🇪")`)
}})
```

`get_flag("XK")` inapaswa kurudisha `"🇽🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "🇽🇰")`)
}})
```

`get_flag("KG")` inapaswa kurudisha `"🇰🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "🇰🇬")`)
}})
```

`get_flag("LV")` inapaswa kurudisha `"🇱🇻"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "🇱🇻")`)
}})
```

`get_flag("LB")` inapaswa kurudisha `"🇱🇧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "🇱🇧")`)
}})
```

`get_flag("LI")` inapaswa kurudisha `"🇱🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "🇱🇮")`)
}})
```

`get_flag("LT")` inapaswa kurudisha `"🇱🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "🇱🇹")`)
}})
```

`get_flag("LU")` inapaswa kurudisha `"🇱🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "🇱🇺")`)
}})
```

`get_flag("MG")` inapaswa kurudisha `"🇲🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "🇲🇬")`)
}})
```

`get_flag("MY")` inapaswa kurudisha `"🇲🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "🇲🇾")`)
}})
```

`get_flag("MT")` inapaswa kurudisha `"🇲🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "🇲🇹")`)
}})
```

`get_flag("MX")` inapaswa kurudisha `"🇲🇽"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "🇲🇽")`)
}})
```

`get_flag("MD")` inapaswa kurudisha `"🇲🇩"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "🇲🇩")`)
}})
```

`get_flag("MC")` inapaswa kurudisha `"🇲🇨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "🇲🇨")`)
}})
```

`get_flag("MN")` inapaswa kurudisha `"🇲🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "🇲🇳")`)
}})
```

`get_flag("ME")` inapaswa kurudisha `"🇲🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "🇲🇪")`)
}})
```

`get_flag("MA")` inapaswa kurudisha `"🇲🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "🇲🇦")`)
}})
```

`get_flag("NL")` inapaswa kurudisha `"🇳🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "🇳🇱")`)
}})
```

`get_flag("NZ")` inapaswa kurudisha `"🇳🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "🇳🇿")`)
}})
```

`get_flag("NG")` inapaswa kurudisha `"🇳🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "🇳🇬")`)
}})
```

`get_flag("MK")` inapaswa kurudisha `"🇲🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "🇲🇰")`)
}})
```

`get_flag("NO")` inapaswa kurudisha `"🇳🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "🇳🇴")`)
}})
```

`get_flag("PK")` inapaswa kurudisha `"🇵🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "🇵🇰")`)
}})
```

`get_flag("PH")` inapaswa kurudisha `"🇵🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "🇵🇭")`)
}})
```

`get_flag("PL")` inapaswa kurudisha `"🇵🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "🇵🇱")`)
}})
```

`get_flag("PT")` inapaswa kurudisha `"🇵🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "🇵🇹")`)
}})
```

`get_flag("PR")` inapaswa kurudisha `"🇵🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "🇵🇷")`)
}})
```

`get_flag("RO")` inapaswa kurudisha `"🇷🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "🇷🇴")`)
}})
```

`get_flag("SM")` inapaswa kurudisha `"🇸🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "🇸🇲")`)
}})
```

`get_flag("SA")` inapaswa kurudisha `"🇸🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "🇸🇦")`)
}})
```

`get_flag("RS")` inapaswa kurudisha `"🇷🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "🇷🇸")`)
}})
```

`get_flag("SG")` inapaswa kurudisha `"🇸🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "🇸🇬")`)
}})
```

`get_flag("SK")` inapaswa kurudisha `"🇸🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "🇸🇰")`)
}})
```

`get_flag("SI")` inapaswa kurudisha `"🇸🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "🇸🇮")`)
}})
```

`get_flag("ZA")` inapaswa kurudisha `"🇿🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "🇿🇦")`)
}})
```

`get_flag("KR")` inapaswa kurudisha `"🇰🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "🇰🇷")`)
}})
```

`get_flag("ES")` inapaswa kurudisha `"🇪🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "🇪🇸")`)
}})
```

`get_flag("SE")` inapaswa kurudisha `"🇸🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "🇸🇪")`)
}})
```

`get_flag("CH")` inapaswa kurudisha `"🇨🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "🇨🇭")`)
}})
```

`get_flag("TH")` inapaswa kurudisha `"🇹🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "🇹🇭")`)
}})
```

`get_flag("TT")` inapaswa kurudisha `"🇹🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "🇹🇹")`)
}})
```

`get_flag("TR")` inapaswa kurudisha `"🇹🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "🇹🇷")`)
}})
```

`get_flag("UA")` inapaswa kurudisha `"🇺🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "🇺🇦")`)
}})
```

`get_flag("AE")` inapaswa kurudisha `"🇦🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "🇦🇪")`)
}})
```

`get_flag("US")` inapaswa kurudisha `"🇺🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "🇺🇸")`)
}})
```

`get_flag("UY")` inapaswa kurudisha `"🇺🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "🇺🇾")`)
}})
```

`get_flag("UZ")` inapaswa kurudisha `"🇺🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "🇺🇿")`)
}})
```

`get_flag("VE")` inapaswa kurudisha `"🇻🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "🇻🇪")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
