---
id: 697a49e9860d24853adef67f
title: "Zoezi la 194: Michezo ya Majira ya Baridi Siku ya 15: Kuteleza kwa Freestyle"
challengeType: 29
dashedName: challenge-194
---

# --description--

Ikiwa jina la mbinu lina maneno mawili, thibitisha kama ni jina halali la mbinu ya kuteleza kwa freestyle.

Mbinu ni halali ikiwa neno la kwanza liko kwenye orodha ya maneno ya kwanza halali, na neno la pili liko kwenye orodha ya maneno ya pili halali.

- Maneno mawili yatatenganishwa na nafasi moja tu.

Maneno ya kwanza halali:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

Maneno ya pili halali:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`is_valid_trick("Polar Vortex")` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Polar Vortex"), True)`)
}})
```

`is_valid_trick("Solar Icequake")` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Solar Icequake"), True)`)
}})
```

`is_valid_trick("Thunder Blizzard")` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Thunder Blizzard"), True)`)
}})
```

`is_valid_trick("Phantom Frostbite")` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Phantom Frostbite"), True)`)
}})
```

`is_valid_trick("Ghost Avalanche")` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Ghost Avalanche"), True)`)
}})
```

`is_valid_trick("Snowstorm Shadow")` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Snowstorm Shadow"), False)`)
}})
```

`is_valid_trick("Solar Sky")` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Solar Sky"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_trick(trick_name):

    return trick_name
```

# --solutions--

```py
def is_valid_trick(trick_name):
    valid_first = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"]
    valid_second = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"]

    words = trick_name.split(" ")
    first, second = words
    return first in valid_first and second in valid_second
```
