---
id: 698a1a73ade5ac0e19180fa9
title: "Challenge 206: Playing Card Values"
challengeType: 29
dashedName: challenge-206
---

# --description--

Ukitoa safu ya data ya kadi za mchezo, rudisha safu mpya ya data yenye thamani ya nambari ya kila kadi.

Card Values:

- Ace (`"A"`) ina thamani ya 1.
- Kadi zilizo na nambari (`"2"` - `"10"`) zina thamani sawa na nambari zilizo mbele yao: 2 - 10, mtawalia.
- Kadi za uso: Jack (`"J"`), Queen (`"Q"`), na King (`"K"`) kila moja ina thamani ya 10.

Suits:

- Kila kadi ina aina: Spades (`"S"`), Clubs (`"C"`), Diamonds (`"D"`), au Hearts (`"H"`).

Card Format:

- Kila kadi inawakilishwa kama mfuatano wa herufi: `"valueSuit"`. Kwa mfano: `"AS"` ni Ace wa Spades, `"10H"` ni Kumi wa Hearts, na `"QC"` ni Queen wa Clubs.

# --hints--

`card_values(["3H", "4D", "5S"])` should return `[3, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` should return `[1, 10, 10, 6, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` should return `[8, 10, 2, 10, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` should return `[1, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` should return `[10, 10, 10, 10, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```
