---
id: 69a890af247de743333bd4d1
title: "Challenge 228: Movie Night"
challengeType: 29
dashedName: challenge-228
---

# --description--

Ukipata mfuatano wa herufi unaoonyesha siku ya wiki, mfuatano mwingine wa herufi unaoonyesha saa ya kuonyesha filamu, na nambari kamili ya tiketi, rudisha jumla ya gharama ya tiketi za filamu kwa onyesho hilo.

Siku iliyotolewa itakuwa moja kati ya:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

Saa ya kuonyesha itatolewa kwa muundo `"H:MMam"` au `"H:MMpm"`. Kwa mfano `"10:00am"` au `"10:00pm"`.

Rudisha jumla ya gharama kwa muundo `"$D.CC"` ukitumia sheria hizi:

- Mwisho wa wiki (Ijumaa - Jumapili): $12.00 kwa tiketi moja.
- Siku za wiki (Jumatatu - Alhamisi): $10.00 kwa tiketi moja.
- Onyesho la mchana (kabla ya saa 11 jioni): pungua $2.00 kwa tiketi moja (isipokuwa siku ya Jumanne).
- Jumatano: tiketi zote ni $5.00 kila moja.

# --hints--

`get_movie_night_cost("Saturday", "10:00pm", 1)` should return `"$12.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Saturday", "10:00pm", 1), "$12.00")`)
}})
```

`get_movie_night_cost("Sunday", "10:00am", 1)` should return `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Sunday", "10:00am", 1), "$10.00")`)
}})
```

`get_movie_night_cost("Tuesday", "7:20pm", 2)` should return `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "7:20pm", 2), "$10.00")`)
}})
```

`get_movie_night_cost("Wednesday", "5:40pm", 3)` should return `"$30.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Wednesday", "5:40pm", 3), "$30.00")`)
}})
```

`get_movie_night_cost("Monday", "11:50am", 4)` should return `"$32.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Monday", "11:50am", 4), "$32.00")`)
}})
```

`get_movie_night_cost("Friday", "4:30pm", 5)` should return `"$50.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Friday", "4:30pm", 5), "$50.00")`)
}})
```

`get_movie_night_cost("Tuesday", "11:30am", 1)` should return `"$5.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "11:30am", 1), "$5.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):

    return day
```

# --solutions--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):
    if day == "Tuesday":
        price = 5
    else:
        if day in ["Friday", "Saturday", "Sunday"]:
            price = 12
        else:
            price = 10

        hour = int(showtime.split(":")[0])
        period = showtime[-2:]

        hour24 = hour

        if period == "pm" and hour != 12:
            hour24 += 12
        if period == "am" and hour == 12:
            hour24 = 0

        if hour24 < 17:
            price -= 2

    total = price * number_of_tickets

    return f"${total:.2f}"
```
