---
id: 69b559d2903b9e4afe9075f8
title: "Challenge 238: Digit Rotation Escape"
challengeType: 29
dashedName: challenge-238
---

# --description--

Ikiwa nambari kamili chanya imetolewa, thibitisha kama nambari hiyo, au mizunguko yoyote ya nambari hiyo, inaweza kugawika sawasawa na idadi ya tarakimu zake.

Mzunguko unamaanisha kusogeza tarakimu ya kwanza mwisho. Kwa mfano, baada ya mzunguko mmoja, 123 inakuwa 231.

- Kagua mzunguko `0` (nambari iliyotolewa) kwanza.
- Nambari zilizotolewa hazitakuwa na sifuri yoyote.
- Rudisha nambari ya mzunguko wa kwanza ikiwa ipo, au `"none"` ikiwa hakuna.

# --hints--

`get_rotation(123)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(123), 0)`)
}})
```

`get_rotation(13579)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(13579), 3)`)
}})
```

`get_rotation(24681)` should return `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(24681), "none")`)
}})
```

`get_rotation(84138789345)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(84138789345), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rotation(n):

    return n
```

# --solutions--

```py
def get_rotation(n):
    s = str(n)
    digit_count = len(s)
    current = s

    for i in range(digit_count):
        if int(current) % digit_count == 0:
            return i
        current = current[1:] + current[0]

    return "none"
```
