---
id: 69c5f3d787b1725d5f00c8bd
title: "Challenge 259: FizzBuzz Explosion"
challengeType: 29
dashedName: challenge-259
---

# --description--

Ukitoa nambari kamili, rudisha idadi ya hatua zinazohitajika kubadilisha neno `"fizzbuzz"` kuwa mfuatano wa herufi wenye angalau idadi ya `"z"` kama ilivyoainishwa kwa kutumia sheria zifuatazo:

- Anza na mfuatano wa herufi `"fizzbuzz"`.
- Kila hatua, tumia sheria za kawaida za FizzBuzz kwa kutumia nafasi ya herufi katika mfuatano (herufi ya kwanza `"f"` iko katika nafasi ya 1).
  - Ikiwa nafasi ya herufi inagawika kwa 3, badilisha herufi hiyo na `"fizz"`
  - Ikiwa inagawika kwa 5, badilisha herufi hiyo na `"buzz"`
  - Ikiwa inagawika kwa 3 na 5, badilisha herufi hiyo na `"fizzbuzz"`

Hivyo baada ya hatua 1, `"fizzbuzz"` hubadilika kuwa `"fifizzzbuzzfizzzz"`, ambayo ina `"z"` tisa.

# --hints--

`explode_fizzbuzz(9)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` should return `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` should return `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
