---
id: 69cfca90e8a0a6d4d6871c53
title: "Challenge 269: Allergen Friendly Meals"
challengeType: 29
dashedName: challenge-269
---

# --description--

Kutolewa safu ya data ya milo na safu ya data ya vitu vinavyosababisha mzio vya kuepuka, rudisha majina ya milo yote ambayo haina vitu vyovyote vya mzio vilivyotajwa.

- Kila mlo uko katika muundo `[meal, allergens]`, ambapo `meal` ni jina la mlo, na `allergens` ni safu ya data ya vitu vinavyosababisha mzio vilivyomo kwenye mlo huo. Kwa mfano, `["pasta", ["wheat", "milk"]]`.
- Vitu vinavyosababisha mzio vya kuepuka vitakuwa safu ya data ya mfuatano wa herufi.

Rudisha majina ya milo salama kwa mpangilio ule ule uliopewa. Ikiwa hakuna mlo salama, rudisha safu ya data tupu.

# --hints--

`get_allergen_friendly_meals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"])` should return `["salad"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"]), ["salad"])`)
}})
```

`get_allergen_friendly_meals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"])` should return `["fried rice", "chicken parmesan"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"]), ["fried rice", "chicken parmesan"])`)
}})
```

`get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"])` should return `["oatmeal", "granola", "toast"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"]), ["oatmeal", "granola", "toast"])`)
}})
```

`get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"])` should return `["granola", "yogurt", "eggs"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"]), ["granola", "yogurt", "eggs"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_allergen_friendly_meals(meals, allergens):

    return meals
```

# --solutions--

```py
def get_allergen_friendly_meals(meals, allergens):
    return [name for name, meal_allergens in meals if not any(a in meal_allergens for a in allergens)]
```
