---
id: 69e2383af7832c8032603b8e
title: "Challenge 272: Transposed Matrix"
challengeType: 29
dashedName: challenge-272
---

# --description--

Ikiwa matrix (safu ya data ya safu za data), rudisha toleo lililobadilishwa safu na nguzo.

Ili kubadilisha matrix, badilisha safu na nguzo. Mfano: thamani iliyoko kwenye kielezo `[0, 1]` inapaswa kuhama kwenda kielezo `[1, 0]`.

Kwa mfano, ukipata:

```json
[
  [1, 2, 3],
  [4, 5, 6]
]
```

Return:

```json
[
  [1, 4],
  [2, 5],
  [3, 6]
]
```

# --hints--

`transpose([[1, 2, 3], [4, 5, 6]])` should return `[[1, 4], [2, 5], [3, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(transpose([[1, 2, 3], [4, 5, 6]]), [[1, 4], [2, 5], [3, 6]])`)
}})
```

`transpose([[1, 2], [3, 4], [5, 6]])` should return `[[1, 3, 5], [2, 4, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(transpose([[1, 2], [3, 4], [5, 6]]), [[1, 3, 5], [2, 4, 6]])`)
}})
```

`transpose([[1, 2], [3, 4], [5, 6], [7, 8]])` should return `[[1, 3, 5, 7], [2, 4, 6, 8]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(transpose([[1, 2], [3, 4], [5, 6], [7, 8]]), [[1, 3, 5, 7], [2, 4, 6, 8]])`)
}})
```

`transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]])` should return `[["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]]), [["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]])`)
}})
```

`transpose([[True, False, True, False], [False, True, False, True], [True, True, False, False], [False, False, True, True], [True, False, False, True]])` should return `[[True, False, True, False, True], [False, True, True, False, False], [True, False, False, True, False], [False, True, False, True, True]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(transpose([[True, False, True, False], [False, True, False, True], [True, True, False, False], [False, False, True, True], [True, False, False, True]]), [[True, False, True, False, True], [False, True, True, False, False], [True, False, False, True, False], [False, True, False, True, True]])`)
}})
```

# --seed--

## --seed-contents--

```py
def transpose(matrix):

    return matrix
```

# --solutions--

```py
def transpose(matrix):
    return [[row[i] for row in matrix] for i in range(len(matrix[0]))]
```
