---
id: 69f35a5bb823ed620fcb7cbb
title: "Challenge 282: Sleep Debt"
challengeType: 29
dashedName: challenge-282
---

# --description--

Ukizingatia safu ya data ya saa ulizolala kila usiku hadi leo, na nambari lengwa ya saa kwa usiku, rudisha ni saa ngapi za usingizi unazohitaji usiku wa leo kufuta deni lako la usingizi.

- Jumuisha saa za usiku wa leo katika jumla ya muda unaohitajika kufikia malengo.
- Ikiwa umelala vya kutosha kufikia au kuzidi lengo la usiku wa leo, rudisha `0`.

# --hints--

`sleep_debt([6, 6, 6, 6, 6, 6], 8)` should return `20`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 6, 6, 6, 6, 6], 8), 20)`)
}})
```

`sleep_debt([6, 7, 8, 4, 8, 6], 7)` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 7, 8, 4, 8, 6], 7), 10)`)
}})
```

`sleep_debt([10, 10, 9, 10, 9, 11], 9)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([10, 10, 9, 10, 9, 11], 9), 4)`)
}})
```

`sleep_debt([8, 7, 6, 7, 6, 8], 6)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 7, 6, 7, 6, 8], 6), 0)`)
}})
```

`sleep_debt([8, 9, 10, 9, 10, 7], 7)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 9, 10, 9, 10, 7], 7), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sleep_debt(hours_slept, target_hours):

    return hours_slept
```

# --solutions--

```py
def sleep_debt(hours_slept, target_hours):
    debt = (len(hours_slept) + 1) * target_hours - sum(hours_slept)
    return max(0, debt)
```
