---
id: 69f8c998d78ad3171a0713be
title: "Challenge 291: FizzBuzz Count"
challengeType: 29
dashedName: challenge-291
---

# --description--

Ukitoa nambari ya kuanzia na ya mwisho, hesabu idadi ya maonyesho ya fizz na buzz katika safu hiyo (ikiwa ni pamoja na nambari hizo).

- Nambari zinazogawika kwa 3 zinahesabiwa kama fizz.
- Nambari zinazogawika kwa 5 zinahesabiwa kama buzz.
- Nambari zinazogawika kwa 3 na 5 zinahesabiwa kama fizz na buzz zote mbili.

Rudisha kitu au kamusi yenye hesabu katika muundo: `{ fizz, buzz }`.

# --hints--

`fizz_buzz_count(1, 11)` should return `{"fizz": 3, "buzz": 2}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(1, 11), {"fizz": 3, "buzz": 2})`)
}})
```

`fizz_buzz_count(14, 41)` should return `{"fizz": 9, "buzz": 6}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(14, 41), {"fizz": 9, "buzz": 6})`)
}})
```

`fizz_buzz_count(24, 100)` should return `{"fizz": 26, "buzz": 16}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(24, 100), {"fizz": 26, "buzz": 16})`)
}})
```

`fizz_buzz_count(-635, -14)` should return `{"fizz": 207, "buzz": 125}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-635, -14), {"fizz": 207, "buzz": 125})`)
}})
```

`fizz_buzz_count(-5432, 6789)` should return `{"fizz": 4074, "buzz": 2444}`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-5432, 6789), {"fizz": 4074, "buzz": 2444})`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_count(start, end):

    return start
```

# --solutions--

```py
def fizz_buzz_count(start, end):
    fizz = sum(1 for i in range(start, end + 1) if i % 3 == 0)
    buzz = sum(1 for i in range(start, end + 1) if i % 5 == 0)
    return {"fizz": fizz, "buzz": buzz}
```
