---
id: 69f8c998d78ad3171a0713bd
title: "Challenge 290: Pizza Party"
challengeType: 29
dashedName: challenge-290
---

# --description--

Kutolewa safu ya data ya saa zilizofanyakaziwa leo kwa kila mtu, rudisha idadi ya pizza za kuagiza kwa ajili ya sherehe ya pizza.

- Gawanya saa zilizofanyakaziwa na kila mtu kwa 3 kupata idadi ya vipande vya pizza anavyopaswa.
- Huwezi kula kipande kidogo, kwa hivyo zungusha idadi ya vipande vya kila mtu hadi nambari nzima inayofuata.
- Kila mtu anapata angalau vipande viwili.
- Kila pizza ina vipande 8. Zungusha jumla ya pizza hadi pizza nzima inayofuata.

# --hints--

`get_pizzas_to_order([8, 8, 8])` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8]), 2)`)
}})
```

`get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10])` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10]), 3)`)
}})
```

`get_pizzas_to_order([1, 2, 3, 4, 5])` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([1, 2, 3, 4, 5]), 2)`)
}})
```

`get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8])` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8]), 3)`)
}})
```

`get_pizzas_to_order([9, 9, 6])` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([9, 9, 6]), 1)`)
}})
```

`get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0])` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_pizzas_to_order(hours_worked):

    return hours_worked
```

# --solutions--

```py
import math

def get_pizzas_to_order(hours_worked):
    total_slices = sum(max(math.ceil(h / 3), 2) for h in hours_worked)
    return math.ceil(total_slices / 8)
```
