---
id: 69f8c998d78ad3171a0713b9
title: "Challenge 286: Open Issues"
challengeType: 29
dashedName: challenge-286
---

# --description--

Ukizingatia safu ya nambari za masuala na safu nyingine ya nambari za ombi la kuvuta mabadiliko (PR), rudisha safu ya masuala ambayo bado yako wazi baada ya maombi yote ya PR kuunganishwa.

- Ombi la PR linafungua suala ikiwa nambari zao ni mzunguko wa kila mmoja. Kwa mfano, suala 123 lingefungwa na PR 231 au 312.
- Ombi la PR halifungi suala lenye nambari sawa kabisa. Kwa mfano, PR 123 halifungi suala 123. Hivyo suala lenye nambari zote sawa haliwezi kufungwa.
- Nambari yoyote inaweza kuwa na sifuri za mbele zilizokatwa. Kwa mfano, PR 201 lingefunga suala 12 (012, mzunguko wa 201). Vivyo hivyo, suala 201 lingefungwa na PR 12.

Rudisha masuala yaliyo wazi yaliyobaki kwa mpangilio waliotolewa.

# --hints--

`get_open_issues([123, 234], [231])` should return `[234]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` should return `[345, 16]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` should return `[456, 332, 15]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` should return `[115, 296, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` should return `[95, 395, 754, 296, 709, 237, 1802]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```
