---
id: 69f90c1329a94b37e2a2086c
title: "Zoezi la 293: Mkono bora"
challengeType: 29
dashedName: challenge-293
---

# --description--

Ikiwa umepewa safu ya data ya mfuatano wa herufi mitano inayowakilisha kadi za kucheza, rudisha jina la mkono bora.

- Kadi kila moja inawakilishwa kama mfuatano wa herufi wa herufi mbili: cheo kifuatiwa na rangi, `"2h"` kwa mfano.
  - Cheo, kutoka chini hadi juu, ni: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"`, na `"A"`.
  - Rangi ni: `"h"`, `"d"`, `"c"`, na `"s"`.
- Aces (`"A"`) zinaweza kutumika kama cheo cha juu au cha chini katika mfululizo.

Mikono, kwa mpangilio kutoka mbaya hadi bora, ni:

| Jina | Maelezo |
| - | - |
| `"High Card"` | Hakuna jozi au bora zaidi |
| `"Pair"` | Jozi moja ya cheo kimoja |
| `"Two Pair"` | Jozi mbili za cheo tofauti |
| `"Three of a Kind"` | Tatu za cheo kimoja |
| `"Straight"` | Cheo tano mfululizo |
| `"Flush"` | Cheo tano za rangi moja |
| `"Full House"` | Tatu za cheo kimoja, na jozi moja ya cheo kingine |
| `"Four of a Kind"` | Nne za cheo kimoja |
| `"Straight Flush"` | Cheo tano mfululizo za rangi moja |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` za rangi moja |

Rudisha jina la mkono bora.

# --hints--

`get_best_hand(["7s", "7h", "7d", "2c", "5h"])` inapaswa kurudisha `"Three of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind")`)
}})
```

`get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"])` inapaswa kurudisha `"Full House"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House")`)
}})
```

`get_best_hand(["2h", "5h", "7h", "9h", "Jh"])` inapaswa kurudisha `"Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["2h", "5h", "7h", "9h", "Jh"]), "Flush")`)
}})
```

`get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"])` inapaswa kurudisha `"Four of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind")`)
}})
```

`get_best_hand(["Ts", "Th", "9d", "9c", "8h"])` inapaswa kurudisha `"Two Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair")`)
}})
```

`get_best_hand(["9c", "8c", "7c", "6c", "5c"])` inapaswa kurudisha `"Straight Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush")`)
}})
```

`get_best_hand(["As", "Kh", "Jd", "8c", "5h"])` inapaswa kurudisha `"High Card"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Kh", "Jd", "8c", "5h"]), "High Card")`)
}})
```

`get_best_hand(["As", "2h", "3d", "4c", "5h"])` inapaswa kurudisha `"Straight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "2h", "3d", "4c", "5h"]), "Straight")`)
}})
```

`get_best_hand(["Ts", "Th", "7c", "6d", "5h"])` inapaswa kurudisha `"Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "7c", "6d", "5h"]), "Pair")`)
}})
```

`get_best_hand(["As", "Ks", "Qs", "Js", "Ts"])` inapaswa kurudisha `"Royal Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_best_hand(cards):

    return cards
```

# --solutions--

```py
def get_best_hand(cards):
    rank_order = "23456789TJQKA"

    ranks = [c[0] for c in cards]
    suits = [c[1] for c in cards]

    is_flush = len(set(suits)) == 1
    rank_indices = sorted(rank_order.index(r) for r in ranks)

    def is_straight(indices):
        is_regular = all(indices[i] == indices[i - 1] + 1 for i in range(1, len(indices)))
        is_low_ace = indices == [0, 1, 2, 3, 12]
        return is_regular or is_low_ace

    straight = is_straight(rank_indices)

    counts = {}
    for r in ranks:
        counts[r] = counts.get(r, 0) + 1
    freq = sorted(counts.values(), reverse=True)

    if is_flush and rank_indices == [8, 9, 10, 11, 12]:
        return "Royal Flush"
    if is_flush and straight:
        return "Straight Flush"
    if freq[0] == 4:
        return "Four of a Kind"
    if freq[0] == 3 and freq[1] == 2:
        return "Full House"
    if is_flush:
        return "Flush"
    if straight:
        return "Straight"
    if freq[0] == 3:
        return "Three of a Kind"
    if freq[0] == 2 and freq[1] == 2:
        return "Two Pair"
    if freq[0] == 2:
        return "Pair"
    return "High Card"
```
