---
id: 6a0dcd03ee4e68698080ef68
title: "Zoezi la 304: Mpangilio wa Ratiba ya Safari"
challengeType: 29
dashedName: challenge-304
---

# --description--

Kutolewa safu ya data yenye vituo vya hiari vya angalau viwili kwa safari ya siku moja, rudisha idadi ya mipangilio halali ya ratiba.

Ratiba daima inajumuisha `"breakfast"`, `"lunch"`, na `"dinner"`, hizi hazitapokelewa kama hoja. Vituo vya hiari vinaweza kuwekwa mahali popote katika ratiba, kwa kufuata sheria zifuatazo:

- `"breakfast"` daima huwa wa kwanza, na angalau kituo kimoja kabla ya `"lunch"`.
- `"lunch"` lazima ionekane kabla ya `"dinner"`, na angalau kituo kimoja kati yao.
- Zaidi ya yote, kituo kimoja cha hiari kinaweza kuonekana baada ya `"dinner"`.

Rudisha idadi ya mipangilio halali.

# --hints--

`get_itinerary_count(["library", "park"])` inapaswa kurudisha `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` inapaswa kurudisha `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` inapaswa kurudisha `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` inapaswa kurudisha `840`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` inapaswa kurudisha `55440`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
