---
id: 6a0dcc730cb92a616f86f0c2
title: "Zoezi la 298: Kagua muundo sehemu ya 4"
challengeType: 29
dashedName: challenge-298
---

# --description--

Ikiwa umepewa kitu (JavaScript) au kamusi (Python), amua kama kinakidhi muundo ufuatao:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean
}
```

- Alama ya bomba (`|`) ina maana ya "au". `role` lazima iwe mojawapo ya thamani zilizoorodheshwa za `Roles`.
- Alama ya swali (`?`) baada ya `supporter` ina maana kuwa sehemu hiyo ni hiari, lakini ni aina iliyobainishwa ikiwa ipo.
- Funguo za ziada zinaruhusiwa

# --hints--

`is_valid_schema({"username": "vivian", "posts": 1, "verified": False, "role": "user", "supporter": True})` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "vivian", "posts": 1, "verified": False, "role": "user", "supporter": True}), True)`)
}})
```

`is_valid_schema({"username": "rudolph", "posts": 15, "verified": True, "role": "creator"})` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "rudolph", "posts": 15, "verified": True, "role": "creator"}), True)`)
}})
```

`is_valid_schema({"username": "hernandez", "posts": 35, "verified": True, "role": "moderator", "supporter": False, "followers": 55})` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "hernandez", "posts": 35, "verified": True, "role": "moderator", "supporter": False, "followers": 55}), True)`)
}})
```

`is_valid_schema({"username": "julia", "posts": 50, "verified": True, "role": "admin", "supporter": "true"})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "julia", "posts": 50, "verified": True, "role": "admin", "supporter": "true"}), False)`)
}})
```

`is_valid_schema({"username": "bernard", "posts": 0, "verified": True, "role": "friend", "supporter": True})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bernard", "posts": 0, "verified": True, "role": "friend", "supporter": True}), False)`)
}})
```

`is_valid_schema({"username": "felix", "posts": 40, "verified": "yes", "role": "staff", "supporter": False})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "felix", "posts": 40, "verified": "yes", "role": "staff", "supporter": False}), False)`)
}})
```

`is_valid_schema({"username": "jimmy", "posts": True, "verified": False, "role": "creator", "supporter": True})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "jimmy", "posts": True, "verified": False, "role": "creator", "supporter": True}), False)`)
}})
```

`is_valid_schema({"username": True, "posts": 30, "verified": True, "role": "moderator", "supporter": False})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": True, "posts": 30, "verified": True, "role": "moderator", "supporter": False}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    roles = ["user", "creator", "moderator", "staff", "admin"]
    return (
        isinstance(obj.get("username"), str) and
        isinstance(obj.get("posts"), int) and not isinstance(obj.get("posts"), bool) and
        isinstance(obj.get("verified"), bool) and
        obj.get("role") in roles and
        ("supporter" not in obj or isinstance(obj["supporter"], bool))
    )
```
