---
id: 6a0dcd03ee4e68698080ef69
title: "Zoezi la 305: Kuweka mawazo kwa vipaumbele"
challengeType: 29
dashedName: challenge-305
---

# --description--

Imetolewa safu ya data ya pande mbili ambapo kila safu ndogo ina (kwa mpangilio huu) jina la wazo, makadirio ya matumaini, makadirio ya halisi, na makadirio ya kukata tamaa (kwa siku), rudisha safu ya majina ya mawazo yaliyopangwa kwa mujibu wa muda unaotarajiwa wa kukamilika, kuanzia mfupi zaidi.

Hesabu muda unaotarajiwa wa kukamilika kwa kila wazo kwa kutumia fomula ifuatayo:

- `expected = ((optimistic + 4 * realistic + pessimistic) / 6) * length of idea name`

# --hints--

`analyze_ideas([["Add logging", 2, 5, 15], ["SEO optimization", 4, 8, 20], ["Fix bug", 1, 3, 5]])` inapaswa kurudisha `["Fix bug", "Add logging", "SEO optimization"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(analyze_ideas([["Add logging", 2, 5, 15], ["SEO optimization", 4, 8, 20], ["Fix bug", 1, 3, 5]]), ["Fix bug", "Add logging", "SEO optimization"])`)
}})
```

`analyze_ideas([["Dark mode", 1, 3, 8], ["Real-time collaboration feature", 6, 12, 20], ["Add tooltip", 1, 2, 4]])` inapaswa kurudisha `["Add tooltip", "Dark mode", "Real-time collaboration feature"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(analyze_ideas([["Dark mode", 1, 3, 8], ["Real-time collaboration feature", 6, 12, 20], ["Add tooltip", 1, 2, 4]]), ["Add tooltip", "Dark mode", "Real-time collaboration feature"])`)
}})
```

`analyze_ideas([["Update user profile page", 3, 7, 14], ["Add pagination", 2, 5, 10], ["Add tags", 2, 3, 6], ["Fix login bug", 1, 4, 8]])` inapaswa kurudisha `["Add tags", "Fix login bug", "Add pagination", "Update user profile page"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(analyze_ideas([["Update user profile page", 3, 7, 14], ["Add pagination", 2, 5, 10], ["Add tags", 2, 3, 6], ["Fix login bug", 1, 4, 8]]), ["Add tags", "Fix login bug", "Add pagination", "Update user profile page"])`)
}})
```

`analyze_ideas([["Migrate database", 14, 25, 40], ["Add chat assistant", 8, 15, 24], ["Redesign onboarding flow", 3, 7, 13], ["Add language support", 6, 11, 18]])` inapaswa kurudisha `["Redesign onboarding flow", "Add language support", "Add chat assistant", "Migrate database"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(analyze_ideas([["Migrate database", 14, 25, 40], ["Add chat assistant", 8, 15, 24], ["Redesign onboarding flow", 3, 7, 13], ["Add language support", 6, 11, 18]]), ["Redesign onboarding flow", "Add language support", "Add chat assistant", "Migrate database"])`)
}})
```

`analyze_ideas([["Add email notifications", 3, 7, 10], ["Migrate deployment flow", 6, 10, 16], ["Add push notifications", 2, 6, 10], ["Optimize continuous integration", 5, 8, 15], ["Analyze user patterns", 5, 10, 18], ["Create onboarding curriculum", 6, 15, 25]])` inapaswa kurudisha `["Add push notifications", "Add email notifications", "Analyze user patterns", "Migrate deployment flow", "Optimize continuous integration", "Create onboarding curriculum"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(analyze_ideas([["Add email notifications", 3, 7, 10], ["Migrate deployment flow", 6, 10, 16], ["Add push notifications", 2, 6, 10], ["Optimize continuous integration", 5, 8, 15], ["Analyze user patterns", 5, 10, 18], ["Create onboarding curriculum", 6, 15, 25]]), ["Add push notifications", "Add email notifications", "Analyze user patterns", "Migrate deployment flow", "Optimize continuous integration", "Create onboarding curriculum"])`)
}})
```

# --seed--

## --seed-contents--

```py
def analyze_ideas(ideas):

    return ideas
```

# --solutions--

```py
def analyze_ideas(ideas):
    def expected(idea):
        name, optimistic, realistic, pessimistic = idea
        return ((optimistic + 4 * realistic + pessimistic) / 6) * len(name)
    return [idea[0] for idea in sorted(ideas, key=expected)]
```
