---
id: 6a0dcc730cb92a616f86f0c0
title: "Zoezi la 296: Kihakiki cha Mchoro Sehemu ya 2"
challengeType: 29
dashedName: challenge-296
---

# --description--

Ikiwa umepewa kitu (JavaScript) au kamusi (Python), amua kama kinakidhi mchoro ufuatao:

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- Funguo za ziada zinaruhusiwa

# --hints--

`is_valid_schema({"username": "alice", "posts": 10, "verified": False})` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "alice", "posts": 10, "verified": False}), True)`)
}})
```

`is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25})` inapaswa kurudisha `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25}), True)`)
}})
```

`is_valid_schema({"username": "frank", "posts": "21", "verified": True})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "frank", "posts": "21", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "sam", "posts": 17, "verified": "false"})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "sam", "posts": 17, "verified": "false"}), False)`)
}})
```

`is_valid_schema({"username": "bill", "verified": True})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bill", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "fred", "verified": True})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "fred", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": 5, "posts": 10, "verified": True})` inapaswa kurudisha `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 5, "posts": 10, "verified": True}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return (
        isinstance(obj.get("username"), str) and
        isinstance(obj.get("posts"), int) and not isinstance(obj.get("posts"), bool) and
        isinstance(obj.get("verified"), bool)
    )
```
