---
id: 6a0dcd03ee4e68698080ef66
title: "Zoezi la 302: Kuathirika na mabadiliko ya saa za ndege"
challengeType: 29
dashedName: challenge-302
---

# --description--

Ukizingatia jiji la kuondoka, jiji la kufikia, muda wa safari kwa saa, na mwelekeo wa kusafiri, rudisha idadi ya saa za kuathirika na mabadiliko ya saa za ndege (jet lag) ambazo msafiri anazipata.

Miji iliyotolewa itakuwa kutoka kwenye orodha ifuatayo inayojumuisha mbali zao za saa za UTC:

| Jiji | Mbali |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Ili kuhesabu saa za kuathirika na mabadiliko ya saa za ndege:

1. Tafuta tofauti ya saa za mkoa kati ya miji miwili.
2. Amua kigezo cha mwelekeo. Ikiwa unasafiri `"east"`, ni 1.5, vinginevyo ni 1.0.
3. Pata saa za kuathirika kwa kutumia fomula: tofauti ya saa za mkoa + (muda wa safari * 0.1) * kigezo cha mwelekeo

Rudisha saa za kuathirika zikiwa zimezungushwa hadi sehemu moja ya desimali.

# --hints--

`get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east")` inapaswa kurudisha `6.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east"), 6.5)`)
}})
```

`get_jet_lag_hours("London", "New York", 8, "west")` inapaswa kurudisha `5.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("London", "New York", 8, "west"), 5.8)`)
}})
```

`get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east")` inapaswa kurudisha `1.6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east"), 1.6)`)
}})
```

`get_jet_lag_hours("Dubai", "London", 7, "west")` inapaswa kurudisha `4.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Dubai", "London", 7, "west"), 4.7)`)
}})
```

`get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west")` inapaswa kurudisha `17.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west"), 17.5)`)
}})
```

`get_jet_lag_hours("Tokyo", "Dubai", 9, "west")` inapaswa kurudisha `5.9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Tokyo", "Dubai", 9, "west"), 5.9)`)
}})
```

`get_jet_lag_hours("New York", "Istanbul", 10, "east")` inapaswa kurudisha `9.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("New York", "Istanbul", 10, "east"), 9.5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):

    return departure_city
```

# --solutions--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):
    offsets = {
        "Los Angeles": -8,
        "New York": -5,
        "London": 0,
        "Istanbul": 3,
        "Dubai": 4,
        "Hong Kong": 8,
        "Tokyo": 9
    }
    timezone_diff = abs(offsets[arrival_city] - offsets[departure_city])
    multiplier = 1.5 if direction == "east" else 1.0
    score = timezone_diff + (flight_duration * 0.1) * multiplier
    return round(score, 1)
```
