---
id: 6a0dcd03ee4e68698080ef6b
title: "Zoezi la 307: Kanuni za Mipangilio ya Miji"
challengeType: 29
dashedName: challenge-307
---

# --description--

Ikiwa umepewa gridi ya 2D (safu ya safu za data) inayowakilisha mpangilio wa majengo ya jiji, rudisha kuratibu za majengo yote yanayokiuka kanuni za mipangilio ya miji.

Kila seli katika gridi ina moja ya lebo kutoka kwenye jedwali hapa chini. Jengo linakiuka kanuni ikiwa yoyote ya majirani zake (hadi nne), wima au mlalo, ni aina ambayo haliwezi kuwa jirani nayo.

| Lebo | Aina | Haiwezi kuwa jirani na |
| - | - | - |
| `"i"` | viwanda | `"R"`, `"I"` |
| `"A"` | kilimo | `"C"` |
| `"R"` | makazi | `"i"`, `"C"` |
| `"I"` | taasisi | `"i"` |
| `"C"` | biashara | `"R"`, `"A"` |
| `""` (mfuatano wa herufi tupu) | haijajengwa | hakuna vikwazo |

Rudisha kuratibu za seli zote zinazokiuka kama safu ya jozi za `[row, col]`, kwa mpangilio wowote. Ikiwa hakuna ukiukaji wowote, rudisha safu tupu.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` inapaswa kurudisha `[[0, 0], [0, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` inapaswa kurudisha `[[0, 1], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` inapaswa kurudisha `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` inapaswa kurudisha `[[0, 1], [0, 2], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` inapaswa kurudisha `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
