---
id: 6a151d32d772271dd2448c2d
title: "Zoezi la 311: Mchawi wa maneno"
challengeType: 29
dashedName: challenge-311
---

# --description--

Kutolewa mfuatano wa misimbo ya maneno unayoitumia, hesabu jumla ya alama.

Kila herufi katika mfuatano inawakilisha neno:

| Msimbo | Neno | Kategoria | Alama ya msingi |
| - | - | - | - |
| `"f"` | Moto | Uharibifu | 3 |
| `"l"` | Umeme | Uharibifu | 3 |
| `"i"` | Barafu | Udhibiti | 2 |
| `"w"` | Upepo | Udhibiti | 2 |
| `"h"` | Tiba | Urejeshaji | 1 |
| `"s"` | Mlinzi | Urejeshaji | 1 |

Kiongezaji cha mchanganyiko kinawekwa kulingana na ni maneno mangapi mfululizo yameitumiwa kutoka katika makategoria tofauti:

- Neno la kwanza daima linapata alama ya msingi.
- Kila neno linalofuata kutoka kategoria tofauti na lile lililotangulia huongeza kiongezaji kwa 1.
- Kutumia neno kutoka kategoria ile ile kama lile lililotangulia kunarejesha kiongezaji kuwa 1.
- Alama kwa kila neno ni alama yake ya msingi ikizidishwa na kiongezaji cha sasa.

Rudisha jumla ya alama kutoka kwa mfululizo wa maneno.

# --hints--

`cast("fihwl")` inapaswa kurudisha `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` inapaswa kurudisha `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` inapaswa kurudisha `37`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` inapaswa kurudisha `50`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` inapaswa kurudisha `101`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
