---
id: 6a15cadf5f240d05a264955b
title: "Zoezi la 317: Kihesabu cha BMI"
challengeType: 29
dashedName: challenge-317
---

# --description--

Ukichukulia uzito kwa pauni na urefu kwa inchi, rudisha BMI (Kielezo cha Uzito wa Mwili) kilichozungushwa hadi sehemu moja ya desimali.

Ili kupata BMI: gawanya uzito kwa urefu uliyozidishwa na mwenyewe, kisha zidisha matokeo na 703.

# --hints--

`calculate_bmi(180, 70)` inapaswa kurudisha `25.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(180, 70), 25.8)`)
}})
```

`calculate_bmi(140, 64)` inapaswa kurudisha `24.0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(140, 64), 24.0)`)
}})
```

`calculate_bmi(160, 76)` inapaswa kurudisha `19.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(160, 76), 19.5)`)
}})
```

`calculate_bmi(200, 60)` inapaswa kurudisha `39.1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(200, 60), 39.1)`)
}})
```

`calculate_bmi(150, 68)` inapaswa kurudisha `22.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(150, 68), 22.8)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_bmi(weight, height):

    return weight
```

# --solutions--

```py
def calculate_bmi(weight, height):
    return round((weight / (height * height)) * 703, 1)
```
