---
id: 6a15cadf5f240d05a264955f
title: "Changamoto 321: Ustafishaji wa Kipindi"
challengeType: 29
dashedName: challenge-321
---

# --description--

Ipe neno, amua kama linaweza kuandikwa kwa kutumia alama za vipengele kutoka kwenye jedwali la vipengele vya kipindi.

- Puuza tofauti za herufi kubwa na ndogo wakati wa kuandika neno. `"neon"` linaweza kuandikwa kwa alama `"Ne"`, `"O"`, na `"N"`.

Hapa kuna orodha kamili ya alama za vipengele:

```json
["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];
```

Rudisha safu ya data ya vipengele vilivyotumika kuandika neno, kwa herufi zao za asili na kwa mpangilio wa kuandika neno. Au, safu ya data tupu ikiwa haiwezi kuandikwa.

# --hints--

`get_periodic_spelling("neon")` inapaswa kurudisha `["Ne", "O", "N"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("neon"), ["Ne", "O", "N"])`)
}})
```

`get_periodic_spelling("rational")` inapaswa kurudisha `["Ra", "Ti", "O", "N", "Al"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("rational"), ["Ra", "Ti", "O", "N", "Al"])`)
}})
```

`get_periodic_spelling("yarn")` inapaswa kurudisha `["Y", "Ar", "N"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("yarn"), ["Y", "Ar", "N"])`)
}})
```

`get_periodic_spelling("carbon")` inapaswa kurudisha `["C", "Ar", "B", "O", "N"]` au `["Ca", "Rb", "O", "N"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("carbon"), [["C", "Ar", "B", "O", "N"], ["Ca", "Rb", "O", "N"]])`)
}})
```

`get_periodic_spelling("noisy")` inapaswa kurudisha `["N", "O", "I", "S", "Y"]` au `["No", "I", "S", "Y"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("noisy"), [["N", "O", "I", "S", "Y"], ["No", "I", "S", "Y"]])`)
}})
```

`get_periodic_spelling("bicycles")` inapaswa kurudisha `["B", "I", "C", "Y", "Cl", "Es"]` au `["Bi", "C", "Y", "Cl", "Es"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("bicycles"), [["B", "I", "C", "Y", "Cl", "Es"], ["Bi", "C", "Y", "Cl", "Es"]])`)
}})
```

`get_periodic_spelling("optics")` inapaswa kurudisha `["O", "P", "Ti", "C", "S"]`, `["O", "P", "Ti", "Cs"]`, `["O", "Pt", "I", "C", "S"]`, au `["O", "Pt", "I", "Cs"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("optics"), [["O", "P", "Ti", "C", "S"], ["O", "P", "Ti", "Cs"], ["O", "Pt", "I", "C", "S"], ["O", "Pt", "I", "Cs"]])`)
}})
```

`get_periodic_spelling("value")` inapaswa kurudisha `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("value"), [])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_periodic_spelling(word):

    return word
```

# --solutions--

```py
def get_periodic_spelling(word):
    elements = ["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"]
    lower = {e.lower(): e for e in elements}

    def spell(w, i):
        if i == len(w):
            return []
        two = w[i:i + 2]
        one = w[i:i + 1]
        if two in lower:
            rest = spell(w, i + 2)
            if rest is not None:
                return [lower[two]] + rest
        if one in lower:
            rest = spell(w, i + 1)
            if rest is not None:
                return [lower[one]] + rest
        return None

    result = spell(word.lower(), 0)
    return result if result is not None else []
```
