---
id: 6a15cadf5f240d05a2649558
title: "Changamoto 314: Gawanyio msingi"
challengeType: 29
dashedName: challenge-314
---

# --description--

Ikiwa umepewa nambari kamili kubwa kuliko 1, rudisha gawanyio msingi wake kama safu ya nambari kwa mpangilio wa kuongezeka.

Gawanyio msingi ni seti ya namba msingi zinazozidiana pamoja kutoa nambari kamili iliyotolewa. Kila nambari ina seti moja tu. Kwa mfano, gawanyio msingi la 20 ni `[2, 2, 5]` kwa sababu 2 * 2 * 5 = 20.

Ikiwa nambari kamili iliyotolewa ni namba msingi yenyewe, irudishe katika safu yenye kipengele kimoja.

# --hints--

`prime_factorization(20)` inapaswa kurudisha `[2, 2, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(20), [2, 2, 5])`)
}})
```

`prime_factorization(17)` inapaswa kurudisha `[17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(17), [17])`)
}})
```

`prime_factorization(15)` inapaswa kurudisha `[3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(15), [3, 5])`)
}})
```

`prime_factorization(35)` inapaswa kurudisha `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(35), [5, 7])`)
}})
```

`prime_factorization(999)` inapaswa kurudisha `[3, 3, 3, 37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(999), [3, 3, 3, 37])`)
}})
```

`prime_factorization(360)` inapaswa kurudisha `[2, 2, 2, 3, 3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(360), [2, 2, 2, 3, 3, 5])`)
}})
```

`prime_factorization(510510)` inapaswa kurudisha `[2, 3, 5, 7, 11, 13, 17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(510510), [2, 3, 5, 7, 11, 13, 17])`)
}})
```

# --seed--

## --seed-contents--

```py
def prime_factorization(n):

    return n
```

# --solutions--

```py
def prime_factorization(n):
    factors = []
    divisor = 2
    while divisor * divisor <= n:
        while n % divisor == 0:
            factors.append(divisor)
            n //= divisor
        divisor += 1
    if n > 1:
        factors.append(n)
    return factors
```
