---
id: 6a15cadf5f240d05a264955c
title: "Changamoto 318: Mabadiliko ya DNA"
challengeType: 29
dashedName: challenge-318
---

# --description--

Ukizingatia safu mbili za DNA zenye urefu sawa, rudisha safu ya data ya vielezo ambavyo safu hizo zinatofautiana (mabadiliko).

- Safu za DNA ni mfuatano wa herufi zilizoundwa na herufi `"A"`, `"T"`, `"C"`, na `"G"`
- Rudisha vielezo kwa mpangilio wa kuongezeka
- Ikiwa hakuna mabadiliko, rudisha safu ya data tupu

# --hints--

`detect_mutations("ATCG", "ATGG")` inapaswa kurudisha `[2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ATCG", "ATGG"), [2])`)
}})
```

`detect_mutations("ATGCGTACGTTAGC", "ATGCATACGATTGC")` inapaswa kurudisha `[4, 9, 11]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ATGCGTACGTTAGC", "ATGCATACGATTGC"), [4, 9, 11])`)
}})
```

`detect_mutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG")` inapaswa kurudisha `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG"), [])`)
}})
```

`detect_mutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG")` inapaswa kurudisha `[15, 16, 17, 18]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG"), [15, 16, 17, 18])`)
}})
```

`detect_mutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT")` inapaswa kurudisha `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28])`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_mutations(strand1, strand2):

    return strand1
```

# --solutions--

```py
def detect_mutations(strand1, strand2):
    return [i for i, (a, b) in enumerate(zip(strand1, strand2)) if a != b]
```
