---
id: 6a15cadf5f240d05a2649557
title: "Changamoto 313: Gharama ya Kodi"
challengeType: 29
dashedName: challenge-313
---

# --description--

Ukitoa tarehe na saa ya kukodisha, tarehe na saa ya kurudisha, na kiwango cha kodi, rudisha jumla ya gharama ya kodi ikijumuisha ada zozote za kuchelewa.

- Tarehe na saa zilizotolewa ni mfuatano wa herufi wa ISO za UTC, kwa mfano: `"2026-06-18T18:30:00Z"`.
- Kiwango cha kodi ni idadi ya siku kabla ya kodi kurudishwa: `1`, `3`, au `7`.
- Kodi zinapaswa kurudishwa ifikapo saa 12:00 PM UTC au kabla yake siku ya mwisho ya kipindi cha kodi. Kwa mfano, kodi ya siku 1 iliyochukuliwa wakati wowote tarehe 15 Machi inapaswa kurudishwa ifikapo saa 12:00 PM UTC tarehe 16 Machi.
- Kila siku inayopita baada ya tarehe na saa ya kurudisha husababisha ada ya kuchelewa.

Bei ni kama ifuatavyo:

| Kiwango | Gharama ya msingi | Ada ya kuchelewa kwa siku |
| - | - | - |
| Siku 1 | $4.99 | $3.99 |
| Siku 3 | $3.99 | $2.99 |
| Siku 7 | $2.99 | $0.99 |

Rudisha jumla ya gharama iliyozungushwa hadi sehemu mbili za desimali kwa muundo `"$D.CC"`.

# --hints--

`get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` inapaswa kurudisha `"$4.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99")`)
}})
```

`get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` inapaswa kurudisha `"$12.97"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97")`)
}})
```

`get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` inapaswa kurudisha `"$3.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99")`)
}})
```

`get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` inapaswa kurudisha `"$9.97"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97")`)
}})
```

`get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` inapaswa kurudisha `"$2.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99")`)
}})
```

`get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` inapaswa kurudisha `"$358.40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rental_cost(rented, returned, tier):

    return rented
```

# --solutions--

```py
from datetime import datetime, timedelta
import math

def get_rental_cost(rented, returned, tier):
    pricing = {
        1: {"base": 4.99, "late": 3.99},
        3: {"base": 3.99, "late": 2.99},
        7: {"base": 2.99, "late": 0.99},
    }
    rented_dt = datetime.fromisoformat(rented.replace("Z", "+00:00"))
    returned_dt = datetime.fromisoformat(returned.replace("Z", "+00:00"))
    due = rented_dt + timedelta(days=tier)
    due = due.replace(hour=12, minute=0, second=0, microsecond=0)
    base = pricing[tier]["base"]
    late = pricing[tier]["late"]
    if returned_dt <= due:
        return f"${base:.2f}"
    days_late = math.ceil((returned_dt - due).total_seconds() / 86400)
    return f"${base + days_late * late:.2f}"
```
