---
id: 6a15cadf5f240d05a264955d
title: "Zoezi la 319: Mchambuzi wa Frontmatter"
challengeType: 29
dashedName: challenge-319
---

# --description--

Ukitoa mfuatano wa herufi unaowakilisha kipande cha frontmatter, chambua na urudishe kitu (JavaScript) au kamusi (Python) yenye funguo na thamani.

Frontmatter imefungwa ndani ya vizingiti vya `---` na ina jozi za `key: value` ndani yake, moja kwa mstari. Kwa mfano:

```md
---
title: My Post
draft: false
views: 100
---
```

Inapaswa kurudisha:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- Nambari, thamani za kweli au siyo kweli, na mfuatano wa herufi zote zinapaswa kurudishwa kama aina zao husika.
- Mfuatano uliotolewa utakuwa na mistari mipya iliyotenganishwa na tabia ya mistari mipya (`"\n"`). Mfano ulio juu utatolewa kama: `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`.

# --hints--

`parse_frontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` inapaswa kurudisha `{ title: "My Post", draft: False, views: 100 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\ntitle: My Post\\ndraft: false\\nviews: 100\\n---"), {"title": "My Post", "draft": False, "views": 100})`)
}})
```

`parse_frontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` inapaswa kurudisha `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: False }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nid: 6a174db57256a112f932195c\\ntitle: My Book\\nlocale: en\\nwordCount: 10000\\npublished: false\\n---"), {"id": "6a174db57256a112f932195c", "title": "My Book", "locale": "en", "wordCount": 10000, "published": False})`)
}})
```

`parse_frontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` inapaswa kurudisha `{ version: "1.0.0", url: "https://example.com", private: True }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nversion: 1.0.0\\nurl: https://example.com\\nprivate: true\\n---"), {"version": "1.0.0", "url": "https://example.com", "private": True})`)
}})
```

`parse_frontmatter("---\nrating: 4.5\nprice: 9.99\n---")` inapaswa kurudisha `{ rating: 4.5, price: 9.99 }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nrating: 4.5\\nprice: 9.99\\n---"), {"rating": 4.5, "price": 9.99})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_frontmatter(s):

    return s
```

# --solutions--

```py
def parse_frontmatter(s):
    lines = [line for line in s.split("\n") if line != "---" and line != ""]
    result = {}
    for line in lines:
        key, *rest = line.split(":")
        value = ":".join(rest).strip()
        if value == "true":
            result[key.strip()] = True
        elif value == "false":
            result[key.strip()] = False
        else:
            try:
                n = float(value)
                result[key.strip()] = int(n) if n == int(n) else n
            except ValueError:
                result[key.strip()] = value
    return result
```
