---
id: 6a15cadf5f240d05a2649556
title: "Changamoto 312: Gharama ya Kuangalia Moja kwa Moja"
challengeType: 29
dashedName: challenge-312
---

# --description--

Kutolewa safu ya data inayowakilisha filamu zilizo kwenye kikapu cha huduma yako ya kuangalia moja kwa moja, na mfuatano wa herufi kwa ngazi yako ya usajili, rudisha jumla ya gharama za filamu hizo.

Kila kipengee katika kikapu ni kitu chenye `"format"` (`"HD"` au `"4K"`) na `"type"` (`"rent"` au `"buy"`). Gharama zao ni:

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

Tumia punguzo zifuatazo za ngazi ya usajili:

- `"none"`: bei kamili
- `"basic"`: punguzo la 10%
- `"premium"`: punguzo la 25%

Rudisha jumla ya gharama iliyozungushwa hadi sehemu mbili za desimali kwa muundo `"$D.CC"`.

# --hints--

`get_streaming_bill([{ "format": "HD", "type": "rent" }], "none")` inapaswa kurudisha `"$3.99"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}], "none"), "$3.99")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "buy" }], "premium")` inapaswa kurudisha `"$9.74"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "buy"}], "premium"), "$9.74")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "HD", "type": "rent" }, { "format": "HD", "type": "buy" }], "basic")` inapaswa kurudisha `"$18.87"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "HD", "type": "rent"}, {"format": "HD", "type": "buy"}], "basic"), "$18.87")`)
}})
```

`get_streaming_bill([{ "format": "4K", "type": "buy" }, { "format": "4K", "type": "buy" }, { "format": "4K", "type": "buy" }], "premium")` inapaswa kurudisha `"$44.98"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "4K", "type": "buy"}, {"format": "4K", "type": "buy"}, {"format": "4K", "type": "buy"}], "premium"), "$44.98")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "4K", "type": "rent" }, { "format": "HD", "type": "buy" }, { "format": "4K", "type": "buy" }], "none")` inapaswa kurudisha `"$42.96"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}], "none"), "$42.96")`)
}})
```

`get_streaming_bill([{ "format": "HD", "type": "rent" }, { "format": "4K", "type": "rent" }, { "format": "HD", "type": "buy" }, { "format": "4K", "type": "buy" }], "basic")` inapaswa kurudisha `"$38.66"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}], "basic"), "$38.66")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_streaming_bill(cart, subscription):

    return cart
```

# --solutions--

```py
def get_streaming_bill(cart, subscription):
    pricing = {
        "HD": {"rent": 3.99, "buy": 12.99},
        "4K": {"rent": 5.99, "buy": 19.99},
    }
    discounts = {"none": 1, "basic": 0.9, "premium": 0.75}
    multiplier = discounts[subscription]
    total = sum(pricing[item["format"]][item["type"]] * multiplier for item in cart)
    return f"${total:.2f}"
```
