---
id: 6a19dd69388e5a6e59e4f2ef
title: "Changamoto 326: Faida Kuu"
challengeType: 29
dashedName: challenge-326
---

# --description--

Ukizingatia safu ya data ya bei za hisa kila siku na bajeti (katika dola), hesabu faida kubwa zaidi unayoweza kupata kwa kununua na kuuza hisa katika kipindi kilichotolewa.

- Unaweza kuuza tu baada ya kununua.
- Unaweza kufanya ununuzi mmoja na uuzaji mmoja tu. Mara tu unapouza, huwezi kununua tena.
- Unaweza kununua hisa kamili tu.
- Rudisha faida kubwa zaidi inayowezekana kama mfuatano wa herufi, ukizungusha chini hadi senti iliyo karibu na kuipanga kwa sehemu mbili za desimali.

# --hints--

`get_max_profit([5, 6], 50)` inapaswa kurudisha `"10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([5, 6], 50), "10.00")`)
}})
```

`get_max_profit([8, 2, 5, 10], 20)` inapaswa kurudisha `"80.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([8, 2, 5, 10], 20), "80.00")`)
}})
```

`get_max_profit([4, 5, 3, 6], 20)` inapaswa kurudisha `"18.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([4, 5, 3, 6], 20), "18.00")`)
}})
```

`get_max_profit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200)` inapaswa kurudisha `"8.31"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200), "8.31")`)
}})
```

`get_max_profit([15.38, 15.01, 14.99, 14.62, 14.28], 80)` inapaswa kurudisha `"0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([15.38, 15.01, 14.99, 14.62, 14.28], 80), "0.00")`)
}})
```

`get_max_profit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25)` inapaswa kurudisha `"73.80"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25), "73.80")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_max_profit(prices, budget):

    return prices
```

# --solutions--

```py
def get_max_profit(prices, budget):
    max_profit = 0

    for buy in range(len(prices) - 1):
        for sell in range(buy + 1, len(prices)):
            shares = int(budget / prices[buy])
            profit = shares * (prices[sell] - prices[buy])
            if profit > max_profit:
                max_profit = profit

    return f"{int(max_profit * 100) / 100:.2f}"
```
