---
id: 6a19b1062d1b153d8ac76d73
title: "Zoezi la 325: Nambari ya Bahati"
challengeType: 29
dashedName: challenge-325
---

# --description--

Kutolewa mfuatano wa herufi wa jina la kwanza na la mwisho la mtu, hesabu nambari yake ya bahati kwa kutumia sheria zifuatazo:

- Majina ya kwanza na ya mwisho yamegawanyika kwa nafasi
- Tafuta idadi ya vokali na konsonanti kwa kila jina
- Zidisha idadi ndogo ya vokali na konsonanti kwa kila mmoja kisha zidisha na urefu wa jina dogo
- Fanya vivyo hivyo kwa idadi kubwa mbili na jina kubwa
- Toa thamani ndogo kutoka kwa kubwa kupata nambari yao ya bahati

Ikiwa thamani ya mwisho ni sifuri (`0`), rudisha `13`.

# --hints--

`get_lucky_number("John Doe")` inapaswa kurudisha `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("John Doe"), 21)`)
}})
```

`get_lucky_number("Olivia Lewis")` inapaswa kurudisha `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Olivia Lewis"), 52)`)
}})
```

`get_lucky_number("James Wilson")` inapaswa kurudisha `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("James Wilson"), 18)`)
}})
```

`get_lucky_number("Elizabeth Hernandez")` inapaswa kurudisha `81`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Elizabeth Hernandez"), 81)`)
}})
```

`get_lucky_number("Mike Walker")` inapaswa kurudisha `32`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Mike Walker"), 32)`)
}})
```

`get_lucky_number("Chloe Perez")` inapaswa kurudisha `13`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Chloe Perez"), 13)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_lucky_number(name):

    return name
```

# --solutions--

```py
def get_lucky_number(name):
    vowels = "aeiouAEIOU"

    def name_stats(s):
        vowel_count = sum(1 for c in s if c in vowels)
        consonant_count = sum(1 for c in s if c.isalpha() and c not in vowels)
        return vowel_count, consonant_count, len(s)

    first, last = name.split(" ")
    first_vowels, first_consonants, first_len = name_stats(first)
    last_vowels, last_consonants, last_len = name_stats(last)

    larger_vowels = max(first_vowels, last_vowels)
    smaller_vowels = min(first_vowels, last_vowels)
    larger_consonants = max(first_consonants, last_consonants)
    smaller_consonants = min(first_consonants, last_consonants)
    larger_length = max(first_len, last_len)
    smaller_length = min(first_len, last_len)

    larger_value = larger_vowels * larger_consonants * larger_length
    smaller_value = smaller_vowels * smaller_consonants * smaller_length
    result = larger_value - smaller_value

    return 13 if result == 0 else result
```
