---
id: 6a19b1062d1b153d8ac76d70
title: "Changamoto 322: Unganisha 3"
challengeType: 29
dashedName: challenge-322
---

# --description--

Kutolewa matriki ya mfuatano wa herufi zinazowakilisha vipande kwenye gridi ya mchezo, amua kama mchezaji yeyote ana mfuatano wa tatu mfululizo.

- Kila seli ina `"R"`, `"Y"`, au `""` (mfuatano wa herufi tupu).
- Mfuatano wa tatu mfululizo unamaanisha seli tatu zisizo tupu za aina ile ile kwa mlalo, wima, au mwinuko.

Rudisha:

- Safu ya data tambarare yenye mshindi na kuratibu za seli zao tatu za ushindi kwa muundo: `["R", [0,2], [1,3], [2,4]]`. Kuratibu hurudishwa kutoka juu kwenda chini, kisha kushoto kwenda kulia.
- Safu ya data tupu ikiwa hakuna mshindi.

# --hints--

`connect_three([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]])` inapaswa kurudisha `["R", [3, 1], [3, 2], [3, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]]), ["R", [3, 1], [3, 2], [3, 3]])`)
}})
```

`connect_three([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]])` inapaswa kurudisha `["Y", [1, 1], [2, 1], [3, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]]), ["Y", [1, 1], [2, 1], [3, 1]])`)
}})
```

`connect_three([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]])` inapaswa kurudisha `["R", [0, 3], [1, 2], [2, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]]), ["R", [0, 3], [1, 2], [2, 1]])`)
}})
```

`connect_three([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]])` inapaswa kurudisha `["Y", [0, 1], [1, 2], [2, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]]), ["Y", [0, 1], [1, 2], [2, 3]])`)
}})
```

`connect_three([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]])` inapaswa kurudisha `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]]), [])`)
}})
```

# --seed--

## --seed-contents--

```py
def connect_three(matrix):

    return matrix
```

# --solutions--

```py
def connect_three(matrix):
    rows = len(matrix)
    cols = len(matrix[0])

    directions = [[0, 1], [1, 0], [1, 1], [1, -1]]

    for r in range(rows):
        for c in range(cols):
            piece = matrix[r][c]
            if not piece:
                continue

            for dr, dc in directions:
                cells = [[r, c]]

                for i in range(1, 3):
                    nr = r + dr * i
                    nc = c + dc * i
                    if nr < 0 or nr >= rows or nc < 0 or nc >= cols:
                        break
                    if matrix[nr][nc] != piece:
                        break
                    cells.append([nr, nc])

                if len(cells) == 3:
                    cells.sort(key=lambda x: (x[0], x[1]))
                    return [piece] + cells

    return []
```
