---
id: 6a1d9f98e819ed70a0e994dc
title: "Zoezi la 330: maneno ya herufi ndogo"
challengeType: 29
dashedName: challenge-330
---

# --description--

Ukipata mfuatano wa herufi, rudisha maneno ambayo yote ni ya herufi ndogo tu, kwa mpangilio wake wa awali na ukiwa na nafasi kati ya kila neno.

# --hints--

`get_lowercase_words("hello GOOD world")` inapaswa kurudisha `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("hello GOOD world"), "hello world")`)
}})
```

`get_lowercase_words("these are all lowercase")` inapaswa kurudisha `"these are all lowercase"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("these are all lowercase"), "these are all lowercase")`)
}})
```

`get_lowercase_words("less is NoT more")` inapaswa kurudisha `"less is more"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("less is NoT more"), "less is more")`)
}})
```

`get_lowercase_words("DonT eat pizza every OTHER day")` inapaswa kurudisha `"eat pizza every day"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("DonT eat pizza every OTHER day"), "eat pizza every day")`)
}})
```

`get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` inapaswa kurudisha `"the quick brown fox jumped over the lazy dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_lowercase_words(s):

    return s
```

# --solutions--

```py
def get_lowercase_words(s):
    return " ".join(word for word in s.split(" ") if word == word.lower())
```
