---
id: 6a1d9f98e819ed70a0e994df
title: "Changamoto 333: Uainishaji wa Masuala 2"
challengeType: 29
dashedName: challenge-333
---

# --description--

Kutolewa kichwa cha suala na safu ya lebo za sasa, rudisha safu iliyosasishwa ya lebo kulingana na sheria zifuatazo:

Kama suala halina lebo yoyote, ongeza:

- `"bug"` na `"needs triage"` ikiwa kichwa kina `"error"` au `"bug"`
- `"enhancement"` na `"discussing"` ikiwa kichwa kina `"feature"` au `"add"`

Vinginevyo, kama lebo zilizotolewa zina:

- `"needs triage"` na kichwa kina `"simple"` au `"easy"`, ondoa `"needs triage"` na ongeza `"good first issue"`
- `"discussing"` na kichwa kina `"planned"` au `"next"`, ondoa `"discussing"` na ongeza `"on the roadmap"`
- Vinginevyo, kama `"needs triage"` au `"discussing"` ipo, iondoe na ongeza `"help wanted"`

Kama kichwa kina:

- `"security"`, ongeza lebo ya `"critical"`

# --hints--

`triage_issue("app crashes with error", [])` inapaswa kurejesha `["bug", "needs triage"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", [])), sorted(["bug", "needs triage"]))`)
}})
```

`triage_issue("app crashes with error", ["bug", "needs triage"])` inapaswa kurejesha `["bug", "help wanted"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", ["bug", "needs triage"])), sorted(["bug", "help wanted"]))`)
}})
```

`triage_issue("add dark mode", [])` inapaswa kurudisha `["enhancement", "discussing"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", [])), sorted(["enhancement", "discussing"]))`)
}})
```

`triage_issue("add dark mode", ["enhancement", "discussing"])` inapaswa kurudisha `["enhancement", "help wanted"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted"]))`)
}})
```

`triage_issue("xss security bug", [])` inapaswa kurejesha `["bug", "needs triage", "critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("xss security bug", [])), sorted(["bug", "needs triage", "critical"]))`)
}})
```

`triage_issue("security vulnerability in auth", [])` inapaswa kurudisha `["critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("security vulnerability in auth", [])), sorted(["critical"]))`)
}})
```

`triage_issue("easy a11y fix", ["bug", "needs triage"])` inapaswa kurejesha `["bug", "good first issue"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("easy a11y fix", ["bug", "needs triage"])), sorted(["bug", "good first issue"]))`)
}})
```

`triage_issue("planned api migration", ["enhancement", "discussing"])` inapaswa kurudisha `["enhancement", "on the roadmap"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("planned api migration", ["enhancement", "discussing"])), sorted(["enhancement", "on the roadmap"]))`)
}})
```

`triage_issue("improve security", ["enhancement", "discussing"])` inapaswa kurudisha `["enhancement", "help wanted", "critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("improve security", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted", "critical"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(title, labels):

    return title
```

# --solutions--

```py
def triage_issue(title, labels):
    result = list(labels)
    t = title.lower()

    def has(label):
        return label in result

    def add(label):
        if not has(label):
            result.append(label)

    def remove(label):
        if label in result:
            result.remove(label)

    if len(result) == 0:
        if "error" in t or "bug" in t:
            add("bug")
            add("needs triage")
        elif "feature" in t or "add" in t:
            add("enhancement")
            add("discussing")
    else:
        if has("needs triage") and ("simple" in t or "easy" in t):
            remove("needs triage")
            add("good first issue")
        elif has("discussing") and ("planned" in t or "next" in t):
            remove("discussing")
            add("on the roadmap")
        elif has("needs triage") or has("discussing"):
            remove("needs triage")
            remove("discussing")
            add("help wanted")

    if "security" in t:
        add("critical")

    return result
```
