---
id: 6a1d9f98e819ed70a0e994dd
title: "Zoezi la 331: Mara ya Karibu Zaidi"
challengeType: 29
dashedName: challenge-331
---

# --description--

Ukizingatia nambari kamili mbili, punguza ya kwanza hadi mara ya karibu zaidi ya ya pili.

# --hints--

`round_to_nearest_multiple(5, 3)` inapaswa kurudisha `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(5, 3), 6)`)
}})
```

`round_to_nearest_multiple(17, 4)` inapaswa kurudisha `16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(17, 4), 16)`)
}})
```

`round_to_nearest_multiple(43, 5)` inapaswa kurudisha `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(43, 5), 45)`)
}})
```

`round_to_nearest_multiple(38, 11)` inapaswa kurudisha `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(38, 11), 33)`)
}})
```

`round_to_nearest_multiple(93, 12)` inapaswa kurudisha `96`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(93, 12), 96)`)
}})
```

# --seed--

## --seed-contents--

```py
def round_to_nearest_multiple(num, multiple):

    return num
```

# --solutions--

```py
def round_to_nearest_multiple(num, multiple):
    return round(num / multiple) * multiple
```
