---
id: 6a1d9f98e819ed70a0e994e2
title: "Zoezi la 336: Ulinganifu wa Horoscope"
challengeType: 29
dashedName: challenge-336
---

# --description--

Ukitoa mfuatano wa herufi wa alama mbili za nyota, rudisha asilimia yao ya ulinganifu.

Alama zimepangwa katika gurudumu lenye nafasi 12 kwa mpangilio huu: `"Aries"`, `"Taurus"`, `"Gemini"`, `"Cancer"`, `"Leo"`, `"Virgo"`, `"Libra"`, `"Scorpio"`, `"Sagittarius"`, `"Capricorn"`, `"Aquarius"`, `"Pisces"`, zikizunguka kurudi `"Aries"` baada ya `"Pisces"`. Tafuta umbali mfupi zaidi kati ya alama hizo mbili na rudisha ulinganifu:

| Umbali | Ulinganifu |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscope_match("Libra", "Sagittarius")` inapaswa kurudisha `"80%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Libra", "Sagittarius"), "80%")`)
}})
```

`horoscope_match("Gemini", "Scorpio")` inapaswa kurudisha `"20%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Gemini", "Scorpio"), "20%")`)
}})
```

`horoscope_match("Pisces", "Aries")` inapaswa kurudisha `"40%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Pisces", "Aries"), "40%")`)
}})
```

`horoscope_match("Capricorn", "Cancer")` inapaswa kurudisha `"50%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Capricorn", "Cancer"), "50%")`)
}})
```

`horoscope_match("Aquarius", "Aquarius")` inapaswa kurudisha `"100%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Aquarius", "Aquarius"), "100%")`)
}})
```

`horoscope_match("Virgo", "Taurus")` inapaswa kurudisha `"90%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Virgo", "Taurus"), "90%")`)
}})
```

`horoscope_match("Leo", "Scorpio")` inapaswa kurudisha `"30%"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(horoscope_match("Leo", "Scorpio"), "30%")`)
}})
```

# --seed--

## --seed-contents--

```py
def horoscope_match(sign1, sign2):

    return sign1
```

# --solutions--

```py
def horoscope_match(sign1, sign2):
    signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"]
    a = signs.index(sign1)
    b = signs.index(sign2)
    diff = abs(a - b)
    distance = min(diff, 12 - diff)
    compatibility = {0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%"}
    return compatibility[distance]
```
