---
id: 6a1d9f98e819ed70a0e994e0
title: "Changamoto 334: Mabadiliko Sahihi"
challengeType: 29
dashedName: challenge-334
---

# --description--

Ukitoa kiasi kamili cha nambari kamili katika senti, rudisha idadi ya njia tofauti za kutengeneza mabadiliko sahihi kwa kutumia pennies (1 senti), nickels (5 senti), dimes (10 senti), na quarters (25 senti).

# --hints--

`exact_change(3)` inapaswa kurudisha `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(3), 1)`)
}})
```

`exact_change(9)` inapaswa kurudisha `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(9), 2)`)
}})
```

`exact_change(17)` inapaswa kurudisha `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(17), 6)`)
}})
```

`exact_change(39)` inapaswa kurudisha `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(39), 24)`)
}})
```

`exact_change(61)` inapaswa kurudisha `73`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(61), 73)`)
}})
```

`exact_change(99)` inapaswa kurudisha `213`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(exact_change(99), 213)`)
}})
```

# --seed--

## --seed-contents--

```py
def exact_change(amount):

    return amount
```

# --solutions--

```py
def exact_change(amount):
    coins = [1, 5, 10, 25]
    dp = [0] * (amount + 1)
    dp[0] = 1
    for coin in coins:
        for i in range(coin, amount + 1):
            dp[i] += dp[i - coin]
    return dp[amount]
```
