---
id: 6a1d9f98e819ed70a0e994de
title: "Zoezi la 332: Usimamizi wa Masuala"
challengeType: 29
dashedName: challenge-332
---

# --description--

Kutolewa idadi ya millisekunde tangu chapisho la mwisho kuhusu suala fulani, na ujumbe wa mwisho uliotumwa kuhusu suala hilo, amua unapaswa kufanya nini na suala hilo kulingana na sheria hizi:

- Ikiwa ujumbe wa mwisho ni chini ya siku 7 zilizopita, rudisha `"leave it"`
- Ikiwa ujumbe wa mwisho ni siku 7 au zaidi zilizopita na maudhui yake yana `"bump"` (bila kujali herufi kubwa au ndogo), rudisha `"close it"`
- Vinginevyo, rudisha `"bump it"`

# --hints--

`triage_issue(86400000, "Lets fix it")` inapaswa kurudisha `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(86400000, "Lets fix it"), "leave it")`)
}})
```

`triage_issue(1209600000, "still waiting")` inapaswa kurudisha `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(1209600000, "still waiting"), "bump it")`)
}})
```

`triage_issue(864000000, "bump")` inapaswa kurudisha `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(864000000, "bump"), "close it")`)
}})
```

`triage_issue(604800000, "Do we still want this?")` inapaswa kurudisha `"bump it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Do we still want this?"), "bump it")`)
}})
```

`triage_issue(604800000, "Bumping this")` inapaswa kurudisha `"close it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Bumping this"), "close it")`)
}})
```

`triage_issue(345600000, "I'll make a PR")` inapaswa kurudisha `"leave it"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(345600000, "I'll make a PR"), "leave it")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(ms, message):

    return ms
```

# --solutions--

```py
def triage_issue(ms, message):
    seven_days = 7 * 24 * 60 * 60 * 1000
    if ms < seven_days:
        return "leave it"
    if "bump" in message.lower():
        return "close it"
    return "bump it"
```
