---
id: 6a1d9f98e819ed70a0e994da
title: "Zoezi la 328: Utaratibu wa Kaprekar"
challengeType: 29
dashedName: challenge-328
---

# --description--

Ukipata nambari yenye tarakimu 4, rudisha idadi ya mara unazohitaji kutumia utaratibu wa Kaprekar hadi kufikia 6174.

Utaratibu wa Kaprekar hufanya kazi kama ifuatavyo:

- Panga tarakimu kwa mpangilio wa kushuka ili kuunda nambari kubwa zaidi
- Panga tarakimu kwa mpangilio wa kuongezeka ili kuunda nambari ndogo zaidi (ongeza sifuri za mbele ikiwa ni lazima)
- Toa nambari ndogo kutoka kwa nambari kubwa
- Rudia kwa nambari mpya

# --hints--

`kaprekar(1234)` inapaswa kurudisha `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(1234), 3)`)
}})
```

`kaprekar(2025)` inapaswa kurudisha `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(2025), 6)`)
}})
```

`kaprekar(7173)` inapaswa kurudisha `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(7173), 4)`)
}})
```

`kaprekar(3164)` inapaswa kurudisha `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(3164), 7)`)
}})
```

`kaprekar(8082)` inapaswa kurudisha `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(8082), 2)`)
}})
```

# --seed--

## --seed-contents--

```py
def kaprekar(n):

    return n
```

# --solutions--

```py
def kaprekar(n):
    steps = 0
    current = n
    while current != 6174:
        digits = list(str(current).zfill(4))
        desc = int("".join(sorted(digits, reverse=True)))
        asc = int("".join(sorted(digits)))
        current = desc - asc
        steps += 1
    return steps
```
