---
id: 6a2037a68a0bc2aef0006002
title: "Zoezi la 340: Mbadalishaji wa Pig Latin"
challengeType: 29
dashedName: challenge-340
---

# --description--

Kutolewa mfuatano wa herufi, ubadilishe kuwa Pig Latin ukitumia sheria zifuatazo:

- Ikiwa neno linaanza na vokali (`"a"`, `"e"`, `"i"`, `"o"`, au `"u"`), ongeza `"way"` mwishoni. Kwa mfano, `"universe"` hubadilika kuwa `"universeway"`.
- Ikiwa neno linaanza na konsonanti moja au zaidi, hamisha konsonanti hizo mwishoni na ongeza `"ay"`. Kwa mfano, `"hello"` hubadilika kuwa `"ellohay"`.
- Hifadhi herufi ya kwanza kama ilivyo. Kwa mfano, `"Hello"` hubadilika kuwa `"Ellohay"`.

# --hints--

`pig_latin("universe")` inapaswa kurudisha `"universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("universe"), "universeway")`)
}})
```

`pig_latin("hello")` inapaswa kurudisha `"ellohay"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello"), "ellohay")`)
}})
```

`pig_latin("hello universe")` inapaswa kurudisha `"ellohay universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello universe"), "ellohay universeway")`)
}})
```

`pig_latin("Hello universe")` inapaswa kurudisha `"Ellohay universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Hello universe"), "Ellohay universeway")`)
}})
```

`pig_latin("Pig Latin is fun")` inapaswa kurudisha `"Igpay Atinlay isway unfay"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Pig Latin is fun"), "Igpay Atinlay isway unfay")`)
}})
```

`pig_latin("The quick brown fox jumped over the lazy dog")` inapaswa kurudisha `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday")`)
}})
```

# --seed--

## --seed-contents--

```py
def pig_latin(s):

    return s
```

# --solutions--

```py
def pig_latin(s):
    vowels = "aeiou"

    def convert_word(word):
        lower_word = word.lower()
        if lower_word[0] in vowels:
            return word + "way"
        i = 0
        while i < len(word) and lower_word[i] not in vowels:
            i += 1
        consonants = word[:i]
        rest = word[i:]
        result = rest + consonants + "ay"
        if word[0].isupper():
            return result[0].upper() + result[1:].lower()
        return result

    return " ".join(convert_word(word) for word in s.split(" "))
```
