---
id: 6a2037a68a0bc2aef0005fff
title: "Zoezi la 337: Kihesabu cha Tally"
challengeType: 29
dashedName: challenge-337
---

# --description--

Kutolewa mfuatano wa herufi wa alama za tally, rudisha jumla ya hesabu inayowakilishwa.

- Kila bomba `"|"` wakilisha hesabu moja.
- Kila alama ya tano inawakilishwa kama mstari wa mbele `"/"`, ukikamilisha kundi la tano (`"||||/"`).
- Makundi yanatenganishwa kwa nafasi.

# --hints--

`get_tally_count("||||")` inapaswa kurudisha `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||"), 4)`)
}})
```

`get_tally_count("||||/")` inapaswa kurudisha `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/"), 5)`)
}})
```

`get_tally_count("||||/ |||")` inapaswa kurudisha `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ |||"), 8)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||")` inapaswa kurudisha `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||"), 17)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` inapaswa kurudisha `41`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_tally_count(s):

    return s
```

# --solutions--

```py
def get_tally_count(s):
    return len(s.replace(" ", ""))
```
