---
id: 6a2037a68a0bc2aef0006001
title: "Changamoto 339: Vipande vya Safu ya Data"
challengeType: 29
dashedName: challenge-339
---

# --description--

Ukitoa safu ya data na ukubwa wa kipande, rudisha safu ya data iliyogawanyika katika safu ndogo za ukubwa huo.

- Kipande cha mwisho kinaweza kuwa kidogo zaidi ikiwa safu ya data haigawani kwa usawa.

# --hints--

`chunk_array([1, 2, 3, 4, 5, 6], 3)` inapaswa kurudisha `[[1, 2, 3], [4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]])`)
}})
```

`chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` inapaswa kurudisha `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]])`)
}})
```

`chunk_array([1, 2, 3, 4, 5], 3)` inapaswa kurudisha `[[1, 2, 3], [4, 5]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]])`)
}})
```

`chunk_array(["a", "b", "c", "d", "e"], 1)` inapaswa kurudisha `[["a"], ["b"], ["c"], ["d"], ["e"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]])`)
}})
```

`chunk_array([1, 2, 3], 5)` inapaswa kurudisha `[[1, 2, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3], 5), [[1, 2, 3]])`)
}})
```

# --seed--

## --seed-contents--

```py
def chunk_array(arr, size):

    return arr
```

# --solutions--

```py
def chunk_array(arr, size):
    return [arr[i:i + size] for i in range(0, len(arr), size)]
```
