---
id: 6a2037a68a0bc2aef0006003
title: "Zoezi la 341: Hesabu ya siku za kuzaliwa"
challengeType: 29
dashedName: challenge-341
---

# --description--

Ukizingatia tarehe ya leo na siku ya kuzaliwa, rudisha idadi ya siku hadi siku ya kuzaliwa inayofuata ya mtu huyo.

- Tarehe ya leo hutolewa kama mfuatano wa herufi katika muundo wa `"YYYY-MM-DD"`, ukiwa na sifuri za kuongoza, kwa mfano: `"2026-07-16"`.
- Siku ya kuzaliwa hutolewa kama mfuatano wa herufi katika muundo wa `"M/D"`, bila sifuri za kuongoza, kwa mfano: `"9/7"`.
- Ikiwa leo ni siku yao ya kuzaliwa, rudisha idadi ya siku hadi siku yao inayofuata ya kuzaliwa (si `0`).
- Miaka ya kurukia inapaswa kuzingatiwa.

# --hints--

`days_until_birthday("2026-07-16", "9/7")` inapaswa kurudisha `53`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "9/7"), 53)`)
}})
```

`days_until_birthday("2026-07-16", "3/22")` inapaswa kurudisha `249`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "3/22"), 249)`)
}})
```

`days_until_birthday("2026-07-16", "7/16")` inapaswa kurudisha `365`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "7/16"), 365)`)
}})
```

`days_until_birthday("2024-02-28", "3/1")` inapaswa kurudisha `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-02-28", "3/1"), 2)`)
}})
```

`days_until_birthday("2023-04-24", "12/30")` inapaswa kurudisha `250`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2023-04-24", "12/30"), 250)`)
}})
```

`days_until_birthday("2024-03-01", "2/29")` inapaswa kurudisha `1460`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-03-01", "2/29"), 1460)`)
}})
```

`days_until_birthday("2096-03-01", "2/29")` inapaswa kurudisha `2920`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2096-03-01", "2/29"), 2920)`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_birthday(today, birthday):

    return today
```

# --solutions--

```py
def days_until_birthday(today, birthday):
    from datetime import date

    year, month, day = [int(x) for x in today.split("-")]
    b_month, b_day = [int(x) for x in birthday.split("/")]

    today_date = date(year, month, day)

    def is_leap_year(y):
        return (y % 4 == 0 and y % 100 != 0) or y % 400 == 0

    def next_valid_birthday(from_year):
        for y in range(from_year, from_year + 9):
            if b_month == 2 and b_day == 29 and not is_leap_year(y):
                continue
            d = date(y, b_month, b_day)
            if d > today_date:
                return d

    return (next_valid_birthday(year) - today_date).days
```
