---
id: 6a22d77ddf034bc4e35b1d56
title: "Changamoto 349: Ishara ya Simu ya Mkononi"
challengeType: 29
dashedName: challenge-349
---

# --description--

Kutolewa gridi yenye vipimo vitatu vya mnara wa simu, amua eneo la simu.

- Kila seli katika gridi ni `0` (hakuna mnara) au nambari kamili chanya inayowakilisha idadi ya seli hadi simu, iliyopimwa kwa mstari wa moja kwa moja: wima, mlalo, au mwinuko.
- Rudisha `[row, col]` ya seli ambayo ina idadi sahihi ya seli kutoka kwa minara yote mitatu.
- Daima kuna suluhisho moja tu.

# --hints--

`find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` inapaswa kurudisha `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2])`)
}})
```

`find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` inapaswa kurudisha `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1])`)
}})
```

`find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` inapaswa kurudisha `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2])`)
}})
```

`find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` inapaswa kurudisha `[3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4])`)
}})
```

`find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` inapaswa kurudisha `[3, 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_signal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_signal(grid):

    return grid
```

# --solutions--

```py
def find_signal(grid):
    rows = len(grid)
    cols = len(grid[0])

    towers = []
    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                towers.append((r, c, grid[r][c]))

    def is_valid(r, c, tr, tc, dist):
        dr = abs(r - tr)
        dc = abs(c - tc)
        return (dr == dist or dc == dist) and (dr == 0 or dc == 0 or dr == dc)

    for r in range(rows):
        for c in range(cols):
            if grid[r][c] != 0:
                continue
            if all(is_valid(r, c, tr, tc, dist) for tr, tc, dist in towers):
                return [r, c]
```
