---
id: 6a22d77ddf034bc4e35b1d54
title: "Changamoto 347: Nadharia ya Michezo"
challengeType: 29
dashedName: challenge-347
---

# --description--

Kutolewa mfuatano wa herufi wenye urefu sawa unaowakilisha mikakati ya wachezaji wawili kwa mchezo, rudisha alama kama safu ya data `[player1, player2]`.

- Mfuatano wa herufi uliotolewa utakuwa na herufi moja kati ya mbili tu: `"C"` (shirikiana) au `"D"` (kataa).
- Kila herufi inawakilisha raundi moja, na alama hutolewa kama ifuatavyo:
  - Ikiwa wachezaji wote wawili watashirikiana, kila mmoja atapata alama 3.
  - Ikiwa wachezaji wote wawili watakata, kila mmoja atapata alama 1.
  - Ikiwa mchezaji mmoja atakata na mwingine atashirikiana, mkatakata atapata alama 5 na mshirikaji atapata alama 0.

# --hints--

`play_game("CCCC", "CCCC")` inapaswa kurudisha `[12, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` inapaswa kurudisha `[4, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` inapaswa kurudisha `[5, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` inapaswa kurudisha `[24, 34]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` inapaswa kurudisha `[66, 21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
