---
id: 6a22d77ddf034bc4e35b1d5b
title: "Changamoto 354: Kiwango cha Utofauti 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

Kutolewa safu mbili za data zinazowakilisha thamani za RGB na thamani ya kweli au siyo kweli inayoonyesha kama maandishi ni makubwa, rudisha kiwango cha utofauti cha WCAG kwa kutumia njia ifuatayo:

Kwanza, badilisha kila thamani ya RGB kuwa mwangaza wa jamaa:

- Gawanya kila chaneli `[R, G, B]` kwa 255 kupata thamani kati ya `0` na `1`
- Tumia fomula ya kurekebisha gamma kwa kila chaneli:
  - Ikiwa thamani ya chaneli ni ndogo au sawa na `0.04045`: `channel / 12.92`
  - Vinginevyo: `((channel + 0.055) / 1.055) ^ 2.4`
- Hesabu mwangaza: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Kisha, hesabu uwiano wa utofauti kwa kuongeza 0.05 kwa kila thamani ya mwangaza, kisha gawanya ile yenye mwanga zaidi kwa ile yenye giza zaidi. Ile yenye mwanga zaidi itakuwa hoja ya kwanza kila mara.

Rudisha kiwango kulingana na uwiano wa utofauti kwa kutumia jedwali lifuatalo:

| Kiwango | Maandishi ya Kawaida | Maandishi Makubwa |
|--------|--------------------|------------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | chini ya 4.5 | chini ya 3.0 |

# --hints--

`get_contrast_rating([255, 255, 255], [0, 0, 0], False)` inapaswa kurudisha `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

`get_contrast_rating([215, 188, 188], [55, 55, 55], False)` inapaswa kurudisha `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

`get_contrast_rating([143, 144, 210], [46, 47, 61], False)` inapaswa kurudisha `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

`get_contrast_rating([167, 167, 210], [53, 10, 53], True)` inapaswa kurudisha `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

`get_contrast_rating([135, 147, 155], [60, 70, 90], True)` inapaswa kurudisha `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

`get_contrast_rating([125, 210, 195], [105, 130, 90], True)` inapaswa kurudisha `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
