---
id: 6a22d77ddf034bc4e35b1d57
title: "Zoezi la 350: Umbali wa Herufi"
challengeType: 29
dashedName: challenge-350
---

# --description--

Ukitoa mfuatano wa herufi miwili wenye urefu sawa, rudisha jumla ya umbali mfupi zaidi kati ya kila jozi ya herufi.

- Ingizo litakuwa na herufi ndogo tu
- Alfabeti hutendewa kama mduara, kwa hivyo umbali kati ya `a` na `z` ni 1.

# --hints--

`letter_distance("abc", "bcd")` inapaswa kurudisha `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "bcd"), 3)`)
}})
```

`letter_distance("abc", "xyz")` inapaswa kurudisha `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "xyz"), 9)`)
}})
```

`letter_distance("encrypt", "decrypt")` inapaswa kurudisha `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("encrypt", "decrypt"), 10)`)
}})
```

`letter_distance("algorithm", "codeblock")` inapaswa kurudisha `43`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("algorithm", "codeblock"), 43)`)
}})
```

`letter_distance("lobster", "penguin")` inapaswa kurudisha `47`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("lobster", "penguin"), 47)`)
}})
```

`letter_distance("alligator", "crocodile")` inapaswa kurudisha `55`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("alligator", "crocodile"), 55)`)
}})
```

# --seed--

## --seed-contents--

```py
def letter_distance(str1, str2):

    return str1
```

# --solutions--

```py
def letter_distance(str1, str2):
    total = 0
    for a, b in zip(str1, str2):
        diff = abs(ord(a) - ord(b))
        total += min(diff, 26 - diff)
    return total
```
