---
id: 6a22d77ddf034bc4e35b1d5a
title: "Changamoto 353: Kiwango cha Utofauti wa Rangi 2"
challengeType: 29
dashedName: challenge-353
---

# --description--

Ukizingatia thamani mbili za mwangaza wa uhusiano na thamani ya kweli au siyo kweli inayoonyesha kama maandishi ni makubwa, rudisha kiwango cha utofauti wa rangi cha WCAG kwa kutumia njia ifuatayo:

Hesabu uwiano wa utofauti kwa kuongeza 0.05 kwa kila thamani ya mwangaza, kisha gawanya ile yenye mwangaza zaidi kwa ile yenye mwangaza mdogo. Ile yenye mwangaza zaidi itakuwa hoja ya kwanza kila mara.

Rudisha kiwango kulingana na uwiano wa utofauti kwa kutumia jedwali lifuatalo:

| Kiwango | Maandishi ya Kawaida | Maandishi Makubwa |
|--------|--------------------|------------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | chini ya 4.5 | chini ya 3.0 |

# --hints--

`get_contrast_rating(1.0, 0.0, False)` inapaswa kurudisha `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(1.0, 0.0, False), "AAA")`)
}})
```

`get_contrast_rating(0.9015, 0.1364, False)` inapaswa kurudisha `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.9015, 0.1364, False), "AA")`)
}})
```

`get_contrast_rating(0.8965, 0.1628, False)` inapaswa kurudisha `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.8965, 0.1628, False), "Fail")`)
}})
```

`get_contrast_rating(0.7469, 0.0957, True)` inapaswa kurudisha `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7469, 0.0957, True), "AAA")`)
}})
```

`get_contrast_rating(0.7489, 0.2018, True)` inapaswa kurudisha `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7489, 0.2018, True), "AA")`)
}})
```

`get_contrast_rating(0.6571, 0.1974, True)` inapaswa kurudisha `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.6571, 0.1974, True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(l1, l2, is_large_text):

    return l1
```

# --solutions--

```py
def get_contrast_rating(l1, l2, is_large_text):
    ratio = (l1 + 0.05) / (l2 + 0.05)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
