---
id: 6a26df3e2988bcdded204893
title: "Changamoto 357: Mnyororo wa Chakula"
challengeType: 29
dashedName: challenge-357
---

# --description--

Kutolewa safu ya data ya jozi za `[predator, prey]`, rudisha mnyororo wa chakula kutoka mnyama mwangizi wa juu hadi chini kabisa.

- Mnyama mwangizi wa juu ni mnyama ambaye hakuwahi kuwa mtego wa mnyama mwingine.
- Rudisha mnyororo huo kama safu ya data ya mfuatano wa herufi.

# --hints--

`get_food_chain([["cat", "mouse"]])` inapaswa kurudisha `["cat", "mouse"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["cat", "mouse"]]), ["cat", "mouse"])`)
}})
```

`get_food_chain([["wolf", "deer"], ["deer", "grass"]])` inapaswa kurudisha `["wolf", "deer", "grass"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"])`)
}})
```

`get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` inapaswa kurudisha `["hawk", "snake", "frog", "fly"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"])`)
}})
```

`get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` inapaswa kurudisha `["eagle", "fox", "rabbit", "grass"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"])`)
}})
```

`get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` inapaswa kurudisha `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_food_chain(pairs):

    return pairs
```

# --solutions--

```py
def get_food_chain(pairs):
    prey_set = {p for _, p in pairs}
    chain_map = dict(pairs)
    apex = next(predator for predator, _ in pairs if predator not in prey_set)
    chain = [apex]
    while chain[-1] in chain_map:
        chain.append(chain_map[chain[-1]])
    return chain
```
