---
id: 6a26df95efa55a2524399745
title: "Changamoto 364: Kati ya Mifuko Miwili"
challengeType: 29
dashedName: challenge-364
---

# --description--

Ukizingatia mifuko miwili ya rangi, kila moja ikiwa na rangi ya RGB na kiwango cha ujazo, rudisha rangi iliyochanganywa ya RGB kama safu ya data ya nambari kamili tatu.

- Kila mfuko ni kitu (JavaScript) au kamusi (Python) yenye sifa ya `color` (safu ya data ya nambari kamili tatu `[r, g, b]`) na sifa ya `fullness` (0–100).
- Rangi iliyochanganywa ni wastani wa uzito wa kila chaneli katika rangi hizo mbili kulingana na kiwango cha ujazo, ambapo kila chaneli inakaribia nambari kamili inayokaribu zaidi.

# --hints--

`mix_paint({"color": [250, 250, 250], "fullness": 50}, {"color": [0, 0, 0], "fullness": 50})` inapaswa kurudisha `[125, 125, 125]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [250, 250, 250], "fullness": 50}, {"color": [0, 0, 0], "fullness": 50}), [125, 125, 125])`)
}})
```

`mix_paint({"color": [250, 250, 250], "fullness": 80}, {"color": [0, 0, 0], "fullness": 20})` inapaswa kurudisha `[200, 200, 200]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [250, 250, 250], "fullness": 80}, {"color": [0, 0, 0], "fullness": 20}), [200, 200, 200])`)
}})
```

`mix_paint({"color": [100, 150, 200], "fullness": 30}, {"color": [100, 150, 200], "fullness": 70})` inapaswa kurudisha `[100, 150, 200]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [100, 150, 200], "fullness": 30}, {"color": [100, 150, 200], "fullness": 70}), [100, 150, 200])`)
}})
```

`mix_paint({"color": [143, 143, 101], "fullness": 45}, {"color": [100, 204, 204], "fullness": 90})` inapaswa kurudisha `[114, 184, 170]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [143, 143, 101], "fullness": 45}, {"color": [100, 204, 204], "fullness": 90}), [114, 184, 170])`)
}})
```

`mix_paint({"color": [15, 134, 249], "fullness": 29}, {"color": [97, 178, 55], "fullness": 54})` inapaswa kurudisha `[68, 163, 123]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [15, 134, 249], "fullness": 29}, {"color": [97, 178, 55], "fullness": 54}), [68, 163, 123])`)
}})
```

# --seed--

## --seed-contents--

```py
def mix_paint(bucket1, bucket2):

    return bucket1
```

# --solutions--

```py
def mix_paint(bucket1, bucket2):
    total = bucket1['fullness'] + bucket2['fullness']
    return [
        round((c1 * bucket1['fullness'] + c2 * bucket2['fullness']) / total)
        for c1, c2 in zip(bucket1['color'], bucket2['color'])
    ]
```
