---
id: 6a26df3e2988bcdded204897
title: "Changamoto 361: Muda Unaosemwa"
challengeType: 29
dashedName: challenge-361
---

# --description--

Ukizingatia pembe za mkono wa saa na mkono wa dakika wa saa ya analog kwa digrii (kuelekea saa 12 kwa mwelekeo wa saa), rudisha muda kwa Kiingereza kinachosemwa.

Badilisha pembe ya mkono wa dakika kuwa dakika (360° = dakika 60), kisha tumia sheria zifuatazo:

| Dakika | Inasemwa |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (isipokuwa 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (isipokuwa 45) | `"X minutes to Z"` (ambapo `X` ni 60 - dakika) |

Ambapo `Y` ni saa ya sasa na `Z` ni saa inayofuata, zote zikichukuliwa kutoka kwa pembe ya mkono wa saa (360° = saa 12).

Kumbuka: Pembe za mikono zinaweza zisipatikane hasa kwenye nambari, fikiria kuzizungusha kwa namna fulani.

# --hints--

`get_spoken_time(90, 0)` inapaswa kurudisha `"3 o'clock"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(90, 0), "3 o'clock")`)
}})
```

`get_spoken_time(160, 120)` inapaswa kurudisha `"20 minutes past 5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(160, 120), "20 minutes past 5")`)
}})
```

`get_spoken_time(255, 180)` inapaswa kurudisha `"half past 8"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(255, 180), "half past 8")`)
}})
```

`get_spoken_time(67.5, 92)` inapaswa kurudisha `"quarter past 2"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(67.5, 92), "quarter past 2")`)
}})
```

`get_spoken_time(200, 240)` inapaswa kurudisha `"20 minutes to 7"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(200, 240), "20 minutes to 7")`)
}})
```

`get_spoken_time(322.5, 273)` inapaswa kurudisha `"quarter to 11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(322.5, 273), "quarter to 11")`)
}})
```

`get_spoken_time(117.5, 335)` inapaswa kurudisha `"5 minutes to 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(117.5, 335), "5 minutes to 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_time(hour_angle, minute_angle):

    return hour_angle
```

# --solutions--

```py
def get_spoken_time(hour_angle, minute_angle):
    minutes = int(minute_angle / 6)
    hour = int(hour_angle / 30) or 12
    next_hour = (hour % 12) + 1

    if minutes == 0:
        return f"{hour} o'clock"
    if minutes == 15:
        return f"quarter past {hour}"
    if minutes == 30:
        return f"half past {hour}"
    if minutes == 45:
        return f"quarter to {next_hour}"
    if minutes < 30:
        return f"{minutes} minutes past {hour}"
    return f"{60 - minutes} minutes to {next_hour}"
```
