---
id: 66ec4c8e9878d8441956516f
title: Crie uma tabela de catálogo de livros
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**Objetivo:** atender às histórias de usuário abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve criar um elemento `table` que liste informações de livros.
1. A tabela deve ter um elemento de cabeçalho com uma linha.
1. A linha no elemento de cabeçalho da tabela (head) deve ter quatro elementos de cabeçalho de tabela (header), com o texto `Title`, `Author`, `Genre` e `Publication Year`, nessa ordem.
1. A tabela deve ter um elemento de corpo de tabela com pelo menos cinco linhas.
1. Cada linha no corpo da tabela deve ter quatro elementos de dados que exibam o Título, Autor, Gênero e Ano de Publicação do livro.
1. A tabela deve ter um elemento de rodapé com uma linha.
1. The row in your table footer element should have a table data element that spans four columns and has the text `Total Books: N` where `N` should be replaced by the number of books in your table.

# --hints--

Você deve ter um elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Você deve ter um elemento `thead` dentro do elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

Você deve ter um elemento `tr` dentro do elemento `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

Você deve ter quatro elementos `th` dentro da linha do elemento `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

Os quatro elementos `th` devem ter o texto `Title`, `Author`, `Genre` e `Publication Year`, nessa ordem.

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent, 'Title');
assert.equal(ths[1]?.textContent, 'Author');
assert.equal(ths[2]?.textContent, 'Genre');
assert.equal(ths[3]?.textContent, 'Publication Year');
```

Você deve ter um elemento `tbody` dentro do elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

O elemento `tbody` deve ter pelo menos cinco linhas.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

Cada linha no elemento `tbody` deve ter exatamente quatro elementos `td` como filhos.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

Cada elemento `td` no corpo da tabela deve ter texto com informações do livro.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.length, 1);
});
```

Você deve ter um elemento `tfoot` dentro do elemento `table`.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

Você deve ter exatamente um elemento `tr` no elemento `tfoot`.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

O elemento `td` na linha do elemento `tfoot` deve ter o atributo `colspan` definido como `4`.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

The `td` element in your `tfoot` element's row should have the text `Total Books: N` where `N` is the number of books in your table.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
