---
id: 66a83e5e491625454b6f62c3
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

フォームはユーザーがデータを入力できる`inputs`で構成されています。関連する入力をまとめるには、`fieldset`要素を使うことができます。

`fieldset`要素を使った例を示します。

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

`form`要素の中に、`fieldset`要素を追加してください。

# --hints--

開始の`fieldset`タグがあるはずです。

```js
assert.match(code, /<fieldset>/i);
```

終了の`fieldset`タグがあるはずです。

```js
assert.match(code, /<\/fieldset>/i);
```

`fieldset`要素は`form`要素の中にあるはずです。

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
