---
id: 66a83601cd819e37f0dccd14
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

このワークショップでは、ホテルのフィードバックフォームを作成しながらHTMLフォームの操作を練習します。

導入文として、メインタイトルの後にフィードバックを残すための短いメモを表示してください。

`body` 要素の中に `header` 要素を追加してください。

`header` 要素の中に `h1` 要素を追加してください。`h1` 要素のテキストは `Hotel Feedback Form` であるべきです。

`h1` 要素の下に `p` 要素を追加してください。`p` 要素のテキストは `Thank you for staying with us. Please provide feedback on your recent stay.` であるべきです。

# --hints--

`header` の開始タグがあるはずです。

```js
assert.match(code, /<header>/i);
```

`header` の終了タグがあるはずです。

```js
assert.match(code, /<\/header>/i);
```

`h1` の開始タグがあるはずです。

```js
assert.match(code, /<h1>/i);
```

`h1` の終了タグがあるはずです。

```js
assert.match(code, /<\/h1>/i);
```

`h1` 要素のテキストは `Hotel Feedback Form` であるべきです。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

`h1` 要素は `header` 要素の中にあるべきです。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

`p` の開始タグがあるはずです。

```js
assert.match(code, /<p>/i);
```

`p` の終了タグがあるはずです。

```js
assert.match(code, /<\/p>/i);
```

段落要素のテキストは `Thank you for staying with us. Please provide feedback on your recent stay.` であるべきです。

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

段落要素は `header` の中にあるべきです。

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

段落要素は `h1` の後に来るべきです。

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
