---
id: 66a83fec026a7a4631e084d2
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`fieldset` 要素を扱うとき、入力のグループを説明するためにキャプションを使うことがよくあります。これには `legend` 要素を使えます。

以下は `legend` 要素を使った例です。

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

`fieldset` 要素の中に、テキスト `legend` を持つ `Personal Information` 要素を追加してください。

# --hints--

開く `legend` タグがあるはずです。

```js
assert.match(code, /<legend>/);
```

閉じる `legend` タグがあるはずです。

```js
assert.match(code, /<\/legend>/);
```

`legend` 要素は `fieldset` 要素の中にあるはずです。

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

`legend` 要素にはテキスト `Personal Information` が含まれているはずです。

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
