---
id: 66a83bdcf425e7446900b7c4
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

前のレッスンでは、`form` 要素の使い方を学びました。

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

`action` 属性は、フォームが提出されたときにフォームデータを送信する先を指定するために使います。

`method` 属性は、フォームデータを送信するときに使う HTTP メソッドを指定するために使います。最も一般的なメソッドは `GET` と `POST` です。

**注意**：HTTP メソッドの仕組みについては、後のモジュールで学習します。

`main` 要素の中に、`form` 要素を追加し、`action` 属性を `"https://hotel-feedback.freecodecamp.org"` に、`method` 属性を `"POST"` に設定してください。

# --hints--

`form` の開始タグがあるはずです。

```js
assert.match(code, /<form.*>/i);
```

`form` の終了タグがあるはずです。

```js
assert.match(code, /<\/form>/i);
```

`form` 要素には `action` 属性が `"https://hotel-feedback.freecodecamp.org"` に設定されているはずです。

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

`form` 要素には `method` 属性が `"POST"` に設定されているはずです。

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
