---
id: 66a84111965a0c46df6bbd0a
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

前のレッスンでは、`label` 要素を `input` と次のように関連付ける方法を学びました。

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

`for` 要素の `label` 属性は、`id` 要素の `input` 属性と一致している必要があります。これは明示的な関連付けとして知られています。

`legend` 要素の下に、テキストが `label` の `Name (required):` 要素を追加してください。その `for` 属性を `"full-name"` の値に設定します。

次に、`label` 要素の下に属性のない `input` 要素を追加してください。次のステップで必要な属性を追加します。

# --hints--

開く `label` タグがあるはずです。

```js
assert.match(code, /<label.*>/);
```

閉じる `label` タグがあるはずです。

```js
assert.match(code, /<\/label>/);
```

`label` 要素は `legend` 要素の中の `fieldset` の下に配置されているはずです。

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

`label` 要素は `for` の値を持つ `"full-name"` 属性があるはずです。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

`label` 要素は `Name (required):` のテキストがあるはずです。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

`input` 要素の下に `label` 要素があるはずです。

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
