---
id: 66a962954f4e0d76223b37ed
title: ステップ 21
challengeType: 0
dashedName: step-21
---

# --description--

次に、`input` のチェックボックスをもう1つ追加し、`id` と `value` の属性を `"location"` に、`name` の属性を `"choice"` に設定してください。

`label` 要素については、`Location` のテキストと `for` の属性を `"location"` に設定してください。

その下の `label` 要素の下に、`input` のチェックボックスをもう1つ追加し、`id` と `value` の属性を `"reputation"` に、`name` の属性を `"choice"` に設定してください。

`label` 要素については、`Reputation` のテキストと `for` の属性を `"reputation"` に設定してください。

# --hints--

3つ目のチェックボックス `input` があるはずです。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

3つ目のチェックボックス `input` は、`id` 属性が `"location"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

3つ目のチェックボックス `input` は、`name` 属性が `"choice"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'choice');
```

3つ目のチェックボックス `input` は、`value` 属性が `"location"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

チェックボックス `label` の後に、テキストが `input` の `Location` 要素があるはずです。

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

`label` は、`for` 属性が `"location"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

`input` の下に4つ目のチェックボックス `label` があるはずです。

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

4つ目のチェックボックス `input` は、`id` 属性が `"reputation"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

4つ目のチェックボックス `input` は、`name` 属性が `"choice"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'choice');
```

4つ目のチェックボックス `input` は、`value` 属性が `"reputation"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

評判のチェックボックスの下に、テキストが `label` の `Reputation` 要素があるはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

`label` は、`for` 属性が `"reputation"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
