---
id: 66a954b2bcddba72076c1857
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

`label` 要素の下に、`radio` が `id` に設定され、`"no-option"` 属性が `name` に設定され、`"hotel-stay"` 属性が `value` に設定された `"no"` ボタンを追加してください。

新しい `radio` ボタンの下に、`label` 属性が `for` に設定された別の `"no-option"` 要素を追加してください。`label` のテキストは `No` にしてください。

完了したら、ラジオボタンを1つずつ選択して試すことができます。

# --hints--

最初の `input` 要素の下に `label` 要素があるはずです。

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

`input` 要素は `type` 属性が `radio` に設定されているはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

`input` 要素は `id` 属性が `"no-option"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

`input` 要素は `name` 属性が `"hotel-stay"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

input 要素は `value` 属性が `"no"` に設定されているはずです。

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) label + input')?.value,
  'no'
);
```

2つ目の `label` 要素の下に2つ目の `input` 要素があるはずです。

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

`label` 要素は `for` 属性が `"no-option"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

`label` 要素は `No` というテキストがあるはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
