---
id: 66a9577022877d72d8f43b4f
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

フォームの次のセクションでは、ユーザーにその特定のホテルを選んだ理由を尋ねます。ユーザーは複数の選択肢を選べるようになります。

まず、別の`fieldset`要素を追加してください。

`fieldset`要素の中に、テキスト`legend`を持つ`Why did you choose to stay at our hotel? (Check all that apply)`要素を追加してください。

# --hints--

`fieldset`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 3);
```

`legend`要素の中に`fieldset`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(3) legend'), 1);
```

`legend`要素には`Why did you choose to stay at our hotel? (Check all that apply)`のテキストがあるはずです。スペルミスやスペースの問題がないか再確認してください。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) legend')?.innerText.trim(), 'Why did you choose to stay at our hotel? (Check all that apply)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        --fcc-editable-region--
        
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
