---
id: 66a9419e2d18476c645ce693
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

数値の `input` は数値入力フィールドを作成するために使われます。

数値入力フィールドの例を示します：

```html
<input type="number" id="age" name="age" min="18" max="100">
```

`min` と `max` の属性は、入力フィールドに入力できる最小値と最大値を設定するために使われます。

`label` 要素の下に、`type` 属性を `"number"` に設定し、`id` を `"age"` に設定した入力を追加してください。

`name` 属性は `"age"` に設定し、`min` 属性は `"3"` に設定し、`max` 属性は `"100"` に設定してください。

# --hints--

`input` 要素の下に数値の `label` があるはずです。

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

数値の `input` は `id` の `age` を持っているはずです。

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

数値の `input` は `name` に設定された `age` 属性を持っているはずです。

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

数値の `input` は `min` に設定された `"3"` 属性を持っているはずです。

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

数値の `input` は `max` に設定された `"100"` 属性を持っているはずです。

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
