---
id: 66a9401c9d660d6bb15993e2
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

ホテルのフィードバックフォームに、ユーザーが年齢を入力できるオプションを追加してください。

まず、フォームにテキスト `label` を持つ `Age (optional):` 要素を追加しましょう。

`for` 属性は `"age"` に設定する必要があります。

# --hints--

`label` 要素が必要です。

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

`label` 要素には `for` 属性を `"age"` に設定してください。

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

`label` 要素にはテキスト `Age (optional):` を含めてください。

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
