---
id: 66a930b20f589b6664c51cb0
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

ユーザーがフルネームを入力すると、`input` はプレーンテキストを受け入れます。

前のレッスンでは、`type` 属性を次のように扱う方法を学びました。

```html
<input type="text">
```

既存の `input` 要素に、`type` 属性を `"text"` に設定して追加してください。

また、`input` 要素に `id` 属性を追加し、その値を `"full-name"` にしてください。

# --hints--

既存の `input` 要素は削除しないでください。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

新しい `input` 要素を作成しないでください。

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

`input` 要素には、`type` 属性を `"text"` の値で設定してください。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

`input` 要素には、`id` 属性を `"full-name"` の値で設定してください。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
