---
id: 66a95d0eff8168747805f1f3
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

ユーザーにリストから複数の選択肢を選んでもらいたい場合は、チェックボックスを使うことができます。

ここでは、食べ物の選択肢を扱うチェックボックスの使い方の例を示します。

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

`value` 属性は、フォームが提出されたときにサーバーに送信される値を指定するために使います。

`legend` 要素の下に、`input` と `id` 属性が `value` に設定され、`"ads"` 属性が `name` に設定されたチェックボックス `"choice"` を追加してください。

チェックボックス `input` の下に、テキスト `label` を持つ `Social Media Ads` 要素を追加してください。`for` 属性は `"ads"` に設定してください。

# --hints--

`input` 要素の下にチェックボックス `legend` があるはずです。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

チェックボックス `input` は `id` 属性が `"ads"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

チェックボックス `input` は `name` 属性が `"choice"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "choice");
```

チェックボックス `input` は `value` 属性が `"ads"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

チェックボックス `label` の下に `input` 要素があるはずです。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

`label` 要素にはテキスト `Social Media Ads` があるはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

`label` 要素は `for` 属性が `"ads"` に設定されているはずです。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
