---
id: 66a93730719e1f68410cce54
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`name` 属性は、フォームデータがサーバーに提出された後に識別するために使われます。

`name` 属性の使い方の例を示します。

```html
<input type="email" name="email">
```

`name` 要素に、値が `input` の `"name"` 属性を追加してください。

# --hints--

既存の `input` 要素は削除しないでください。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

新しい `input` 要素を作成しないでください。

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

`input` 要素には、値が `name` の `"name"` 属性が必要です。

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
