---
id: 66a972137acd1179fa3fe8a0
title: ステップ 26
challengeType: 0
dashedName: step-26
---

# --description--

`select` 要素の中に、次の5つの `option` 要素を追加してください。それぞれの `option` テキストと `value` 属性には以下の値を設定します。

**Value 属性:**

- poor
- satisfactory
- good
- very-good
- excellent

**Option 要素のテキスト:**

- Poor
- Satisfactory
- Good
- Very Good
- Excellent

# --hints--

`option` 要素があり、その value が `"poor"` に設定されているはずです。

```js
assert.exists(document.querySelector('option[value="poor"]'));
```

`option` 要素で `value` が `"poor"` のものは、テキストが `Poor` であるはずです。

```js
assert.strictEqual(document.querySelector('option[value="poor"]')?.textContent.trim(), 'Poor');
```

`option` 要素があり、その `value` が `"satisfactory"` に設定されているはずです。

```js
assert.exists(document.querySelector('option[value="satisfactory"]'));
```

`option` 要素で `value` が `"satisfactory"` のものは、テキストが `Satisfactory` であるはずです。

```js
assert.strictEqual(document.querySelector('option[value="satisfactory"]')?.textContent.trim(), 'Satisfactory');
```

`option` 要素があり、その `value` が `"good"` に設定されているはずです。

```js
assert.exists(document.querySelector('option[value="good"]'));
```

`option` 要素で `value` が `"good"` のものは、テキストが `Good` であるはずです。

```js

assert.strictEqual(document.querySelector('option[value="good"]')?.textContent.trim(), 'Good');
```

`option` 要素があり、その value が `"very-good"` に設定されているはずです。

```js
assert.exists(document.querySelector('option[value="very-good"]'));
```

`option` 要素で `value` が `"very-good"` のものは、テキストが `Very Good` であるはずです。

```js
assert.strictEqual(document.querySelector('option[value="very-good"]')?.textContent.trim(), 'Very Good');
```

`option` 要素があり、その value が `"excellent"` に設定されているはずです。

```js
assert.exists(document.querySelector('option[value="excellent"]'));
```

`option` 要素で `value` が `"excellent"` のものは、テキストが `Excellent` であるはずです。

```js
assert.strictEqual(document.querySelector('option[value="excellent"]')?.textContent.trim(), 'Excellent');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
          --fcc-editable-region--
            
          --fcc-editable-region--
          </select>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
