---
id: 66a9521bc70162712caf118d
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

ユーザーに複数の選択肢から1つだけ選んでもらいたい場合は、ラジオボタンのセットを使うことができます。

ここに2つのラジオボタンの例があります：

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

この例では、ラジオボタンは同じ`name`属性値を使ってグループ化されています。これは、一度に選択できるラジオボタンが1つだけであることを意味します。

`legend`要素の下に、`radio`が`id`に設定され、`"yes-option"`属性が`name`に設定され、`"hotel-stay"`属性が`value`に設定された`"yes"`ボタンを追加してください。

`radio`ボタンの下に、テキストが`label`で`Yes`属性が`for`に設定された`"yes-option"`要素を追加してください。

# --hints--

`radio`要素の下に`legend`ボタンがあるはずです。

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

`radio`ボタンには`id`属性が`"yes-option"`に設定されているはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

`radio`ボタンには`name`属性が`"hotel-stay"`に設定されているはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

ラジオボタンには`value`属性が`"yes"`に設定されているはずです。

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

`label`ボタンの下に`radio`要素があるはずです。

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

`label`要素には`for`属性が`"yes-option"`に設定されているはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

`label`要素にはテキストが`Yes`であるはずです。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
