---
id: 66a937e74920ba68ebe5e86d
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

前のレッスンでは、`placeholder` と `required` 属性の使い方を学びました。

```html
<input type="text" placeholder="e.g., John Doe" required>
```

既存の `input` 要素に、値が `placeholder` の `"e.g., John Doe"` 属性を追加してください。

また、`required` 要素に `input` 属性を追加してください。

# --hints--

既存の `input` 要素は削除しないでください。

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

`input` には、値が `placeholder` の `"e.g., John Doe"` 属性が必要です。

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

`input` には `required` 属性が必要です。

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
