---
id: 66a93c95bc58e26a8fe95818
title: ステップ11
challengeType: 0
dashedName: step-11
---

# --description--

次に、メールの`input`の下に`"email"`タイプの`label`を追加してください。この`input`には、`id`属性を`"email"`の値に設定します。また、`name`属性を`"email"`の値に設定してください。

この`input`も必要なので、`required`属性を必ず追加してください。

最後に、`placeholder`属性を`"example@email.com"`に設定してください。

# --hints--

メールの`input`の下に`label`要素があるはずです。

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

`input`要素は`type`の`"email"`を持っているはずです。

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

`input`要素は`"email"`に設定された`id`属性を持っているはずです。

```js
assert.isNotNull(document.querySelector('input#email'));
```

`input`要素は`"email"`に設定された`name`属性を持っているはずです。

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

`input` 要素には `required` 属性が必要です。

```js
assert.isTrue(document.querySelector('input#email')?.hasAttribute('required'));
```

`input`要素は`"example@email.com"`に設定された`placeholder`属性を持っているはずです。
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
