---
id: 66a93bbe65a26169dbf3bc39
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

ホテルのフィードバックフォームには、ユーザーのメールアドレスも収集する必要があります。

まず、フォームにテキストが`label`の新しい`Email address (required):`要素を追加してください。`label`要素には、`for`の値に設定された`"email"`属性が必要です。

# --hints--

`label`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

`label`要素にはテキスト`Email address (required):`が含まれている必要があります。

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

`label`要素には、値が`for`に設定された`"email"`属性が必要です。

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
