---
id: 67a51d3a8a6fe123b77b0c6e
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`input` 要素は `size` 属性を持つことができます。この属性は、ユーザーが `input` に入力するときに表示される文字数を定義します。`size` の値は正の整数である必要があります。`size` が指定されていないか、無効な値が指定されている場合、`input` はブラウザーによって設定されたデフォルトの幅になります。

```html
<label for="last-name">Last Name:</label>
<input id="last-name" name="last-name" type="text" size="10" />
```

name と email の入力に `size` 属性を値 `"20"` で付けてください。

# --hints--

name `input` に値 `size` の `"20"` 属性を付けてください。

```js
const nameInput = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(nameInput?.getAttribute('size'), '20');
```

email `input` に値 `size` の `"20"` 属性を付けてください。

```js
const emailInput = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(emailInput?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
