---
id: 67fccc2463253b213a000142
title: Construir un reproductor multimedia
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

En las lecciones anteriores, se te presentó cómo trabajar con los elementos `audio` y `video`. En este laboratorio, construirás un reproductor multimedia que mostrará una pista de `audio` y un `video` con una transcripción.

Para el elemento `audio`, necesitará incluir un elemento `source` que se utiliza para especificar el medio que se está utilizando. 

Aquí hay un ejemplo:

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

El elemento `source` también se puede usar en el elemento `video` de esta manera:

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debería tener un elemento `h1` para el título principal de la página.
2. Debería tener tres elementos `section`.
3. Dentro del primer elemento `section`, debería tener un elemento `h2` para el título de la canción que se está reproduciendo.
4. Debajo del elemento `h2`, debería tener un elemento `audio` con el atributo `controls` y un atributo `aria-label`.
5. Dentro del elemento `audio`, debería tener un elemento `source` con un atributo `src` que apunte a un archivo de audio y un atributo `type`. Es libre de usar esta URL de audio si lo desea: `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. Dentro del segundo elemento `section`, debería tener un elemento `h2` para el título del video que se está reproduciendo.
7. Debajo del elemento `h2`, debería tener un elemento `video` con atributos `controls`, `width` y un atributo `aria-label`.
8. Dentro del elemento `video`, debes tener un elemento `source` con un atributo `src` que apunte a un archivo de video y un atributo `type`. Puedes usar esta URL de video si quieres: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. Debajo del elemento `source`, debes tener un elemento `track` con un atributo `src` que apunte a un archivo de subtítulos y un atributo `kind`, un atributo `srclang` y un atributo `label`. Puedes usar esta URL de subtítulos si quieres: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. Dentro del tercer elemento `section`, debes tener un elemento `h2` para el título de la sección, por ejemplo, "Transcript".
11. Debajo del elemento `h2`, debería tener un elemento `p` con la transcripción del video.

# --hints--

Debería tener un elemento `h1`.

```js
assert.isNotNull(document.querySelector("h1"));
```

Su elemento `h1` debería tener algo de texto que represente el título principal de la página.  

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```  

Debería tener tres elementos `section`.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

La primera `section` debería contener un elemento `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

Su `h2` debería tener algo de texto que represente el título de la canción que se está reproduciendo.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

La primera `section` debería contener un elemento `audio`.  

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

Su elemento `audio` debería tener un atributo `controls`.  

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

Su elemento `audio` debería tener un atributo `aria-label` con texto que describa el audio.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

Su elemento `audio` debería tener un elemento `source`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

Su elemento `source` debería tener un atributo `src` que apunte a un archivo de audio.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

Su elemento `source` debería tener un atributo `type` con un valor que especifique el tipo de audio.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

La segunda `section` debería contener un elemento `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

Su `h2` debería tener algo de texto que represente el título del video que se está reproduciendo.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

La segunda `section` debería contener un elemento `video`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

Su elemento `video` debería tener un atributo `controls`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

Su elemento `video` debería tener un atributo `width` con un valor que especifique el ancho del video.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

El elemento `video` debería tener un elemento `source`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

El elemento `source` debería tener un atributo `src` que apunte a un archivo de video.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

El elemento `source` debería tener un atributo `type` con un valor que especifique el tipo de video.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

Su elemento `video` debería tener un elemento `track`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

Su elemento `track` debería tener un atributo `src` que apunte a un archivo de subtítulos.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

Su elemento `track` debería tener un atributo `kind` con texto que describa el tipo de archivo.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

Su elemento `track` debería tener un atributo `srclang` con texto que describa el idioma de los subtítulos.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

Su elemento `track` debería tener un atributo `label` con texto que describa el idioma de los subtítulos.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

La tercera `section` debería contener un elemento `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

Su `h2` debería tener algo de texto que represente el título de la sección.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

La tercera `section` debería contener un elemento `p` para la transcripción.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
