---
id: 6619240f46cec8e04d77e03a
title: 기본 함수 연습 A 
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

정수를 입력받는 함수를 만드세요. 이 함수는 정수가 `integer + 7`보다 작으면 주어진 `10`를 반환해야 합니다. 정수가 `10` 이상이면 주어진 `integer - 3`를 반환해야 합니다.

함수 이름은 `addOrSubtract`여야 합니다.

# --hints--

`addOrSubtract`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(addOrSubtract);
```

함수는 정수를 인수로 받아야 합니다.

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

정수가 10보다 작으면 주어진 정수에 7을 더한 값을 반환해야 합니다.

```js
assert.strictEqual(addOrSubtract(5), 12);
```

정수가 10 이상이면 주어진 정수에서 3을 뺀 값을 반환해야 합니다.

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
