---
id: 661e131f068359c3ccf2f4d6
title: 기본 함수 연습 B
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

두 개의 매개변수를 받아 그 곱을 반환하는 `multiply`라는 이름의 함수를 작성하세요.

# --hints--

`multiply`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(multiply);
```

함수는 두 개의 정수를 인수로 받아야 합니다.

```js
assert.match(multiply.toString(), /\s*multiply\(\s*\w+\s*,\s*\w+\s*\)/);
```

두 정수의 곱을 반환해야 합니다.

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js 
function multiply(a, b) {
  return a * b;
}
```
