---
id: 661e17c6068359c3ccf2f4d8
title: 기본 함수 연습 D
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

문자열을 매개변수로 받아 문자열의 마지막 글자를 반환하는 `lastLetter`라는 함수를 작성하세요.

# --hints--

`lastLetter`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(lastLetter);
```

함수는 문자열을 매개변수로 받아야 합니다.

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

문자열의 마지막 글자를 반환해야 합니다.

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
