---
id: 661e151f068359c3ccf2f4d7
title: 기본 함수 연습 C
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

문자열을 매개변수로 받아 첫 글자를 대문자로 바꾼 새 문자열을 반환하는 `capitalize`라는 함수를 작성하세요.

# --hints--

`capitalize`라는 함수가 있어야 합니다.

```js
assert.isFunction(capitalize);
```

함수는 문자열을 매개변수로 받아야 합니다.

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

함수는 첫 글자가 대문자인 새 문자열을 반환해야 합니다.

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
