---
id: 699d1f8e51fcdc8230f779fb
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop bauen Sie ein Katalogsystem für eine Erb-Bibliothek. Sie werden rohe Buchdaten parsen, Einträge suchen und gruppieren, die Datenqualität validieren und Ergebnisse in verschiedene Formate exportieren.

Jarteintrag wird als einzelner String im Format `"Title | Author | Year | Location"` gespeichert, wobei die Felder durch das Pipe-Zeichen (`|`) getrennt sind.

Sie verwenden im gesamten Workshop ein Array dieser Strings, um Ihre Katalogfunktionen zu erstellen und zu testen. Erstellen Sie ein Array mit dem Namen `rawCatalogCards` und fügen Sie ihm die folgenden zwei Strings hinzu:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

Sie sollten eine Variable mit dem Namen `rawCatalogCards` deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` sollte ein Array sein.

```js
assert.isArray(rawCatalogCards);
```

Das erste Element von `rawCatalogCards` sollte `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"` sein.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

Das zweite Element von `rawCatalogCards` sollte `"The Shining | King, Stephen | 1977 | Shelf K1"` sein.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
