---
id: 699d20623e57e8ffee885459
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

`parseCard` wurde getestet und funktioniert korrekt. Die Variable `cardResult` wird nicht mehr benötigt.

Beginnen Sie damit, die Zeile `const cardResult` und die Zeile `console.log(cardResult)` zu entfernen. Erstellen Sie dann eine Funktion mit dem Namen `parseCatalog`, die `rawCards` als Parameter entgegennimmt. Deklarieren Sie vorerst ein leeres Array mit dem Namen `catalog` innerhalb der Funktion und geben Sie es zurück.

Rufen Sie unterhalb der Funktion `parseCatalog(rawCatalogCards)` auf, weisen Sie das Ergebnis einer Variablen namens `catalog` zu und geben Sie `catalog.length` in der Konsole aus.

# --hints--

`parseCatalog` sollte eine Funktion sein, die ein Array zurückgibt.

```js
assert.isArray(parseCatalog(rawCatalogCards));
```

`catalog` sollte definiert und auf `parseCatalog(rawCatalogCards)` gesetzt sein.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+catalog\s*=\s*parseCatalog\s*\(/);
```

Sie sollten die Variable `cardResult` und den zugehörigen Aufruf `console.log()` löschen.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bcardResult\b/);
```

Sie sollten `catalog.length` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*catalog\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

--fcc-editable-region--
const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
--fcc-editable-region--
```
