---
id: 699d20623e57e8ffee885455
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Wie Sie an der Ausgabe in der Konsole sehen können, hat jeder Teil nach dem Aufteilen immer noch zusätzliche führende und/oder nachfolgende Leerzeichen. Um jedes Element zu bereinigen, müssen Sie jedes einzeln verarbeiten.

Dies ist ein perfekter Anwendungsfall für eine `for`-Schleife, die einen Codeblock einmal für jedes Element in einem iterierbaren Objekt wie einem Array wiederholt:

```js
for (let i = 0; i < array.length; i++) {
  // code to run on each iteration
}
```

`i` beginnt bei `0` (dem ersten Index), die Schleife läuft, solange `i` kleiner als die Länge des Arrays ist, und `i++` erhöht `i` nach jeder Iteration um eins.

Deklarieren Sie innerhalb von `parseCard` ein leeres Array mit dem Namen `trimmedParts`. Fügen Sie dann eine `for`-Schleife hinzu, die `parts` durchläuft. Lassen Sie den Schleifenkörper vorerst leer, da Sie ihn im nächsten Schritt ausfüllen werden.

# --hints--

Sie sollten eine Variable namens `trimmedParts` deklarieren, die auf ein leeres Array gesetzt ist.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+trimmedParts\s*=\s*\[\s*\]/);
```

Sie sollten eine `for`-Schleife verwenden, um `parts` zu durchlaufen.

```js
assert.match(parseCard.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*parts\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--
  
--fcc-editable-region--
  return parts;
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
