---
id: 699d20623e57e8ffee885450
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Der Rest der Katalogkarten wurde hinzugefügt. Beachten Sie die drei Einträge gegen Ende mit fehlenden Feldern; diese werden Sie in einem späteren Schritt bearbeiten.

`parseCard` nimmt einen rohen Katalog-String und gibt ein strukturiertes Objekt mit vier Eigenschaften zurück: `title`, `author`, `year` und `location`. Erstellen Sie vorerst eine Funktion `parseCard` mit einem Parameter namens `rawString`.

# --hints--

Sie sollten eine Funktion mit dem Namen `parseCard` haben.

```js
assert.isFunction(parseCard);
```

`parseCard` sollte eine Funktion mit einem Parameter namens `rawString` sein.

```js
const regex = __helpers.functionRegex('parseCard', ['rawString']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

--fcc-editable-region--

--fcc-editable-region--
```
