---
id: 699d20623e57e8ffee885458
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Das Jahr-Feld kommt als String an, aber Sie benötigen es als Zahl für eine spätere Berechnung. Verwenden Sie `parseInt()`, um es umzuwandeln.

Das Jahr kann auch leer sein, daher brauchen Sie einen Fallback. Der ternäre Operator `condition ? valueIfTrue : valueIfFalse` gibt basierend auf einer Bedingung einen von zwei Werten zurück.

Aktualisieren Sie die `year`-Eigenschaft in der return-Anweisung zu:

```js
year: year ? parseInt(year) : "Unknown"
```

# --hints--

Wenn die Funktion mit einer Karte aufgerufen wird, die ein gültiges Jahr hat, sollte die `year`-Eigenschaft des zurückgegebenen Objekts eine Zahl und kein String sein.

```js
assert.isNumber(parseCard("The Shining | King, Stephen | 1977 | Shelf K1").year);
```

Wenn die Funktion mit `"The Shining | King, Stephen | 1977 | Shelf K1"` aufgerufen wird, sollte die `year`-Eigenschaft `1977` sein.

```js
assert.strictEqual(parseCard("The Shining | King, Stephen | 1977 | Shelf K1").year, 1977);
```

Wenn die Funktion mit einem String aufgerufen wird, dessen Jahr-Feld leer ist, sollte die `year`-Eigenschaft `"Unknown"` sein.

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "Unknown");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year,
    location: location || "Unknown"
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
