---
id: 699d20623e57e8ffee885451
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Die `parseCard`-Funktion nimmt einen rohen Katalog-String und gibt ein strukturiertes Objekt zurück. Bevor Sie Ihre Funktion ausbauen, richten Sie eine Möglichkeit ein, um die Ausgabe während der Arbeit zu sehen.

Rufen Sie `parseCard` mit `rawCatalogCards[2]` als Argument auf, weisen Sie das Ergebnis einer Variablen namens `cardResult` zu, und geben Sie `cardResult` in der Konsole aus.

Im Moment hat `parseCard` keinen Funktionskörper, daher gibt sie `undefined` zurück. Während Sie in den nächsten Schritten Code zur Funktion hinzufügen, zeigt Ihnen die ausgegebene Ausgabe genau, was die Funktion jeweils zurückgibt.

# --hints--

Sie sollten `parseCard` mit `rawCatalogCards[2]` als Argument aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)/);
```

Sie sollten das Ergebnis des Aufrufs von `parseCard` mit `rawCatalogCards[2]` einer Variablen namens `cardResult` zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+cardResult\s*=\s*parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)\s*;?/);
```

Sie sollten `cardResult` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*cardResult\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {}

--fcc-editable-region--

--fcc-editable-region--
```
