---
id: 699d20623e57e8ffee885454
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Der erste Schritt beim Parsen einer Katalogkarte besteht darin, die vier Felder zu trennen. Die `.split()`-Methode zerlegt einen String in ein Array von Teilstrings bei jedem Vorkommen des von Ihnen angegebenen Trennzeichens.

Zum Beispiel:

```js
"The Shining | King, Stephen | 1977 | Shelf K1".split("|")
// ["The Shining ", " King, Stephen ", " 1977 ", " Shelf K1"]
```

Rufen Sie innerhalb von `parseCard` die `.split("|")`-Methode auf `rawString` auf und geben Sie das Ergebnis zurück. Nach diesem Schritt sollte die Konsole ein Array mit vier Strings anzeigen.

**Hinweis:** Möglicherweise bemerken Sie zusätzliche Leerzeichen um die Werte. Diese werden Sie im nächsten Schritt bereinigen.

# --hints--

Sie sollten `.split("|")` auf `rawString` innerhalb Ihrer `parseCard`-Funktion aufrufen.

```js
assert.match(parseCard.toString(), /rawString\s*\.\s*split\s*\(\s*('|"|`)\|\1\s*\)\s*;?/);
```

`parseCard` sollte ein Array zurückgeben.

```js
assert.isArray(parseCard("A | B | C | D"));
```

Wenn `"A | B | C | D"` mit `parseCard` aufgerufen wird, sollte es ein Array mit vier Elementen zurückgeben.

```js
assert.strictEqual(parseCard("A | B | C | D").length, 4);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
--fcc-editable-region--
  
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
