---
id: 699d20623e57e8ffee88545d
title: Schritt 19
challengeType: 1
dashedName: step-19
---

# --description--

Nach der `"Unknown"`-Prüfung fügen Sie die Logik zur Gruppierung nach Jahrzehnten für Bücher mit numerischen Jahreszahlen hinzu. Deklarieren Sie `const decade = Math.floor(book.year / 10) * 10`. Zum Beispiel ergibt `Math.floor(1977 / 10)` `197`, multipliziert mit `10` ergibt `1970`.

Deklarieren Sie dann `const decadeKey =` `` `${decade}s` `` to produce a string like `"1970s"`.

Verwenden Sie Klammernotation, um das Buch zu gruppieren: Prüfen Sie, ob `grouped[decadeKey]` noch nicht existiert, und initialisieren Sie es als leeres Array, dann fügen Sie `book` hinzu. Hier ist Klammernotation erforderlich, weil der Schlüssel in einer Variablen gespeichert ist.

# --hints--

Sie sollten eine Variable namens `decade` deklarieren, die auf `Math.floor(book.year / 10) * 10` gesetzt ist.

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+decade\s*=\s*Math\.floor\s*\(\s*book\.year\s*\/\s*10\s*\)\s*\*\s*10/);
```

Sie sollten eine Variable namens `decadeKey` deklarieren, die `decade` und den Buchstaben `"s"` enthält.

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+decadeKey\s*=.*decade/s);
assert.match(groupByDecade.toString(), /decadeKey.*s/s);
```

Sie sollten prüfen, ob `grouped[decadeKey]` noch nicht existiert, und es als leeres Array initialisieren.

```js
assert.match(groupByDecade.toString(), /!\s*grouped\s*\[\s*decadeKey\s*\]/);
assert.match(groupByDecade.toString(), /grouped\s*\[\s*decadeKey\s*\]\s*=\s*\[\s*\]/);
```

Sie sollten `book` in `grouped[decadeKey]` hinzufügen.

```js
assert.match(groupByDecade.toString(), /grouped\s*\[\s*decadeKey\s*\]\s*\.push\s*\(\s*book\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

function groupByDecade(catalog) {
  const grouped = {};
  for (let i = 0; i < catalog.length; i++) {
    const book = catalog[i];
    if (book.year === "Unknown") {
      if (!grouped["Unknown"]) {
        grouped["Unknown"] = [];
      }
      grouped["Unknown"].push(book);
      continue;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
