---
id: 699d20623e57e8ffee88545a
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Nachdem der Katalog geparst wurde, können Sie ihn nach Feldern durchsuchen. `findByAuthor` filtert Einträge, deren Autor einen Suchbegriff enthält.

Erstellen Sie eine Funktion namens `findByAuthor`, die `catalog` und `author` als Parameter entgegennimmt. Deklarieren Sie darin eine Variable namens `searchTerm` und setzen Sie sie auf die Kleinbuchstaben-Version von `author`, um später eine Groß-/Kleinschreibungs-unabhängige Übereinstimmung zu ermöglichen.

Deklarieren Sie außerdem eine Variable namens `results` und weisen Sie ihr ein leeres Array zu, um die Treffer zu speichern.

# --hints--

Sie sollten eine Funktion namens `findByAuthor` mit `catalog` und `author` als Parametern haben.

```js
const regex = __helpers.functionRegex('findByAuthor', ['catalog', 'author']);
assert.match(__helpers.removeJSComments(code), regex);
```

Sie sollten eine Variable namens `searchTerm` deklarieren, die auf `author.toLowerCase()` gesetzt ist.

```js
assert.match(findByAuthor.toString(), /searchTerm\s*=\s*author\.toLowerCase\s*\(\s*\)/);
```

Sie sollten eine Variable namens `results` deklarieren, die auf ein leeres Array gesetzt ist.

```js
assert.match(findByAuthor.toString(), /(const|let|var)\s+results\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);
--fcc-editable-region--

--fcc-editable-region--
```
