---
id: 69aa11c2584523e6e991330c
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Die `.trim()`-Methode entfernt Leerzeichen an beiden Enden eines Strings. Zum Beispiel gibt `" hello ".trim()` `"hello"` zurück.

Innerhalb des Schleifenkörpers fügen Sie `parts[i].trim()` in `trimmedParts` ein. Aktualisieren Sie dann die bestehende `return`-Anweisung, sodass sie `trimmedParts` anstelle von `parts` zurückgibt.

Nach diesem Schritt sollte die protokollierte Ausgabe saubere Werte ohne zusätzliche Leerzeichen zeigen.

# --hints--

Sie sollten `parts[i].trim()` innerhalb der Schleife in `trimmedParts` einfügen.

```js
assert.match(parseCard.toString(), /trimmedParts\s*\.\s*push\s*\(\s*parts\s*\[\s*\w+\s*\]\s*\.\s*trim\s*\(\s*\)\s*\)/);
```

`parseCard` sollte `trimmedParts` zurückgeben.

```js
const explorer = await __helpers.Explorer(code);
const { parseCard } = explorer.allFunctions;
assert.isTrue(parseCard.hasReturn("trimmedParts"));
```

`parseCard` sollte ein Array zurückgeben, bei dem jedes Element keine führenden oder nachfolgenden Leerzeichen hat.

```js
assert.isArray(parseCard("  A  |  B  |  C  |  D  "));
```

Wenn `"  A  |  B  |  C  |  D  "` aufgerufen wird, sollte das erste Element des zurückgegebenen Arrays `"A"` sein.

```js
assert.strictEqual(parseCard("  A  |  B  |  C  |  D  ")[0], "A");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
--fcc-editable-region--
    
  }
  return parts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
