---
id: 66bb6a9c2dd58b73cd759034
title: Створіть застосунок для журналу оцінок
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з назвою `getAverage`, яка приймає масив оцінок за тести як параметр і повертає середній бал.
1. У вас має бути функція з назвою `getGrade`, яка приймає оцінку студента як параметр і повертає рядок, що представляє буквений бал на основі оцінки. Ось діапазони оцінок і відповідні їм буквені бали:

    | Діапазон оцінок | Бал   |
    | ----------- | ------- |
    | `100`       | `"A+"`  |
    | `90 - 99`   | `"A"`   |
    | `80 - 89`   | `"B"`   |
    | `70 - 79`   | `"C"`   |
    | `60 - 69`   | `"D"`   |
    | `0 - 59`    | `"F"`   |

1. У вас має бути функція з назвою `hasPassingGrade`, яка приймає оцінку як параметр і повертає або `true`, або `false` залежно від того, чи відповідає оцінка прохідному балу.
1. Функція `hasPassingGrade` має використовувати функцію `getGrade` для отримання буквенного балу і на її основі визначати, чи є оцінка прохідною. Прохідний бал — це будь-який, крім `"F"`.
1. У вас має бути функція з назвою `studentMsg`, яка приймає масив оцінок і оцінку студента як параметри. Функція має повертати рядок у форматі:

    - `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`, якщо студент склав курс.
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`, якщо студент не склав курс.

    Замініть `average-goes-here` на середній бал усіх оцінок, а `grade-goes-here` — на буквений бал студента.
    Використовуйте `getAverage` для отримання середнього балу і `getGrade` для отримання буквенного балу студента.

# --hints--

У вас має бути функція з назвою `getAverage`.

```js
assert.isFunction(getAverage);
```

Ваша функція `getAverage` має повертати число.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` має повертати `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` має повертати `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` має повертати `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` має повертати `47.25`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` має повертати `47.125`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

Ваша функція `getAverage` має повертати середній бал тестових оцінок.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

У вас має бути функція з назвою `getGrade`.

```js
assert.isFunction(getGrade);
```

Ваша функція `getGrade` має повертати `"A+"`, якщо оцінка є `100`.

```js
assert.strictEqual(getGrade(100), "A+");
```

Ваша функція `getGrade` має повертати `"A"`, якщо оцінка знаходиться між `90` і `99`.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

Ваша функція `getGrade` має повертати `"B"`, якщо оцінка знаходиться між `80` і `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Ваша функція `getGrade` має повертати `"C"`, якщо оцінка знаходиться між `70` і `79`.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

Ваша функція `getGrade` має повертати `"D"`, якщо оцінка знаходиться між `60` і `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Ваша функція `getGrade` має повертати `"F"`, якщо оцінка знаходиться між `0` і `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


У вас має бути функція з назвою `hasPassingGrade`.

```js
assert.isFunction(hasPassingGrade);
```

Ваша функція `hasPassingGrade` має повертати булеве значення.

```js
assert.isBoolean(hasPassingGrade(100));
```

Ваша функція `hasPassingGrade` має повертати `true`, якщо бал є `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Ваша функція `hasPassingGrade` має повертати `false`, якщо бал є `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Ваша функція `hasPassingGrade` має повертати `true` для всіх прохідних балів.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

У вас має бути функція з назвою `studentMsg`.

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` має повертати таке повідомлення: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` має повертати таке повідомлення: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` має повертати таке повідомлення: `"Class average: 48.25. Your grade: B. You passed the course."`.

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` має повертати таке повідомлення: `"Class average: 45.625. Your grade: C. You passed the course."`.

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

Ваша функція `studentMsg` має повертати правильне повідомлення на основі оцінки студента та середнього балу класу.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
