---
id: 65576ff7888f9e96f52a4be1
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

تخزن القواميس البيانات على شكل أزواج *مفتاح*-*قيمة*. يُفصل المفتاح عن القيمة المقابلة بنقطتين. ويُفصل كل زوج مفتاح-قيمة عن الزوج التالي بفاصلة:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

أضف زوج مفتاح-قيمة جديد إلى قاموسك. استخدم السلسلة النصية `'species'` كمفتاح، والسلسلة النصية `'guinea pig'` كقيمة.

# --hints--

يجب أن يحتوي قاموس `copper` على مفتاح `'species'` بقيمة `'guinea pig'`.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

يجب أن يحتوي قاموس `copper` على زوج مفتاح-قيمة واحد فقط.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
