---
id: 65578cb031cd93b77a285db2
title: الخطوة 38
challengeType: 20
dashedName: step-38
---

# --description--

ستستكشف الدالة جميع العقد المتصلة بالعقدة الابتدائية. ستحسب أقصر المسارات لكل منها. ثم ستزيل العقدة الابتدائية من العقد غير المزارة.

بعد ذلك، ستُزار أقرب عقدة مجاورة وسيُكرر هذا الإجراء حتى تُزار جميع العقد.

من الآن فصاعدًا، ستعمل على الحلقة الرئيسية التي تستكشف العقد في الرسم البياني. لتجنب مشاكل تشغيل حلقة لا نهائية أثناء تطوير الخوارزمية، حوّل استدعاء الدالة إلى تعليق.

# --hints--

يجب أن تحوّل استدعاء الدالة إلى تعليق.

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
