---
id: 655774d01daeeaa1978b99d5
title: الخطوة 29
challengeType: 20
dashedName: step-29
---

# --description--

في البداية، تُعتبر جميع العقد الأخرى في الرسم البياني على مسافة لا نهائية من العقدة المصدر، لأن المسافة لم تُحدد بعد.

أنشئ بند `else` وعيّن قيمة لا نهائية للعقدة في قاموس `distances`. استخدم لذلك الدالة `float()` مع السلسلة النصية `'inf'` كمعلمة لتوليد رقم عشري يمثل اللانهاية الموجبة.

# --hints--

يجب أن يكون لديك بند `else`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

يجب تعيين `float('inf')` إلى `distances[node]` داخل بند `else` الجديد.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
