---
id: 655771d889132f9ccd341060
title: الخطوة 20
challengeType: 20
dashedName: step-20
---

# --description--

تُسمى الرسم البياني *موزونًا* عندما ترتبط حوافها بأوزان تمثل مسافة أو زمنًا أو قيمة كمية أخرى.

في حالتك، ستكون هذه الأوزان هي المسافات بين كل عقدة أو نقطة في الفضاء. لتمثيل رسم بياني موزون يمكنك تعديل القاموس الخاص بك باستخدام قائمة من الصفوف المرتبة لكل قيمة.

العنصر الأول في الصف المرتب سيكون العقدة المتصلة، والعنصر الثاني سيكون عددًا صحيحًا يشير إلى المسافة.

عدّل `my_graph['A']` إلى قائمة من الصفوف المرتبة، مع اعتبار أن المسافة بين `A-B` هي `3` والمسافة بين `A-D` هي `1`.

# --hints--

يجب أن تكون `my_graph["A"]` قائمة تحتوي على الصفوف المرتبة `('B', 3)` و `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

يجب أن يحتوي `my_graph` على 4 مفاتيح مسماة `'A'` و `'B'` و `'C'` و `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
