---
id: 655774955b097ea14897db12
title: الخطوة 28
challengeType: 20
dashedName: step-28
---

# --description--

المسافة من العقدة البداية تساوي صفرًا، لأن الخوارزمية تبدأ تقييمها من هناك مباشرة.

بعد إضافة `node` إلى `unvisited` في الحلقة، أنشئ جملة `if` تُنفذ إذا كانت العقدة تساوي العقدة البداية. ثم عيّن القيمة `0` لتلك العقدة داخل قاموس `distances`.

# --hints--

يجب أن تنشئ جملة `if` تُنفذ عندما تكون `node` مساوية لـ `start`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

داخل جملة `if` الجديدة، يجب أن تعيّن القيمة `0` للعقدة في قاموس `distances`.

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
