---
id: 65578d4fc3afc3b8f554c882
title: الخطوة 41
challengeType: 20
dashedName: step-41
---

# --description--

تأخذ الدالة `min()` أيضًا معلمة خاصة بالكلمة المفتاحية فقط. بتمرير دالة كمعلمة إضافية إلى `min()`، يمكنك تعديل طريقة مقارنة عناصر القائمة.

النتيجة من السطر الذي كتبته في الخطوة السابقة هي العقدة التي تأتي أولاً حسب الترتيب الأبجدي. بدلًا من ذلك، تريد اختيار العقدة غير المزارة التي تملك أصغر مسافة من العقدة البداية.

مرر `key=distances.get` كالمعلمة الثانية لاستدعاء `min()` الخاص بك. بهذه الطريقة، ستتم المقارنة بناءً على القيمة التي يملكها كل عنصر من قائمة `unvisited` داخل قاموس `distances`.

# --hints--

يجب أن تمرر `key=distances.get` كالمعلمة الثانية لاستدعاء `min()` الخاص بك.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
