---
id: 6557709b0aee699a6a00528c
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

على سبيل المثال، يمكن استخدام رسم بياني لتمثيل نقطتين في الفضاء، `A` و `B`، متصلتين بمسار. مثل هذا الرسم البياني يتكون من عقدتين متصلتين بحافة.

استبدل المفتاح الموجود `'species'` بالسلسلة النصية `'A'`. ثم استبدل القيمة المقابلة بالسلسلة النصية `'B'` لتمثيل الاتصال بين العقدتين `'A'` و `'B'`.

# --hints--

يجب أن يحتوي القاموس على مفتاح `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

يجب أن يكون للمفتاح `'A'` القيمة `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

يجب أن يحتوي القاموس على مفتاحين.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
