---
id: 65579228c669fcbebffd01d5
title: الخطوة 52
challengeType: 20
dashedName: step-52
---

# --description--

الخطأ البرمجي الآخر دقيق. عندما يُعثر على مسافة أقصر لعقدة مجاورة، يُعيَّن `paths[current]` إلى مسار العقدة المجاورة، `paths[node]`.

هذا يعني أن المتغيرين يشيران إلى نفس القائمة. بما أن القوائم قابلة للتغيير، عند إضافة العقدة المجاورة إلى مسارها، يتغير كل من `paths[node]` و`paths[current]` لأنهما نفس القائمة. هذا يؤدي إلى مسارات خاطئة، رغم أن المسافات صحيحة.

يمكنك إصلاح هذا الخطأ البرمجي بتعيين نسخة من `paths[current]` إلى مسار العقدة المجاورة. لهذا يمكنك استخدام صيغة القطع (slice):

```py
my_list[:]
```

حيث `my_list` هي القائمة التي تريد نسخها. عدل التعيين الحالي `paths[node] = paths[current]` داخل كتلة `if` الخاصة بك باستخدام القطع على `paths[current]`.

# --hints--

يجب أن تعيّن `paths[current][:]` إلى مسار العقدة المجاورة.

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
