---
id: 6557712d77ce2d9bd7e63afd
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

أضف عقدة أخرى متصلة بـ `B` إلى الرسم البياني وسمها `C`.

عدّل القاموس الحالي لتمثيل هذا الترتيب: أضف مفتاحًا جديدًا `'C'` إلى `my_graph` وأعطه القيمة السلسلة النصية `'B'`.

غيّر أيضًا قيمة المفتاح الموجود `'B'` إلى القائمة `['A', 'C']` لتمثيل الاتصالات المتعددة لعقدة `'B'` الخاصة بك.

# --hints--

يجب أن يحتوي قاموسك على 3 مفاتيح — `'A'` و `'B'` و `'C'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

يجب أن تكون قيمة `my_graph['A']` هي السلسلة النصية `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

يجب أن تكون `my_graph['B']` قائمة.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

يجب أن تكون قيمة `my_graph['B']` قائمة تحتوي على `'A'` و `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

يجب أن تكون قيمة `my_graph['C']` هي السلسلة النصية `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
