---
id: 65578d0f6c78a0b868a43b9c
title: الخطوة 40
challengeType: 20
dashedName: step-40
---

# --description--

داخل حلقة `while`، أول شيء تفعله هو تعريف العقدة الحالية التي ستزور. لهذا يمكنك استخدام الدالة `min()`. تُرجع أصغر عنصر من القابل للتكرار الممرر كمعلمة.

احذف `pass`، ثم أنشئ متغيرًا اسمه `current` وقم بتعيين له `min(unvisited)`.

# --hints--

يجب أن تنشئ متغير `current` داخل حلقة `while` الخاصة بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

يجب أن تعيّن `min(unvisited)` إلى متغير `current` الخاص بك. تذكر حذف `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
