---
id: 655777060d8ddea6741be1b1
title: الخطوة 32
challengeType: 20
dashedName: step-32
---

# --description--

جميع المسافات في `distances` مضبوطة على ما لا نهاية، باستثناء عقدة البداية. تحتوي قائمة `unvisited` على جميع العقد في الرسم البياني الخاص بك. لكن في الواقع، لا تحتاج إلى حلقة `for` تلك لتحقيق هذه النتيجة.

احذف حلقة `for` كاملة مع محتواها.

# --hints--

يجب حذف حلقة `for` وكل الشفرة المسكنة بداخلها.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)distances\s*=\s*\{\s*\}\s*\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
