---
id: 655773f8b8b5db9fc6d0ae76
title: الخطوة 25
challengeType: 20
dashedName: step-25
---

# --description--

لمتابعة العقد التي تمت زيارتها، تحتاج إلى قائمة بكل العقد في الرسم البياني. بمجرد زيارة عقدة، تُحذف من تلك القائمة.

الآن، استبدل الكلمة المفتاحية `pass` بمتغير اسمه `unvisited` وقم بتعيين قائمة فارغة له.

# --hints--

يجب أن يكون لديك متغير اسمه `unvisited` داخل الدالة `shortest_path`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

يجب تعيين قائمة فارغة للمتغير `unvisited`. تذكر حذف `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
