---
id: 65577791ad8c26a7705e2919
title: الخطوة 34
challengeType: 20
dashedName: step-34
---

# --description--

باستخدام تعبير القاموس، يمكنك إنشاء قاموس يبدأ من قاموس موجود بالفعل:

```py
{key: val for key in dict}
```

في المثال أعلاه، `val` هي القيمة التي سيحصل عليها `key` في القاموس الجديد، و`dict` هو القاموس الموجود.

تريد تتبع المسارات بين العقدة البداية وكل عقدة أخرى.

بعد المتغير `distances`، أنشئ متغيرًا باسم `paths` وعيّنه إلى قاموس يحتوي على كل المفاتيح من `graph`. عيّن قائمة فارغة لكل مفتاح واستخدم تعبير القاموس لبناء القاموس الخاص بك.

# --hints--

يجب أن يكون لديك متغير باسم `paths`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

يجب أن يستخدم متغير `paths` صيغة تعبير القاموس لتعيين قائمة فارغة لكل عقدة في graph.

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
