---
id: 655773b0591c5f9f4045883e
title: الخطوة 24
challengeType: 20
dashedName: step-24
---

# --description--

تبدأ الخوارزمية عند عقدة محددة. ثم تستكشف الرسم البياني لإيجاد أقصر مسار بين العقدة البداية، أو *المصدر*، وجميع العقد الأخرى.

لهذا تحتاج دالتك إلى معلمتين: `graph` و `start`. أضفهما إلى إعلان الدالة.

# --hints--

يجب أن تأخذ دالتك `graph` و `start` كمعلمتين، بهذا الترتيب.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
