---
id: 65577236b056379d5dbc7000
title: الخطوة 23
challengeType: 20
dashedName: step-23
---

# --description--

الآن ستبدأ في تطوير الخوارزمية لحساب أقصر مسار بين كل عقدة في الرسم البياني الجديد.

أعلن دالة فارغة باسم `shortest_path`. استخدم الكلمة المفتاحية `pass` لملء جسم الدالة.

# --hints--

يجب أن تعرف دالة باسم `shortest_path`.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
