---
id: 6566195b0a021bb660b2b4b1
title: الخطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

عدّل `my_graph['B']` إلى قائمة من الصفوف المرتبة، حيث يكون العنصر الأول في الصف المرتب هو العقدة المتصلة، والعنصر الثاني هو المسافة. المسافة بين `B-C` هي `4`.

# --hints--

يجب أن تكون `my_graph['B']` قائمة تحتوي على الصفوف المرتبة `('A', 3)` و `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

يجب أن يحتوي `my_graph` على 4 مفاتيح مسماة `'A'` و `'B'` و `'C'` و `'D'`.

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
