---
id: 65661b72d6745ebec6a96923
title: الخطوة 22
challengeType: 20
dashedName: step-22
---

# --description--

بنفس الطريقة، عدّل القائمتين المتبقيتين مع الأخذ في الاعتبار أن المسافة بين `C-D` هي `7`.

# --hints--

يجب أن تكون `my_graph['C']` قائمة تحتوي على الصفوف المرتبة `('B', 4)` و `('D', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

يجب أن تكون `my_graph['D']` قائمة تحتوي على الصفوف المرتبة `('A', 1)` و `('C', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

يجب أن يحتوي `my_graph` على 4 مفاتيح مسماة `'A'` و `'B'` و `'C'` و `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

يجب أن تكون `my_graph['A']` قائمة تحتوي على الصفوف المرتبة `('B', 3)` و `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

يجب أن تكون `my_graph['B']` قائمة تحتوي على الصفوف المرتبة `('A', 3)` و `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
