---
id: 6578b57361f2f132a02e2a18
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

يمكنك الوصول إلى البيانات المخزنة في القاموس عبر مفاتيحه:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

بعد قاموسك، اتبع المثال أعلاه للوصول إلى المفتاح `'species'` في `copper` وطباعة النتيجة.

# --hints--

لا يجب تعديل تعيين قاموسك.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

يجب استخدام `copper['species']` للوصول إلى قيمة المفتاح `'species'`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

يجب استدعاء `print()` مع تمرير `copper['species']` كمعلمة.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
