---
id: 65789506b30453080f77470c
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

حتى الآن، تعرفت على أنواع بيانات مختلفة:

- أنواع بيانات غير قابلة للتغيير، مثل الأعداد الصحيحة، والسلاسل النصية، والصفوف المرتبة، والقيم المنطقية.
- أنواع بيانات قابلة للتغيير، مثل القوائم والقواميس.

يُعرف القاموس باستخدام زوج من الأقواس المتعرجة، `{}`.

ابدأ بإنشاء متغير اسمه `copper` وقم بتعيين قاموس فارغ له باستخدام زوج من الأقواس المتعرجة، بنفس الطريقة التي تنشئ بها قائمة فارغة باستخدام زوج من الأقواس المربعة.

# --hints--

يجب أن يكون لديك متغير اسمه `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

يجب أن يحتوي متغير `copper` على قيمة قاموس فارغ. استخدم زوجًا من الأقواس المتعرجة لذلك.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
