---
id: 6578b13757611e2825beb8a5
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

يجب أن تكون المفاتيح فريدة داخل القاموس، ويجوز أن تكون فقط من أنواع البيانات غير القابلة للتغيير. هذا يعني أنه لا يمكنك استخدام قائمة أو قاموس آخر كمفاتيح.

أضف مفتاحًا آخر `'age'` إلى قاموسك وأعطه القيمة العدد الصحيح `2`.

# --hints--

يجب أن يكون لديك مفتاح `'age'` بقيمة `2` مباشرة بعد `'species': 'guinea pig'` داخل قاموس `copper`. لا تنسَ الفاصلة.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

يجب أن يحتوي قاموس `copper` على زوجين من أزواج المفتاح-القيمة.

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
