---
id: 6579dd49fa8a8e1fd06b85a9
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

الرسوم البيانية هي هياكل بيانات تمثل العلاقات بين أزواج من العناصر.
هذه العناصر، التي تُسمى *عقدًا*، يمكن أن تكون أشياء من الحياة الواقعية، كيانات، نقاط في الفضاء أو غيرها. الاتصالات بين العقد تُسمى *الحواف*.

إليك تمثيلًا بصريًا للرسم البياني:

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


أعد تسمية القاموس `copper` إلى `my_graph`. هذا سيمثل الرسم البياني لاختبار خوارزميتك.

# --hints--

يجب أن تعيد تسمية قاموس `copper` إلى `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

يجب أن تكون متغير `my_graph` قاموسًا.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
