---
id: 6579cd5f6dd62c189e53ddbb
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

الآن بعد مراجعتك الجوانب الأساسية للقواميس، يمكنك المتابعة لبناء خوارزمية أقصر مسار.

احذف كل سطر شفرة بعد الإعلان عن القاموس `copper`.

# --hints--

يجب حذف كل أسطر الشفرة بعد الإعلان عن قاموسك.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
