---
id: 6579ca0923cfa7162089d2f0
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

يمكن استخدام نفس الصيغة النحوية لتغيير قيمة مفتاح موجود.

قبل استدعاء `print()` مباشرة، وصول إلى المفتاح `'species'` وأعد تعيين قيمته إلى `'Cavia porcellus'`.

# --hints--

يجب إعادة تعيين `copper['species']` إلى السلسلة النصية `'Cavia porcellus'` قبل استدعاء `print()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
