---
id: 6579cbab9825b8170974c69a
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

يمكنك حذف زوج مفتاح-قيمة من القاموس باستخدام الكلمة المفتاحية `del`:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

قبل حلقة `for` مباشرة، استخدم الكلمة المفتاحية `del` لحذف المفتاح `'age'` وقيمته من `copper`.

# --hints--

يجب استخدام الكلمة المفتاحية `del` لحذف `copper['age']` قبل حلقة `for`.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
