---
id: 6579717f0920131304286804
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

لإضافة زوج مفتاح-قيمة جديد بعد الإعلان عن القاموس، يمكنك تحديد المفتاح بنفس الطريقة التي تصل بها إلى مفتاح موجود، وتعيين قيمة المفتاح الجديد باستخدام معامل التعيين:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

احذف استدعاء `print()` الخاص بك. ثم، بعد الإعلان عن `copper`، أضف المفتاح `'food'` إلى القاموس وحدد قيمته إلى `'hay'`.

# --hints--

يجب ألا يحتوي شفرتك على `print(copper['age'])`.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

يجب إضافة المفتاح `'food'` إلى `copper` بعد الإعلان عن القاموس.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

يجب تعيين `copper['food']` إلى `'hay'` بعد الإعلان عن القاموس.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
