---
id: 65affe602a2a828b0382a058
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

استبدل المفتاح `'age'` بالسلسلة النصية `'B'` واضبط قيمته على السلسلة النصية `'A'` لتمثيل الاتصال بين العقد في كلا الاتجاهين.

# --hints--

يجب أن يحتوي القاموس على مفتاح `'B'`.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

يجب أن تكون قيمة المفتاح `'B'` هي `'A'`.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

يجب أن يحتوي القاموس على مفتاح `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

يجب أن تكون قيمة المفتاح `'A'` هي `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

يجب أن يحتوي القاموس على مفتاحين.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
