---
id: 65b7cf140d34058d7ea8935f
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

إذا أردت التكرار على قيم مفاتيح القاموس، فطريقة واحدة هي استخدام الطريقة `.values()`.

عدّل حلقة `for` لتكرار `copper.values()` بدلاً من `copper` واطلع على الناتج.

# --hints--

يجب تعديل حلقة `for` لتكرار `copper.values()` بدلاً من `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
