---
id: 65b7d3319ebec69b983fb91b
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

كما ترى من الناتج، تنشئ `.items()` هيكل بيانات يخزن كل زوج مفتاح-قيمة في صف مرتب مميز. للتكرار على العناصر في تلك الصفوف المرتبة يمكنك إضافة متغير حلقة ثاني:

```py
for i, j in dict.items():
    print(i, j)
```

عدّل حلقة `for` لتأخذ متغيري حلقة واطبع كلاهما داخل جسم الحلقة.

# --hints--

يجب تعديل حلقة `for` لاستخدام متغيرين.

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

يجب طباعة متغيري الحلقة داخل حلقة `for`.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
