---
id: 65b7d25992879698180e6a71
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

أخيرًا، إذا أردت أن تتمكن من التكرار على أزواج المفتاح-القيمة، يمكنك استخدام طريقة `.items()`.

عدّل حلقة `for` لتكرار `copper.items()` بدلًا من `copper.values()`.

# --hints--

يجب تعديل حلقة `for` لتكرار `copper.items()` بدلًا من `copper.values()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
