---
id: 65b7cd2b7bd9a684ccf1dc16
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

للتكرار على مفاتيح القاموس، يمكنك ببساطة وضع القاموس داخل حلقة `for`. الشفرة أدناه ستطبع كل مفتاح في القاموس `dict`:

```py
for i in dict:
   print(i)
```

استبدل استدعاء `print()` بحلقة `for` تتكرر على `copper` وتطبع كل مفتاح.

# --hints--

يجب ألا يحتوي شفرتك على `print(copper)`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

يجب إنشاء حلقة `for` للتكرار على قاموس `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

يجب طباعة كل مفتاح من `copper` داخل حلقة `for` الخاصة بك.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
