---
id: 61696ef7ac756c829f9e4048
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

In this project you'll work with an external CSS file to style the page. We've already created a `styles.css` file for you. But before you can use it, you'll need to link it to the page.

Вкладіть елемент `link` в елемент `head`. Надайте йому атрибут `rel` зі значенням `"stylesheet"` та атрибут `href` зі значенням `"styles.css"`.

# --hints--

Ви повинні мати один елемент `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Елемент `link` має бути порожнім елементом, він не повинен мати кінцевий тег `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Елемент `link` повинен мати атрибут `rel` зі значенням `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Елемент `link` повинен мати атрибут `href` зі значенням `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Елемент `link` має бути в межах елемента `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
