---
id: 61695d1fbc003856628bf561
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Атрибут `charset` визначає кодування символів, яке використовується в документі. `utf-8` (формат перетворення Юнікоду – 8-біт) — це стандарт кодування символів, який використовується для електронного зв’язку.

В межах елемента `head` вкладіть елемент `meta` з атрибутом `charset` зі значенням `"utf-8"`.

# --hints--

Ви повинні мати один елемент `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

Елемент `meta` має бути порожнім елементом, він не повинен мати кінцевий тег `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Елемент `meta` повинен мати атрибут `charset` зі значенням `"utf-8"`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Елемент `meta` має бути в межах елемента `head`.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
