---
id: 616965351e74d4689eb6de30
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Ви можете мати декілька елементів `meta` на вебсторінці. Кожен елемент `meta` додає інформацію про сторінку, яка не може бути виражена іншими елементами HTML.

Додайте ще один елемент `meta` в межах `head`. Надайте йому атрибут `name` зі значенням `"viewport"` та атрибут `content` зі значенням `"width=device-width, initial-scale=1.0"`, щоб сторінка мала однаковий вигляд на всіх пристроях.

# --hints--

Ви повинні мати два елементи `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

Новий елемент `meta` має бути порожнім елементом, він не повинен мати кінцевий тег `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```


Новий елемент `meta` повинен мати атрибут `name` зі значенням `"viewport"` та атрибут `content` зі значенням `"width=device-width, initial-scale=1.0"`.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

Your two `meta` elements should be inside the `head` element.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
