---
id: 616d4a84b756d9c4b8255093
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

It's time to add some color to the marker. Remember that one way to add color to an element is to use a <dfn>color keyword</dfn> like `black`, `cyan`, or `yellow`.

Нагадуємо, як націлити клас `freecodecamp`:

```css
.freecodecamp {

}
```

Створіть нове CSS-правило, націлене на клас `marker`, та встановіть значення його властивості `background-color` на `red`.

**Примітка:** після додавання CSS ви не побачите жодних змін.

# --hints--

Ви повинні створити селектор класу, щоб націлити клас `marker`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Ваше CSS-правило `.marker` повинне мати властивість `background-color` зі значенням `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
