---
id: 618a0927005553b74bfa05ae
title: Крок 36
challengeType: 0
dashedName: step-36
---

# --description--

A color wheel is a circle where similar colors, or <dfn>hues</dfn>, are near each other, and different ones are further apart. For example, pure red is between the hues rose and orange.

Два кольори, протилежні один одному на колірному колі, називають <dfn>доповнюючими кольорами</dfn>. Якщо поєднати два доповнюючі кольори, вони створять сірий. Але коли вони розташовані поруч, то створюють сильний візуальний контраст та виглядають яскравішими.

В функції `rgb` для CSS-правила `.one` встановіть максимальне значення червоного на `255`, щоб отримати чистий червоний. В функції `rgb` для CSS-правила `.two` встановіть максимальне значення для зеленого та синього на `255`, щоб отримати блакитний.

# --hints--

Ваше CSS-правило `.one` повинне мати властивість `background-color` зі значенням `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

Ваше CSS-правило `.two` повинне мати властивість `background-color` зі значенням `rgb(0, 255, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
