---
id: 67d83df6f82eda3868dd2a84
title: Construir un Personaje de RPG
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

En este laboratorio, practicarás los fundamentos de Python construyendo una pequeña aplicación que crea un personaje para una aventura RPG.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función llamada `create_character`.
1. La función debe aceptar, en orden, un nombre de personaje seguido por tres estadísticas: fuerza, inteligencia y carisma.
1. The character name should be validated:
   - Si el nombre del personaje no es una cadena, la función debe devolver `The character name should be a string`.
   - Si el nombre del personaje tiene más de 10 caracteres, la función debe devolver `The character name is too long`.
   - Si el nombre del personaje contiene espacios, la función debe devolver `The character name should not contain spaces`.
1. The stats should also be validated:
   - Si una o más estadísticas no son enteros, la función debe devolver `All stats should be integers`.
   - Si una o más estadísticas son menores que 1, la función debe devolver `All stats should be no less than 1`.
   - Si una o más estadísticas son mayores que 4, la función debe devolver `All stats should be no more than 4`.
   - Si la suma de todas las estadísticas es diferente de 7, la función debe devolver `The character should start with 7 points`.
1. If all values pass the verification, the function should return a string with four lines:
   - la primera línea debe contener el nombre del personaje
   - las líneas 2-4 deben comenzar con la abreviatura de la estadística, `STR`, `INT` o `CHA` (en ese orden), luego un espacio, y luego un número de puntos llenos (`●`) igual al valor de la estadística, y un número de puntos vacíos (`○`) para llegar a 10. Ejemplo: si el valor de fuerza es 3 debe haber 3 puntos llenos seguidos de 7 puntos vacíos. Los puntos se dan en el editor.

Aquí está la cadena que `create_character("ren", 4, 2, 1)` debería devolver:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

NOTA: aunque `str` e `int` son abreviaturas comunes para las estadísticas, recuerda que son palabras clave reservadas en Python y no deben usarse como nombres de variables.

# --hints--

Debes tener una función llamada `create_character`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

Cuando `create_character` se llama con un primer argumento que no es una cadena, debería devolver `The character name should be a string`.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {"a": 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'

# even if wrong numbers, wrong type has precedence

assert create_character(1, 1, 1, 1) == 'The character name should be a string'
`
    )
})
```

Cuando `create_character` se llama con un primer argumento que tiene más de 10 caracteres, debería devolver `The character name is too long`.

```js
({
    test: () => runPython(
        `
assert create_character("aaabbbcccdd", 4, 2, 1) == 'The character name is too long'
`
    )
})
```

Cuando `create_character` se llama con un primer argumento que contiene un espacio debería devolver `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character("cha cha", 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

Cuando `create_character` se llama con un segundo, tercer o cuarto argumento que no es un entero, debería devolver `All stats should be integers`.

```js
({
    test: () => runPython(
        `
wrong_types = ["3", "5", {"a": 3}, 3.1, [1, 2], (1, 2), {1, 2}, "str", "friend"]
for wrong_type in wrong_types:
    assert create_character("friend", wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character("friend", 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character("friend", 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

Cuando `create_character` se llama con un segundo, tercer o cuarto argumento que sea menor que `1`, debería devolver `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
expected = "All stats should be no less than 1"

assert create_character("ren", 0, 4, 3) == expected
assert create_character("ren", 4, 0, 3) == expected
assert create_character("ren", 4, 3, 0) == expected
assert create_character("ren", -1, 4, 4) == expected
assert create_character("ren", 4, -1, 4) == expected
assert create_character("ren", 4, 4, -1) == expected
`
    )
})
```

Cuando `create_character` se llama con un segundo, tercer o cuarto argumento que sea mayor que `4`, debería devolver `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character("ren", 1, 5, 1) == 'All stats should be no more than 4'
assert create_character("ren", 5, 1, 1) == 'All stats should be no more than 4'
assert create_character("ren", 1, 1, 5) == 'All stats should be no more than 4'
assert create_character("ren", 1, 1, 7) == 'All stats should be no more than 4'
assert create_character("ren", 1, 7, 1) == 'All stats should be no more than 4'
assert create_character("ren", 7, 1, 1) == 'All stats should be no more than 4'
`
    )
})
```

Cuando `create_character` se llama con un segundo, tercer o cuarto argumento que no suma `7`, debería devolver `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character("ren", 4, 4, 4) == 'The character should start with 7 points'
assert create_character("ren", 1, 1, 1) == 'The character should start with 7 points'
assert create_character("ren", 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

`create_character("ren", 4, 2, 1)` debería devolver `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`.

```js
({
    test: () => runPython(
        `
assert create_character("ren", 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

Cuando `create_character` se llama con valores válidos debería mostrar las estadísticas del personaje como se requiere.

```js
({
    test: () => runPython(
        `
assert create_character("Bob", 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character("aaabbbcccd", 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return "The character name should be a string"
    if len(character_name) > 10:
        return "The character name is too long"
    if ' ' in character_name:
        return "The character name should not contain spaces"

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return "All stats should be integers"
    if strength < 1 or intelligence < 1 or charisma < 1:
        return "All stats should be no less than 1"
    if strength > 4 or intelligence > 4 or charisma > 4:
        return "All stats should be no more than 4"
    if strength + intelligence + charisma != 7:
        return "The character should start with 7 points"

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
