---
id: 690ddb5e6ee94eb73ed172b0
title: 调试捐赠表单
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

一个局部的慈善机构建立了一个捐赠表单网站，但有几个问题需要修复。该表单无法访问，并且存在一些超文本标记语言语法错误。

你的工作是修复所有误差，使表单能够正确工作并且对所有用户可访问。完成下面用户故事中的各项，然后点击 `Run the Tests` 查看你是否修复了所有误差。

**用户故事：**

1. `input` 元素是空元素，不应有闭合标签。请从形式中删除所有 `</input>` 闭合标签。
2. 为每个表单输入字段添加 `label` 元素，以便用户知道每个字段的用途。标签文本应匹配当前每个输入旁边的内容。
3. `Email Address:` 输入类型应为 `email` 而不是 `text`。
4. 你应该使用 `label` 上的 `for` 属性和 `input` 上匹配的 `id` 属性，将每个 `label` 元素与其对应的 `input` 元素关联起来。
5. 为文本、电子邮件和数字输入字段添加 `required` 属性（但不包括复选框或提交按钮），以确保用户填写必填信息。

# --hints--

你的代码中不应有任何 `</input>` 关闭标签。

```js
assert.notMatch(code, /<\/input>/);
```

你的形式中应该有正好五个 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

你的 `form` 中应有正好四个 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

你的第一个 `label` 应该有文本 `Full Name:`。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

你的第一个 `input` 应该有一个 `required` 属性。

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

你的第二个 `label` 应该有文本 `Email Address:`。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

你的第二个 `input` 应该有一个 `required` 属性。

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

你的第三个 `label` 应该有文本 `Donation Amount ($):`。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

你的第三个 `input` 应该有一个 `required` 属性。

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

你的第四个 `label` 应该有文本 `Subscribe`。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

你应该使用 `for` 和 `id` 属性将每个 `input` 元素与 `label` 元素关联起来。

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

你的 `input` 用于电子邮件 `address` 应该是 `email` 类型。

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

你的复选框 `input` 不应有 `required` 属性。

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

你的提交 `button` 不应有 `required` 属性。

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
