---
id: acda2fb1324d9b0fa741e6b5
title: Construye una herramienta para confirmar el final
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

En este laboratorio, implementarás una función que verifica si una cadena termina con la cadena objetivo dada.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una función llamada `confirmEnding` que tome dos parámetros: la cadena a verificar y la cadena con la que se verifica.
2. La función debe devolver `true` si la primera cadena termina con la segunda cadena, y `false` en caso contrario.
3. No debes usar el método `.endsWith()`; en su lugar, usa uno de los métodos de subcadena de JavaScript para lograr esto.

# --hints--

Debes crear una función llamada `confirmEnding`.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` debería tomar 2 parámetros.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` debe devolver `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` debe devolver `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` debe devolver `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` debe devolver `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` debe devolver `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` debe devolver `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` debe devolver `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` debe devolver `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` debería devolver `false`.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` debe devolver `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Tu código no debe usar el método incorporado `.endsWith()` para resolver el laboratorio.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
