---
id: 619665c9abd72906f3ad30f9
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Vas a construir un feliz Pingüino Flappy, y además practicarás las transformaciones y animaciones CSS en el proceso.

Comienza vinculando tu hoja de estilo a la página.

# --hints--

Tu código debe tener un elemento `link`.

```js
assert.exists(document.querySelector('link'));
```

Tu elemento `link` debe estar dentro de tu elemento `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Tu elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
