---
id: 587d7b89367417b2b2512b4b
title: Tumia mgawo wa ufumbuzi wa muundo kuweka thamani kwa vigezo kutoka katika safu za data
challengeType: 1
forumTopicId: 301213
dashedName: use-destructuring-assignment-to-assign-variables-from-arrays
---

# --description--

ES6 hufanya ufumbuzi wa muundo wa safu za data kuwa rahisi kama ufumbuzi wa muundo wa vitu.

Tofauti kuu kati ya opereta ya spread na ufumbuzi wa muundo wa safu ya data ni kwamba opereta ya spread hufungua maudhui yote ya safu ya data kuwa orodha iliyotenganishwa kwa koma. Kwa hiyo, huwezi kuchagua vipengele unavyotaka kuweka thamani kwa vigezo.

Ufumbuzi wa muundo wa safu ya data huturuhusu kufanya hasa hivyo:

```js
const [a, b] = [1, 2, 3, 4, 5, 6];
console.log(a, b);
```

Konsoli itaonyesha thamani za `a` na `b` kama `1, 2`.

Kigezo `a` kimepewa thamani ya kwanza ya safu ya data, na `b` kimepewa thamani ya pili ya safu ya data. Pia tunaweza kupata thamani katika kielezo chochote cha safu ya data kwa ufumbuzi wa muundo kwa kutumia koma kufikia kielezo kinachotakiwa:

```js
const [a, b,,, c] = [1, 2, 3, 4, 5, 6];
console.log(a, b, c);
```

Konsoli itaonyesha thamani za `a`, `b`, na `c` kama `1, 2, 5`.

# --instructions--

Tumia mgawo wa ufumbuzi wa muundo kubadilisha thamani za `a` na `b` ili `a` ipokee thamani iliyohifadhiwa katika `b`, na `b` ipokee thamani iliyohifadhiwa katika `a`.

# --hints--

Thamani ya `a` inapaswa kuwa `6`, baada ya kubadilisha.

```js
assert(a === 6);
```

Thamani ya `b` inapaswa kuwa `8`, baada ya kubadilisha.

```js
assert(b === 8);
```

Unapaswa kutumia ufumbuzi wa muundo wa safu ya data kubadilisha `a` na `b`.

```js
assert(/\[\s*(\w)\s*,\s*(\w)\s*\]\s*=\s*\[\s*\2\s*,\s*\1\s*\]/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let a = 8, b = 6;
// Only change code below this line
```

# --solutions--

```js
let a = 8, b = 6;
[a, b] = [b, a];
```
