---
id: 587d7b8c367417b2b2512b57
title: Tumia * kuleta kila kitu kutoka kwenye faili
challengeType: 1
forumTopicId: 301210
dashedName: use--to-import-everything-from-a-file
---

# --description--

Fikiria una faili na unataka kuleta kila yaliyomo ndani yake katika faili ya sasa. Hii inaweza kufanywa kwa kutumia sintaksia ya `import * as`. Hapa kuna mfano ambapo yaliyomo ya faili liitwalo `math_functions.js` yameletwa katika faili moja katika saraka ile ile:

```js
import * as myMathModule from "./math_functions.js";
```

Sentensi ya `import` hapo juu itaunda kitu kinachoitwa `myMathModule`. Hii ni jina la kigezo tu, unaweza kulipa jina lolote. Kitu hicho kitakuwa na kila export kutoka `math_functions.js` ndani yake, hivyo unaweza kupata upatikanaji wa kitendakazi kama unavyofanya kwa sifa nyingine yoyote ya kitu. Hapa ni jinsi unavyoweza kutumia kitendakazi cha `add` na `subtract` ambacho kimeletwa:

```js
myMathModule.add(2,3);
myMathModule.subtract(5,3);
```

# --instructions--

Msimbo katika faili hii unahitaji yaliyomo ya faili: `string_functions.js`, ambayo iko katika saraka ile ile na faili ya sasa. Tumia sintaksia ya `import * as` kuleta kila kitu kutoka kwenye faili ndani ya kitu kinachoitwa `stringFunctions`.

# --hints--

Msimbo wako unapaswa kutumia sintaksia ya `import * as` ipasavyo.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*\*\s*as\s+stringFunctions\s+from\s*('|")\.\/string_functions\.js\1/g
  )
);
```

# --seed--

## --seed-contents--

```js

// Only change code above this line

stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```

# --solutions--

```js
import * as stringFunctions from "./string_functions.js";

// add code above this line
stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```
