---
id: 587d7b88367417b2b2512b46
title: Weka vigezo vya chaguo-msingi kwa vitendakazi vyako
challengeType: 1
forumTopicId: 301209
dashedName: set-default-parameters-for-your-functions
---

# --description--

Ili kutusaidia kuunda vitendakazi vyenye kubadilika zaidi, ES6 inaleta <dfn>vigezo vya chaguo-msingi</dfn> kwa vitendakazi.

Tazama msimbo huu:

```js
const greeting = (name = "Anonymous") => "Hello " + name;

console.log(greeting("John"));
console.log(greeting());
```

Konsoli itaonyesha mfuatano wa herufi `Hello John` na `Hello Anonymous`.

Kigezo cha chaguo-msingi kinaanza kufanya kazi wakati hoja haijatajwa (ni undefined). Kama unavyoona katika mfano hapo juu, kigezo `name` kitapokea thamani yake ya chaguo-msingi `Anonymous` wakati hutoi thamani kwa kigezo hicho. Unaweza kuongeza thamani za chaguo-msingi kwa vigezo vingi kadri unavyotaka.

# --instructions--

Badilisha kitendakazi `increment` kwa kuongeza vigezo vya chaguo-msingi ili kiweze kuongeza 1 kwa `number` ikiwa `value` haitatajwa.

# --hints--

Matokeo ya `increment(5, 2)` yanapaswa kuwa `7`.

```js
assert(increment(5, 2) === 7);
```

Matokeo ya `increment(5)` yanapaswa kuwa `6`.

```js
assert(increment(5) === 6);
```

Thamani ya kigezo cha chaguo-msingi ya `1` inapaswa kutumika kwa `value`.

```js
assert(__helpers.removeJSComments(code).match(/value\s*=\s*1/g));
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const increment = (number, value) => number + value;
// Only change code above this line
```

# --solutions--

```js
const increment = (number, value = 1) => number + value;
```
