---
id: 587d7b8a367417b2b2512b4e
title: Tengeneza mfuatano wa herufi kwa kutumia herufi za kiolezo
challengeType: 1
forumTopicId: 301200
dashedName: create-strings-using-template-literals
---

# --description--

Kipengele kipya cha ES6 ni <dfn>herufi za kiolezo</dfn>. Hii ni aina maalum ya mfuatano wa herufi inayorahisisha kutengeneza mfuatano wa herufi tata.

Herufi za kiolezo zinakuwezesha kutengeneza mfuatano wa herufi wenye mistari mingi na kutumia vipengele vya uingizaji wa thamani kuunda mfuatano wa herufi.

Tafakari msimbo ulio hapa chini:

```js
const person = {
  name: "Zodiac Hasbro",
  age: 56
};

const greeting = `Hello, my name is ${person.name}!
I am ${person.age} years old.`;

console.log(greeting);
```

Konsoli itaonyesha mfuatano wa herufi `Hello, my name is Zodiac Hasbro!` na `I am 56 years old.`.

Mambo mengi yalitokea hapo. Kwanza, mfano huu unatumia alama za kupigia (`` ` ``), not quotes (`'` or `"`), to wrap the string. Secondly, notice that the string is multi-line, both in the code and the output. This saves inserting `\n` within strings. The `${variable}` syntax used above is a placeholder. Basically, you won't have to use concatenation with the `+` operator anymore. To add variables to strings, you just drop the variable in a template string and wrap it with `${` and `}`. Similarly, you can include other expressions in your string literal, for example `${a + b}`). Njia hii mpya ya kutengeneza mfuatano wa herufi inakupa uhuru zaidi wa kutengeneza mfuatano imara.

# --instructions--

Tumia sintaksia ya herufi za kiolezo na alama za kupigia kutengeneza safu ya data ya mfuatano wa herufi wa kipengele cha orodha (`li`). Kila kipengele cha orodha kiwe na maandishi mojawapo ya vipengele vya safu ya data kutoka kwenye sifa `failure` kwenye kitu `result` na kiwe na sifa ya kipengele `class` yenye thamani `text-warning`. Kitendakazi `makeList` kinapaswa kurudisha safu ya data ya mfuatano wa herufi wa kipengele cha orodha.

Tumia njia ya kurudia (aina yoyote ya mzunguko) kupata matokeo yanayotarajiwa (yaliyoonyeshwa hapa chini).

```js
[
  '<li class="text-warning">no-var</li>',
  '<li class="text-warning">var-on-top</li>',
  '<li class="text-warning">linebreak</li>'
]
```

# --hints--

`failuresList` inapaswa kuwa safu ya data yenye ujumbe `result failure`.

```js
assert(
  typeof makeList(result.failure) === 'object' && failuresList.length === 3
);
```

`failuresList` inapaswa kuwa sawa na matokeo yaliyotajwa.

```js
assert(
  makeList(result.failure).every(
    (v, i) =>
      v === `<li class="text-warning">${result.failure[i]}</li>` ||
      v === `<li class='text-warning'>${result.failure[i]}</li>`
  )
);
```

Tumia mfuatano wa herufi za kiolezo na uingizaji wa thamani.

```js
assert.match(code, /(`.*\${.*}.*`)/);
```

Tumia njia ya kurudia.

```js
assert(__helpers.removeJSComments(code).match(/for|map|reduce|forEach|while/));
```

# --seed--

## --seed-contents--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  // Only change code below this line
  const failureItems = [];
  // Only change code above this line

  return failureItems;
}

const failuresList = makeList(result.failure);
```

# --solutions--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  return arr.map(val => `<li class="text-warning">${val}</li>`);
}

const failuresList = makeList(result.failure);
```
