---
id: 587d7b8c367417b2b2512b56
title: Tumia export kushiriki kipande cha msimbo
challengeType: 1
forumTopicId: 301219
dashedName: use-export-to-share-a-code-block
---

# --description--

Fikiria faili liitwalo `math_functions.js` ambalo lina kitendakazi kadhaa kinachohusiana na shughuli za kihesabu. Kimoja kati yao kimehifadhiwa katika kigezo, `add`, kinachopokea nambari mbili na kurudisha jumla yao. Unataka kutumia kitendakazi hiki katika mafaili tofauti ya JavaScript. Ili kushiriki na mafaili haya mengine, kwanza unahitaji `export`.

```js
export const add = (x, y) => {
  return x + y;
}
```

Hii ni njia ya kawaida ya ku-export kitendakazi kimoja, lakini unaweza kufanikisha jambo lile lile hivi:

```js
const add = (x, y) => {
  return x + y;
}

export { add };
```

Unapofanya export ya kigezo au kitendakazi, unaweza kuimport katika faili nyingine na kukitumia bila kuandika msimbo tena. Unaweza ku-export vitu vingi kwa kurudia mfano wa kwanza kwa kila kitu unachotaka ku-export, au kwa kuviweka vyote katika tamko la export la mfano wa pili, kama hivi:

```js
export { add, subtract };
```

# --instructions--

Kuna kitendakazi viwili vinavyohusiana na mfuatano wa herufi katika mhariri. Export zote mbili kwa kutumia njia unayopendelea.

# --hints--

Unapaswa ku-export `uppercaseString` ipasavyo.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+uppercaseString|export\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)})/g
  )
);
```

Unapaswa ku-export `lowercaseString` ipasavyo.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+lowercaseString|export\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)})/g
  )
);
```

# --seed--

## --seed-contents--

```js
const uppercaseString = (string) => {
  return string.toUpperCase();
}

const lowercaseString = (string) => {
  return string.toLowerCase()
}
```

# --solutions--

```js
export const uppercaseString = (string) => {
  return string.toUpperCase();
}

export const lowercaseString = (string) => {
  return string.toLowerCase()
}
```
