---
id: 587d7b8b367417b2b2512b50
title: Andika kitendakazi cha tamko fupi kwa ES6
challengeType: 1
forumTopicId: 301224
dashedName: write-concise-declarative-functions-with-es6
---

# --description--

Wakati wa kufafanua vitendakazi ndani ya vitu katika ES5, tunapaswa kutumia neno kuu `function` kama ifuatavyo:

```js
const person = {
  name: "Taylor",
  sayHello: function() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

Kwa ES6, unaweza kuondoa kabisa neno kuu `function` na alama ya nukta wawili wakati wa kufafanua vitendakazi ndani ya vitu. Hapa kuna mfano wa sintaksia hii:

```js
const person = {
  name: "Taylor",
  sayHello() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

# --instructions--

Badilisha kitendakazi `setGear` ndani ya kitu `bicycle` ili kitumie sintaksia fupi iliyotajwa hapo juu.

# --hints--

Usemi wa kitendakazi wa kawaida hautumiki.

```js
assert(!__helpers.removeJSComments(code).match(/function/));
```

`setGear` inapaswa kuwa kitendakazi cha tamko.

```js
assert(
  typeof bicycle.setGear === 'function' && __helpers.removeJSComments(code).match(/setGear\s*\(.+\)\s*\{/)
);
```

`bicycle.setGear(48)` inapaswa kubadilisha thamani ya `gear` kuwa 48.

```js
bicycle.setGear(48);
assert(bicycle.gear === 48);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const bicycle = {
  gear: 2,
  setGear: function(newGear) {
    this.gear = newGear;
  }
};
// Only change code above this line
bicycle.setGear(3);
console.log(bicycle.gear);
```

# --solutions--

```js
const bicycle = {
  gear: 2,
  // setGear: function(newGear) {
  setGear(newGear) {
    this.gear = newGear;
  }
};
bicycle.setGear(3);
```
