---
id: 587d7b8c367417b2b2512b58
title: Unda Njia Mbadala ya Export kwa kutumia export default
challengeType: 1
forumTopicId: 301199
dashedName: create-an-export-fallback-with-export-default
---

# --description--

Katika funzo la `export`, ulijifunza kuhusu sintaksia inayojulikana kama <dfn>export yenye jina</dfn>. Hii ilikuruhusu kufanya vitendakazi na vigezo vingi viweze kutumika katika mafaili mengine.

Kuna sintaksia nyingine ya `export` unayohitaji kujua, inayojulikana kama <dfn>export default</dfn>. Kawaida utatumia sintaksia hii ikiwa thamani moja tu inatolewa kutoka kwenye faili. Pia hutumika kuunda thamani ya njia mbadala kwa faili au moduli.

Hapo chini ni mifano inayotumia `export default`:

```js
export default function add(x, y) {
  return x + y;
}

export default function(x, y) {
  return x + y;
}
```

Ya kwanza ni kitendakazi chenye jina, na ya pili ni kitendakazi kisicho na jina.

Kwa kuwa `export default` hutumika kutangaza thamani ya njia mbadala kwa moduli au faili, unaweza kuwa na thamani moja tu kama export default katika kila moduli au faili. Zaidi ya hayo, huwezi kutumia `export default` pamoja na `var`, `let`, au `const`

# --instructions--

Kitendakazi kilicho hapa chini kinapaswa kuwa thamani ya njia mbadala kwa moduli. Tafadhali ongeza msimbo unaohitajika kufanya hivyo.

# --hints--

Msimbo wako unapaswa kutumia njia mbadala ya `export`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /export\s+default\s+function(\s+subtract\s*|\s*)\(\s*x,\s*y\s*\)\s*{/g
  )
);
```

# --seed--

## --seed-contents--

```js
function subtract(x, y) {
  return x - y;
}
```

# --solutions--

```js
export default function subtract(x, y) {
  return x - y;
}
```
