---
id: 587d7b8c367417b2b2512b55
title: Tumia tena msimbo wa JavaScript kwa kutumia import
challengeType: 1
forumTopicId: 301208
dashedName: reuse-javascript-code-using-import
---

# --description--

`import` inakuwezesha kuchagua sehemu gani za faili au moduli unazotaka kupakia. Katika funzo lililopita, mifano ilitumia export kupeleka `add` kutoka faili la `math_functions.js`. Hapa ni jinsi unavyoweza kuimport ili kuitumia katika faili nyingine:

```js
import { add } from './math_functions.js';
```

Hapa, `import` itatafuta `add` katika `math_functions.js`, kuleta kitendakazi hicho pekee ili utumie, na kupuuza vingine vyote. `./` inaeleza import kutafuta faili la `math_functions.js` katika folda ile ile kama faili ya sasa. Njia ya faili kwa uhusiano (`./`) na kiambatisho cha faili (`.js`) vinahitajika wakati unapotumia import kwa njia hii.

Unaweza kuleta vitu zaidi ya kimoja kutoka faili kwa kuviweka katika tamko la `import` kama ifuatavyo:

```js
import { add, subtract } from './math_functions.js';
```

# --instructions--

Ongeza tamko sahihi la `import` litakaloruhusu faili ya sasa kutumia kitendakazi cha `uppercaseString` na `lowercaseString` ulizozitumia export katika funzo lililopita. Vitendakazi hivi viko katika faili liitwalo `string_functions.js`, ambalo liko katika saraka ile ile kama faili ya sasa.

# --hints--

Unapaswa kuimport `uppercaseString` kwa usahihi.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

Unapaswa kuimport `lowercaseString` kwa usahihi.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

uppercaseString("hello");
lowercaseString("WORLD!");
```

# --solutions--

```js
import { uppercaseString, lowercaseString } from './string_functions.js';

uppercaseString("hello");
lowercaseString("WORLD!");
```
