---
id: 587d7b88367417b2b2512b44
title: Andika kitendakazi cha mshale chenye vigezo
challengeType: 1
forumTopicId: 301223
dashedName: write-arrow-functions-with-parameters
---

# --description--

Kama kitendakazi cha kawaida, unaweza kupitisha hoja ndani ya kitendakazi cha mshale.

```js
const doubler = (item) => item * 2;
doubler(4);
```

`doubler(4)` itarejesha thamani ya `8`.

Ikiwa kitendakazi cha mshale kina kigezo kimoja tu, mabano ya kawaida yanayozunguka kigezo yanaweza kuachwa.

```js
const doubler = item => item * 2;
```

Inawezekana kupitisha hoja zaidi ya moja ndani ya kitendakazi cha mshale.

```js
const multiplier = (item, multi) => item * multi;
multiplier(4, 2);
```

`multiplier(4, 2)` itarejesha thamani ya `8`.

# --instructions--

Andika upya kitendakazi cha `myConcat` ambacho kinaongeza yaliyomo ya `arr2` kwenye `arr1` ili kitendakazi hicho kitumie sintaksia ya kitendakazi cha mshale.

# --hints--

Unapaswa kubadilisha neno la `var`.

```js
assert.notMatch(code, /var/g);
```

`myConcat` inapaswa kuwa kigezo cha thabiti (kwa kutumia `const`).

```js
assert.match(code, /const\s+myConcat/g);
```

`myConcat` inapaswa kuwa kitendakazi cha mshale chenye vigezo viwili

```js
assert(
  /myConcat=\(\w+,\w+\)=>/.test(__helpers.removeJSComments(code).replace(/\s/g, '')) &&
    typeof myConcat === 'function'
);
```

`myConcat()` inapaswa kurudisha `[1, 2, 3, 4, 5]`.

```js
assert.deepEqual(myConcat([1, 2], [3, 4, 5]), [1, 2, 3, 4, 5]);
```

Neno la `function` halipaswi kutumika.

```js
assert.notMatch(code, /function/g);
```

# --seed--

## --seed-contents--

```js
var myConcat = function(arr1, arr2) {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```

# --solutions--

```js
const myConcat = (arr1, arr2) =>  {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```
