---
id: 587d7b8a367417b2b2512b4f
title: Andika tamko fupi za vitu vya literal kwa kutumia kifupi cha sifa za vitu
challengeType: 1
forumTopicId: 301225
dashedName: write-concise-object-literal-declarations-using-object-property-shorthand
---

# --description--

ES6 inaongeza msaada mzuri wa kuunda vitu vya literal kwa urahisi.

Tafakari msimbo ufuatao:

```js
const getMousePosition = (x, y) => ({
  x: x,
  y: y
});
```

`getMousePosition` ni kitendakazi rahisi kinachorejesha kitu chenye sifa mbili. ES6 hutoa sukari ya sintaksia kuondoa rudufu ya kuandika `x: x`. Unaweza tu kuandika `x` mara moja, na itabadilishwa kuwa `x: x` (au kitu kinachofanana) chini ya hood. Hapa ni kitendakazi kile kile kilichoandikwa tena kwa kutumia sintaksia hii mpya:

```js
const getMousePosition = (x, y) => ({ x, y });
```

# --instructions--

Tumia kifupi cha sifa za vitu na vitu vya literal kuunda na kurudisha kitu chenye sifa za `name`, `age` na `gender`.

# --hints--

`createPerson("Zodiac Hasbro", 56, "male")` inapaswa kurudisha `{name: "Zodiac Hasbro", age: 56, gender: "male"}`.

```js
assert.deepEqual(
  { name: 'Zodiac Hasbro', age: 56, gender: 'male' },
  createPerson('Zodiac Hasbro', 56, 'male')
);
```

Msimbo wako haupaswi kutumia `key:value`.

```js
assert(!__helpers.removeJSComments(code).match(/:/g))
```

# --seed--

## --seed-contents--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  return {
    name: name,
    age: age,
    gender: gender
  };
  // Only change code above this line
};
```

# --solutions--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  /*return {
    name: name,
    age: age,
    gender: gender
  };*/
  return {
    name,
    age,
    gender
  };
  // Only change code above this line
};
```
