---
id: 587d7b8a367417b2b2512b4d
title: Tumia mgawo wa ufumbuzi wa muundo kupitisha kitu kama vigezo vya kitendakazi
challengeType: 1
forumTopicId: 301217
dashedName: use-destructuring-assignment-to-pass-an-object-as-a-functions-parameters
---

# --description--

Katika baadhi ya hali, unaweza kufumbua muundo wa kitu ndani ya hoja ya kitendakazi yenyewe.

Tazama msimbo ulio hapa chini:

```js
const profileUpdate = (profileData) => {
  const { name, age, nationality, location } = profileData;

}
```

Hii inafumbua muundo wa kitu kinachotumwa ndani ya kitendakazi kwa ufanisi. Hii pia inaweza kufanywa mahali hapo hapo:

```js
const profileUpdate = ({ name, age, nationality, location }) => {

}
```

Wakati `profileData` inapopita kwenye kitendakazi kilicho juu, thamani zinatolewa kutoka kwenye kigezo cha kitendakazi kwa matumizi ndani ya kitendakazi hicho.

# --instructions--

Tumia mgawo wa ufumbuzi wa muundo ndani ya hoja ya kitendakazi `half` ili kupitisha tu `max` na `min` ndani ya kitendakazi.

# --hints--

`stats` inapaswa kuwa `object`.

```js
assert(typeof stats === 'object');
```

`half(stats)` inapaswa kuwa `28.015`

```js
assert(half(stats) === 28.015);
```

Tumia ufumbuzi wa muundo.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/half=\({\w+,\w+}\)/));
```

Tumia kigezo kilichofumbuliwa.

```js
assert(!__helpers.removeJSComments(code).match(/stats\.max|stats\.min/));
```

# --seed--

## --seed-contents--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

// Only change code below this line
const half = (stats) => (stats.max + stats.min) / 2.0; 
// Only change code above this line
```

# --solutions--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

const half = ( {max, min} ) => (max + min) / 2.0;
```
