---
id: 587d7b8d367417b2b2512b59
title: Import export chaguo-msingi
challengeType: 1
forumTopicId: 301205
dashedName: import-a-default-export
---

# --description--

Katika changamoto iliyopita, ulijifunza kuhusu `export default` na matumizi yake. Ili ku-import export chaguo-msingi, unahitaji kutumia sintaksia tofauti ya `import`. Katika mfano ufuatao, `add` ni export chaguo-msingi ya faili ya `math_functions.js`. Hapa ni jinsi ya ku-import:

```js
import add from "./math_functions.js";
```

Sintaksia inatofautiana mahali moja muhimu. Thamani inayoinuliwa, `add`, haizungushwi na mabano ya mviringo (`{}`). `add` hapa ni jina la kigezo kwa chochote kinachotolewa kama export chaguo-msingi ya faili ya `math_functions.js`. Unaweza kutumia jina lolote hapa unapo-import export chaguo-msingi.

# --instructions--

Katika msimbo ufuatao, import export chaguo-msingi kutoka faili ya `math_functions.js`, iliyopo katika saraka ile ile kama faili hii. Toa import jina `subtract`.

# --hints--

Unapaswa ku-import `subtract` ipasavyo kutoka `math_functions.js`.

```js
assert(__helpers.removeJSComments(code).match(/import\s+subtract\s+from\s+('|")\.\/math_functions\.js\1/g));
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

subtract(7,4);
```

# --solutions--

```js
import subtract from "./math_functions.js";

subtract(7,4);
```
