---
id: 598f48a36c8c40764b4e52b3
title: Zuia mabadiliko ya kitu
challengeType: 1
forumTopicId: 301207
dashedName: prevent-object-mutation
---

# --description--

Kama ilivyoonekana katika changamoto iliyopita, tamko la `const` pekee halilindi data yako dhidi ya mabadiliko. Ili kuhakikisha data yako haibadiliki, JavaScript hutoa kitendakazi `Object.freeze` ili kuzuia mabadiliko ya data.

Jaribio lolote la kubadilisha kitu litakataliwa, na hitilafu itatupwa ikiwa script inaendesha katika hali ya strict mode.

```js
let obj = {
  name:"FreeCodeCamp",
  review:"Awesome"
};
Object.freeze(obj);
obj.review = "bad";
obj.newProp = "Test";
console.log(obj); 
```

Mgawo wa `obj.review` na `obj.newProp` utasababisha makosa, kwa sababu mhariri wetu unaendesha katika strict mode kwa chaguo-msingi, na konsoli itaonyesha thamani `{ name: "FreeCodeCamp", review: "Awesome" }`.

# --instructions--

Katika changamoto hii utatumia `Object.freeze` kuzuia constants za kihesabu zisibadilike. Unahitaji kufunga kitu cha `MATH_CONSTANTS` ili mtu yeyote asiweze kubadilisha thamani ya `PI`, kuongeza, au kufuta vigezo.

# --hints--

Haupaswi kubadilisha neno la `const`.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`MATH_CONSTANTS` inapaswa kuwa variable ya constant (kwa kutumia `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+MATH_CONSTANTS/g));
```

Haupaswi kubadilisha tamko la awali la `MATH_CONSTANTS`.

```js
assert(__helpers.removeJSComments(code).match(
   /const\s+MATH_CONSTANTS\s+=\s+{\s+PI:\s+3.14\s+};/g
));
```

`PI` inapaswa kuwa sawa na `3.14`.

```js
assert(PI === 3.14);
```

# --seed--

## --seed-contents--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  // Only change code below this line


  // Only change code above this line
  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```

# --solutions--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  Object.freeze(MATH_CONSTANTS);

  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```
