---
id: 5cdafbb0291309899753167f
title: Tengeneza ahadi ya JavaScript
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

Ahadi katika JavaScript ni kama inavyosikika - unaitumia kutoa ahadi ya kufanya jambo fulani, kawaida isiyo ya wakati mmoja. Wakati zoezi linapokamilika, au unatekeleza ahadi yako au unashindwa kufanya hivyo. `Promise` ni kitendakazi cha kijenzi, kwa hivyo unahitaji kutumia neno kuu la `new` kuunda moja. Inachukua kitendakazi, kama hoja yake, yenye vigezo viwili - `resolve` na `reject`. Hizi ni njia zinazotumika kuamua matokeo ya ahadi. Sintaksia inaonekana kama hii:

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

Tengeneza ahadi mpya iitwayo `makeServerRequest`. Pitia kitendakazi chenye vigezo `resolve` na `reject` kwa kijenzi.

# --hints--

Unapaswa weka ahadi kwenye kigezo kilichotangazwa kiitwacho `makeServerRequest`.

```js
assert(makeServerRequest instanceof Promise);
```

Ahadi yako inapaswa kupokea kitendakazi chenye vigezo `resolve` na `reject`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
