---
id: 5cddbfd622f1a59093ec611d
title: Unda script ya moduli
challengeType: 6
forumTopicId: 301198
dashedName: create-a-module-script
---

# --description--

JavaScript ilianza na jukumu dogo katika mtandao uliojaa hasa HTML. Leo, ni kubwa sana, na baadhi ya tovuti zinajengwa karibu kabisa kwa kutumia JavaScript pekee. Ili kufanya JavaScript iwe moduli zaidi, safi, na rahisi kudumishwa; ES6 ilianzisha njia ya kushiriki msimbo kwa urahisi kati ya mafaili ya JavaScript. Hii inahusisha kusafirisha sehemu za faili kwa matumizi katika faili moja au zaidi, na kuleta sehemu unazohitaji, mahali unazohitaji. Ili kufaidika na kipengele hiki, unahitaji kuunda script katika hati yako ya HTML yenye `type` ya `module`. Hapa kuna mfano:

```html
<script type="module" src="filename.js"></script>
```

Script inayotumia aina hii ya `module` sasa inaweza kutumia vipengele vya `import` na `export` utakavyojifunza katika changamoto zijazo.

# --instructions--

Ongeza script katika hati ya HTML yenye aina ya `module` na uipe faili chanzo la `index.js`

# --hints--

Unapaswa kuunda lebo ya `script`.

```js
assert(__helpers.removeJSComments(code).match(/<\s*script[^>]*>\s*<\/\s*script\s*>/g));
```

Leboa yako ya `script` inapaswa kuwa na sifa ya `type` yenye thamani ya `module`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^t]*type\s*=\s*('|")module\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

Leboa yako ya `script` inapaswa kuwa na `src` ya `index.js`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^s]*src\s*=\s*('|")index\.js\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <!-- Only change code below this line -->

    <!-- Only change code above this line -->
  </body>
</html>
```

# --solutions--

```html
<html>
  <body>
    <script type="module" src="index.js"></script>
  </body>
</html>
```
