---
id: 69461e4f98594012939bc1c6
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、簡単な成績表プリンターを作成しながらPythonのデータ型を練習します。

前のレッスンで学んだように、変数は次のように代入されます。

```py
greeting = 'Hello World'
```

上の例では、変数`greeting`に文字列の値が代入されています。文字列は引用符で囲まれた文字のシーケンスです。

`name`という変数を作成し、文字列`Alice`を代入してください。値は例のようにシングルクォートか二重引用符で囲むことを忘れないでください。

**重要:** Pythonではコードを整理するためにインデント（行の先頭のスペース）を使います。このワークショップでは、コードの各行の先頭に余分なスペースがないことを必ず確認してください。余分なスペースがあると`IndentationError`が発生し、コードが実行されなくなります。

# --hints--

`name`という変数があるはずです。

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

変数`name`には値`'Alice'`が代入されている必要があります。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
