---
id: 694648acde178bb8202d9515
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

次に、生徒のスコアを追加しましょう。

`score` という名前の変数を宣言し、値 `80.5` を代入してください。

`age` と `score` はどちらも数値ですが、同じ種類とは限りません。Python にはこれを確認するための `isinstance()` という関数があります。

```py
x = 10
print(isinstance(x, int)) # Output: True
```

`isinstance()` を使って `score` が `int` かどうかを確認し、上の例のように結果をターミナルに出力してください。

# --hints--

`score` という変数を用意してください。

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

変数 `score` には値 `80.5` を代入してください。値は引用符で囲まないでください。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

`isinstance(score, int)` の結果を出力してください。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
