---
id: 694648acde178bb8202d9517
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

出力は `True` で、`score` が `float` であることを確認しています。

`score` の値とそのデータ型を1つの `print()` 文で出力して、成績表を完成させてください。

# --hints--

`score` と `type(score)` はカンマ区切りで出力してください。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
