---
id: 694648acde178bb8202d9512
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

成績表には、生徒が現在在籍しているかどうかも表示する必要があります。これはブール値を使って表すことができます。

ブール値は「はい」か「いいえ」の状態を表し、コード内で判断を行う際によく使われます。ブール値は2つだけで、`True` と `False` です。

`is_student` という名前の変数を宣言し、値として `True` を代入してください。

# --hints--

`is_student` という変数があるはずです。

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

変数 `is_student` は値 `True` をストアする必要があります。値を引用符で囲まないでください。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
