---
id: 694648abde178bb8202d9511
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

ターミナルに学生の名前が出力されているはずです。

Pythonには、値の型を確認するために使える`type()`という関数があります。

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

上の例では、出力の`<class 'str'>`は、`type()`関数に渡された変数が文字列であることを意味しています。

`type()`関数を`name`を引数にして呼び出し、その出力を例のように出力してください。ターミナルの出力で`name`が`str`（文字列）の型であることを確認しましょう。

# --hints--

`type(name)`を呼び出した結果を出力してください。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
