---
id: 694648acde178bb8202d9514
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

今度は生徒の年齢を成績表に追加する必要があります。そのために、Pythonの数値データ型の一つである整数を使います。

`age` という名前の変数を宣言し、整数値 `20` を代入してください。

次に、`age` の値とデータ型をカンマで区切って一緒に出力してください。ターミナルの出力で、`age` の値と、その型が `int`（整数）であることを確認しましょう。

# --hints--

`age` という変数があるはずです。

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

変数 `age` は値 `20` をストアしている必要があります。値をクォートで囲まないでください。

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

`age` と `type(age)` をカンマ区切りで出力してください。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
