---
id: 694648acde178bb8202d9516
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

出力は `False` で、`score` が `int` ではないことを示しています。

Pythonでよく使われるもう一つの数値の種類は `float` で、小数点を含む数値を表します。既存の `int` 呼び出しで `float` を `isinstance()` に置き換えて、これを確認してください。

# --hints--

既存の `int` 呼び出しで `float` を `isinstance(score, int)` に置き換えてください。

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
