---
id: 694648acde178bb8202d9513
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

生徒の名前は他の詳細と同じ形式にしてください。

`name` 変数の以前の出力を削除してください。それから、前のステップのように `name` と `type(name)` をカンマで区切って1行で出力してください。

# --hints--

コードに `print(name)` を含めてはいけません。

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

コードに `print(type(name))` を含めてはいけません。

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

`name` と `type(name)` を同じ `print()` の呼び出しでカンマ区切りで出力してください。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
