---
id: 6947a0c2b6de035dd8d3eea1
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

`print()` 関数を使って変数の値を出力できます。

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

関数については今後のレッスンで詳しく学びます。ここでは、関数とは再利用可能なコードの塊で、呼び出して実行でき、引数を渡すことができるものだと理解してください。

上の例では、`print(greeting)` が関数の呼び出しで、`greeting` が関数の引数です。

例を参考にして、`name` 変数を出力してください。ターミナルの出力を確認しましょう。

# --hints--

`print()` を `name` を引数として呼び出してください。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
