---
id: 6947a3e516bb8a4a681274a7
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`print()` 関数は一度に複数の値を表示できます。値をカンマ（`,`）で区切ると、同じ行に出力されます。

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

例のように、カンマ `is_student` を使って `type(is_student)` と `,` の両方を同じ行に出力してください。その後、`is_student` の値と、その型が `bool`（ブール値）であることをターミナルの出力で確認してください。

# --hints--

`is_student` と `type(is_student)` をカンマ区切りで出力してください。

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
