---
id: 68e76b1ce8216249d04df670
title: Eine Planet-Klasse erstellen
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Klasse namens `Planet` erstellen.
1. Die `Planet`-Klasse sollte eine `__init__`-Methode haben, die:
   - Vier Parameter hat: `self`, `name`, `planet_type` und `star`.
   - Eine `TypeError` mit der Nachricht `name, planet type, and star must be strings` auslöst, wenn eines der übergebenen Argumente kein String ist.
   - Eine `ValueError` mit der Nachricht `name, planet_type, and star must be non-empty strings` auslöst, wenn eines der übergebenen Argumente ein leerer String ist.
   - Die übergebenen Werte den Instanzattributen `name`, `planet_type` und `star` zuweist.
1. Die `Planet`-Klasse sollte eine `orbit`-Methode haben, die einen String im Format `{name} is orbiting around {star}...` zurückgibt.
1. Die `Planet`-Klasse sollte eine `__str__`-Methode haben, die einen String im Format `Planet: {name} | Type: {planet_type} | Star: {star}` zurückgibt.
1. Sie sollten drei Instanzen der `Planet`-Klasse mit den Namen `planet_1`, `planet_2` und `planet_3` erstellen.
1. Sie sollten jedes Planetenobjekt ausgeben, um die Ausgabe der `__str__`-Methode zu sehen.
1. Sie sollten die `orbit`-Methode bei jedem Planetenobjekt aufrufen und das Ergebnis ausgeben.

# --hints--

Sie sollten eine Klasse namens `Planet` erstellen.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

Die `Planet`-Klasse sollte eine `__init__`-Methode haben.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

Die `__init__`-Methode sollte vier Parameter haben: `self`, `name`, `planet_type` und `star` in dieser Reihenfolge.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

Die `__init__`-Methode sollte eine `TypeError` mit der Nachricht `name, planet type, and star must be strings` auslösen, wenn das erste übergebene Argument kein String ist.

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Die `__init__`-Methode sollte eine `TypeError` mit der Nachricht `name, planet type, and star must be strings` auslösen, wenn das zweite übergebene Argument kein String ist.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Die `__init__`-Methode sollte eine `TypeError` mit der Nachricht `name, planet type, and star must be strings` auslösen, wenn das dritte übergebene Argument kein String ist.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Die `__init__`-Methode sollte eine `ValueError` mit der Nachricht `name, planet_type, and star must be non-empty strings` auslösen, wenn das erste Argument ein leerer String ist.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Die `__init__`-Methode sollte eine `ValueError` mit der Nachricht `name, planet_type, and star must be non-empty strings` auslösen, wenn das zweite Argument ein leerer String ist.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Die `__init__`-Methode sollte eine `ValueError` mit der Nachricht `name, planet_type, and star must be non-empty strings` auslösen, wenn das dritte Argument ein leerer String ist.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Sie sollten `name` in der `self.name`-Methode `__init__` zuweisen.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

Sie sollten `planet_type` in der `self.planet_type`-Methode `__init__` zuweisen.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

Sie sollten `star` in der `self.star`-Methode `__init__` zuweisen.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

Die `Planet`-Klasse sollte eine `orbit`-Methode haben.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

Die `orbit`-Methode sollte einen String im Format `{name} is orbiting around {star}...` zurückgeben.

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

Die `Planet`-Klasse sollte eine `__str__`-Methode haben.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

Die `__str__`-Methode sollte einen String im Format `Planet: {name} | Type: {planet_type} | Star: {star}` zurückgeben.

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

Sie sollten drei Instanzen der `Planet`-Klasse mit den Namen `planet_1`, `planet_2` und `planet_3` erstellen.

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

Sie sollten jedes Planetenobjekt ausgeben, um die Ausgabe der `__str__`-Methode zu sehen.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

Sie sollten die `orbit`-Methode bei jedem Planetenobjekt aufrufen und das Ergebnis ausgeben.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
