---
id: 60f027c87bc98f050395c139
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Тепер додайте початкові та кінцеві теги `head` й `body` в межах елементу `html`.

# --hints--

Ви повинні мати початковий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви повинні мати кінцевий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви повинні мати початковий тег `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви повинні мати кінцевий тег `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` та `body` мають бути братськими.

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

Елемент `head` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Елемент `body` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
