---
id: 60f0286404aefb0562a4fdf9
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Додайте елементи `title` та `meta` до елемента `head`. Назвіть свій проєкт `Registration Form` та додайте атрибут `charset` зі значенням `utf-8` до елемента `meta`.

# --hints--

Ваш код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Елемент `title` має бути в межах елемента `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Проєкт повинен мати титул `Registration Form`. Пам’ятайте, що регістр та орфографія мають значення.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Створіть елемент `meta` в межах елемента `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Надайте елементу `meta` атрибут `charset` зі значенням `utf-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Елемент `meta` має бути порожнім елементом, він не має кінцевий тег `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
