---
id: 62cc5b1779e4d313466f73c5
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Вкладіть елемент `link` в елемент `head`. Надайте йому атрибут `rel` зі значенням `stylesheet` та атрибут `href` зі значенням `styles.css`.

# --hints--

Ваш код повинен мати елемент `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Ви не повинні змінювати наявні теґи `head`. Переконайтеся, що не видалили кінцевий теґ.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

Ви повинні мати один елемент `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Елемент `link` має бути порожнім елементом, він не повинен мати кінцевий тег `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Елемент `link` має бути в межах елемента `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Ваш елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
