---
id: 5a858944d96184f06fd60d61
title: 첫 번째 CSS 그리드 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid
---

# --description--

어떤 HTML 요소든 `display` 속성을 `grid`로 설정해 그리드 컨테이너로 만드세요. 이렇게 하면 CSS Grid와 관련된 모든 다른 속성을 사용할 수 있습니다.

**참고:** CSS Grid에서 부모 요소를 <dfn>컨테이너</dfn>라고 하며, 자식 요소들은 <dfn>아이템</dfn>이라고 부릅니다.

# --instructions--

`container` 클래스를 가진 div의 display를 `grid`로 변경하세요.

# --hints--

`container` 클래스는 `display` 속성을 `grid` 값으로 가져야 합니다.

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```
