---
id: 5a9036ee38fddaf9a66b5d34
title: CSS Grid 단위를 사용해 열과 행 크기 변경하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cvE8phd'
forumTopicId: 301134
dashedName: use-css-grid-units-to-change-the-size-of-columns-and-rows
---

# --description--

CSS Grid에서 `px`와 `em` 같은 절대 단위와 상대 단위를 사용해 행과 열의 크기를 정의할 수 있습니다. 다음 단위도 사용할 수 있습니다:

`fr`: 사용 가능한 공간의 분수만큼 열이나 행 크기를 설정합니다,

`auto`: 열이나 행 크기를 내용의 너비나 높이에 맞게 자동으로 설정합니다,

`%`: 열이나 행 크기를 컨테이너의 백분율 너비에 맞게 조정합니다.

아래 코드는 미리보기에서 출력되는 내용을 생성합니다:

```css
grid-template-columns: auto 50px 10% 2fr 1fr;
```

이 코드 조각은 다섯 개의 열을 만듭니다. 첫 번째 열은 내용 크기만큼 넓고, 두 번째 열은 50px, 세 번째 열은 컨테이너 너비의 10%이며, 마지막 두 열은 남은 공간을 세 부분으로 나누어 네 번째 열에 두 부분, 다섯 번째 열에 한 부분을 할당합니다.

# --instructions--

너비가 각각 1fr, 100px, 2fr인 세 개의 열로 구성된 그리드를 만드세요.

# --hints--

`container` 클래스에 너비가 각각 `grid-template-columns`, `1fr`, `100px`인 세 개의 열을 가진 `2fr` 속성이 있어야 합니다.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?1fr\s*?100px\s*?2fr\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: auto 50px 10% 2fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 50px 50px;
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 1fr 100px 2fr;
  }
</style>
```
