---
id: 5a90376038fddaf9a66b5d3c
title: justify-items를 사용해 모든 아이템을 수평으로 정렬하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--

때때로 CSS 그리드의 모든 아이템을 같은 정렬로 맞추고 싶을 때가 있습니다. 이전에 배운 속성들을 사용해 각각 개별적으로 정렬할 수도 있지만, 그리드 컨테이너에서 `justify-items` 속성을 사용해 한 번에 모든 아이템을 수평으로 정렬할 수도 있습니다. 이 속성은 이전 두 도전 과제에서 배운 것과 같은 모든 값을 받을 수 있으며, 차이점은 그리드 내 **모든** 아이템을 원하는 정렬로 이동시킨다는 점입니다.

# --instructions--

이 속성을 사용해 모든 아이템을 수평으로 가운데 정렬하세요.

# --hints--

`container` 클래스에는 `justify-items` 속성이 있고, 그 값은 `center`이어야 합니다.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```
