---
id: 5a9036d038fddaf9a66b5d32
title: grid-template-columns로 열 추가하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c7NzDHv'
forumTopicId: 301117
dashedName: add-columns-with-grid-template-columns
---

# --description--

그리드 요소를 단순히 만드는 것만으로는 충분하지 않습니다. 그리드의 구조도 정의해야 합니다. 그리드에 열을 추가하려면 아래 예시처럼 그리드 컨테이너에 `grid-template-columns` 속성을 사용하세요:

```css
.container {
  display: grid;
  grid-template-columns: 50px 50px;
}
```

이렇게 하면 너비가 각각 50px인 두 개의 열이 그리드에 생깁니다. `grid-template-columns` 속성에 주어진 매개변수 개수는 그리드의 열 개수를 나타내고, 각 매개변수의 값은 각 열의 너비를 나타냅니다.

# --instructions--

그리드 컨테이너에 너비가 각각 `100px`인 세 개의 열을 만드세요.

# --hints--

`container` 클래스에는 `grid-template-columns` 속성이 있어야 하며, 세 개의 `100px` 단위가 포함되어야 합니다.

```js
const templateColumns = new __helpers.CSSHelp(document).getStyle(
  '.container'
)?.gridTemplateColumns;
assert.strictEqual(templateColumns, '100px 100px 100px');
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 100px 100px 100px;
  }
</style>
```
