---
id: 5a9036ee38fddaf9a66b5d37
title: grid-gap으로 간격을 더 빠르게 추가하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ca2qVtv'
forumTopicId: 301118
dashedName: add-gaps-faster-with-grid-gap
---

# --description--

`grid-gap`는 이전 두 도전 과제에서 사용한 `grid-row-gap`과 `grid-column-gap`의 단축 속성으로, 더 편리하게 사용할 수 있습니다. `grid-gap`에 값이 하나만 있으면 모든 행과 열 사이에 간격을 만듭니다. 그러나 값이 두 개 있으면 첫 번째 값은 행 사이 간격에, 두 번째 값은 열 사이 간격에 사용됩니다.

# --instructions--

`grid-gap`를 사용해 행 사이에 `10px` 간격을, 열 사이에 `20px` 간격을 만드세요.

# --hints--

`container` 클래스에는 행 사이에 `grid-gap` 간격을, 열 사이에 `10px` 간격을 만드는 `20px` 속성이 있어야 합니다.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-gap\s*?:\s*?10px\s+?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>
<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-gap: 10px 20px;
  }
</style>
```
