---
id: 5a90374338fddaf9a66b5d3a
title: justify-self로 항목을 수평 정렬하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpKHq'
forumTopicId: 301122
dashedName: align-an-item-horizontally-using-justify-self
---

# --description--

CSS Grid에서 각 항목의 내용은 <dfn>셀</dfn>이라고 하는 박스 안에 위치합니다. 그리드 항목에 `justify-self` 속성을 사용해 셀 내에서 내용의 수평 위치를 정렬할 수 있습니다. 기본적으로 이 속성은 `stretch` 값을 가지며, 이 경우 내용이 셀의 전체 너비를 채우게 됩니다. 이 CSS Grid 속성은 다음과 같은 다른 값도 허용합니다:

`start`: 내용을 셀의 왼쪽에 정렬,

`center`: 내용을 셀의 가운데에 정렬,

`end`: 내용을 셀의 오른쪽에 정렬.

# --instructions--

클래스가 `justify-self`인 항목을 가운데 정렬하려면 `item2` 속성을 사용하세요.

# --hints--

`item2` 클래스에 `justify-self` 속성을 추가하고 그 값을 `center`로 설정하세요.

```js
assert.match(
  code,
  /.item2\s*?{[\s\S]*justify-self\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }

  .item2 {
    background: LightSalmon;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item2 {
    justify-self: center;
  }
</style>
```
