---
id: 5a9036ee38fddaf9a66b5d35
title: grid-column-gap을 사용해 열 간격 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cVZ8vfD'
forumTopicId: 301124
dashedName: create-a-column-gap-using-grid-column-gap
---

# --description--

지금까지 만든 그리드에서는 열들이 모두 서로 딱 붙어 있었습니다. 때로는 열 사이에 간격을 두고 싶을 때가 있습니다. 열 사이에 간격을 추가하려면 `grid-column-gap` 속성을 이렇게 사용하세요:

```css
grid-column-gap: 10px;
```

이렇게 하면 모든 열 사이에 10px의 빈 공간이 생깁니다.

# --instructions--

그리드의 열들 사이에 `20px` 간격을 주세요.

# --hints--

`container` 클래스에는 값이 `grid-column-gap`인 `20px` 속성이 있어야 합니다.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-column-gap\s*?:\s*?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

   
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-column-gap: 20px;
  }
</style>
```
