---
id: 5a9036ee38fddaf9a66b5d36
title: grid-row-gap을 사용해 행 간격 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cPbJ2Cv'
forumTopicId: 301125
dashedName: create-a-row-gap-using-grid-row-gap
---

# --description--

이전 도전 과제에서 열 사이에 간격을 추가한 것처럼 `grid-row-gap`을 사용해 그리드 행 사이에 간격을 추가할 수 있습니다.

# --instructions--

행 사이에 `5px` 높이의 간격을 만드세요.

# --hints--

`container` 클래스에 `grid-row-gap` 속성이 있고 그 값이 `5px`여야 합니다.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-row-gap\s*?:\s*?5px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-row-gap: 5px;
  }
</style>
```
