---
id: 5a90375238fddaf9a66b5d3b
title: align-self를 사용해 아이템을 수직으로 정렬하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cmzd4fz'
forumTopicId: 301123
dashedName: align-an-item-vertically-using-align-self
---

# --description--

아이템을 수평으로 정렬할 수 있듯이, 아이템을 수직으로 정렬하는 방법도 있습니다. 이를 위해 아이템에 `align-self` 속성을 사용하세요. 이 속성은 지난 도전 과제에서 사용한 `justify-self`과 동일한 값을 모두 받을 수 있습니다.

# --instructions--

클래스가 `item3`인 아이템을 `end`에서 수직으로 정렬하세요.

# --hints--

`item3` 클래스에는 값이 `align-self`인 `end` 속성이 있어야 합니다.

```js
assert.match(code, /.item3\s*?{[\s\S]*align-self\s*?:\s*?end\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }

  .item3 {
    background: PaleTurquoise;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    align-self: end;
  }
</style>
```
