---
id: 5a94fe5469fb03452672e461
title: auto-fill을 사용해 유연한 레이아웃 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cmzdycW'
forumTopicId: 301126
dashedName: create-flexible-layouts-using-auto-fill
---

# --description--

repeat 함수에는 <dfn>auto-fill</dfn>이라는 옵션이 있습니다. 이 옵션은 컨테이너 크기에 따라 원하는 크기의 행이나 열을 가능한 한 많이 자동으로 삽입할 수 있게 해줍니다. `auto-fill`와 `minmax`을 결합하면 다음과 같이 유연한 레이아웃을 만들 수 있습니다:

```css
repeat(auto-fill, minmax(60px, 1fr));
```

컨테이너 크기가 변경되면 이 설정은 60px 열을 계속 삽입하고, 또 하나를 삽입할 수 있을 때까지 늘립니다. **참고:** 컨테이너가 모든 항목을 한 행에 맞출 수 없으면 항목을 아래 새 행으로 이동합니다.

# --instructions--

첫 번째 그리드에서 `auto-fill`와 `repeat`을 사용해 최솟값이 `60px`이고 최댓값이 `1fr`인 열로 그리드를 채우세요. 그런 다음 미리보기 크기를 조절해 auto-fill이 작동하는 모습을 확인하세요.

# --hints--

`container` 클래스에는 `grid-template-columns` 속성이 있어야 하며, `repeat`와 `auto-fill`를 사용해 최솟값이 `60px`이고 최댓값이 `1fr`인 열로 그리드를 채울 수 있습니다.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?auto-fill\s*?,\s*?minmax\s*?\(\s*?60px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    grid-template-columns: repeat(3, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>
<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    grid-template-columns: repeat(3, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>
<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```
