---
id: 5a94fe0569fb03452672e45c
title: 그리드를 영역 템플릿으로 나누기
challengeType: 0
forumTopicId: 301130
dashedName: divide-the-grid-into-an-area-template
---

# --description--

그리드의 셀들을 하나의 <dfn>영역</dfn>으로 묶고 그 영역에 사용자 지정 이름을 붙일 수 있습니다. 컨테이너에 `grid-template-areas`를 사용하여 이렇게 하세요:

```css
grid-template-areas:
  'header header header'
  'advert content content'
  'advert footer footer';
```

위 코드는 그리드의 셀들을 네 개의 영역인 `header`, `advert`, `content`, 그리고 `footer`로 묶습니다. 각 단어는 셀 하나를 나타내고, 각 쌍의 따옴표는 한 행을 나타냅니다.

# --instructions--

템플릿을 변경하여 `footer` 영역이 맨 아래 행 전체를 차지하도록 하세요. 영역을 정의하는 것은 지금은 시각적으로 아무런 효과가 없습니다. 나중에 아이템이 영역을 사용하도록 만들어서 어떻게 작동하는지 확인할 것입니다.

# --hints--

`container` 클래스는 예시와 비슷한 `grid-template-areas` 속성을 가져야 하지만 `footer` 영역이 맨 아래 행 전체를 차지하도록 하세요.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.container\s*?{[\s\S]*grid-template-areas\s*?:\s*?["|']\s*?header\s*?header\s*?header\s*?["|']\s*?["|']\s*?advert\s*?content\s*?content\s*?["|']\s*?["|']\s*?footer\s*?footer\s*?footer\s*?["|']\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
    /* Only change code below this line */
      'header header header'
      'advert content content'
      'advert footer footer';
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;

    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```
