---
id: 5a94fe4469fb03452672e460
title: minmax 함수로 아이템 크기 제한하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cD97RTv'
forumTopicId: 301131
dashedName: limit-item-size-using-the-minmax-function
---

# --description--

`grid-template-columns`와 `grid-template-rows`과 함께 사용할 수 있는 또 다른 내장 함수가 `minmax`입니다. 이 함수는 그리드 컨테이너 크기가 변할 때 아이템 크기를 제한하는 데 사용합니다. 이를 위해 아이템의 허용 가능한 크기 범위를 지정해야 합니다. 예시는 다음과 같습니다:

```css
grid-template-columns: 100px minmax(50px, 200px);
```

위 코드에서 `grid-template-columns`는 두 개의 열을 만들도록 설정되어 있습니다. 첫 번째 열은 100px 너비이고, 두 번째 열은 최솟값 50px, 최댓값 200px 너비를 가집니다.

# --instructions--

`minmax` 함수를 사용하여 `1fr` 함수 내 `repeat`를 최솟값이 `90px`이고 최댓값이 `1fr`인 열 크기로 바꾸고, 미리보기 패널 크기를 조절하여 변화를 확인하세요.

# --hints--

`container` 클래스에는 최솟값이 `grid-template-columns`이고 최댓값이 `90px`인 3개의 열을 반복하도록 설정된 `1fr` 속성이 있어야 합니다.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?minmax\s*?\(\s*?90px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, 1fr);

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, minmax(90px, 1fr));
  }
</style>
```
