---
id: 5a94fe7769fb03452672e463
title: 미디어 쿼리를 사용해 반응형 레이아웃 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cMbqeHk'
forumTopicId: 301138
dashedName: use-media-queries-to-create-responsive-layouts
---

# --description--

CSS Grid는 미디어 쿼리를 사용해 그리드 영역을 재배치하고, 그리드 크기를 변경하며, 아이템 배치를 조정해 사이트를 더 반응형으로 만드는 쉬운 방법이 될 수 있습니다.

미리보기에서 뷰포트 너비가 300px 이상일 때 열 수가 1에서 2로 변경됩니다. 광고 영역은 왼쪽 열을 완전히 차지합니다.

# --instructions--

뷰포트 너비가 `400px` 이상일 때 헤더 영역이 상단 행을 완전히 차지하고 푸터 영역이 하단 행을 완전히 차지하도록 만드세요.

# --hints--

뷰포트가 `400px` 이상일 때 `container` 클래스는 헤더와 푸터 영역이 각각 상단과 하단 행을 차지하고 광고와 콘텐츠가 중간 행의 왼쪽과 오른쪽 열을 차지하는 `grid-template-areas` 속성을 가져야 합니다.

```js
assert.match(
  __helpers.removeCssComments(code),
  /@media\s*?\(\s*?min-width\s*?:\s*?400px\s*?\)[\s\S]*.container\s*?{[\s\S]*grid-template-areas\s*?:\s*?["|']\s*?header\s*?header\s*?["|']\s*?["|']\s*?advert\s*?content\s*?["|']\s*?["|']\s*?footer\s*?footer\s*?["|']\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
    grid-area: header;
  }

  .item2 {
    background: LightSalmon;
    grid-area: advert;
  }

  .item3 {
    background: PaleTurquoise;
    grid-area: content;
  }

  .item4 {
    background: lightpink;
    grid-area: footer;
  }

  .container {
    font-size: 1.5em;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr;
    grid-template-rows: 50px auto 1fr auto;
    grid-gap: 10px;
    grid-template-areas:
      'header'
      'advert'
      'content'
      'footer';
  }

  @media (min-width: 300px) {
    .container {
      grid-template-columns: auto 1fr;
      grid-template-rows: auto 1fr auto;
      grid-template-areas:
        'advert header'
        'advert content'
        'advert footer';
    }
  }

  @media (min-width: 400px) {
    .container {
      grid-template-areas:
      /* Only change code below this line */
        'advert header'
        'advert content'
        'advert footer';
      /* Only change code above this line */
    }
  }
</style>

<div class="container">
  <div class="item1">header</div>
  <div class="item2">advert</div>
  <div class="item3">content</div>
  <div class="item4">footer</div>
</div>
```

# --solutions--

```html
<style>
  .item1 {
    background: LightSkyBlue;
    grid-area: header;
  }

  .item2 {
    background: LightSalmon;
    grid-area: advert;
  }

  .item3 {
    background: PaleTurquoise;
    grid-area: content;
  }

  .item4 {
    background: lightpink;
    grid-area: footer;
  }

  .container {
    font-size: 1.5em;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr;
    grid-template-rows: 50px auto 1fr auto;
    grid-gap: 10px;
    grid-template-areas:
      'header'
      'advert'
      'content'
      'footer';
  }

  @media (min-width: 300px) {
    .container {
      grid-template-columns: auto 1fr;
      grid-template-rows: auto 1fr auto;
      grid-template-areas:
        'advert header'
        'advert content'
        'advert footer';
    }
  }

  @media (min-width: 400px) {
    .container {
      grid-template-areas:
        'header header'
        'advert content'
        'footer footer';
    }
  }
</style>

<div class="container">
  <div class="item1">header</div>
  <div class="item2">advert</div>
  <div class="item3">content</div>
  <div class="item4">footer</div>
</div>
```
