---
id: 5a94fe3669fb03452672e45f
title: repeat 함수로 반복 줄이기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--

`grid-template-columns`와 `grid-template-rows`을 사용해 그리드 구조를 정의할 때, 생성한 각 행이나 열에 대해 값을 입력했습니다.

높이가 같은 100개의 행으로 이루어진 그리드를 만든다고 가정해 보세요. 100개의 값을 하나씩 입력하는 것은 실용적이지 않습니다. 다행히 더 좋은 방법이 있습니다. `repeat` 함수를 사용해 열이나 행을 반복할 횟수를 지정하고, 쉼표 뒤에 반복할 값을 적으면 됩니다.

다음 예시는 100개의 행으로 이루어진 그리드를 만들며, 각 행의 높이는 50px입니다.

```css
grid-template-rows: repeat(100, 50px);
```

repeat 함수를 사용해 여러 값을 반복할 수도 있고, 그리드 구조를 정의할 때 다른 값들 사이에 이 함수를 넣을 수도 있습니다. 다음은 그 예시입니다:

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

이는 다음과 같습니다:

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**참고:** `1fr 50px`가 두 번 반복된 뒤 20px가 옵니다.

# --instructions--

`repeat`를 사용해 `grid-template-columns` 속성에서 반복을 제거하세요.

# --hints--

`container` 클래스는 너비가 `grid-template-columns`인 3개의 열을 반복하도록 설정된 `1fr` 속성을 가져야 합니다.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```
