---
id: 5a94fe1369fb03452672e45d
title: grid-area 속성을 사용해 그리드 영역에 아이템 배치하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cRrqmtV'
forumTopicId: 301132
dashedName: place-items-in-grid-areas-using-the-grid-area-property
---

# --description--

이전 도전 과제에서 보여준 것처럼 그리드 컨테이너에 영역 템플릿을 만든 후, 지정한 이름을 참조해 사용자 정의 영역에 아이템을 배치할 수 있습니다. 이렇게 하려면 아이템에 `grid-area` 속성을 다음과 같이 사용하세요:

```css
.item1 {
  grid-area: header;
}
```

이렇게 하면 그리드가 `item1` 클래스를 `header`라는 이름의 영역에 배치하도록 알립니다. 이 경우, 아이템은 전체 상단 행을 차지하는데, 그 행 전체가 `header` 영역으로 이름 붙여져 있기 때문입니다.

# --instructions--

`item5` 속성을 사용해 `footer` 클래스를 `grid-area` 영역에 배치하세요.

# --hints--

`item5` 클래스에 `grid-area` 속성을 추가하고 값으로 `footer`를 지정하세요.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?footer\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: footer;
  }
</style>
```
