---
id: 5a94fe2669fb03452672e45e
title: 영역 템플릿 없이 grid-area 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c6N7VhK'
forumTopicId: 301135
dashedName: use-grid-area-without-creating-an-areas-template
---

# --description--

지난 도전 과제에서 배운 `grid-area` 속성은 다른 방식으로도 사용할 수 있습니다. 만약 그리드에 참조할 영역 템플릿이 없다면, 다음과 같이 항목을 배치할 영역을 즉석에서 만들 수 있습니다:

```css
item1 {
  grid-area: 1/1/2/4;
}
```

위 예제는 앞서 배운 선 번호를 사용해 이 항목의 영역이 어디에 위치할지 정의하는 것입니다. 위 예제의 숫자는 다음 값을 나타냅니다:

```css
grid-area: horizontal line to start at / vertical line to start at / horizontal
  line to end at / vertical line to end at;
```

따라서 예제의 항목은 1번과 2번 행 사이, 그리고 1번과 4번 열 사이를 차지하게 됩니다.

# --instructions--

`grid-area` 속성을 사용해 `item5` 클래스를 가진 요소를 세 번째와 네 번째 가로선 사이, 그리고 첫 번째와 네 번째 세로선 사이에 배치하세요.

# --hints--

`item5` 클래스에는 세 번째와 네 번째 가로선 사이, 그리고 첫 번째와 네 번째 세로선 사이 전체 영역을 채우도록 `grid-area` 속성이 있어야 합니다.

```js
assert.match(
  code,
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?3\s*?\/\s*?1\s*?\/\s*?4\s*?\/\s*?4\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: 3/1/4/4;
  }
</style>
```
