---
id: 5a94fe6269fb03452672e462
title: auto-fit을 사용해 유연한 레이아웃 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c3dPph8'
forumTopicId: 301127
dashedName: create-flexible-layouts-using-auto-fit
---

# --description--

`auto-fit`는 `auto-fill`와 거의 동일하게 작동합니다. 유일한 차이점은 컨테이너 크기가 모든 아이템 크기를 합친 것보다 클 때, `auto-fill`는 빈 행이나 열을 계속 삽입해 아이템을 옆으로 밀어내는 반면, `auto-fit`는 그 빈 행이나 열을 접어 아이템을 컨테이너 크기에 맞게 늘린다는 점입니다.

**참고:** 컨테이너가 모든 아이템을 한 행에 담지 못하면, 아이템을 아래 새 행으로 이동합니다.

# --instructions--

두 번째 그리드에서 `auto-fit`와 `repeat`를 사용해 최소 너비가 `60px`이고 최대 너비가 `1fr`인 열들로 그리드를 채우세요. 그런 다음 미리보기 크기를 조절해 차이를 확인하세요.

# --hints--

`container2` 클래스에는 `grid-template-columns` 속성이 있어야 하며, `repeat`과 `auto-fit`를 사용해 최소 너비가 `60px`이고 최대 너비가 `1fr`인 열들로 그리드를 채울 수 있습니다.

```js
assert.match(
  code,
  /.container2\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?auto-fit\s*?,\s*?minmax\s*?\(\s*?60px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));
  }
  .container2 {
    grid-template-columns: repeat(auto-fit, minmax(60px, 1fr));
  }
</style>
```
