---
id: 5a94fe8569fb03452672e464
title: 그리드 안에 그리드 만들기
challengeType: 0
forumTopicId: 301128
dashedName: create-grids-within-grids
---

# --description--

요소를 그리드로 만들면 직접 자손 요소의 동작에만 영향을 줍니다. 그래서 직접 자손 요소를 그리드로 만들면 그리드 안에 그리드가 생깁니다.

예를 들어, `display` 클래스가 있는 요소의 `grid-template-columns`와 `item3` 속성을 설정하면 그리드 안에 그리드를 만들 수 있습니다.

# --instructions--

`item3` 클래스가 있는 요소를 `auto`와 `1fr`를 사용해 너비가 `display`와 `grid-template-columns`인 두 개의 열을 가진 그리드로 만드세요.

# --hints--

`item3` 클래스는 `grid-template-columns` 속성을 가지고, 값으로 `auto`과 `1fr`을 가져야 합니다.

```js
assert.match(
  code,
  /.item3\s*?{[\s\S]*grid-template-columns\s*?:\s*?auto\s*?1fr\s*?;[\s\S]*}/gi
);
```

`item3` 클래스는 `display` 속성을 가지고, 값으로 `grid`를 가져야 합니다.

```js
assert.match(code, /.item3\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .container {
    font-size: 1.5em;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: auto 1fr;
    grid-template-rows: auto 1fr auto;
    grid-gap: 10px;
    grid-template-areas:
      'advert header'
      'advert content'
      'advert footer';
  }
  .item1 {
    background: LightSkyBlue;
    grid-area: header;
  }

  .item2 {
    background: LightSalmon;
    grid-area: advert;
  }

  .item3 {
    background: PaleTurquoise;
    grid-area: content;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .item4 {
    background: lightpink;
    grid-area: footer;
  }

  .itemOne {
    background: PaleGreen;
  }

  .itemTwo {
    background: BlanchedAlmond;
  }
</style>

<div class="container">
  <div class="item1">header</div>
  <div class="item2">advert</div>
  <div class="item3">
    <div class="itemOne">paragraph1</div>
    <div class="itemTwo">paragraph2</div>
  </div>
  <div class="item4">footer</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    grid-template-columns: auto 1fr;
    display: grid;
  }
</style>
```
