---
id: 699f52220de88c8091032076
title: ステップ 18
challengeType: 1
dashedName: step-18
---

# --description--

コンソールを開くと、`NaN` の結果が表示されます。これは `exponent` の2番目の引数が提供されていなかったためです。この `raiseTo` 関数にはいくつか修正が必要です。

最初の変更は `base ** 2` を返すことです。

次のステップでは、オプションの `exponent` パラメータの扱い方を学びます。

# --hints--

`raiseTo` 関数は数値を返す必要があります。

```js
assert.isNumber(raiseTo(4));
```

`raiseTo` 関数は `base` を `2` 乗した値を返す必要があります。

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
