---
id: 699f3cb8d057f29954205d7d
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

これまでのレッスンで学んだように、関数のパラメータにこのように型注釈を追加できます。

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

この例では、`add` 関数を数値以外の引数で呼び出そうとすると、TypeScript がエラーを表示します。

`square` 関数の `num` パラメータに型注釈を追加してください。型は `number` にしてください。

# --hints--

`num` 関数の `square` パラメータに型注釈を追加してください。型は `number` であるべきです。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
