---
id: 699f412114c87221f16a99e3
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

TypeScriptを使うもう一つの利点は、関数呼び出しに渡す引数が多すぎたり少なすぎたりするとエラーを検出してくれることです。

これを示すために、`square` 関数呼び出しを更新して、`5` と `10` の2つの引数を取るようにしてください。

# --hints--

`square` 関数呼び出しは、`5` と `10` の2つの引数を持つ必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
