---
id: 699f1d72e906e4616155290d
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、型注釈と関数での使い方について引き続き学びながら、型安全な数学ツールキットを作成していきます。

まず、パラメータを受け取り、数値の二乗を返す`square`という関数を作成してください。

# --hints--

`square`という関数があるはずです。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

`square`関数はパラメータを1つ持つべきです。

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

`square`関数は数値の二乗を返すべきです。

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
