---
id: 699f434e294a9e11cf4f65f0
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

次は、ツールキットプロジェクトにもう一つ数学のヘルパーを追加する時間です。

まず、1つのパラメータを受け取る`getAverage`という関数を作成してください。関数は数値の配列の平均を返す必要があります。

次のいくつかのステップで、この関数に型注釈を追加していきます。

# --hints--

`getAverage`という関数があるはずです。

```js
assert.isFunction(getAverage);
```

関数は1つのパラメータを受け取る必要があります。

```js
assert.strictEqual(getAverage.length, 1);
```

関数は数値を返す必要があります。

```js
assert.isNumber(getAverage([1, 2, 3]));
```

関数はどんな数値の配列に対しても正しい結果を返す必要があります。

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
