---
id: 699f4ec939f10da62fb56c3e
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

いよいよ関数をテストする時間です。

`raisedResult` という変数を作成し、`raiseTo` を `2` と `3` の引数で呼び出した結果を代入してください。

その後、`raisedResult` をコンソールに出力してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`raisedResult` という変数があるはずです。

```js
assert.isDefined(raisedResult);
```

`raisedResult` 変数は数値であるべきです。

```js
assert.isNumber(raisedResult);
```

`raisedResult` をコンソールに出力してください。

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
