---
id: 699f4861dd5b5f164e1bda4a
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

TypeScriptを使うもう一つの利点は、戻り値の型が強制され、期待通りであることを保証できる点です。TypeScriptは関数の戻り値の型を推論できる場合もありますが、明示的に指定することもできます。

これを示すために、`getAverage` 関数の戻り値の型を文字列の`"something"`に変更してください。

# --hints--

`getAverage` 関数は文字列を返す必要があります。

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

`getAverage` 関数は文字列の`"something"`を返す必要があります。

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
