---
id: 699f4a3e7b4c6bfa538cc4f7
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

コンソールを開くと、次のエラーメッセージが表示されます。

```md
Type 'string' is not assignable to type 'number'.
```

このエラーを解決するには、戻り値を `sum / numbers.length` に変更する必要があります。

# --hints--

`getAverage` 関数は数値を返す必要があります。

```js
assert.isNumber(getAverage([1, 2, 3]));
```

`getAverage` 関数は配列内の数値の平均を返す必要があります。

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```
