---
id: 699f41b8da20f816b5c7dece
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

コンソールを開くと、次のエラーメッセージが表示されているはずです。

```md
Expected 1 arguments, but got 2.
```

このエラーを解決するには、`square` 関数の呼び出しを `5` の引数1つだけに更新してください。

# --hints--

`square` 関数の呼び出しは、`5` の引数1つだけにしてください。

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
