---
id: 699f3738921fb79a651ded56
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

いよいよ関数を呼び出す時です。

`result` という変数を作成し、`square` を `5` という引数で呼び出した結果を代入してください。

その下に、`result` の値をコンソールに出力してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`result` という変数があるはずです。

```js
assert.isDefined(result);
```

`result` 変数は数値である必要があります。

```js
assert.isNumber(result);
```

`result` 変数は `5` の二乗である必要があります。

```js
assert.strictEqual(result, 25);
```

`result` 変数をコンソールに出力してください。

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
