---
id: 699f52fac083214b6bf51de6
title: ステップ 19
challengeType: 1
dashedName: step-19
---

# --description--

この最後のステップでは、`exponent` 関数の `raiseTo` パラメータに対する条件チェックを追加する必要があります。

`exponent` 引数が提供された場合、関数は `base` を `exponent` 乗した値を返すべきです。`exponent` 引数が提供されていない場合は、`base ** 2` を返すべきです。

この最後のステップで、型安全な数学ツールキットが完成します！

# --hints--

`raiseTo` 関数は、`base ** exponent` 引数が提供されたときに `exponent` を返すべきです。

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

`raiseTo` 関数は、`base ** 2` 引数が提供されていないときに `exponent` を返すべきです。

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
