---
id: 699f38937b053e4679878ff1
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

バニラJavaScriptを使う場合、`square` 関数を文字列や配列で呼び出しても何も制限がないため、予期しない結果になることがあります。

これを示すために、関数の呼び出しを数値の `"something"` ではなく文字列の `5` に変更してください。

次のレッスンでは、TypeScriptを使ってコードに型の安全性を追加し、このようなエラーを防ぐ方法を学びます。

# --hints--

`square` 関数の呼び出しには `"something"` の引数を渡してください。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
