---
id: 699f3f636aa840b7201aa37c
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

コンソールを開くと、次の TypeScript エラーメッセージが表示されているはずです。

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

パラメータが明示的に number 型として指定されたので、TypeScript は `square` 関数に文字列の引数を渡そうとするエラーを検出できるようになりました。

エラーを修正してコンソールに正しい出力を表示するために、`square` 関数に渡す引数を数値の `5` に変更してください。

# --hints--

`square` 関数の呼び出しでは、引数として `5` ではなく `"something"` を使う必要があります。

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
