---
id: 699f5063d2393c674b1f45dc
title: ステップ 16
challengeType: 1
dashedName: step-16
---

# --description--

前のステップで学んだように、関数を間違った数の引数で呼び出すと、TypeScript はエラーをスローします。

パラメータ名の後に `?` を追加すると、そのパラメータをオプションとしてマークできます。これにより、パラメータより少ない引数で関数を呼び出せるようになります。

例を示します。

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

この例では、`greeting` 関数を呼び出す際に `greet` パラメータが指定されなければ、`undefined` となり、関数はデフォルトのあいさつを返します。

`raiseTo` 関数を更新して、`exponent` パラメータをオプションにしてください。

# --hints--

`raiseTo` 関数の `exponent` パラメータはオプションとしてマークされている必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
