---
id: 699f49669ea51a31d2ba89d8
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

これまでのレッスンで学んだように、関数の戻り値の型を明示的に指定することができます。

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

この例では、`circleArea` 関数の戻り値の型は `number` です。異なる型を返そうとすると、TypeScript はエラーを表示します。

`getAverage` 関数の戻り値の型を `number` として明示的に指定してください。

# --hints--

`getAverage` 関数の戻り値の型は `number` である必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
