---
id: 699f458e47fb88c756fdba81
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

現在、`numbers` パラメータは暗黙的に `any` 型になっています。これを修正するには、パラメータを数値の配列として明示的に型指定できます。

前のレッスンで学んだように、数値の配列は次のように型指定できます。

```ts
const numbers: number[] = [1, 2, 3];
```

`getAverage` 関数を更新して、`numbers` パラメータを数値の配列として明示的に型指定してください。

# --hints--

`getAverage` 関数は、数値の配列として明示的に型指定された `numbers` という名前のパラメータを持つ必要があります。

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
