---
id: 699f4b903bf5427a599f791d
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

最後の数ステップでは、関数の型注釈の練習を続けるために、もう1つ関数を作成します。

まず、`raiseTo` という名前の関数を作成してください。この関数は `base` と `exponent` という2つのパラメータを受け取ります。両方のパラメータは `number` 型である必要があります。関数は `base` を `exponent` 乗した結果を返すようにしてください。関数の戻り値の型も `number` として注釈を付けてください。

# --hints--

`raiseTo` 関数があるはずです。

```js
assert.isFunction(raiseTo);
```

`raiseTo` 関数は2つのパラメータ、`base` と `exponent` を持つべきです。

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*=>/
);
```

`base` と `exponent` のパラメータは両方とも `number` 型であるべきです。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

`raiseTo` 関数は数値を返すべきです。

```js
assert.isNumber(raiseTo(2, 3));
```

`raiseTo` 関数は `base` を `exponent` 乗した結果を返すべきです。

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

`raiseTo` 関数は戻り値の型注釈として `number` を持つべきです。

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
