---
id: 699f4639ae0016caaf81b7c6
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

関数をテストする時間です。

まず、`avgResult` という変数を作成し、`getAverage` の関数呼び出しを代入してください。その関数呼び出しの引数は `[2, 14, 26, 8]` にしてください。

次に、`avgResult` の値をコンソールに出力してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`avgResult` という変数を作成しているはずです。

```js
assert.isDefined(avgResult);
```

`getAverage([2, 14, 26, 8])` 変数に `avgResult` を代入してください。

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

`avgResult` 変数は数値である必要があります。

```js
assert.isNumber(avgResult);
```

`avgResult` をコンソールに出力してください。

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
