---
id: 699f513e01eb1768cc5df726
title: ステップ 17
challengeType: 1
dashedName: step-17
---

# --description--

`exponent` パラメータがオプションとしてマークされたので、引数を1つだけ渡して関数を呼び出せます。

`raiseTo` 関数の呼び出しを更新して、数値の `3` のみを渡してください。

# --hints--

`raiseTo` 関数は `3` の引数1つだけで呼び出す必要があります。

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
